fof(f9648, plain, $false, inference(avatar_sat_refutation, [], [f329, f333, f341, f346, f351, f513, f1124, f2974, f4766, f4770, f4911, f4914, f5087, f5099, f6909, f7526, f7544, f7591, f7613, f7732, f7734, f9647])).
fof(f9647, plain, (~ spl50_34 | spl50_86 | ~ spl50_166), inference(avatar_contradiction_clause, [], [f9646])).
fof(f9646, plain, ($false | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(subsumption_resolution, [], [f9645, f337])).
fof(f337, plain, (sP1(sK40) | ~ spl50_34), inference(avatar_component_clause, [], [f335])).
fof(f335, plain, (spl50_34 <=> sP1(sK40)), introduced(avatar_definition, [new_symbols(naming, [spl50_34])])).
fof(f9645, plain, (~ sP1(sK40) | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(duplicate_literal_removal, [], [f9644])).
fof(f9644, plain, (~ sP1(sK40) | ~ sP1(sK40) | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(resolution, [], [f9003, f120])).
fof(f120, plain, ! [X0] : (r1(X0, sK24(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f55, plain, ! [X0] : ((! [X1] : ((p2(sK22(X1)) & (~ p2(sK23(X1)) & r1(sK22(X1), sK23(X1))) & r1(X1, sK22(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK25(X0), X6)) & ~ p2(sK25(X0)) & r1(sK24(X0), sK25(X0))) & r1(X0, sK24(X0)))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22, sK23, sK24, sK25])], [f50, f54, f53, f52, f51])).
fof(f51, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK22(X1)) & ? [X3] : (~ p2(X3) & r1(sK22(X1), X3)) & r1(X1, sK22(X1)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK22(X1), X3)) => (~ p2(sK23(X1)) & r1(sK22(X1), sK23(X1)))), introduced(choice_axiom, [])).
fof(f53, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK24(X0), X5)) & r1(X0, sK24(X0)))), introduced(choice_axiom, [])).
fof(f54, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK24(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK25(X0), X6)) & ~ p2(sK25(X0)) & r1(sK24(X0), sK25(X0)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP1(X0)), inference(rectify, [], [f49])).
fof(f49, plain, ! [X33] : ((! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59))) | ~ sP1(X33)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X33] : ((! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59))) | ~ sP1(X33)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X33] : (sP1(X33) <=> (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9003, plain, (! [X0] : (~ r1(X0, sK24(sK40)) | ~ sP1(X0)) | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(subsumption_resolution, [], [f9002, f631])).
fof(f631, plain, (~ p2(sK24(sK40)) | spl50_86), inference(avatar_component_clause, [], [f629])).
fof(f629, plain, (spl50_86 <=> p2(sK24(sK40))), introduced(avatar_definition, [new_symbols(naming, [spl50_86])])).
fof(f9002, plain, (! [X0] : (p2(sK24(sK40)) | ~ r1(X0, sK24(sK40)) | ~ sP1(X0)) | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(resolution, [], [f8805, f126])).
fof(f126, plain, ! [X0, X1] : (~ p2(sK23(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f8805, plain, (p2(sK23(sK24(sK40))) | (~ spl50_34 | spl50_86 | ~ spl50_166)), inference(subsumption_resolution, [], [f8804, f337])).
fof(f8804, plain, (p2(sK23(sK24(sK40))) | ~ sP1(sK40) | (spl50_86 | ~ spl50_166)), inference(subsumption_resolution, [], [f8795, f631])).
fof(f8795, plain, (p2(sK23(sK24(sK40))) | p2(sK24(sK40)) | ~ sP1(sK40) | ~ spl50_166), inference(resolution, [], [f4963, f876])).
fof(f876, plain, ! [X3] : (r1(sK22(sK24(X3)), sK23(sK24(X3))) | p2(sK24(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f862])).
fof(f862, plain, ! [X3] : (p2(sK24(X3)) | r1(sK22(sK24(X3)), sK23(sK24(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f125, f120])).
fof(f125, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK22(X1), sK23(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f4963, plain, (! [X3] : (~ r1(sK22(sK24(sK40)), X3) | p2(X3)) | ~ spl50_166), inference(avatar_component_clause, [], [f4962])).
fof(f4962, plain, (spl50_166 <=> ! [X3] : (p2(X3) | ~ r1(sK22(sK24(sK40)), X3))), introduced(avatar_definition, [new_symbols(naming, [spl50_166])])).
fof(f7734, plain, (~ spl50_34 | spl50_86 | spl50_166 | ~ spl50_100), inference(avatar_split_clause, [], [f7733, f1121, f4962, f629, f335])).
fof(f1121, plain, (spl50_100 <=> ! [X7, X6] : (~ p2(X6) | p2(X7) | ~ r1(X6, X7) | ~ r1(sK24(sK40), X6))), introduced(avatar_definition, [new_symbols(naming, [spl50_100])])).
fof(f7733, plain, (! [X3] : (p2(X3) | ~ r1(sK22(sK24(sK40)), X3) | p2(sK24(sK40)) | ~ sP1(sK40)) | ~ spl50_100), inference(subsumption_resolution, [], [f4918, f717])).
fof(f717, plain, ! [X3] : (p2(sK22(sK24(X3))) | p2(sK24(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f703])).
fof(f703, plain, ! [X3] : (p2(sK24(X3)) | p2(sK22(sK24(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f127, f120])).
fof(f127, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK22(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f4918, plain, (! [X3] : (p2(X3) | ~ r1(sK22(sK24(sK40)), X3) | ~ p2(sK22(sK24(sK40))) | p2(sK24(sK40)) | ~ sP1(sK40)) | ~ spl50_100), inference(resolution, [], [f1122, f805])).
fof(f805, plain, ! [X3] : (r1(sK24(X3), sK22(sK24(X3))) | p2(sK24(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f791])).
fof(f791, plain, ! [X3] : (p2(sK24(X3)) | r1(sK24(X3), sK22(sK24(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f124, f120])).
fof(f124, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK22(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f1122, plain, (! [X6, X7] : (~ r1(sK24(sK40), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6)) | ~ spl50_100), inference(avatar_component_clause, [], [f1121])).
fof(f7732, plain, (spl50_36 | ~ spl50_37 | ~ spl50_266), inference(avatar_contradiction_clause, [], [f7731])).
fof(f7731, plain, ($false | (spl50_36 | ~ spl50_37 | ~ spl50_266)), inference(subsumption_resolution, [], [f7730, f350])).
fof(f350, plain, (r1(sK28, sK40) | ~ spl50_37), inference(avatar_component_clause, [], [f348])).
fof(f348, plain, (spl50_37 <=> r1(sK28, sK40)), introduced(avatar_definition, [new_symbols(naming, [spl50_37])])).
fof(f7730, plain, (~ r1(sK28, sK40) | (spl50_36 | ~ spl50_37 | ~ spl50_266)), inference(subsumption_resolution, [], [f7729, f345])).
fof(f345, plain, (~ p2(sK40) | spl50_36), inference(avatar_component_clause, [], [f343])).
fof(f343, plain, (spl50_36 <=> p2(sK40)), introduced(avatar_definition, [new_symbols(naming, [spl50_36])])).
fof(f7729, plain, (p2(sK40) | ~ r1(sK28, sK40) | (spl50_36 | ~ spl50_37 | ~ spl50_266)), inference(resolution, [], [f7686, f142])).
fof(f142, plain, ! [X35] : (~ p2(sK45(X35)) | p2(X35) | ~ r1(sK28, X35)), inference(cnf_transformation, [], [f84])).
fof(f84, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK29(X1), X3)) & ~ p2(sK29(X1)) & r1(X1, sK29(X1))) | p2(X1) | ~ r1(sK28, X1)) & ((! [X6] : ((r1(X6, sK31(X6)) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK30, X6)) & r1(sK30, sK32) & ~ p1(sK30) & r1(sK28, sK30)) | ! [X11] : ~ r1(sK28, X11) | p1(sK28)) & ((! [X13] : ((r1(X13, sK34(X13)) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK33, X13)) & r1(sK33, sK35) & ~ p1(sK33) & ~ p2(sK33) & r1(sK28, sK33)) | ! [X18] : ~ r1(sK28, X18) | p1(sK28) | p2(sK28)) & ((sP7(sK36) & r1(sK36, sK37) & ~ p1(sK36) & ~ p2(sK36) & ~ p3(sK36) & r1(sK28, sK36)) | ! [X21] : ~ r1(sK28, X21) | p1(sK28) | p2(sK28) | p3(sK28)) & ((sP6(sK38) & r1(sK38, sK39) & ~ p1(sK38) & ~ p2(sK38) & ~ p3(sK38) & ~ p4(sK38) & r1(sK28, sK38)) | ! [X24] : ~ r1(sK28, X24) | p1(sK28) | p2(sK28) | p3(sK28) | p4(sK28)) & ((! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(sK40, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK40, X29)) & ~ p2(sK40)) | sP1(sK40)) & r1(sK28, sK40)) | sP5(sK28)) & ! [X31] : ((p1(sK41(X31)) & (~ p1(sK42(X31)) & r1(sK41(X31), sK42(X31))) & r1(X31, sK41(X31))) | p1(X31) | ~ r1(sK28, X31)) & (~ p1(sK43) & r1(sK28, sK43)) & ! [X35] : ((p2(sK44(X35)) & (~ p2(sK45(X35)) & r1(sK44(X35), sK45(X35))) & r1(X35, sK44(X35))) | p2(X35) | ~ r1(sK28, X35)) & (~ p2(sK46) & r1(sK28, sK46)) & ! [X39] : ((p3(sK47(X39)) & (~ p3(sK48(X39)) & r1(sK47(X39), sK48(X39))) & r1(X39, sK47(X39))) | p3(X39) | ~ r1(sK28, X39)) & (~ p3(sK49) & r1(sK28, sK49))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29, sK30, sK31, sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49])], [f61, f83, f82, f81, f80, f79, f78, f77, f76, f75, f74, f73, f72, f71, f70, f69, f68, f67, f66, f65, f64, f63, f62])).
fof(f62, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP7(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP6(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP1(X25)) & r1(X0, X25)) | sP5(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & ! [X35] : (? [X36] : (p2(X36) & ? [X37] : (~ p2(X37) & r1(X36, X37)) & r1(X35, X36)) | p2(X35) | ~ r1(X0, X35)) & ? [X38] : (~ p2(X38) & r1(X0, X38)) & ! [X39] : (? [X40] : (p3(X40) & ? [X41] : (~ p3(X41) & r1(X40, X41)) & r1(X39, X40)) | p3(X39) | ~ r1(X0, X39)) & ? [X42] : (~ p3(X42) & r1(X0, X42))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK28, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK28, X5)) | ! [X11] : ~ r1(sK28, X11) | p1(sK28)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK28, X12)) | ! [X18] : ~ r1(sK28, X18) | p1(sK28) | p2(sK28)) & (? [X19] : (sP7(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK28, X19)) | ! [X21] : ~ r1(sK28, X21) | p1(sK28) | p2(sK28) | p3(sK28)) & (? [X22] : (sP6(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK28, X22)) | ! [X24] : ~ r1(sK28, X24) | p1(sK28) | p2(sK28) | p3(sK28) | p4(sK28)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP1(X25)) & r1(sK28, X25)) | sP5(sK28)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(sK28, X31)) & ? [X34] : (~ p1(X34) & r1(sK28, X34)) & ! [X35] : (? [X36] : (p2(X36) & ? [X37] : (~ p2(X37) & r1(X36, X37)) & r1(X35, X36)) | p2(X35) | ~ r1(sK28, X35)) & ? [X38] : (~ p2(X38) & r1(sK28, X38)) & ! [X39] : (? [X40] : (p3(X40) & ? [X41] : (~ p3(X41) & r1(X40, X41)) & r1(X39, X40)) | p3(X39) | ~ r1(sK28, X39)) & ? [X42] : (~ p3(X42) & r1(sK28, X42)))), introduced(choice_axiom, [])).
fof(f63, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK29(X1), X3)) & ~ p2(sK29(X1)) & r1(X1, sK29(X1)))), introduced(choice_axiom, [])).
fof(f64, plain, (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK28, X5)) => (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK30, X6)) & ? [X10] : r1(sK30, X10) & ~ p1(sK30) & r1(sK28, sK30))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X6] : (? [X7] : r1(X6, X7) => r1(X6, sK31(X6))), introduced(choice_axiom, [])).
fof(f66, plain, (? [X10] : r1(sK30, X10) => r1(sK30, sK32)), introduced(choice_axiom, [])).
fof(f67, plain, (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK28, X12)) => (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK33, X13)) & ? [X17] : r1(sK33, X17) & ~ p1(sK33) & ~ p2(sK33) & r1(sK28, sK33))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X13] : (? [X14] : r1(X13, X14) => r1(X13, sK34(X13))), introduced(choice_axiom, [])).
fof(f69, plain, (? [X17] : r1(sK33, X17) => r1(sK33, sK35)), introduced(choice_axiom, [])).
fof(f70, plain, (? [X19] : (sP7(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK28, X19)) => (sP7(sK36) & ? [X20] : r1(sK36, X20) & ~ p1(sK36) & ~ p2(sK36) & ~ p3(sK36) & r1(sK28, sK36))), introduced(choice_axiom, [])).
fof(f71, plain, (? [X20] : r1(sK36, X20) => r1(sK36, sK37)), introduced(choice_axiom, [])).
fof(f72, plain, (? [X22] : (sP6(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK28, X22)) => (sP6(sK38) & ? [X23] : r1(sK38, X23) & ~ p1(sK38) & ~ p2(sK38) & ~ p3(sK38) & ~ p4(sK38) & r1(sK28, sK38))), introduced(choice_axiom, [])).
fof(f73, plain, (? [X23] : r1(sK38, X23) => r1(sK38, sK39)), introduced(choice_axiom, [])).
fof(f74, plain, (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP1(X25)) & r1(sK28, X25)) => (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(sK40, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(sK40, X29)) & ~ p2(sK40)) | sP1(sK40)) & r1(sK28, sK40))), introduced(choice_axiom, [])).
fof(f75, plain, ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) => (p1(sK41(X31)) & ? [X33] : (~ p1(X33) & r1(sK41(X31), X33)) & r1(X31, sK41(X31)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X31] : (? [X33] : (~ p1(X33) & r1(sK41(X31), X33)) => (~ p1(sK42(X31)) & r1(sK41(X31), sK42(X31)))), introduced(choice_axiom, [])).
fof(f77, plain, (? [X34] : (~ p1(X34) & r1(sK28, X34)) => (~ p1(sK43) & r1(sK28, sK43))), introduced(choice_axiom, [])).
fof(f78, plain, ! [X35] : (? [X36] : (p2(X36) & ? [X37] : (~ p2(X37) & r1(X36, X37)) & r1(X35, X36)) => (p2(sK44(X35)) & ? [X37] : (~ p2(X37) & r1(sK44(X35), X37)) & r1(X35, sK44(X35)))), introduced(choice_axiom, [])).
fof(f79, plain, ! [X35] : (? [X37] : (~ p2(X37) & r1(sK44(X35), X37)) => (~ p2(sK45(X35)) & r1(sK44(X35), sK45(X35)))), introduced(choice_axiom, [])).
fof(f80, plain, (? [X38] : (~ p2(X38) & r1(sK28, X38)) => (~ p2(sK46) & r1(sK28, sK46))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X39] : (? [X40] : (p3(X40) & ? [X41] : (~ p3(X41) & r1(X40, X41)) & r1(X39, X40)) => (p3(sK47(X39)) & ? [X41] : (~ p3(X41) & r1(sK47(X39), X41)) & r1(X39, sK47(X39)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X39] : (? [X41] : (~ p3(X41) & r1(sK47(X39), X41)) => (~ p3(sK48(X39)) & r1(sK47(X39), sK48(X39)))), introduced(choice_axiom, [])).
fof(f83, plain, (? [X42] : (~ p3(X42) & r1(sK28, X42)) => (~ p3(sK49) & r1(sK28, sK49))), introduced(choice_axiom, [])).
fof(f61, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP7(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP6(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (! [X26] : ((! [X27] : (~ p2(X27) | ! [X28] : (p2(X28) | ~ r1(X27, X28)) | ~ r1(X26, X27)) & ~ p2(X26)) | sP3(X26) | sP4(X26) | ~ r1(X25, X26)) & ((! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X25, X29)) & ~ p2(X25)) | sP1(X25)) & r1(X0, X25)) | sP5(X0)) & ! [X31] : (? [X32] : (p1(X32) & ? [X33] : (~ p1(X33) & r1(X32, X33)) & r1(X31, X32)) | p1(X31) | ~ r1(X0, X31)) & ? [X34] : (~ p1(X34) & r1(X0, X34)) & ! [X35] : (? [X36] : (p2(X36) & ? [X37] : (~ p2(X37) & r1(X36, X37)) & r1(X35, X36)) | p2(X35) | ~ r1(X0, X35)) & ? [X38] : (~ p2(X38) & r1(X0, X38)) & ! [X39] : (? [X40] : (p3(X40) & ? [X41] : (~ p3(X41) & r1(X40, X41)) & r1(X39, X40)) | p3(X39) | ~ r1(X0, X39)) & ? [X42] : (~ p3(X42) & r1(X0, X42))), inference(rectify, [], [f16])).
fof(f16, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP7(X19) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (sP6(X26) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | sP3(X34) | sP4(X34) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | sP1(X33)) & r1(X0, X33)) | sP5(X0)) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79)) & ! [X80] : (? [X81] : (p3(X81) & ? [X82] : (~ p3(X82) & r1(X81, X82)) & r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) & ? [X83] : (~ p3(X83) & r1(X0, X83))), inference(definition_folding, [], [f7, e15, e14, e13, e12, e11, e10, e9, e8])).
fof(f8, plain, ! [X0] : (! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)) | ~ sP0(X0)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X0] : (sP0(X0) <=> ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X44] : (! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)) | ~ sP2(X44)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X44] : (sP2(X44) <=> ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f11, plain, ! [X34] : ((! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ~ sP3(X34)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X34] : (sP3(X34) <=> (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f12, plain, ! [X34] : (! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP2(X44))) | ~ r1(X34, X44)) | ~ sP4(X34)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X34] : (sP4(X34) <=> ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP2(X44))) | ~ r1(X34, X44))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f13, plain, ! [X0] : (((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP0(X0))) | ~ sP5(X0)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X0] : (sP5(X0) <=> ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f14, plain, ! [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ~ sP6(X26)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X26] : (sP6(X26) <=> ! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f15, plain, ! [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ~ sP7(X19)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X19] : (sP7(X19) <=> ! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f7, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79)) & ! [X80] : (? [X81] : (p3(X81) & ? [X82] : (~ p3(X82) & r1(X81, X82)) & r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) & ? [X83] : (~ p3(X83) & r1(X0, X83))), inference(flattening, [], [f6])).
fof(f6, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : (((! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p2(X34)) | (! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40)))) | ! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | ! [X50] : (! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) & ((! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) & ~ p2(X33)) | (! [X56] : (? [X57] : (p2(X57) & ? [X58] : (~ p2(X58) & r1(X57, X58)) & r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) & ? [X59] : (? [X60] : (! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) & ~ p2(X60) & r1(X59, X60)) & r1(X33, X59)))) & r1(X0, X33)) | ((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | ! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) & ! [X72] : (? [X73] : (p1(X73) & ? [X74] : (~ p1(X74) & r1(X73, X74)) & r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) & ? [X75] : (~ p1(X75) & r1(X0, X75)) & ! [X76] : (? [X77] : (p2(X77) & ? [X78] : (~ p2(X78) & r1(X77, X78)) & r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) & ? [X79] : (~ p2(X79) & r1(X0, X79)) & ! [X80] : (? [X81] : (p3(X81) & ? [X82] : (~ p3(X82) & r1(X81, X82)) & r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) & ? [X83] : (~ p3(X83) & r1(X0, X83))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79)) | ~ ! [X80] : (~ ! [X81] : (~ p3(X81) | ! [X82] : (p3(X82) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) | ! [X83] : (p3(X83) | ~ r1(X0, X83))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79)) | ~ ! [X80] : (~ ! [X81] : (~ p3(X81) | ! [X82] : (p3(X82) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) | ! [X83] : (p3(X83) | ~ r1(X0, X83))), inference(true_and_false_elimination, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ($false | ~ r1(X6, X7)) | p1(X6)) | ! [X8] : (! [X9] : ($false | ~ r1(X8, X9)) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ($false | ~ r1(X5, X10)) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ($false | ~ r1(X0, X11)) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ($false | ~ r1(X13, X14)) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ($false | ~ r1(X15, X16)) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ($false | ~ r1(X12, X17)) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ($false | ~ r1(X0, X18)) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ($false | ~ r1(X20, X21)) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ($false | ~ r1(X22, X23)) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ($false | ~ r1(X19, X24)) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ($false | ~ r1(X0, X25)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ($false | ~ r1(X27, X28)) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ($false | ~ r1(X29, X30)) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ($false | ~ r1(X26, X31)) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ($false | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p2(X34)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X34, X40)))) | ! [X44] : (((~ ! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p2(X44)) & (~ ! [X47] : (~ ! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | p2(X47) | ~ r1(X44, X47)) | ! [X50] : (! [X51] : (~ ! [X52] : (~ p2(X52) | ! [X53] : (p2(X53) | ~ r1(X52, X53)) | ~ r1(X51, X52)) | p2(X51) | ~ r1(X50, X51)) | ~ r1(X44, X50)))) | ~ r1(X34, X44)) | ~ r1(X33, X34)) | ((~ ! [X54] : (~ p2(X54) | ! [X55] : (p2(X55) | ~ r1(X54, X55)) | ~ r1(X33, X54)) | p2(X33)) & (~ ! [X56] : (~ ! [X57] : (~ p2(X57) | ! [X58] : (p2(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p2(X56) | ~ r1(X33, X56)) | ! [X59] : (! [X60] : (~ ! [X61] : (~ p2(X61) | ! [X62] : (p2(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | p2(X60) | ~ r1(X59, X60)) | ~ r1(X33, X59)))) | ~ r1(X0, X33)) | ((~ ! [X63] : (~ p2(X63) | ! [X64] : (p2(X64) | ~ r1(X63, X64)) | ~ r1(X0, X63)) | p2(X0)) & (~ ! [X65] : (~ ! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) | p2(X65) | ~ r1(X0, X65)) | ! [X68] : (! [X69] : (~ ! [X70] : (~ p2(X70) | ! [X71] : (p2(X71) | ~ r1(X70, X71)) | ~ r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)))))) | ~ ! [X72] : (~ ! [X73] : (~ p1(X73) | ! [X74] : (p1(X74) | ~ r1(X73, X74)) | ~ r1(X72, X73)) | p1(X72) | ~ r1(X0, X72)) | ! [X75] : (p1(X75) | ~ r1(X0, X75)) | ~ ! [X76] : (~ ! [X77] : (~ p2(X77) | ! [X78] : (p2(X78) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | p2(X76) | ~ r1(X0, X76)) | ! [X79] : (p2(X79) | ~ r1(X0, X79)) | ~ ! [X80] : (~ ! [X81] : (~ p3(X81) | ! [X82] : (p3(X82) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | p3(X80) | ~ r1(X0, X80)) | ! [X83] : (p3(X83) | ~ r1(X0, X83))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL642+1.005.p', main)).
fof(f7686, plain, (p2(sK45(sK40)) | (spl50_36 | ~ spl50_37 | ~ spl50_266)), inference(subsumption_resolution, [], [f7685, f350])).
fof(f7685, plain, (p2(sK45(sK40)) | ~ r1(sK28, sK40) | (spl50_36 | ~ spl50_266)), inference(subsumption_resolution, [], [f7675, f345])).
fof(f7675, plain, (p2(sK45(sK40)) | p2(sK40) | ~ r1(sK28, sK40) | ~ spl50_266), inference(resolution, [], [f7612, f141])).
fof(f141, plain, ! [X35] : (r1(sK44(X35), sK45(X35)) | p2(X35) | ~ r1(sK28, X35)), inference(cnf_transformation, [], [f84])).
fof(f7612, plain, (! [X6] : (~ r1(sK44(sK40), X6) | p2(X6)) | ~ spl50_266), inference(avatar_component_clause, [], [f7611])).
fof(f7611, plain, (spl50_266 <=> ! [X6] : (~ r1(sK44(sK40), X6) | p2(X6))), introduced(avatar_definition, [new_symbols(naming, [spl50_266])])).
fof(f7613, plain, (spl50_36 | spl50_266 | ~ spl50_35 | ~ spl50_37 | ~ spl50_65), inference(avatar_split_clause, [], [f7609, f510, f348, f339, f7611, f343])).
fof(f339, plain, (spl50_35 <=> ! [X29, X30] : (~ p2(X29) | ~ r1(sK40, X29) | ~ r1(X29, X30) | p2(X30))), introduced(avatar_definition, [new_symbols(naming, [spl50_35])])).
fof(f510, plain, (spl50_65 <=> p2(sK44(sK40))), introduced(avatar_definition, [new_symbols(naming, [spl50_65])])).
fof(f7609, plain, (! [X6] : (~ r1(sK44(sK40), X6) | p2(X6) | p2(sK40)) | (~ spl50_35 | ~ spl50_37 | ~ spl50_65)), inference(subsumption_resolution, [], [f7608, f350])).
fof(f7608, plain, (! [X6] : (~ r1(sK44(sK40), X6) | p2(X6) | p2(sK40) | ~ r1(sK28, sK40)) | (~ spl50_35 | ~ spl50_65)), inference(subsumption_resolution, [], [f4907, f512])).
fof(f512, plain, (p2(sK44(sK40)) | ~ spl50_65), inference(avatar_component_clause, [], [f510])).
fof(f4907, plain, (! [X6] : (~ p2(sK44(sK40)) | ~ r1(sK44(sK40), X6) | p2(X6) | p2(sK40) | ~ r1(sK28, sK40)) | ~ spl50_35), inference(resolution, [], [f340, f140])).
fof(f140, plain, ! [X35] : (r1(X35, sK44(X35)) | p2(X35) | ~ r1(sK28, X35)), inference(cnf_transformation, [], [f84])).
fof(f340, plain, (! [X30, X29] : (~ r1(sK40, X29) | ~ p2(X29) | ~ r1(X29, X30) | p2(X30)) | ~ spl50_35), inference(avatar_component_clause, [], [f339])).
fof(f7591, plain, ~ spl50_185, inference(avatar_contradiction_clause, [], [f7590])).
fof(f7590, plain, ($false | ~ spl50_185), inference(subsumption_resolution, [], [f7589, f138])).
fof(f138, plain, r1(sK28, sK46), inference(cnf_transformation, [], [f84])).
fof(f7589, plain, (~ r1(sK28, sK46) | ~ spl50_185), inference(subsumption_resolution, [], [f7588, f139])).
fof(f139, plain, ~ p2(sK46), inference(cnf_transformation, [], [f84])).
fof(f7588, plain, (p2(sK46) | ~ r1(sK28, sK46) | ~ spl50_185), inference(resolution, [], [f5185, f181])).
fof(f181, plain, ! [X1] : (~ p2(sK29(X1)) | p2(X1) | ~ r1(sK28, X1)), inference(cnf_transformation, [], [f84])).
fof(f5185, plain, (p2(sK29(sK46)) | ~ spl50_185), inference(avatar_component_clause, [], [f5183])).
fof(f5183, plain, (spl50_185 <=> p2(sK29(sK46))), introduced(avatar_definition, [new_symbols(naming, [spl50_185])])).
fof(f7544, plain, (spl50_121 | ~ spl50_261), inference(avatar_split_clause, [], [f7543, f7524, f2972])).
fof(f2972, plain, (spl50_121 <=> ! [X1] : (~ r1(X1, sK46) | ~ sP0(X1))), introduced(avatar_definition, [new_symbols(naming, [spl50_121])])).
fof(f7524, plain, (spl50_261 <=> ! [X1, X0] : (~ r1(X0, sK29(sK46)) | ~ sP0(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl50_261])])).
fof(f7543, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK46)) | ~ spl50_261), inference(subsumption_resolution, [], [f7542, f138])).
fof(f7542, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK46) | ~ r1(sK28, sK46)) | ~ spl50_261), inference(subsumption_resolution, [], [f7541, f139])).
fof(f7541, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK46) | p2(sK46) | ~ r1(sK28, sK46)) | ~ spl50_261), inference(resolution, [], [f7525, f180])).
fof(f180, plain, ! [X1] : (r1(X1, sK29(X1)) | p2(X1) | ~ r1(sK28, X1)), inference(cnf_transformation, [], [f84])).
fof(f7525, plain, (! [X0, X1] : (~ r1(X0, sK29(sK46)) | ~ sP0(X1) | ~ r1(X1, X0)) | ~ spl50_261), inference(avatar_component_clause, [], [f7524])).
fof(f7526, plain, (spl50_261 | spl50_185 | ~ spl50_120 | ~ spl50_141 | ~ spl50_235), inference(avatar_split_clause, [], [f7522, f6906, f4763, f2968, f5183, f7524])).
fof(f2968, plain, (spl50_120 <=> p2(sK26(sK29(sK46)))), introduced(avatar_definition, [new_symbols(naming, [spl50_120])])).
fof(f4763, plain, (spl50_141 <=> r1(sK29(sK46), sK26(sK29(sK46)))), introduced(avatar_definition, [new_symbols(naming, [spl50_141])])).
fof(f6906, plain, (spl50_235 <=> r1(sK26(sK29(sK46)), sK27(sK29(sK46)))), introduced(avatar_definition, [new_symbols(naming, [spl50_235])])).
fof(f7522, plain, (! [X0, X1] : (p2(sK29(sK46)) | ~ r1(X0, sK29(sK46)) | ~ r1(X1, X0) | ~ sP0(X1)) | (~ spl50_120 | ~ spl50_141 | ~ spl50_235)), inference(resolution, [], [f7135, f130])).
fof(f130, plain, ! [X2, X0, X1] : (~ p2(sK27(X2)) | p2(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ! [X0] : (! [X1] : (! [X2] : ((p2(sK26(X2)) & (~ p2(sK27(X2)) & r1(sK26(X2), sK27(X2))) & r1(X2, sK26(X2))) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27])], [f57, f59, f58])).
fof(f58, plain, ! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) => (p2(sK26(X2)) & ? [X4] : (~ p2(X4) & r1(sK26(X2), X4)) & r1(X2, sK26(X2)))), introduced(choice_axiom, [])).
fof(f59, plain, ! [X2] : (? [X4] : (~ p2(X4) & r1(sK26(X2), X4)) => (~ p2(sK27(X2)) & r1(sK26(X2), sK27(X2)))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f56])).
fof(f56, plain, ! [X0] : (! [X68] : (! [X69] : (? [X70] : (p2(X70) & ? [X71] : (~ p2(X71) & r1(X70, X71)) & r1(X69, X70)) | p2(X69) | ~ r1(X68, X69)) | ~ r1(X0, X68)) | ~ sP0(X0)), inference(nnf_transformation, [], [f8])).
fof(f7135, plain, (p2(sK27(sK29(sK46))) | (~ spl50_120 | ~ spl50_141 | ~ spl50_235)), inference(resolution, [], [f6908, f5177])).
fof(f5177, plain, (! [X0] : (~ r1(sK26(sK29(sK46)), X0) | p2(X0)) | (~ spl50_120 | ~ spl50_141)), inference(subsumption_resolution, [], [f5176, f138])).
fof(f5176, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK29(sK46)), X0) | ~ r1(sK28, sK46)) | (~ spl50_120 | ~ spl50_141)), inference(subsumption_resolution, [], [f5175, f139])).
fof(f5175, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK29(sK46)), X0) | p2(sK46) | ~ r1(sK28, sK46)) | (~ spl50_120 | ~ spl50_141)), inference(subsumption_resolution, [], [f5141, f2970])).
fof(f2970, plain, (p2(sK26(sK29(sK46))) | ~ spl50_120), inference(avatar_component_clause, [], [f2968])).
fof(f5141, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK29(sK46)), X0) | ~ p2(sK26(sK29(sK46))) | p2(sK46) | ~ r1(sK28, sK46)) | ~ spl50_141), inference(resolution, [], [f4765, f182])).
fof(f182, plain, ! [X4, X3, X1] : (~ r1(sK29(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK28, X1)), inference(cnf_transformation, [], [f84])).
fof(f4765, plain, (r1(sK29(sK46), sK26(sK29(sK46))) | ~ spl50_141), inference(avatar_component_clause, [], [f4763])).
fof(f6908, plain, (r1(sK26(sK29(sK46)), sK27(sK29(sK46))) | ~ spl50_235), inference(avatar_component_clause, [], [f6906])).
fof(f6909, plain, (spl50_235 | spl50_121), inference(avatar_split_clause, [], [f6904, f2972, f6906])).
fof(f6904, plain, ! [X2] : (~ r1(X2, sK46) | ~ sP0(X2) | r1(sK26(sK29(sK46)), sK27(sK29(sK46)))), inference(subsumption_resolution, [], [f6894, f139])).
fof(f6894, plain, ! [X2] : (~ r1(X2, sK46) | ~ sP0(X2) | p2(sK46) | r1(sK26(sK29(sK46)), sK27(sK29(sK46)))), inference(resolution, [], [f1090, f138])).
fof(f1090, plain, ! [X8, X9] : (~ r1(sK28, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | r1(sK26(sK29(X8)), sK27(sK29(X8)))), inference(subsumption_resolution, [], [f1076, f181])).
fof(f1076, plain, ! [X8, X9] : (p2(sK29(X8)) | r1(sK26(sK29(X8)), sK27(sK29(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK28, X8)), inference(resolution, [], [f129, f180])).
fof(f129, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(sK26(X2), sK27(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f60])).
fof(f5099, plain, (~ spl50_34 | ~ spl50_86 | ~ spl50_33 | spl50_84 | spl50_85), inference(avatar_split_clause, [], [f5098, f625, f621, f331, f629, f335])).
fof(f331, plain, (spl50_33 <=> ! [X26] : (~ p2(X26) | ~ r1(sK40, X26) | sP4(X26) | sP3(X26))), introduced(avatar_definition, [new_symbols(naming, [spl50_33])])).
fof(f621, plain, (spl50_84 <=> sP3(sK24(sK40))), introduced(avatar_definition, [new_symbols(naming, [spl50_84])])).
fof(f625, plain, (spl50_85 <=> sP4(sK24(sK40))), introduced(avatar_definition, [new_symbols(naming, [spl50_85])])).
fof(f5098, plain, (~ p2(sK24(sK40)) | ~ sP1(sK40) | (~ spl50_33 | spl50_84 | spl50_85)), inference(subsumption_resolution, [], [f5097, f622])).
fof(f622, plain, (~ sP3(sK24(sK40)) | spl50_84), inference(avatar_component_clause, [], [f621])).
fof(f5097, plain, (~ p2(sK24(sK40)) | sP3(sK24(sK40)) | ~ sP1(sK40) | (~ spl50_33 | spl50_85)), inference(subsumption_resolution, [], [f583, f626])).
fof(f626, plain, (~ sP4(sK24(sK40)) | spl50_85), inference(avatar_component_clause, [], [f625])).
fof(f583, plain, (~ p2(sK24(sK40)) | sP4(sK24(sK40)) | sP3(sK24(sK40)) | ~ sP1(sK40) | ~ spl50_33), inference(resolution, [], [f332, f120])).
fof(f332, plain, (! [X26] : (~ r1(sK40, X26) | ~ p2(X26) | sP4(X26) | sP3(X26)) | ~ spl50_33), inference(avatar_component_clause, [], [f331])).
fof(f5087, plain, (spl50_58 | ~ spl50_31), inference(avatar_split_clause, [], [f3949, f323, f469])).
fof(f469, plain, (spl50_58 <=> sP0(sK28)), introduced(avatar_definition, [new_symbols(naming, [spl50_58])])).
fof(f323, plain, (spl50_31 <=> sP5(sK28)), introduced(avatar_definition, [new_symbols(naming, [spl50_31])])).
fof(f3949, plain, (~ sP5(sK28) | sP0(sK28)), inference(duplicate_literal_removal, [], [f3948])).
fof(f3948, plain, (~ sP5(sK28) | sP0(sK28) | sP0(sK28) | ~ sP5(sK28)), inference(resolution, [], [f3947, f94])).
fof(f94, plain, ! [X0] : (r1(X0, sK12(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : ((((p2(sK10(X0)) & (~ p2(sK11(X0)) & r1(sK10(X0), sK11(X0))) & r1(X0, sK10(X0))) | p2(X0)) & ((! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK12(X0), X4)) & ~ p2(sK12(X0)) & r1(X0, sK12(X0))) | sP0(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12])], [f26, f29, f28, f27])).
fof(f27, plain, ! [X0] : (? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) => (p2(sK10(X0)) & ? [X2] : (~ p2(X2) & r1(sK10(X0), X2)) & r1(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (? [X2] : (~ p2(X2) & r1(sK10(X0), X2)) => (~ p2(sK11(X0)) & r1(sK10(X0), sK11(X0)))), introduced(choice_axiom, [])).
fof(f29, plain, ! [X0] : (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) => (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK12(X0), X4)) & ~ p2(sK12(X0)) & r1(X0, sK12(X0)))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X0] : (((? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) | p2(X0)) & (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) | sP0(X0))) | ~ sP5(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X0] : (((? [X63] : (p2(X63) & ? [X64] : (~ p2(X64) & r1(X63, X64)) & r1(X0, X63)) | p2(X0)) & (? [X65] : (! [X66] : (~ p2(X66) | ! [X67] : (p2(X67) | ~ r1(X66, X67)) | ~ r1(X65, X66)) & ~ p2(X65) & r1(X0, X65)) | sP0(X0))) | ~ sP5(X0)), inference(nnf_transformation, [], [f13])).
fof(f3947, plain, ! [X0] : (~ r1(sK28, sK12(X0)) | ~ sP5(X0) | sP0(X0)), inference(subsumption_resolution, [], [f3946, f95])).
fof(f95, plain, ! [X0] : (~ p2(sK12(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f30])).
fof(f3946, plain, ! [X0] : (sP0(X0) | ~ sP5(X0) | ~ r1(sK28, sK12(X0)) | p2(sK12(X0))), inference(subsumption_resolution, [], [f3945, f142])).
fof(f3945, plain, ! [X0] : (p2(sK45(sK12(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK28, sK12(X0)) | p2(sK12(X0))), inference(duplicate_literal_removal, [], [f3936])).
fof(f3936, plain, ! [X0] : (p2(sK45(sK12(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK28, sK12(X0)) | p2(sK12(X0)) | ~ r1(sK28, sK12(X0))), inference(resolution, [], [f1048, f141])).
fof(f1048, plain, ! [X12, X13] : (~ r1(sK44(sK12(X13)), X12) | p2(X12) | sP0(X13) | ~ sP5(X13) | ~ r1(sK28, sK12(X13))), inference(subsumption_resolution, [], [f1047, f95])).
fof(f1047, plain, ! [X12, X13] : (p2(X12) | ~ r1(sK44(sK12(X13)), X12) | sP0(X13) | ~ sP5(X13) | p2(sK12(X13)) | ~ r1(sK28, sK12(X13))), inference(subsumption_resolution, [], [f1043, f143])).
fof(f143, plain, ! [X35] : (~ r1(sK28, X35) | p2(X35) | p2(sK44(X35))), inference(cnf_transformation, [], [f84])).
fof(f1043, plain, ! [X12, X13] : (p2(X12) | ~ r1(sK44(sK12(X13)), X12) | ~ p2(sK44(sK12(X13))) | sP0(X13) | ~ sP5(X13) | p2(sK12(X13)) | ~ r1(sK28, sK12(X13))), inference(resolution, [], [f96, f140])).
fof(f96, plain, ! [X4, X0, X5] : (~ r1(sK12(X0), X4) | p2(X5) | ~ r1(X4, X5) | ~ p2(X4) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f30])).
fof(f4914, plain, (~ spl50_34 | ~ spl50_85), inference(avatar_contradiction_clause, [], [f4913])).
fof(f4913, plain, ($false | (~ spl50_34 | ~ spl50_85)), inference(subsumption_resolution, [], [f4912, f337])).
fof(f4912, plain, (~ sP1(sK40) | ~ spl50_85), inference(resolution, [], [f627, f2250])).
fof(f2250, plain, ! [X0] : (~ sP4(sK24(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f2249])).
fof(f2249, plain, ! [X0] : (~ sP4(sK24(X0)) | ~ sP1(X0) | ~ sP4(sK24(X0)) | ~ sP1(X0)), inference(resolution, [], [f2214, f121])).
fof(f121, plain, ! [X0] : (r1(sK24(X0), sK25(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f2214, plain, ! [X0, X1] : (~ r1(X1, sK25(X0)) | ~ sP4(sK24(X0)) | ~ sP1(X0) | ~ sP4(X1)), inference(subsumption_resolution, [], [f2213, f122])).
fof(f122, plain, ! [X0] : (~ p2(sK25(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f2213, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP4(sK24(X0)) | p2(sK25(X0)) | ~ r1(X1, sK25(X0)) | ~ sP4(X1)), inference(resolution, [], [f2177, f106])).
fof(f106, plain, ! [X0, X1] : (~ p2(sK14(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ! [X0] : (! [X1] : ((((p2(sK13(X1)) & (~ p2(sK14(X1)) & r1(sK13(X1), sK14(X1))) & r1(X1, sK13(X1))) | p2(X1)) & ((! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK15(X1), X5)) & ~ p2(sK15(X1)) & r1(X1, sK15(X1))) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15])], [f32, f35, f34, f33])).
fof(f33, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK13(X1)) & ? [X3] : (~ p2(X3) & r1(sK13(X1), X3)) & r1(X1, sK13(X1)))), introduced(choice_axiom, [])).
fof(f34, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK13(X1), X3)) => (~ p2(sK14(X1)) & r1(sK13(X1), sK14(X1)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X1] : (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) => (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK15(X1), X5)) & ~ p2(sK15(X1)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f32, plain, ! [X0] : (! [X1] : (((? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1)) & (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(rectify, [], [f31])).
fof(f31, plain, ! [X34] : (! [X44] : (((? [X45] : (p2(X45) & ? [X46] : (~ p2(X46) & r1(X45, X46)) & r1(X44, X45)) | p2(X44)) & (? [X47] : (! [X48] : (~ p2(X48) | ! [X49] : (p2(X49) | ~ r1(X48, X49)) | ~ r1(X47, X48)) & ~ p2(X47) & r1(X44, X47)) | sP2(X44))) | ~ r1(X34, X44)) | ~ sP4(X34)), inference(nnf_transformation, [], [f12])).
fof(f2177, plain, ! [X0] : (p2(sK14(sK25(X0))) | ~ sP1(X0) | ~ sP4(sK24(X0))), inference(duplicate_literal_removal, [], [f2168])).
fof(f2168, plain, ! [X0] : (~ sP1(X0) | p2(sK14(sK25(X0))) | ~ sP4(sK24(X0)) | ~ sP4(sK24(X0)) | ~ sP1(X0)), inference(resolution, [], [f1249, f833])).
fof(f833, plain, ! [X10] : (r1(sK13(sK25(X10)), sK14(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f822, f122])).
fof(f822, plain, ! [X10] : (p2(sK25(X10)) | r1(sK13(sK25(X10)), sK14(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f105, f121])).
fof(f105, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK13(X1), sK14(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f36])).
fof(f1249, plain, ! [X0, X1] : (~ r1(sK13(sK25(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP4(sK24(X0))), inference(subsumption_resolution, [], [f1248, f654])).
fof(f654, plain, ! [X10] : (p2(sK13(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f643, f122])).
fof(f643, plain, ! [X10] : (p2(sK25(X10)) | p2(sK13(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f107, f121])).
fof(f107, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK13(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f36])).
fof(f1248, plain, ! [X0, X1] : (~ sP4(sK24(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK13(sK25(X0)), X1) | ~ p2(sK13(sK25(X0)))), inference(duplicate_literal_removal, [], [f1228])).
fof(f1228, plain, ! [X0, X1] : (~ sP4(sK24(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK13(sK25(X0)), X1) | ~ p2(sK13(sK25(X0))) | ~ sP1(X0)), inference(resolution, [], [f762, f123])).
fof(f123, plain, ! [X6, X0, X7] : (~ r1(sK25(X0), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6) | ~ sP1(X0)), inference(cnf_transformation, [], [f55])).
fof(f762, plain, ! [X10] : (r1(sK25(X10), sK13(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f751, f122])).
fof(f751, plain, ! [X10] : (p2(sK25(X10)) | r1(sK25(X10), sK13(sK25(X10))) | ~ sP4(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f104, f121])).
fof(f104, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK13(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f36])).
fof(f627, plain, (sP4(sK24(sK40)) | ~ spl50_85), inference(avatar_component_clause, [], [f625])).
fof(f4911, plain, (~ spl50_34 | ~ spl50_84), inference(avatar_contradiction_clause, [], [f4910])).
fof(f4910, plain, ($false | (~ spl50_34 | ~ spl50_84)), inference(subsumption_resolution, [], [f4909, f337])).
fof(f4909, plain, (~ sP1(sK40) | ~ spl50_84), inference(resolution, [], [f623, f2525])).
fof(f2525, plain, ! [X0] : (~ sP3(sK24(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f2524])).
fof(f2524, plain, ! [X0] : (~ sP3(sK24(X0)) | ~ sP1(X0) | ~ sP3(sK24(X0)) | ~ sP1(X0)), inference(resolution, [], [f2498, f121])).
fof(f2498, plain, ! [X0, X1] : (~ r1(X1, sK25(X0)) | ~ sP3(sK24(X0)) | ~ sP1(X0) | ~ sP3(X1)), inference(subsumption_resolution, [], [f2497, f122])).
fof(f2497, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP3(sK24(X0)) | p2(sK25(X0)) | ~ r1(X1, sK25(X0)) | ~ sP3(X1)), inference(resolution, [], [f2470, f114])).
fof(f114, plain, ! [X0, X1] : (~ p2(sK17(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f43])).
fof(f43, plain, ! [X0] : ((! [X1] : ((p2(sK16(X1)) & (~ p2(sK17(X1)) & r1(sK16(X1), sK17(X1))) & r1(X1, sK16(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK19(X0), X6)) & ~ p2(sK19(X0)) & r1(sK18(X0), sK19(X0))) & r1(X0, sK18(X0)))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17, sK18, sK19])], [f38, f42, f41, f40, f39])).
fof(f39, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK16(X1)) & ? [X3] : (~ p2(X3) & r1(sK16(X1), X3)) & r1(X1, sK16(X1)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK16(X1), X3)) => (~ p2(sK17(X1)) & r1(sK16(X1), sK17(X1)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK18(X0), X5)) & r1(X0, sK18(X0)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK18(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK19(X0), X6)) & ~ p2(sK19(X0)) & r1(sK18(X0), sK19(X0)))), introduced(choice_axiom, [])).
fof(f38, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP3(X0)), inference(rectify, [], [f37])).
fof(f37, plain, ! [X34] : ((! [X37] : (? [X38] : (p2(X38) & ? [X39] : (~ p2(X39) & r1(X38, X39)) & r1(X37, X38)) | p2(X37) | ~ r1(X34, X37)) & ? [X40] : (? [X41] : (! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) & ~ p2(X41) & r1(X40, X41)) & r1(X34, X40))) | ~ sP3(X34)), inference(nnf_transformation, [], [f11])).
fof(f2470, plain, ! [X0] : (p2(sK17(sK25(X0))) | ~ sP1(X0) | ~ sP3(sK24(X0))), inference(duplicate_literal_removal, [], [f2461])).
fof(f2461, plain, ! [X0] : (~ sP1(X0) | p2(sK17(sK25(X0))) | ~ sP3(sK24(X0)) | ~ sP3(sK24(X0)) | ~ sP1(X0)), inference(resolution, [], [f1414, f857])).
fof(f857, plain, ! [X10] : (r1(sK16(sK25(X10)), sK17(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f845, f122])).
fof(f845, plain, ! [X10] : (p2(sK25(X10)) | r1(sK16(sK25(X10)), sK17(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f113, f121])).
fof(f113, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK16(X1), sK17(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f43])).
fof(f1414, plain, ! [X0, X1] : (~ r1(sK16(sK25(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP3(sK24(X0))), inference(subsumption_resolution, [], [f1413, f698])).
fof(f698, plain, ! [X10] : (p2(sK16(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f686, f122])).
fof(f686, plain, ! [X10] : (p2(sK25(X10)) | p2(sK16(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f115, f121])).
fof(f115, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK16(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f43])).
fof(f1413, plain, ! [X0, X1] : (~ sP3(sK24(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK16(sK25(X0)), X1) | ~ p2(sK16(sK25(X0)))), inference(duplicate_literal_removal, [], [f1390])).
fof(f1390, plain, ! [X0, X1] : (~ sP3(sK24(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK16(sK25(X0)), X1) | ~ p2(sK16(sK25(X0))) | ~ sP1(X0)), inference(resolution, [], [f786, f123])).
fof(f786, plain, ! [X10] : (r1(sK25(X10), sK16(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f774, f122])).
fof(f774, plain, ! [X10] : (p2(sK25(X10)) | r1(sK25(X10), sK16(sK25(X10))) | ~ sP3(sK24(X10)) | ~ sP1(X10)), inference(resolution, [], [f112, f121])).
fof(f112, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK16(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f43])).
fof(f623, plain, (sP3(sK24(sK40)) | ~ spl50_84), inference(avatar_component_clause, [], [f621])).
fof(f4770, plain, (~ spl50_58 | ~ spl50_121), inference(avatar_split_clause, [], [f4767, f2972, f469])).
fof(f4767, plain, (~ sP0(sK28) | ~ spl50_121), inference(resolution, [], [f2973, f138])).
fof(f2973, plain, (! [X1] : (~ r1(X1, sK46) | ~ sP0(X1)) | ~ spl50_121), inference(avatar_component_clause, [], [f2972])).
fof(f4766, plain, (spl50_141 | spl50_121), inference(avatar_split_clause, [], [f4761, f2972, f4763])).
fof(f4761, plain, ! [X1] : (~ r1(X1, sK46) | ~ sP0(X1) | r1(sK29(sK46), sK26(sK29(sK46)))), inference(subsumption_resolution, [], [f4751, f139])).
fof(f4751, plain, ! [X1] : (~ r1(X1, sK46) | ~ sP0(X1) | p2(sK46) | r1(sK29(sK46), sK26(sK29(sK46)))), inference(resolution, [], [f1032, f138])).
fof(f1032, plain, ! [X8, X9] : (~ r1(sK28, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | r1(sK29(X8), sK26(sK29(X8)))), inference(subsumption_resolution, [], [f1018, f181])).
fof(f1018, plain, ! [X8, X9] : (p2(sK29(X8)) | r1(sK29(X8), sK26(sK29(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK28, X8)), inference(resolution, [], [f128, f180])).
fof(f128, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(X2, sK26(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f60])).
fof(f2974, plain, (spl50_120 | spl50_121), inference(avatar_split_clause, [], [f2966, f2972, f2968])).
fof(f2966, plain, ! [X1] : (~ r1(X1, sK46) | ~ sP0(X1) | p2(sK26(sK29(sK46)))), inference(subsumption_resolution, [], [f2956, f139])).
fof(f2956, plain, ! [X1] : (~ r1(X1, sK46) | ~ sP0(X1) | p2(sK46) | p2(sK26(sK29(sK46)))), inference(resolution, [], [f967, f138])).
fof(f967, plain, ! [X8, X9] : (~ r1(sK28, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | p2(sK26(sK29(X8)))), inference(subsumption_resolution, [], [f953, f181])).
fof(f953, plain, ! [X8, X9] : (p2(sK29(X8)) | p2(sK26(sK29(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK28, X8)), inference(resolution, [], [f131, f180])).
fof(f131, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | p2(sK26(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f60])).
fof(f1124, plain, (~ spl50_34 | spl50_84 | spl50_85 | spl50_100 | ~ spl50_32), inference(avatar_split_clause, [], [f1110, f327, f1121, f625, f621, f335])).
fof(f327, plain, (spl50_32 <=> ! [X27, X26, X28] : (~ p2(X27) | ~ r1(sK40, X26) | sP4(X26) | sP3(X26) | ~ r1(X26, X27) | ~ r1(X27, X28) | p2(X28))), introduced(avatar_definition, [new_symbols(naming, [spl50_32])])).
fof(f1110, plain, (! [X6, X7] : (~ p2(X6) | sP4(sK24(sK40)) | sP3(sK24(sK40)) | ~ r1(sK24(sK40), X6) | ~ r1(X6, X7) | p2(X7) | ~ sP1(sK40)) | ~ spl50_32), inference(resolution, [], [f328, f120])).
fof(f328, plain, (! [X28, X26, X27] : (~ r1(sK40, X26) | ~ p2(X27) | sP4(X26) | sP3(X26) | ~ r1(X26, X27) | ~ r1(X27, X28) | p2(X28)) | ~ spl50_32), inference(avatar_component_clause, [], [f327])).
fof(f513, plain, (spl50_65 | spl50_36 | ~ spl50_37), inference(avatar_split_clause, [], [f498, f348, f343, f510])).
fof(f498, plain, (p2(sK40) | p2(sK44(sK40)) | ~ spl50_37), inference(resolution, [], [f350, f143])).
fof(f351, plain, (spl50_31 | spl50_37), inference(avatar_split_clause, [], [f150, f348, f323])).
fof(f150, plain, (r1(sK28, sK40) | sP5(sK28)), inference(cnf_transformation, [], [f84])).
fof(f346, plain, (spl50_31 | spl50_34 | ~ spl50_36), inference(avatar_split_clause, [], [f151, f343, f335, f323])).
fof(f151, plain, (~ p2(sK40) | sP1(sK40) | sP5(sK28)), inference(cnf_transformation, [], [f84])).
fof(f341, plain, (spl50_31 | spl50_34 | spl50_35), inference(avatar_split_clause, [], [f152, f339, f335, f323])).
fof(f152, plain, ! [X30, X29] : (~ p2(X29) | p2(X30) | ~ r1(X29, X30) | ~ r1(sK40, X29) | sP1(sK40) | sP5(sK28)), inference(cnf_transformation, [], [f84])).
fof(f333, plain, (spl50_31 | spl50_33), inference(avatar_split_clause, [], [f153, f331, f323])).
fof(f153, plain, ! [X26] : (~ p2(X26) | sP3(X26) | sP4(X26) | ~ r1(sK40, X26) | sP5(sK28)), inference(cnf_transformation, [], [f84])).
fof(f329, plain, (spl50_31 | spl50_32), inference(avatar_split_clause, [], [f154, f327, f323])).
fof(f154, plain, ! [X28, X26, X27] : (~ p2(X27) | p2(X28) | ~ r1(X27, X28) | ~ r1(X26, X27) | sP3(X26) | sP4(X26) | ~ r1(sK40, X26) | sP5(sK28)), inference(cnf_transformation, [], [f84])).