fof(f91700, plain, $false, inference(avatar_sat_refutation, [], [f418, f926, f931, f936, f1041, f1058, f5565, f16030, f28355, f28421, f29231, f29443, f29486, f30638, f32049, f32073, f34426, f35063, f35138, f35165, f41336, f42718, f91699])).
fof(f91699, plain, ~ spl50_2241, inference(avatar_contradiction_clause, [], [f91698])).
fof(f91698, plain, ($false | ~ spl50_2241), inference(subsumption_resolution, [], [f91675, f132])).
fof(f132, plain, r1(sK34, sK37), inference(cnf_transformation, [], [f71])).
fof(f71, plain, (! [X1] : ((! [X2] : ((~ p1(sK14(X2)) & r1(X2, sK14(X2))) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & ((! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1))) | ! [X9] : ((~ p1(sK16(X9)) & r1(X9, sK16(X9))) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : ((~ p1(sK17(X13)) & r1(X13, sK17(X13))) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ((! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12))) | ! [X20] : ((~ p1(sK19(X20)) & r1(X20, sK19(X20))) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : ((~ p1(sK20(X25)) & r1(X25, sK20(X25))) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & ((! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24))) | ! [X32] : ((~ p1(sK22(X32)) & r1(X32, sK22(X32))) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : ((~ p1(sK23(X38)) & r1(X38, sK23(X38))) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & ((! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37))) | ! [X45] : ((~ p1(sK25(X45)) & r1(X45, sK25(X45))) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : ((~ p1(sK26(X52)) & r1(X52, sK26(X52))) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & ((! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51))) | ! [X59] : ((~ p1(sK28(X59)) & r1(X59, sK28(X59))) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ((((((! [X67] : ((p1(sK35(X67)) & (~ p1(sK36(X67)) & r1(sK35(X67), sK36(X67))) & r1(X67, sK35(X67))) | p1(X67) | ~ r1(sK34, X67)) & (! [X71] : (p1(X71) | ~ r1(sK37, X71)) & r1(sK34, sK37)) & (~ p1(sK38) & r1(sK34, sK38)) & r1(sK33, sK34)) & r1(sK32, sK33)) & r1(sK31, sK32)) & r1(sK30, sK31)) & r1(sK29, sK30)) & r1(sK13, sK29)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : ((((! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(sK39(X78), X80)) & ~ p1(sK39(X78)) & r1(X78, sK39(X78))) | sP1(X78)) & ! [X82] : ((~ p1(sK40(X82)) & r1(X82, sK40(X82))) | ! [X84] : (! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X82, X84)) | ~ r1(X78, X82)) & sP2(X78) & sP3(X78)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(sK13, X73)) & ! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : ((! [X93] : ((~ p1(sK41(X93)) & r1(X93, sK41(X93))) | ! [X95] : (! [X96] : (p1(X96) | ~ r1(X95, X96)) | ~ r1(X93, X95)) | ~ r1(X92, X93)) & ((! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(sK42(X92), X98)) & ~ p1(sK42(X92)) & r1(X92, sK42(X92))) | ! [X100] : ((~ p1(sK43(X100)) & r1(X100, sK43(X100))) | ~ r1(X92, X100)) | p1(X92))) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(sK13, X86)) & ! [X102] : (! [X103] : (! [X104] : (! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : ((! [X110] : ((~ p1(sK44(X110)) & r1(X110, sK44(X110))) | ! [X112] : (! [X113] : (p1(X113) | ~ r1(X112, X113)) | ~ r1(X110, X112)) | ~ r1(X109, X110)) & ((! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(sK45(X109), X115)) & ~ p1(sK45(X109)) & r1(X109, sK45(X109))) | ! [X117] : ((~ p1(sK46(X117)) & r1(X117, sK46(X117))) | ~ r1(X109, X117)) | p1(X109))) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | ~ r1(X102, X103)) | ~ r1(sK13, X102)) & ! [X119] : (! [X120] : (! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : ((! [X128] : ((~ p1(sK47(X128)) & r1(X128, sK47(X128))) | ! [X130] : (! [X131] : (p1(X131) | ~ r1(X130, X131)) | ~ r1(X128, X130)) | ~ r1(X127, X128)) & ((! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(sK48(X127), X133)) & ~ p1(sK48(X127)) & r1(X127, sK48(X127))) | ! [X135] : ((~ p1(sK49(X135)) & r1(X135, sK49(X135))) | ~ r1(X127, X135)) | p1(X127))) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | ~ r1(sK13, X119))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15, sK16, sK17, sK18, sK19, sK20, sK21, sK22, sK23, sK24, sK25, sK26, sK27, sK28, sK29, sK30, sK31, sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49])], [f33, f70, f69, f68, f67, f66, f65, f64, f63, f62, f61, f60, f59, f58, f57, f56, f55, f54, f53, f52, f51, f50, f49, f48, f47, f46, f45, f44, f43, f42, f41, f40, f39, f38, f37, f36, f35, f34])).
fof(f34, plain, (? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(X0, X61)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | sP1(X78)) & ! [X82] : (? [X83] : (~ p1(X83) & r1(X82, X83)) | ! [X84] : (! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X82, X84)) | ~ r1(X78, X82)) & sP2(X78) & sP3(X78)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) & ! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : ((! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ! [X95] : (! [X96] : (p1(X96) | ~ r1(X95, X96)) | ~ r1(X93, X95)) | ~ r1(X92, X93)) & (? [X97] : (! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X97, X98)) & ~ p1(X97) & r1(X92, X97)) | ! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ~ r1(X92, X100)) | p1(X92))) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X0, X86)) & ! [X102] : (! [X103] : (! [X104] : (! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : ((! [X110] : (? [X111] : (~ p1(X111) & r1(X110, X111)) | ! [X112] : (! [X113] : (p1(X113) | ~ r1(X112, X113)) | ~ r1(X110, X112)) | ~ r1(X109, X110)) & (? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X109, X114)) | ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ~ r1(X109, X117)) | p1(X109))) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | ~ r1(X102, X103)) | ~ r1(X0, X102)) & ! [X119] : (! [X120] : (! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : ((! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ! [X130] : (! [X131] : (p1(X131) | ~ r1(X130, X131)) | ~ r1(X128, X130)) | ~ r1(X127, X128)) & (? [X132] : (! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X132, X133)) & ~ p1(X132) & r1(X127, X132)) | ! [X135] : (? [X136] : (~ p1(X136) & r1(X135, X136)) | ~ r1(X127, X135)) | p1(X127))) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | ~ r1(X0, X119))) => (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(sK13, X61)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | sP1(X78)) & ! [X82] : (? [X83] : (~ p1(X83) & r1(X82, X83)) | ! [X84] : (! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X82, X84)) | ~ r1(X78, X82)) & sP2(X78) & sP3(X78)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(sK13, X73)) & ! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : ((! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ! [X95] : (! [X96] : (p1(X96) | ~ r1(X95, X96)) | ~ r1(X93, X95)) | ~ r1(X92, X93)) & (? [X97] : (! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X97, X98)) & ~ p1(X97) & r1(X92, X97)) | ! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ~ r1(X92, X100)) | p1(X92))) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(sK13, X86)) & ! [X102] : (! [X103] : (! [X104] : (! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : ((! [X110] : (? [X111] : (~ p1(X111) & r1(X110, X111)) | ! [X112] : (! [X113] : (p1(X113) | ~ r1(X112, X113)) | ~ r1(X110, X112)) | ~ r1(X109, X110)) & (? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X109, X114)) | ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ~ r1(X109, X117)) | p1(X109))) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | ~ r1(X102, X103)) | ~ r1(sK13, X102)) & ! [X119] : (! [X120] : (! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : ((! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ! [X130] : (! [X131] : (p1(X131) | ~ r1(X130, X131)) | ~ r1(X128, X130)) | ~ r1(X127, X128)) & (? [X132] : (! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X132, X133)) & ~ p1(X132) & r1(X127, X132)) | ! [X135] : (? [X136] : (~ p1(X136) & r1(X135, X136)) | ~ r1(X127, X135)) | p1(X127))) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | ~ r1(sK13, X119)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK14(X2)) & r1(X2, sK14(X2)))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) => (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) => (~ p1(sK16(X9)) & r1(X9, sK16(X9)))), introduced(choice_axiom, [])).
fof(f38, plain, ! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) => (~ p1(sK17(X13)) & r1(X13, sK17(X13)))), introduced(choice_axiom, [])).
fof(f39, plain, ! [X12] : (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) => (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) => (~ p1(sK19(X20)) & r1(X20, sK19(X20)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) => (~ p1(sK20(X25)) & r1(X25, sK20(X25)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X24] : (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) => (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24)))), introduced(choice_axiom, [])).
fof(f43, plain, ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) => (~ p1(sK22(X32)) & r1(X32, sK22(X32)))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) => (~ p1(sK23(X38)) & r1(X38, sK23(X38)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X37] : (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) => (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37)))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) => (~ p1(sK25(X45)) & r1(X45, sK25(X45)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) => (~ p1(sK26(X52)) & r1(X52, sK26(X52)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X51] : (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) => (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51)))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) => (~ p1(sK28(X59)) & r1(X59, sK28(X59)))), introduced(choice_axiom, [])).
fof(f50, plain, (? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(sK13, X61)) => (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(sK29, X62)) & r1(sK13, sK29))), introduced(choice_axiom, [])).
fof(f51, plain, (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(sK29, X62)) => (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(sK30, X63)) & r1(sK29, sK30))), introduced(choice_axiom, [])).
fof(f52, plain, (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(sK30, X63)) => (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(sK31, X64)) & r1(sK30, sK31))), introduced(choice_axiom, [])).
fof(f53, plain, (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(sK31, X64)) => (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(sK32, X65)) & r1(sK31, sK32))), introduced(choice_axiom, [])).
fof(f54, plain, (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(sK32, X65)) => (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(sK33, X66)) & r1(sK32, sK33))), introduced(choice_axiom, [])).
fof(f55, plain, (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(sK33, X66)) => (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(sK34, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(sK34, X70)) & ? [X72] : (~ p1(X72) & r1(sK34, X72)) & r1(sK33, sK34))), introduced(choice_axiom, [])).
fof(f56, plain, ! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) => (p1(sK35(X67)) & ? [X69] : (~ p1(X69) & r1(sK35(X67), X69)) & r1(X67, sK35(X67)))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X67] : (? [X69] : (~ p1(X69) & r1(sK35(X67), X69)) => (~ p1(sK36(X67)) & r1(sK35(X67), sK36(X67)))), introduced(choice_axiom, [])).
fof(f58, plain, (? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(sK34, X70)) => (! [X71] : (p1(X71) | ~ r1(sK37, X71)) & r1(sK34, sK37))), introduced(choice_axiom, [])).
fof(f59, plain, (? [X72] : (~ p1(X72) & r1(sK34, X72)) => (~ p1(sK38) & r1(sK34, sK38))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X78] : (? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) => (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(sK39(X78), X80)) & ~ p1(sK39(X78)) & r1(X78, sK39(X78)))), introduced(choice_axiom, [])).
fof(f61, plain, ! [X82] : (? [X83] : (~ p1(X83) & r1(X82, X83)) => (~ p1(sK40(X82)) & r1(X82, sK40(X82)))), introduced(choice_axiom, [])).
fof(f62, plain, ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) => (~ p1(sK41(X93)) & r1(X93, sK41(X93)))), introduced(choice_axiom, [])).
fof(f63, plain, ! [X92] : (? [X97] : (! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X97, X98)) & ~ p1(X97) & r1(X92, X97)) => (! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(sK42(X92), X98)) & ~ p1(sK42(X92)) & r1(X92, sK42(X92)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) => (~ p1(sK43(X100)) & r1(X100, sK43(X100)))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X110] : (? [X111] : (~ p1(X111) & r1(X110, X111)) => (~ p1(sK44(X110)) & r1(X110, sK44(X110)))), introduced(choice_axiom, [])).
fof(f66, plain, ! [X109] : (? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X109, X114)) => (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(sK45(X109), X115)) & ~ p1(sK45(X109)) & r1(X109, sK45(X109)))), introduced(choice_axiom, [])).
fof(f67, plain, ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) => (~ p1(sK46(X117)) & r1(X117, sK46(X117)))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) => (~ p1(sK47(X128)) & r1(X128, sK47(X128)))), introduced(choice_axiom, [])).
fof(f69, plain, ! [X127] : (? [X132] : (! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X132, X133)) & ~ p1(X132) & r1(X127, X132)) => (! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(sK48(X127), X133)) & ~ p1(sK48(X127)) & r1(X127, sK48(X127)))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X135] : (? [X136] : (~ p1(X136) & r1(X135, X136)) => (~ p1(sK49(X135)) & r1(X135, sK49(X135)))), introduced(choice_axiom, [])).
fof(f33, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(X0, X61)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | sP1(X78)) & ! [X82] : (? [X83] : (~ p1(X83) & r1(X82, X83)) | ! [X84] : (! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X82, X84)) | ~ r1(X78, X82)) & sP2(X78) & sP3(X78)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) & ! [X86] : (! [X87] : (! [X88] : (! [X89] : (! [X90] : (! [X91] : (! [X92] : ((! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ! [X95] : (! [X96] : (p1(X96) | ~ r1(X95, X96)) | ~ r1(X93, X95)) | ~ r1(X92, X93)) & (? [X97] : (! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X97, X98)) & ~ p1(X97) & r1(X92, X97)) | ! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ~ r1(X92, X100)) | p1(X92))) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | ~ r1(X89, X90)) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | ~ r1(X86, X87)) | ~ r1(X0, X86)) & ! [X102] : (! [X103] : (! [X104] : (! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : ((! [X110] : (? [X111] : (~ p1(X111) & r1(X110, X111)) | ! [X112] : (! [X113] : (p1(X113) | ~ r1(X112, X113)) | ~ r1(X110, X112)) | ~ r1(X109, X110)) & (? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X109, X114)) | ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ~ r1(X109, X117)) | p1(X109))) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | ~ r1(X102, X103)) | ~ r1(X0, X102)) & ! [X119] : (! [X120] : (! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : ((! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ! [X130] : (! [X131] : (p1(X131) | ~ r1(X130, X131)) | ~ r1(X128, X130)) | ~ r1(X127, X128)) & (? [X132] : (! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X132, X133)) & ~ p1(X132) & r1(X127, X132)) | ! [X135] : (? [X136] : (~ p1(X136) & r1(X135, X136)) | ~ r1(X127, X135)) | p1(X127))) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | ~ r1(X0, X119))), inference(rectify, [], [f11])).
fof(f11, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(X0, X61)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | sP1(X78)) & ! [X86] : (? [X87] : (~ p1(X87) & r1(X86, X87)) | ! [X88] : (! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X86, X88)) | ~ r1(X78, X86)) & sP2(X78) & sP3(X78)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) & ! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : ((! [X112] : (? [X113] : (~ p1(X113) & r1(X112, X113)) | ! [X114] : (! [X115] : (p1(X115) | ~ r1(X114, X115)) | ~ r1(X112, X114)) | ~ r1(X111, X112)) & (? [X116] : (! [X117] : (~ p1(X117) | ! [X118] : (p1(X118) | ~ r1(X117, X118)) | ~ r1(X116, X117)) & ~ p1(X116) & r1(X111, X116)) | ! [X119] : (? [X120] : (~ p1(X120) & r1(X119, X120)) | ~ r1(X111, X119)) | p1(X111))) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X0, X105)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : ((! [X129] : (? [X130] : (~ p1(X130) & r1(X129, X130)) | ! [X131] : (! [X132] : (p1(X132) | ~ r1(X131, X132)) | ~ r1(X129, X131)) | ~ r1(X128, X129)) & (? [X133] : (! [X134] : (~ p1(X134) | ! [X135] : (p1(X135) | ~ r1(X134, X135)) | ~ r1(X133, X134)) & ~ p1(X133) & r1(X128, X133)) | ! [X136] : (? [X137] : (~ p1(X137) & r1(X136, X137)) | ~ r1(X128, X136)) | p1(X128))) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X138] : (! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : ((! [X147] : (? [X148] : (~ p1(X148) & r1(X147, X148)) | ! [X149] : (! [X150] : (p1(X150) | ~ r1(X149, X150)) | ~ r1(X147, X149)) | ~ r1(X146, X147)) & (? [X151] : (! [X152] : (~ p1(X152) | ! [X153] : (p1(X153) | ~ r1(X152, X153)) | ~ r1(X151, X152)) & ~ p1(X151) & r1(X146, X151)) | ! [X154] : (? [X155] : (~ p1(X155) & r1(X154, X155)) | ~ r1(X146, X154)) | p1(X146))) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X138, X139)) | ~ r1(X0, X138))), inference(definition_folding, [], [f6, e10, e9, e8, e7])).
fof(f7, plain, ! [X78] : (? [X95] : (! [X96] : ((p1(X96) & ? [X97] : (~ p1(X97) & r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) & p1(X95) & ? [X100] : (~ p1(X100) & r1(X95, X100)) & r1(X78, X95)) | ~ sP0(X78)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X78] : (sP0(X78) <=> ? [X95] : (! [X96] : ((p1(X96) & ? [X97] : (~ p1(X97) & r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) & p1(X95) & ? [X100] : (~ p1(X100) & r1(X95, X100)) & r1(X78, X95))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X78] : (! [X82] : (! [X83] : (? [X84] : (p1(X84) & ? [X85] : (~ p1(X85) & r1(X84, X85)) & r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82)) | ~ sP1(X78)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X78] : (sP1(X78) <=> ! [X82] : (! [X83] : (? [X84] : (p1(X84) & ? [X85] : (~ p1(X85) & r1(X84, X85)) & r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ! [X78] : (? [X90] : (! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) & ~ p1(X90) & r1(X78, X90)) | ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78) | ~ sP2(X78)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X78] : (sP2(X78) <=> (? [X90] : (! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) & ~ p1(X90) & r1(X78, X90)) | ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ! [X78] : (sP0(X78) | ! [X101] : (? [X102] : (p1(X102) & ? [X103] : (~ p1(X103) & r1(X102, X103)) & r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)) | ~ sP3(X78)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X78] : (sP3(X78) <=> (sP0(X78) | ! [X101] : (? [X102] : (p1(X102) & ? [X103] : (~ p1(X103) & r1(X102, X103)) & r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f6, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(X0, X61)) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | ! [X82] : (! [X83] : (? [X84] : (p1(X84) & ? [X85] : (~ p1(X85) & r1(X84, X85)) & r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82))) & ! [X86] : (? [X87] : (~ p1(X87) & r1(X86, X87)) | ! [X88] : (! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X86, X88)) | ~ r1(X78, X86)) & (? [X90] : (! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) & ~ p1(X90) & r1(X78, X90)) | ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78)) & (? [X95] : (! [X96] : ((p1(X96) & ? [X97] : (~ p1(X97) & r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) & p1(X95) & ? [X100] : (~ p1(X100) & r1(X95, X100)) & r1(X78, X95)) | ! [X101] : (? [X102] : (p1(X102) & ? [X103] : (~ p1(X103) & r1(X102, X103)) & r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)))) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) & ! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : ((! [X112] : (? [X113] : (~ p1(X113) & r1(X112, X113)) | ! [X114] : (! [X115] : (p1(X115) | ~ r1(X114, X115)) | ~ r1(X112, X114)) | ~ r1(X111, X112)) & (? [X116] : (! [X117] : (~ p1(X117) | ! [X118] : (p1(X118) | ~ r1(X117, X118)) | ~ r1(X116, X117)) & ~ p1(X116) & r1(X111, X116)) | ! [X119] : (? [X120] : (~ p1(X120) & r1(X119, X120)) | ~ r1(X111, X119)) | p1(X111))) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X0, X105)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : ((! [X129] : (? [X130] : (~ p1(X130) & r1(X129, X130)) | ! [X131] : (! [X132] : (p1(X132) | ~ r1(X131, X132)) | ~ r1(X129, X131)) | ~ r1(X128, X129)) & (? [X133] : (! [X134] : (~ p1(X134) | ! [X135] : (p1(X135) | ~ r1(X134, X135)) | ~ r1(X133, X134)) & ~ p1(X133) & r1(X128, X133)) | ! [X136] : (? [X137] : (~ p1(X137) & r1(X136, X137)) | ~ r1(X128, X136)) | p1(X128))) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X138] : (! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : ((! [X147] : (? [X148] : (~ p1(X148) & r1(X147, X148)) | ! [X149] : (! [X150] : (p1(X150) | ~ r1(X149, X150)) | ~ r1(X147, X149)) | ~ r1(X146, X147)) & (? [X151] : (! [X152] : (~ p1(X152) | ! [X153] : (p1(X153) | ~ r1(X152, X153)) | ~ r1(X151, X152)) & ~ p1(X151) & r1(X146, X151)) | ! [X154] : (? [X155] : (~ p1(X155) & r1(X154, X155)) | ~ r1(X146, X154)) | p1(X146))) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X138, X139)) | ~ r1(X0, X138))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : ((! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ? [X61] : (? [X62] : (? [X63] : (? [X64] : (? [X65] : (? [X66] : (! [X67] : (? [X68] : (p1(X68) & ? [X69] : (~ p1(X69) & r1(X68, X69)) & r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) & ? [X70] : (! [X71] : (p1(X71) | ~ r1(X70, X71)) & r1(X66, X70)) & ? [X72] : (~ p1(X72) & r1(X66, X72)) & r1(X65, X66)) & r1(X64, X65)) & r1(X63, X64)) & r1(X62, X63)) & r1(X61, X62)) & r1(X0, X61))) & ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((? [X79] : (! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) & ~ p1(X79) & r1(X78, X79)) | ! [X82] : (! [X83] : (? [X84] : (p1(X84) & ? [X85] : (~ p1(X85) & r1(X84, X85)) & r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82))) & ! [X86] : (? [X87] : (~ p1(X87) & r1(X86, X87)) | ! [X88] : (! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X86, X88)) | ~ r1(X78, X86)) & (? [X90] : (! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) & ~ p1(X90) & r1(X78, X90)) | ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78)) & (? [X95] : (! [X96] : ((p1(X96) & ? [X97] : (~ p1(X97) & r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) & p1(X95) & ? [X100] : (~ p1(X100) & r1(X95, X100)) & r1(X78, X95)) | ! [X101] : (? [X102] : (p1(X102) & ? [X103] : (~ p1(X103) & r1(X102, X103)) & r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)))) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) & ! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : ((! [X112] : (? [X113] : (~ p1(X113) & r1(X112, X113)) | ! [X114] : (! [X115] : (p1(X115) | ~ r1(X114, X115)) | ~ r1(X112, X114)) | ~ r1(X111, X112)) & (? [X116] : (! [X117] : (~ p1(X117) | ! [X118] : (p1(X118) | ~ r1(X117, X118)) | ~ r1(X116, X117)) & ~ p1(X116) & r1(X111, X116)) | ! [X119] : (? [X120] : (~ p1(X120) & r1(X119, X120)) | ~ r1(X111, X119)) | p1(X111))) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X0, X105)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : ((! [X129] : (? [X130] : (~ p1(X130) & r1(X129, X130)) | ! [X131] : (! [X132] : (p1(X132) | ~ r1(X131, X132)) | ~ r1(X129, X131)) | ~ r1(X128, X129)) & (? [X133] : (! [X134] : (~ p1(X134) | ! [X135] : (p1(X135) | ~ r1(X134, X135)) | ~ r1(X133, X134)) & ~ p1(X133) & r1(X128, X133)) | ! [X136] : (? [X137] : (~ p1(X137) & r1(X136, X137)) | ~ r1(X128, X136)) | p1(X128))) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X138] : (! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : ((! [X147] : (? [X148] : (~ p1(X148) & r1(X147, X148)) | ! [X149] : (! [X150] : (p1(X150) | ~ r1(X149, X150)) | ~ r1(X147, X149)) | ~ r1(X146, X147)) & (? [X151] : (! [X152] : (~ p1(X152) | ! [X153] : (p1(X153) | ~ r1(X152, X153)) | ~ r1(X151, X152)) & ~ p1(X151) & r1(X146, X151)) | ! [X154] : (? [X155] : (~ p1(X155) & r1(X154, X155)) | ~ r1(X146, X154)) | p1(X146))) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X138, X139)) | ~ r1(X0, X138))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ~ ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : (~ ! [X67] : (~ ! [X68] : (~ p1(X68) | ! [X69] : (p1(X69) | ~ r1(X68, X69)) | ~ r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) | ! [X70] : (~ ! [X71] : (p1(X71) | ~ r1(X70, X71)) | ~ r1(X66, X70)) | ! [X72] : (p1(X72) | ~ r1(X66, X72)) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61))) | ~ ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((~ ! [X79] : (~ ! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | p1(X79) | ~ r1(X78, X79)) | ! [X82] : (! [X83] : (~ ! [X84] : (~ p1(X84) | ! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82))) & ! [X86] : (~ ! [X87] : (p1(X87) | ~ r1(X86, X87)) | ! [X88] : (! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X86, X88)) | ~ r1(X78, X86)) & (~ ! [X90] : (~ ! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | p1(X90) | ~ r1(X78, X90)) | ! [X93] : (~ ! [X94] : (p1(X94) | ~ r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78)) & (~ ! [X95] : (~ ! [X96] : (~ (~ p1(X96) | ! [X97] : (p1(X97) | ~ r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) | ~ p1(X95) | ! [X100] : (p1(X100) | ~ r1(X95, X100)) | ~ r1(X78, X95)) | ! [X101] : (~ ! [X102] : (~ p1(X102) | ! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)))) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) | ~ ! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : ((! [X112] : (~ ! [X113] : (p1(X113) | ~ r1(X112, X113)) | ! [X114] : (! [X115] : (p1(X115) | ~ r1(X114, X115)) | ~ r1(X112, X114)) | ~ r1(X111, X112)) & (~ ! [X116] : (~ ! [X117] : (~ p1(X117) | ! [X118] : (p1(X118) | ~ r1(X117, X118)) | ~ r1(X116, X117)) | p1(X116) | ~ r1(X111, X116)) | ! [X119] : (~ ! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X111, X119)) | p1(X111))) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X0, X105)) | ~ ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : ((! [X129] : (~ ! [X130] : (p1(X130) | ~ r1(X129, X130)) | ! [X131] : (! [X132] : (p1(X132) | ~ r1(X131, X132)) | ~ r1(X129, X131)) | ~ r1(X128, X129)) & (~ ! [X133] : (~ ! [X134] : (~ p1(X134) | ! [X135] : (p1(X135) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | p1(X133) | ~ r1(X128, X133)) | ! [X136] : (~ ! [X137] : (p1(X137) | ~ r1(X136, X137)) | ~ r1(X128, X136)) | p1(X128))) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) | ~ ! [X138] : (! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : ((! [X147] : (~ ! [X148] : (p1(X148) | ~ r1(X147, X148)) | ! [X149] : (! [X150] : (p1(X150) | ~ r1(X149, X150)) | ~ r1(X147, X149)) | ~ r1(X146, X147)) & (~ ! [X151] : (~ ! [X152] : (~ p1(X152) | ! [X153] : (p1(X153) | ~ r1(X152, X153)) | ~ r1(X151, X152)) | p1(X151) | ~ r1(X146, X151)) | ! [X154] : (~ ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X146, X154)) | p1(X146))) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X138, X139)) | ~ r1(X0, X138))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ~ ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : (~ ! [X67] : (~ ! [X68] : (~ p1(X68) | ! [X69] : (p1(X69) | ~ r1(X68, X69)) | ~ r1(X67, X68)) | p1(X67) | ~ r1(X66, X67)) | ! [X70] : (~ ! [X71] : (p1(X71) | ~ r1(X70, X71)) | ~ r1(X66, X70)) | ! [X72] : (p1(X72) | ~ r1(X66, X72)) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61))) | ~ ! [X73] : (! [X74] : (! [X75] : (! [X76] : (! [X77] : (! [X78] : (((~ ! [X79] : (~ ! [X80] : (~ p1(X80) | ! [X81] : (p1(X81) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | p1(X79) | ~ r1(X78, X79)) | ! [X82] : (! [X83] : (~ ! [X84] : (~ p1(X84) | ! [X85] : (p1(X85) | ~ r1(X84, X85)) | ~ r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82))) & ! [X86] : (~ ! [X87] : (p1(X87) | ~ r1(X86, X87)) | ! [X88] : (! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X86, X88)) | ~ r1(X78, X86)) & (~ ! [X90] : (~ ! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) | p1(X90) | ~ r1(X78, X90)) | ! [X93] : (~ ! [X94] : (p1(X94) | ~ r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78)) & (~ ! [X95] : (~ ! [X96] : (~ (~ p1(X96) | ! [X97] : (p1(X97) | ~ r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) | ~ p1(X95) | ! [X100] : (p1(X100) | ~ r1(X95, X100)) | ~ r1(X78, X95)) | ! [X101] : (~ ! [X102] : (~ p1(X102) | ! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)))) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X75, X76)) | ~ r1(X74, X75)) | ~ r1(X73, X74)) | ~ r1(X0, X73)) | ~ ! [X105] : (! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : ((! [X112] : (~ ! [X113] : (p1(X113) | ~ r1(X112, X113)) | ! [X114] : (! [X115] : (p1(X115) | ~ r1(X114, X115)) | ~ r1(X112, X114)) | ~ r1(X111, X112)) & (~ ! [X116] : (~ ! [X117] : (~ p1(X117) | ! [X118] : (p1(X118) | ~ r1(X117, X118)) | ~ r1(X116, X117)) | p1(X116) | ~ r1(X111, X116)) | ! [X119] : (~ ! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X111, X119)) | p1(X111))) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X105, X106)) | ~ r1(X0, X105)) | ~ ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : ((! [X129] : (~ ! [X130] : (p1(X130) | ~ r1(X129, X130)) | ! [X131] : (! [X132] : (p1(X132) | ~ r1(X131, X132)) | ~ r1(X129, X131)) | ~ r1(X128, X129)) & (~ ! [X133] : (~ ! [X134] : (~ p1(X134) | ! [X135] : (p1(X135) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | p1(X133) | ~ r1(X128, X133)) | ! [X136] : (~ ! [X137] : (p1(X137) | ~ r1(X136, X137)) | ~ r1(X128, X136)) | p1(X128))) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) | ~ ! [X138] : (! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : ((! [X147] : (~ ! [X148] : (p1(X148) | ~ r1(X147, X148)) | ! [X149] : (! [X150] : (p1(X150) | ~ r1(X149, X150)) | ~ r1(X147, X149)) | ~ r1(X146, X147)) & (~ ! [X151] : (~ ! [X152] : (~ p1(X152) | ! [X153] : (p1(X153) | ~ r1(X152, X153)) | ~ r1(X151, X152)) | p1(X151) | ~ r1(X146, X151)) | ! [X154] : (~ ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X146, X154)) | p1(X146))) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X138, X139)) | ~ r1(X0, X138))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) & ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1))) | ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) & ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1))) | ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL640+1.010.p', main)).
fof(f91675, plain, (~ r1(sK34, sK37) | ~ spl50_2241), inference(resolution, [], [f91651, f129])).
fof(f129, plain, r1(sK33, sK34), inference(cnf_transformation, [], [f71])).
fof(f91651, plain, (! [X0] : (~ r1(sK33, X0) | ~ r1(X0, sK37)) | ~ spl50_2241), inference(resolution, [], [f91629, f128])).
fof(f128, plain, r1(sK32, sK33), inference(cnf_transformation, [], [f71])).
fof(f91629, plain, (! [X6, X7] : (~ r1(sK32, X7) | ~ r1(X6, sK37) | ~ r1(X7, X6)) | ~ spl50_2241), inference(resolution, [], [f91577, f127])).
fof(f127, plain, r1(sK31, sK32), inference(cnf_transformation, [], [f71])).
fof(f91577, plain, (! [X2, X0, X1] : (~ r1(sK31, X1) | ~ r1(X0, sK37) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl50_2241), inference(resolution, [], [f91555, f126])).
fof(f126, plain, r1(sK30, sK31), inference(cnf_transformation, [], [f71])).
fof(f91555, plain, (! [X2, X0, X3, X1] : (~ r1(sK30, X1) | ~ r1(X0, sK37) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl50_2241), inference(resolution, [], [f91533, f125])).
fof(f125, plain, r1(sK29, sK30), inference(cnf_transformation, [], [f71])).
fof(f91533, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK29, X1) | ~ r1(X0, sK37) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X0)) | ~ spl50_2241), inference(resolution, [], [f28420, f124])).
fof(f124, plain, r1(sK13, sK29), inference(cnf_transformation, [], [f71])).
fof(f28420, plain, (! [X6, X4, X2, X7, X5, X3] : (~ r1(sK13, X7) | ~ r1(X2, sK37) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2)) | ~ spl50_2241), inference(avatar_component_clause, [], [f28419])).
fof(f28419, plain, (spl50_2241 <=> ! [X3, X5, X7, X2, X4, X6] : (~ r1(X2, sK37) | ~ r1(sK13, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2))), introduced(avatar_definition, [new_symbols(naming, [spl50_2241])])).
fof(f42718, plain, ~ spl50_2282, inference(avatar_contradiction_clause, [], [f42717])).
fof(f42717, plain, ($false | ~ spl50_2282), inference(subsumption_resolution, [], [f42699, f131])).
fof(f131, plain, ~ p1(sK38), inference(cnf_transformation, [], [f71])).
fof(f42699, plain, (p1(sK38) | ~ spl50_2282), inference(resolution, [], [f29126, f130])).
fof(f130, plain, r1(sK34, sK38), inference(cnf_transformation, [], [f71])).
fof(f29126, plain, (! [X0] : (~ r1(sK34, X0) | p1(X0)) | ~ spl50_2282), inference(avatar_component_clause, [], [f29125])).
fof(f29125, plain, (spl50_2282 <=> ! [X0] : (p1(X0) | ~ r1(sK34, X0))), introduced(avatar_definition, [new_symbols(naming, [spl50_2282])])).
fof(f41336, plain, ~ spl50_429, inference(avatar_contradiction_clause, [], [f41335])).
fof(f41335, plain, ($false | ~ spl50_429), inference(subsumption_resolution, [], [f41317, f129])).
fof(f41317, plain, (~ r1(sK33, sK34) | ~ spl50_429), inference(resolution, [], [f41262, f128])).
fof(f41262, plain, (! [X3] : (~ r1(sK32, X3) | ~ r1(X3, sK34)) | ~ spl50_429), inference(resolution, [], [f41240, f127])).
fof(f41240, plain, (! [X0, X1] : (~ r1(sK31, X1) | ~ r1(X0, sK34) | ~ r1(X1, X0)) | ~ spl50_429), inference(resolution, [], [f41191, f126])).
fof(f41191, plain, (! [X2, X0, X1] : (~ r1(sK30, X1) | ~ r1(X0, sK34) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl50_429), inference(resolution, [], [f41174, f125])).
fof(f41174, plain, (! [X2, X0, X3, X1] : (~ r1(sK29, X1) | ~ r1(X0, sK34) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl50_429), inference(resolution, [], [f5560, f124])).
fof(f5560, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK13, X4) | ~ r1(X0, sK34) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0)) | ~ spl50_429), inference(avatar_component_clause, [], [f5559])).
fof(f5559, plain, (spl50_429 <=> ! [X1, X3, X0, X2, X4] : (~ r1(X0, sK34) | ~ r1(sK13, X4) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl50_429])])).
fof(f35165, plain, (spl50_2282 | ~ spl50_20 | spl50_64 | ~ spl50_63 | ~ spl50_88 | ~ spl50_89 | ~ spl50_2240), inference(avatar_split_clause, [], [f35164, f28353, f934, f929, f521, f525, f275, f29125])).
fof(f275, plain, (spl50_20 <=> p1(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl50_20])])).
fof(f525, plain, (spl50_64 <=> sP0(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl50_64])])).
fof(f521, plain, (spl50_63 <=> sP3(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl50_63])])).
fof(f929, plain, (spl50_88 <=> ! [X12] : (r1(sK4(X12), sK5(X12)) | ~ r1(sK34, X12))), introduced(avatar_definition, [new_symbols(naming, [spl50_88])])).
fof(f934, plain, (spl50_89 <=> ! [X12] : (r1(X12, sK4(X12)) | ~ r1(sK34, X12))), introduced(avatar_definition, [new_symbols(naming, [spl50_89])])).
fof(f28353, plain, (spl50_2240 <=> ! [X3, X2] : (~ r1(sK37, X2) | p1(X3) | ~ r1(X2, X3))), introduced(avatar_definition, [new_symbols(naming, [spl50_2240])])).
fof(f35164, plain, (! [X0] : (sP0(sK34) | ~ p1(sK34) | p1(X0) | ~ r1(sK34, X0)) | (~ spl50_63 | ~ spl50_88 | ~ spl50_89 | ~ spl50_2240)), inference(subsumption_resolution, [], [f29075, f522])).
fof(f522, plain, (sP3(sK34) | ~ spl50_63), inference(avatar_component_clause, [], [f521])).
fof(f29075, plain, (! [X0] : (sP0(sK34) | ~ p1(sK34) | p1(X0) | ~ r1(sK34, X0) | ~ sP3(sK34)) | (~ spl50_88 | ~ spl50_89 | ~ spl50_2240)), inference(resolution, [], [f29042, f132])).
fof(f29042, plain, (! [X0, X1] : (~ r1(X0, sK37) | sP0(X0) | ~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ sP3(X0)) | (~ spl50_88 | ~ spl50_89 | ~ spl50_2240)), inference(resolution, [], [f28976, f74])).
fof(f74, plain, ! [X4, X0, X1] : (~ p1(sK5(X1)) | sP0(X0) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | ~ r1(X0, X4) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (sP0(X0) | ! [X1] : ((p1(sK4(X1)) & (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1))) & r1(X1, sK4(X1))) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f13, f15, f14])).
fof(f14, plain, ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) => (p1(sK4(X1)) & ? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) & r1(X1, sK4(X1)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X1] : (? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) => (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1)))), introduced(choice_axiom, [])).
fof(f13, plain, ! [X0] : (sP0(X0) | ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(rectify, [], [f12])).
fof(f12, plain, ! [X78] : (sP0(X78) | ! [X101] : (? [X102] : (p1(X102) & ? [X103] : (~ p1(X103) & r1(X102, X103)) & r1(X101, X102)) | ~ r1(X78, X101)) | ~ p1(X78) | ! [X104] : (p1(X104) | ~ r1(X78, X104)) | ~ sP3(X78)), inference(nnf_transformation, [], [f10])).
fof(f28976, plain, (p1(sK5(sK37)) | (~ spl50_88 | ~ spl50_89 | ~ spl50_2240)), inference(subsumption_resolution, [], [f28959, f132])).
fof(f28959, plain, (p1(sK5(sK37)) | ~ r1(sK34, sK37) | (~ spl50_88 | ~ spl50_89 | ~ spl50_2240)), inference(resolution, [], [f28685, f930])).
fof(f930, plain, (! [X12] : (r1(sK4(X12), sK5(X12)) | ~ r1(sK34, X12)) | ~ spl50_88), inference(avatar_component_clause, [], [f929])).
fof(f28685, plain, (! [X2] : (~ r1(sK4(sK37), X2) | p1(X2)) | (~ spl50_89 | ~ spl50_2240)), inference(subsumption_resolution, [], [f28670, f132])).
fof(f28670, plain, (! [X2] : (p1(X2) | ~ r1(sK4(sK37), X2) | ~ r1(sK34, sK37)) | (~ spl50_89 | ~ spl50_2240)), inference(resolution, [], [f28354, f935])).
fof(f935, plain, (! [X12] : (r1(X12, sK4(X12)) | ~ r1(sK34, X12)) | ~ spl50_89), inference(avatar_component_clause, [], [f934])).
fof(f28354, plain, (! [X2, X3] : (~ r1(sK37, X2) | p1(X3) | ~ r1(X2, X3)) | ~ spl50_2240), inference(avatar_component_clause, [], [f28353])).
fof(f35138, plain, (spl50_20 | ~ spl50_42 | ~ spl50_101 | ~ spl50_2536), inference(avatar_split_clause, [], [f35137, f34390, f1055, f402, f275])).
fof(f402, plain, (spl50_42 <=> sP2(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl50_42])])).
fof(f1055, plain, (spl50_101 <=> p1(sK6(sK34))), introduced(avatar_definition, [new_symbols(naming, [spl50_101])])).
fof(f34390, plain, (spl50_2536 <=> p1(sK7(sK37))), introduced(avatar_definition, [new_symbols(naming, [spl50_2536])])).
fof(f35137, plain, (p1(sK34) | (~ spl50_42 | ~ spl50_101 | ~ spl50_2536)), inference(subsumption_resolution, [], [f35133, f403])).
fof(f403, plain, (sP2(sK34) | ~ spl50_42), inference(avatar_component_clause, [], [f402])).
fof(f35133, plain, (p1(sK34) | ~ sP2(sK34) | (~ spl50_101 | ~ spl50_2536)), inference(subsumption_resolution, [], [f35131, f132])).
fof(f35131, plain, (~ r1(sK34, sK37) | p1(sK34) | ~ sP2(sK34) | (~ spl50_101 | ~ spl50_2536)), inference(resolution, [], [f1057, f34428])).
fof(f34428, plain, (! [X1] : (~ p1(sK6(X1)) | ~ r1(X1, sK37) | p1(X1) | ~ sP2(X1)) | ~ spl50_2536), inference(resolution, [], [f34392, f79])).
fof(f79, plain, ! [X4, X0] : (~ p1(sK7(X4)) | ~ p1(sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : ((! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0))) | ! [X4] : ((~ p1(sK7(X4)) & r1(X4, sK7(X4))) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7])], [f18, f20, f19])).
fof(f19, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) => (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) => (~ p1(sK7(X4)) & r1(X4, sK7(X4)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) | ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X78] : (? [X90] : (! [X91] : (~ p1(X91) | ! [X92] : (p1(X92) | ~ r1(X91, X92)) | ~ r1(X90, X91)) & ~ p1(X90) & r1(X78, X90)) | ! [X93] : (? [X94] : (~ p1(X94) & r1(X93, X94)) | ~ r1(X78, X93)) | p1(X78) | ~ sP2(X78)), inference(nnf_transformation, [], [f9])).
fof(f34392, plain, (p1(sK7(sK37)) | ~ spl50_2536), inference(avatar_component_clause, [], [f34390])).
fof(f1057, plain, (p1(sK6(sK34)) | ~ spl50_101), inference(avatar_component_clause, [], [f1055])).
fof(f35063, plain, (spl50_20 | ~ spl50_42 | spl50_43 | ~ spl50_2536), inference(avatar_contradiction_clause, [], [f35062])).
fof(f35062, plain, ($false | (spl50_20 | ~ spl50_42 | spl50_43 | ~ spl50_2536)), inference(subsumption_resolution, [], [f35061, f403])).
fof(f35061, plain, (~ sP2(sK34) | (spl50_20 | spl50_43 | ~ spl50_2536)), inference(subsumption_resolution, [], [f35060, f276])).
fof(f276, plain, (~ p1(sK34) | spl50_20), inference(avatar_component_clause, [], [f275])).
fof(f35060, plain, (p1(sK34) | ~ sP2(sK34) | (spl50_43 | ~ spl50_2536)), inference(subsumption_resolution, [], [f34442, f132])).
fof(f34442, plain, (~ r1(sK34, sK37) | p1(sK34) | ~ sP2(sK34) | (spl50_43 | ~ spl50_2536)), inference(resolution, [], [f34427, f407])).
fof(f407, plain, (~ r1(sK34, sK6(sK34)) | spl50_43), inference(avatar_component_clause, [], [f406])).
fof(f406, plain, (spl50_43 <=> r1(sK34, sK6(sK34))), introduced(avatar_definition, [new_symbols(naming, [spl50_43])])).
fof(f34427, plain, (! [X0] : (r1(X0, sK6(X0)) | ~ r1(X0, sK37) | p1(X0) | ~ sP2(X0)) | ~ spl50_2536), inference(resolution, [], [f34392, f77])).
fof(f77, plain, ! [X4, X0] : (~ p1(sK7(X4)) | r1(X0, sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f34426, plain, (spl50_2536 | ~ spl50_45), inference(avatar_split_clause, [], [f30658, f415, f34390])).
fof(f415, plain, (spl50_45 <=> r1(sK37, sK7(sK37))), introduced(avatar_definition, [new_symbols(naming, [spl50_45])])).
fof(f30658, plain, (p1(sK7(sK37)) | ~ spl50_45), inference(resolution, [], [f417, f133])).
fof(f133, plain, ! [X71] : (~ r1(sK37, X71) | p1(X71)), inference(cnf_transformation, [], [f71])).
fof(f417, plain, (r1(sK37, sK7(sK37)) | ~ spl50_45), inference(avatar_component_clause, [], [f415])).
fof(f32073, plain, (~ spl50_43 | spl50_2289 | spl50_20 | ~ spl50_42 | ~ spl50_100 | spl50_2336), inference(avatar_split_clause, [], [f32072, f29440, f1051, f402, f275, f29159, f406])).
fof(f29159, plain, (spl50_2289 <=> ! [X5] : (~ p1(sK7(X5)) | ~ r1(sK34, X5))), introduced(avatar_definition, [new_symbols(naming, [spl50_2289])])).
fof(f1051, plain, (spl50_100 <=> p1(sK35(sK6(sK34)))), introduced(avatar_definition, [new_symbols(naming, [spl50_100])])).
fof(f29440, plain, (spl50_2336 <=> p1(sK36(sK6(sK34)))), introduced(avatar_definition, [new_symbols(naming, [spl50_2336])])).
fof(f32072, plain, (! [X3] : (~ p1(sK7(X3)) | ~ r1(sK34, X3) | ~ r1(sK34, sK6(sK34))) | (spl50_20 | ~ spl50_42 | ~ spl50_100 | spl50_2336)), inference(subsumption_resolution, [], [f29445, f29441])).
fof(f29441, plain, (~ p1(sK36(sK6(sK34))) | spl50_2336), inference(avatar_component_clause, [], [f29440])).
fof(f29445, plain, (! [X3] : (p1(sK36(sK6(sK34))) | ~ p1(sK7(X3)) | ~ r1(sK34, X3) | ~ r1(sK34, sK6(sK34))) | (spl50_20 | ~ spl50_42 | ~ spl50_100)), inference(subsumption_resolution, [], [f29444, f403])).
fof(f29444, plain, (! [X3] : (p1(sK36(sK6(sK34))) | ~ p1(sK7(X3)) | ~ r1(sK34, X3) | ~ sP2(sK34) | ~ r1(sK34, sK6(sK34))) | (spl50_20 | ~ spl50_100)), inference(subsumption_resolution, [], [f29435, f276])).
fof(f29435, plain, (! [X3] : (p1(sK36(sK6(sK34))) | ~ p1(sK7(X3)) | ~ r1(sK34, X3) | p1(sK34) | ~ sP2(sK34) | ~ r1(sK34, sK6(sK34))) | ~ spl50_100), inference(resolution, [], [f1053, f1352])).
fof(f1352, plain, ! [X0, X1] : (~ p1(sK35(sK6(X0))) | p1(sK36(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0))), inference(subsumption_resolution, [], [f1351, f79])).
fof(f1351, plain, ! [X0, X1] : (p1(sK36(sK6(X0))) | ~ p1(sK35(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f1345])).
fof(f1345, plain, ! [X0, X1] : (p1(sK36(sK6(X0))) | ~ p1(sK35(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK34, sK6(X0))), inference(resolution, [], [f616, f135])).
fof(f135, plain, ! [X67] : (r1(sK35(X67), sK36(X67)) | p1(X67) | ~ r1(sK34, X67)), inference(cnf_transformation, [], [f71])).
fof(f616, plain, ! [X4, X5, X3] : (~ r1(sK35(sK6(X4)), X3) | p1(X3) | ~ p1(sK35(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK34, sK6(X4))), inference(subsumption_resolution, [], [f614, f79])).
fof(f614, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK35(sK6(X4)), X3) | ~ p1(sK35(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK34, sK6(X4))), inference(resolution, [], [f81, f134])).
fof(f134, plain, ! [X67] : (r1(X67, sK35(X67)) | p1(X67) | ~ r1(sK34, X67)), inference(cnf_transformation, [], [f71])).
fof(f81, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | ~ p1(sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f1053, plain, (p1(sK35(sK6(sK34))) | ~ spl50_100), inference(avatar_component_clause, [], [f1051])).
fof(f32049, plain, (~ spl50_2287 | ~ spl50_2289), inference(avatar_contradiction_clause, [], [f32048])).
fof(f32048, plain, ($false | (~ spl50_2287 | ~ spl50_2289)), inference(subsumption_resolution, [], [f32044, f132])).
fof(f32044, plain, (~ r1(sK34, sK37) | (~ spl50_2287 | ~ spl50_2289)), inference(resolution, [], [f29160, f30643])).
fof(f30643, plain, (p1(sK7(sK37)) | ~ spl50_2287), inference(subsumption_resolution, [], [f30402, f132])).
fof(f30402, plain, (~ r1(sK34, sK37) | p1(sK7(sK37)) | ~ spl50_2287), inference(resolution, [], [f29150, f133])).
fof(f29150, plain, (! [X5] : (r1(X5, sK7(X5)) | ~ r1(sK34, X5)) | ~ spl50_2287), inference(avatar_component_clause, [], [f29149])).
fof(f29149, plain, (spl50_2287 <=> ! [X5] : (r1(X5, sK7(X5)) | ~ r1(sK34, X5))), introduced(avatar_definition, [new_symbols(naming, [spl50_2287])])).
fof(f29160, plain, (! [X5] : (~ p1(sK7(X5)) | ~ r1(sK34, X5)) | ~ spl50_2289), inference(avatar_component_clause, [], [f29159])).
fof(f30638, plain, (spl50_45 | ~ spl50_2287), inference(avatar_contradiction_clause, [], [f30637])).
fof(f30637, plain, ($false | (spl50_45 | ~ spl50_2287)), inference(subsumption_resolution, [], [f30049, f132])).
fof(f30049, plain, (~ r1(sK34, sK37) | (spl50_45 | ~ spl50_2287)), inference(resolution, [], [f29150, f416])).
fof(f416, plain, (~ r1(sK37, sK7(sK37)) | spl50_45), inference(avatar_component_clause, [], [f415])).
fof(f29486, plain, (spl50_101 | ~ spl50_43 | ~ spl50_2336), inference(avatar_split_clause, [], [f29485, f29440, f406, f1055])).
fof(f29485, plain, (p1(sK6(sK34)) | (~ spl50_43 | ~ spl50_2336)), inference(subsumption_resolution, [], [f29481, f408])).
fof(f408, plain, (r1(sK34, sK6(sK34)) | ~ spl50_43), inference(avatar_component_clause, [], [f406])).
fof(f29481, plain, (p1(sK6(sK34)) | ~ r1(sK34, sK6(sK34)) | ~ spl50_2336), inference(resolution, [], [f29442, f136])).
fof(f136, plain, ! [X67] : (~ p1(sK36(X67)) | p1(X67) | ~ r1(sK34, X67)), inference(cnf_transformation, [], [f71])).
fof(f29442, plain, (p1(sK36(sK6(sK34))) | ~ spl50_2336), inference(avatar_component_clause, [], [f29440])).
fof(f29443, plain, (spl50_2287 | spl50_2336 | spl50_20 | ~ spl50_42 | ~ spl50_43 | ~ spl50_100), inference(avatar_split_clause, [], [f29438, f1051, f406, f402, f275, f29440, f29149])).
fof(f29438, plain, (! [X2] : (p1(sK36(sK6(sK34))) | r1(X2, sK7(X2)) | ~ r1(sK34, X2)) | (spl50_20 | ~ spl50_42 | ~ spl50_43 | ~ spl50_100)), inference(subsumption_resolution, [], [f29437, f408])).
fof(f29437, plain, (! [X2] : (p1(sK36(sK6(sK34))) | r1(X2, sK7(X2)) | ~ r1(sK34, X2) | ~ r1(sK34, sK6(sK34))) | (spl50_20 | ~ spl50_42 | ~ spl50_100)), inference(subsumption_resolution, [], [f29436, f403])).
fof(f29436, plain, (! [X2] : (p1(sK36(sK6(sK34))) | r1(X2, sK7(X2)) | ~ r1(sK34, X2) | ~ sP2(sK34) | ~ r1(sK34, sK6(sK34))) | (spl50_20 | ~ spl50_100)), inference(subsumption_resolution, [], [f29434, f276])).
fof(f29434, plain, (! [X2] : (p1(sK36(sK6(sK34))) | r1(X2, sK7(X2)) | ~ r1(sK34, X2) | p1(sK34) | ~ sP2(sK34) | ~ r1(sK34, sK6(sK34))) | ~ spl50_100), inference(resolution, [], [f1053, f1709])).
fof(f1709, plain, ! [X0, X1] : (~ p1(sK35(sK6(X0))) | p1(sK36(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0))), inference(subsumption_resolution, [], [f1708, f78])).
fof(f78, plain, ! [X4, X0] : (~ p1(sK6(X0)) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f1708, plain, ! [X0, X1] : (p1(sK36(sK6(X0))) | ~ p1(sK35(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f1701])).
fof(f1701, plain, ! [X0, X1] : (p1(sK36(sK6(X0))) | ~ p1(sK35(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK34, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK34, sK6(X0))), inference(resolution, [], [f634, f135])).
fof(f634, plain, ! [X4, X5, X3] : (~ r1(sK35(sK6(X4)), X3) | p1(X3) | ~ p1(sK35(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK34, sK6(X4))), inference(subsumption_resolution, [], [f632, f78])).
fof(f632, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK35(sK6(X4)), X3) | ~ p1(sK35(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK34, sK6(X4))), inference(resolution, [], [f80, f134])).
fof(f80, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f29231, plain, (spl50_2287 | spl50_20 | ~ spl50_42 | ~ spl50_101), inference(avatar_split_clause, [], [f29230, f1055, f402, f275, f29149])).
fof(f29230, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK34, X0)) | (spl50_20 | ~ spl50_42 | ~ spl50_101)), inference(subsumption_resolution, [], [f29229, f403])).
fof(f29229, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK34, X0) | ~ sP2(sK34)) | (spl50_20 | ~ spl50_101)), inference(subsumption_resolution, [], [f29228, f276])).
fof(f29228, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK34, X0) | p1(sK34) | ~ sP2(sK34)) | ~ spl50_101), inference(resolution, [], [f1057, f78])).
fof(f28421, plain, (spl50_2241 | spl50_2240 | ~ spl50_2239), inference(avatar_split_clause, [], [f28417, f28349, f28353, f28419])).
fof(f28349, plain, (spl50_2239 <=> r1(sK37, sK40(sK37))), introduced(avatar_definition, [new_symbols(naming, [spl50_2239])])).
fof(f28417, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (p1(X0) | ~ r1(X1, X0) | ~ r1(sK37, X1) | ~ r1(X2, sK37) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7)) | ~ spl50_2239), inference(resolution, [], [f28381, f120])).
fof(f120, plain, ! [X78, X76, X74, X85, X77, X75, X73, X84, X82] : (~ p1(sK40(X82)) | p1(X85) | ~ r1(X84, X85) | ~ r1(X82, X84) | ~ r1(X78, X82) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | ~ r1(sK13, X73)), inference(cnf_transformation, [], [f71])).
fof(f28381, plain, (p1(sK40(sK37)) | ~ spl50_2239), inference(resolution, [], [f28351, f133])).
fof(f28351, plain, (r1(sK37, sK40(sK37)) | ~ spl50_2239), inference(avatar_component_clause, [], [f28349])).
fof(f28355, plain, (spl50_2239 | spl50_2240), inference(avatar_split_clause, [], [f28324, f28353, f28349])).
fof(f28324, plain, ! [X2, X3] : (~ r1(sK37, X2) | ~ r1(X2, X3) | p1(X3) | r1(sK37, sK40(sK37))), inference(resolution, [], [f28306, f132])).
fof(f28306, plain, ! [X2, X0, X1] : (~ r1(sK34, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | p1(X2) | r1(X0, sK40(X0))), inference(resolution, [], [f11511, f129])).
fof(f11511, plain, ! [X2, X0, X3, X1] : (~ r1(sK33, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | p1(X3) | r1(X1, sK40(X1))), inference(resolution, [], [f5191, f128])).
fof(f5191, plain, ! [X19, X17, X15, X18, X16] : (~ r1(sK32, X17) | ~ r1(X17, X15) | ~ r1(X15, X16) | ~ r1(X16, X18) | ~ r1(X18, X19) | p1(X19) | r1(X16, sK40(X16))), inference(resolution, [], [f3660, f127])).
fof(f3660, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK31, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | p1(X5) | r1(X0, sK40(X0))), inference(resolution, [], [f1828, f126])).
fof(f1828, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK30, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | p1(X1) | r1(X2, sK40(X2))), inference(resolution, [], [f835, f125])).
fof(f835, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK29, X7) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | p1(X0) | r1(X2, sK40(X2))), inference(resolution, [], [f119, f124])).
fof(f119, plain, ! [X78, X76, X74, X85, X77, X75, X73, X84, X82] : (~ r1(sK13, X73) | p1(X85) | ~ r1(X84, X85) | ~ r1(X82, X84) | ~ r1(X78, X82) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | r1(X82, sK40(X82))), inference(cnf_transformation, [], [f71])).
fof(f16030, plain, (~ spl50_64 | ~ spl50_430), inference(avatar_split_clause, [], [f16027, f5562, f525])).
fof(f5562, plain, (spl50_430 <=> sP1(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl50_430])])).
fof(f16027, plain, (~ sP0(sK34) | ~ spl50_430), inference(resolution, [], [f5564, f729])).
fof(f729, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f728])).
fof(f728, plain, ! [X0] : (~ sP0(X0) | ~ sP1(X0) | ~ sP1(X0) | ~ sP0(X0)), inference(resolution, [], [f724, f86])).
fof(f86, plain, ! [X0] : (r1(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((! [X2] : ((p1(X2) & (~ p1(sK11(X2)) & r1(X2, sK11(X2)))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0))) & r1(X0, sK10(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12])], [f28, f31, f30, f29])).
fof(f29, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) => (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & ? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) & r1(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK11(X2)) & r1(X2, sK11(X2)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) => (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f27])).
fof(f27, plain, ! [X78] : (? [X95] : (! [X96] : ((p1(X96) & ? [X97] : (~ p1(X97) & r1(X96, X97))) | ! [X98] : (~ p1(X98) | ! [X99] : (p1(X99) | ~ r1(X98, X99)) | ~ r1(X96, X98)) | ~ r1(X95, X96)) & p1(X95) & ? [X100] : (~ p1(X100) & r1(X95, X100)) & r1(X78, X95)) | ~ sP0(X78)), inference(nnf_transformation, [], [f7])).
fof(f724, plain, ! [X0, X1] : (~ r1(X1, sK10(X0)) | ~ sP0(X0) | ~ sP1(X0) | ~ sP1(X1)), inference(duplicate_literal_removal, [], [f723])).
fof(f723, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | ~ r1(X1, sK10(X0)) | ~ sP1(X1) | ~ sP0(X0)), inference(resolution, [], [f722, f87])).
fof(f87, plain, ! [X0] : (r1(sK10(X0), sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f722, plain, ! [X2, X0, X1] : (~ r1(X1, sK12(X0)) | ~ sP1(X0) | ~ sP0(X0) | ~ r1(X2, X1) | ~ sP1(X2)), inference(subsumption_resolution, [], [f721, f88])).
fof(f88, plain, ! [X0] : (~ p1(sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f721, plain, ! [X2, X0, X1] : (~ sP0(X0) | ~ sP1(X0) | p1(sK12(X0)) | ~ r1(X1, sK12(X0)) | ~ r1(X2, X1) | ~ sP1(X2)), inference(resolution, [], [f720, f84])).
fof(f84, plain, ! [X2, X0, X1] : (~ p1(sK9(X2)) | p1(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (! [X1] : (! [X2] : ((p1(sK8(X2)) & (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2))) & r1(X2, sK8(X2))) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9])], [f23, f25, f24])).
fof(f24, plain, ! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) => (p1(sK8(X2)) & ? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) & r1(X2, sK8(X2)))), introduced(choice_axiom, [])).
fof(f25, plain, ! [X2] : (? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) => (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f22])).
fof(f22, plain, ! [X78] : (! [X82] : (! [X83] : (? [X84] : (p1(X84) & ? [X85] : (~ p1(X85) & r1(X84, X85)) & r1(X83, X84)) | p1(X83) | ~ r1(X82, X83)) | ~ r1(X78, X82)) | ~ sP1(X78)), inference(nnf_transformation, [], [f8])).
fof(f720, plain, ! [X0] : (p1(sK9(sK12(X0))) | ~ sP0(X0) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f712])).
fof(f712, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0) | p1(sK9(sK12(X0))) | ~ sP1(X0)), inference(resolution, [], [f707, f664])).
fof(f664, plain, ! [X0, X1] : (~ r1(sK8(sK12(X0)), X1) | ~ sP0(X0) | p1(X1) | ~ sP1(X0)), inference(subsumption_resolution, [], [f663, f618])).
fof(f618, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f617])).
fof(f617, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f206, f86])).
fof(f206, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | p1(sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f196, f88])).
fof(f196, plain, ! [X4, X5] : (p1(sK12(X4)) | p1(sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f85, f87])).
fof(f85, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | p1(sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f663, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0)))), inference(duplicate_literal_removal, [], [f655])).
fof(f655, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0))) | ~ sP0(X0)), inference(resolution, [], [f651, f573])).
fof(f573, plain, ! [X2, X0, X1] : (~ r1(sK12(X2), X0) | p1(X1) | ~ r1(X0, X1) | ~ p1(X0) | ~ sP0(X2)), inference(subsumption_resolution, [], [f572, f88])).
fof(f572, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2)), inference(duplicate_literal_removal, [], [f569])).
fof(f569, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2) | ~ sP0(X2)), inference(resolution, [], [f92, f87])).
fof(f92, plain, ! [X4, X2, X0, X5] : (~ r1(sK10(X0), X2) | ~ p1(X4) | p1(X5) | ~ r1(X4, X5) | ~ r1(X2, X4) | p1(X2) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f651, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f650])).
fof(f650, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f302, f86])).
fof(f302, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f292, f88])).
fof(f292, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f82, f87])).
fof(f82, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(X2, sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f707, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f706])).
fof(f706, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f431, f86])).
fof(f431, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f421, f88])).
fof(f421, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f83, f87])).
fof(f83, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(sK8(X2), sK9(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f5564, plain, (sP1(sK34) | ~ spl50_430), inference(avatar_component_clause, [], [f5562])).
fof(f5565, plain, (spl50_429 | spl50_430), inference(avatar_split_clause, [], [f5557, f5562, f5559])).
fof(f5557, plain, ! [X4, X2, X0, X3, X1] : (sP1(sK34) | ~ r1(X0, sK34) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(sK13, X4)), inference(subsumption_resolution, [], [f5556, f129])).
fof(f5556, plain, ! [X4, X2, X0, X3, X1] : (sP1(sK34) | ~ r1(sK33, sK34) | ~ r1(X0, sK34) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(sK13, X4)), inference(duplicate_literal_removal, [], [f5328])).
fof(f5328, plain, ! [X4, X2, X0, X3, X1] : (sP1(sK34) | ~ r1(sK33, sK34) | ~ r1(X0, sK34) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(sK13, X4) | sP1(sK34)), inference(resolution, [], [f5316, f1842])).
fof(f1842, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK34, sK39(X0)) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(sK13, X5) | sP1(X0)), inference(subsumption_resolution, [], [f1841, f122])).
fof(f122, plain, ! [X78, X76, X74, X77, X75, X73] : (~ p1(sK39(X78)) | sP1(X78) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | ~ r1(sK13, X73)), inference(cnf_transformation, [], [f71])).
fof(f1841, plain, ! [X4, X2, X0, X5, X3, X1] : (sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(sK13, X5) | ~ r1(sK34, sK39(X0)) | p1(sK39(X0))), inference(subsumption_resolution, [], [f1840, f136])).
fof(f1840, plain, ! [X4, X2, X0, X5, X3, X1] : (p1(sK36(sK39(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(sK13, X5) | ~ r1(sK34, sK39(X0)) | p1(sK39(X0))), inference(duplicate_literal_removal, [], [f1833])).
fof(f1833, plain, ! [X4, X2, X0, X5, X3, X1] : (p1(sK36(sK39(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(sK13, X5) | ~ r1(sK34, sK39(X0)) | p1(sK39(X0)) | ~ r1(sK34, sK39(X0))), inference(resolution, [], [f771, f135])).
fof(f771, plain, ! [X12, X10, X8, X7, X13, X11, X9] : (~ r1(sK35(sK39(X8)), X7) | p1(X7) | sP1(X8) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(sK13, X13) | ~ r1(sK34, sK39(X8))), inference(subsumption_resolution, [], [f770, f122])).
fof(f770, plain, ! [X12, X10, X8, X7, X13, X11, X9] : (p1(X7) | ~ r1(sK35(sK39(X8)), X7) | sP1(X8) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(sK13, X13) | p1(sK39(X8)) | ~ r1(sK34, sK39(X8))), inference(subsumption_resolution, [], [f768, f137])).
fof(f137, plain, ! [X67] : (~ r1(sK34, X67) | p1(X67) | p1(sK35(X67))), inference(cnf_transformation, [], [f71])).
fof(f768, plain, ! [X12, X10, X8, X7, X13, X11, X9] : (p1(X7) | ~ r1(sK35(sK39(X8)), X7) | ~ p1(sK35(sK39(X8))) | sP1(X8) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(sK13, X13) | p1(sK39(X8)) | ~ r1(sK34, sK39(X8))), inference(resolution, [], [f123, f134])).
fof(f123, plain, ! [X80, X78, X76, X74, X81, X77, X75, X73] : (~ r1(sK39(X78), X80) | p1(X81) | ~ r1(X80, X81) | ~ p1(X80) | sP1(X78) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | ~ r1(sK13, X73)), inference(cnf_transformation, [], [f71])).
fof(f5316, plain, ! [X0] : (r1(X0, sK39(X0)) | sP1(X0) | ~ r1(sK33, X0)), inference(resolution, [], [f4322, f128])).
fof(f4322, plain, ! [X6, X7] : (~ r1(sK32, X6) | ~ r1(X6, X7) | sP1(X7) | r1(X7, sK39(X7))), inference(resolution, [], [f2645, f127])).
fof(f2645, plain, ! [X2, X0, X1] : (~ r1(sK31, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP1(X2) | r1(X2, sK39(X2))), inference(resolution, [], [f1265, f126])).
fof(f1265, plain, ! [X2, X0, X3, X1] : (~ r1(sK30, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | sP1(X1) | r1(X1, sK39(X1))), inference(resolution, [], [f652, f125])).
fof(f652, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK29, X4) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | sP1(X0) | r1(X0, sK39(X0))), inference(resolution, [], [f121, f124])).
fof(f121, plain, ! [X78, X76, X74, X77, X75, X73] : (~ r1(sK13, X73) | sP1(X78) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | r1(X78, sK39(X78))), inference(cnf_transformation, [], [f71])).
fof(f1058, plain, (spl50_100 | spl50_101 | ~ spl50_43), inference(avatar_split_clause, [], [f1042, f406, f1055, f1051])).
fof(f1042, plain, (p1(sK6(sK34)) | p1(sK35(sK6(sK34))) | ~ spl50_43), inference(resolution, [], [f408, f137])).
fof(f1041, plain, spl50_42, inference(avatar_split_clause, [], [f1036, f402])).
fof(f1036, plain, sP2(sK34), inference(resolution, [], [f1031, f129])).
fof(f1031, plain, ! [X0] : (~ r1(sK33, X0) | sP2(X0)), inference(resolution, [], [f1022, f128])).
fof(f1022, plain, ! [X0, X1] : (~ r1(sK32, X0) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f1013, f127])).
fof(f1013, plain, ! [X2, X0, X1] : (~ r1(sK31, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP2(X2)), inference(resolution, [], [f1006, f126])).
fof(f1006, plain, ! [X2, X0, X3, X1] : (~ r1(sK30, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | sP2(X3)), inference(resolution, [], [f600, f125])).
fof(f600, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK29, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f118, f124])).
fof(f118, plain, ! [X78, X76, X74, X77, X75, X73] : (~ r1(sK13, X73) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | sP2(X78)), inference(cnf_transformation, [], [f71])).
fof(f936, plain, (~ spl50_63 | spl50_64 | ~ spl50_20 | spl50_89), inference(avatar_split_clause, [], [f932, f934, f275, f525, f521])).
fof(f932, plain, ! [X12] : (r1(X12, sK4(X12)) | ~ r1(sK34, X12) | ~ p1(sK34) | sP0(sK34) | ~ sP3(sK34)), inference(subsumption_resolution, [], [f562, f131])).
fof(f562, plain, ! [X12] : (r1(X12, sK4(X12)) | ~ r1(sK34, X12) | ~ p1(sK34) | p1(sK38) | sP0(sK34) | ~ sP3(sK34)), inference(resolution, [], [f72, f130])).
fof(f72, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(X1, sK4(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f931, plain, (~ spl50_63 | spl50_64 | ~ spl50_20 | spl50_88), inference(avatar_split_clause, [], [f927, f929, f275, f525, f521])).
fof(f927, plain, ! [X12] : (r1(sK4(X12), sK5(X12)) | ~ r1(sK34, X12) | ~ p1(sK34) | sP0(sK34) | ~ sP3(sK34)), inference(subsumption_resolution, [], [f586, f131])).
fof(f586, plain, ! [X12] : (r1(sK4(X12), sK5(X12)) | ~ r1(sK34, X12) | ~ p1(sK34) | p1(sK38) | sP0(sK34) | ~ sP3(sK34)), inference(resolution, [], [f73, f130])).
fof(f73, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(sK4(X1), sK5(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f926, plain, spl50_63, inference(avatar_split_clause, [], [f907, f521])).
fof(f907, plain, sP3(sK34), inference(resolution, [], [f886, f129])).
fof(f886, plain, ! [X0] : (~ r1(sK33, X0) | sP3(X0)), inference(resolution, [], [f866, f128])).
fof(f866, plain, ! [X0, X1] : (~ r1(sK32, X0) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f844, f127])).
fof(f844, plain, ! [X2, X0, X1] : (~ r1(sK31, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP3(X2)), inference(resolution, [], [f838, f126])).
fof(f838, plain, ! [X2, X0, X3, X1] : (~ r1(sK30, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | sP3(X3)), inference(resolution, [], [f597, f125])).
fof(f597, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK29, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f117, f124])).
fof(f117, plain, ! [X78, X76, X74, X77, X75, X73] : (~ r1(sK13, X73) | ~ r1(X77, X78) | ~ r1(X76, X77) | ~ r1(X75, X76) | ~ r1(X74, X75) | ~ r1(X73, X74) | sP3(X78)), inference(cnf_transformation, [], [f71])).
fof(f418, plain, (~ spl50_42 | spl50_20 | spl50_43 | spl50_45), inference(avatar_split_clause, [], [f316, f415, f406, f275, f402])).
fof(f316, plain, (r1(sK37, sK7(sK37)) | r1(sK34, sK6(sK34)) | p1(sK34) | ~ sP2(sK34)), inference(resolution, [], [f76, f132])).
fof(f76, plain, ! [X4, X0] : (~ r1(X0, X4) | r1(X4, sK7(X4)) | r1(X0, sK6(X0)) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).