fof(f422599, plain, $false, inference(avatar_sat_refutation, [], [f427, f445, f464, f484, f505, f526, f622, f624, f627, f631, f633, f636, f640, f645, f647, f650, f654, f659, f5813, f6105, f6597, f6750, f11388, f12871, f14546, f14614, f14662, f14747, f14802, f14803, f14816, f14817, f15848, f15853, f15855, f16000, f16003, f16006, f16009, f16012, f16015, f16018, f16020, f16225, f16321, f16322, f16400, f16564, f16646, f16825, f17521, f31391, f44415, f45441, f46408, f46409, f54279, f65792, f83032, f83217, f98459, f107318, f107397, f107555, f108944, f110474, f110536, f132509, f133656, f133680, f146961, f146967, f159554, f165097, f165098, f166247, f171200, f172574, f173118, f173436, f173437, f173438, f175044, f175962, f175963, f177414, f177731, f179039, f179057, f179150, f179156, f183659, f183677, f192520, f193016, f195542, f198896, f199213, f199214, f199215, f200981, f202849, f338443, f339732, f339733, f339793, f340626, f342382, f342395, f342602, f342603, f343398, f343401, f343889, f350064, f350068, f365020, f365033, f401755, f401764, f402713, f403850, f403871, f403878, f403879, f403883, f404024, f404997, f405002, f405073, f405074, f405585, f405627, f406129, f414084, f414085, f414086, f415064, f415560, f419096, f419597, f419602, f420115, f421107, f421611, f422598])).
fof(f422598, plain, (~ spl92_12 | ~ spl92_17), inference(avatar_contradiction_clause, [], [f422597])).
fof(f422597, plain, ($false | (~ spl92_12 | ~ spl92_17)), inference(subsumption_resolution, [], [f422596, f544])).
fof(f544, plain, sP20(sK91), inference(resolution, [], [f527, f249])).
fof(f249, plain, ! [X0] : (~ sP40(X0) | sP20(X0)), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : ((! [X1] : (~ p101(X1) | ~ p201(X1) | ~ r1(X0, X1)) & ! [X2] : (~ p101(X2) | ~ p301(X2) | ~ r1(X0, X2)) & ! [X3] : (~ p101(X3) | ~ p401(X3) | ~ r1(X0, X3)) & ! [X4] : (~ p101(X4) | ~ p501(X4) | ~ r1(X0, X4)) & ! [X5] : (~ p101(X5) | ~ p601(X5) | ~ r1(X0, X5)) & ! [X6] : (~ p201(X6) | ~ p301(X6) | ~ r1(X0, X6)) & ! [X7] : (~ p201(X7) | ~ p401(X7) | ~ r1(X0, X7)) & ! [X8] : (~ p201(X8) | ~ p501(X8) | ~ r1(X0, X8)) & ! [X9] : (~ p201(X9) | ~ p601(X9) | ~ r1(X0, X9)) & ! [X10] : (~ p301(X10) | ~ p401(X10) | ~ r1(X0, X10)) & ! [X11] : (~ p301(X11) | ~ p501(X11) | ~ r1(X0, X11)) & ! [X12] : (~ p301(X12) | ~ p601(X12) | ~ r1(X0, X12)) & ! [X13] : (~ p401(X13) | ~ p501(X13) | ~ r1(X0, X13)) & ! [X14] : (~ p401(X14) | ~ p601(X14) | ~ r1(X0, X14)) & ! [X15] : (~ p501(X15) | ~ p601(X15) | ~ r1(X0, X15)) & sP39(X0) & sP38(X0) & sP37(X0) & sP36(X0) & sP35(X0) & ! [X16] : (~ p202(X16) | ~ p302(X16) | ~ r1(X0, X16)) & ! [X17] : (~ p202(X17) | ~ p402(X17) | ~ r1(X0, X17)) & ! [X18] : (~ p202(X18) | ~ p502(X18) | ~ r1(X0, X18)) & ! [X19] : (~ p202(X19) | ~ p602(X19) | ~ r1(X0, X19)) & ! [X20] : (~ p302(X20) | ~ p402(X20) | ~ r1(X0, X20)) & ! [X21] : (~ p302(X21) | ~ p502(X21) | ~ r1(X0, X21)) & ! [X22] : (~ p302(X22) | ~ p602(X22) | ~ r1(X0, X22)) & ! [X23] : (~ p402(X23) | ~ p502(X23) | ~ r1(X0, X23)) & ! [X24] : (~ p402(X24) | ~ p602(X24) | ~ r1(X0, X24)) & ! [X25] : (~ p502(X25) | ~ p602(X25) | ~ r1(X0, X25)) & sP9(X0) & sP34(X0) & sP33(X0) & sP32(X0) & sP31(X0) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & ! [X26] : (~ p303(X26) | ~ p403(X26) | ~ r1(X0, X26)) & ! [X27] : (~ p303(X27) | ~ p503(X27) | ~ r1(X0, X27)) & ! [X28] : (~ p303(X28) | ~ p603(X28) | ~ r1(X0, X28)) & ! [X29] : (~ p403(X29) | ~ p503(X29) | ~ r1(X0, X29)) & ! [X30] : (~ p403(X30) | ~ p603(X30) | ~ r1(X0, X30)) & ! [X31] : (~ p503(X31) | ~ p603(X31) | ~ r1(X0, X31)) & sP8(X0) & sP7(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP6(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP19(X0) & sP18(X0) & ! [X32] : (~ p404(X32) | ~ p504(X32) | ~ r1(X0, X32)) & ! [X33] : (~ p404(X33) | ~ p604(X33) | ~ r1(X0, X33)) & ! [X34] : (~ p504(X34) | ~ p604(X34) | ~ r1(X0, X34)) & sP5(X0) & sP4(X0) & sP3(X0) & sP17(X0) & sP16(X0) & sP2(X0) & sP1(X0) & sP15(X0) & sP14(X0) & sP0(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & ! [X35] : (~ p505(X35) | ~ p605(X35) | ~ r1(X0, X35))) | ~ sP40(X0)), inference(rectify, [], [f53])).
fof(f53, plain, ! [X11] : ((! [X12] : (~ p101(X12) | ~ p201(X12) | ~ r1(X11, X12)) & ! [X13] : (~ p101(X13) | ~ p301(X13) | ~ r1(X11, X13)) & ! [X14] : (~ p101(X14) | ~ p401(X14) | ~ r1(X11, X14)) & ! [X15] : (~ p101(X15) | ~ p501(X15) | ~ r1(X11, X15)) & ! [X16] : (~ p101(X16) | ~ p601(X16) | ~ r1(X11, X16)) & ! [X17] : (~ p201(X17) | ~ p301(X17) | ~ r1(X11, X17)) & ! [X18] : (~ p201(X18) | ~ p401(X18) | ~ r1(X11, X18)) & ! [X19] : (~ p201(X19) | ~ p501(X19) | ~ r1(X11, X19)) & ! [X20] : (~ p201(X20) | ~ p601(X20) | ~ r1(X11, X20)) & ! [X21] : (~ p301(X21) | ~ p401(X21) | ~ r1(X11, X21)) & ! [X22] : (~ p301(X22) | ~ p501(X22) | ~ r1(X11, X22)) & ! [X23] : (~ p301(X23) | ~ p601(X23) | ~ r1(X11, X23)) & ! [X24] : (~ p401(X24) | ~ p501(X24) | ~ r1(X11, X24)) & ! [X25] : (~ p401(X25) | ~ p601(X25) | ~ r1(X11, X25)) & ! [X26] : (~ p501(X26) | ~ p601(X26) | ~ r1(X11, X26)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & ! [X37] : (~ p202(X37) | ~ p302(X37) | ~ r1(X11, X37)) & ! [X38] : (~ p202(X38) | ~ p402(X38) | ~ r1(X11, X38)) & ! [X39] : (~ p202(X39) | ~ p502(X39) | ~ r1(X11, X39)) & ! [X40] : (~ p202(X40) | ~ p602(X40) | ~ r1(X11, X40)) & ! [X41] : (~ p302(X41) | ~ p402(X41) | ~ r1(X11, X41)) & ! [X42] : (~ p302(X42) | ~ p502(X42) | ~ r1(X11, X42)) & ! [X43] : (~ p302(X43) | ~ p602(X43) | ~ r1(X11, X43)) & ! [X44] : (~ p402(X44) | ~ p502(X44) | ~ r1(X11, X44)) & ! [X45] : (~ p402(X45) | ~ p602(X45) | ~ r1(X11, X45)) & ! [X46] : (~ p502(X46) | ~ p602(X46) | ~ r1(X11, X46)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & ! [X66] : (~ p303(X66) | ~ p403(X66) | ~ r1(X11, X66)) & ! [X67] : (~ p303(X67) | ~ p503(X67) | ~ r1(X11, X67)) & ! [X68] : (~ p303(X68) | ~ p603(X68) | ~ r1(X11, X68)) & ! [X69] : (~ p403(X69) | ~ p503(X69) | ~ r1(X11, X69)) & ! [X70] : (~ p403(X70) | ~ p603(X70) | ~ r1(X11, X70)) & ! [X71] : (~ p503(X71) | ~ p603(X71) | ~ r1(X11, X71)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & ! [X99] : (~ p404(X99) | ~ p504(X99) | ~ r1(X11, X99)) & ! [X100] : (~ p404(X100) | ~ p604(X100) | ~ r1(X11, X100)) & ! [X101] : (~ p504(X101) | ~ p604(X101) | ~ r1(X11, X101)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & ! [X136] : (~ p505(X136) | ~ p605(X136) | ~ r1(X11, X136))) | ~ sP40(X11)), inference(nnf_transformation, [], [f51])).
fof(f51, plain, ! [X11] : ((! [X12] : (~ p101(X12) | ~ p201(X12) | ~ r1(X11, X12)) & ! [X13] : (~ p101(X13) | ~ p301(X13) | ~ r1(X11, X13)) & ! [X14] : (~ p101(X14) | ~ p401(X14) | ~ r1(X11, X14)) & ! [X15] : (~ p101(X15) | ~ p501(X15) | ~ r1(X11, X15)) & ! [X16] : (~ p101(X16) | ~ p601(X16) | ~ r1(X11, X16)) & ! [X17] : (~ p201(X17) | ~ p301(X17) | ~ r1(X11, X17)) & ! [X18] : (~ p201(X18) | ~ p401(X18) | ~ r1(X11, X18)) & ! [X19] : (~ p201(X19) | ~ p501(X19) | ~ r1(X11, X19)) & ! [X20] : (~ p201(X20) | ~ p601(X20) | ~ r1(X11, X20)) & ! [X21] : (~ p301(X21) | ~ p401(X21) | ~ r1(X11, X21)) & ! [X22] : (~ p301(X22) | ~ p501(X22) | ~ r1(X11, X22)) & ! [X23] : (~ p301(X23) | ~ p601(X23) | ~ r1(X11, X23)) & ! [X24] : (~ p401(X24) | ~ p501(X24) | ~ r1(X11, X24)) & ! [X25] : (~ p401(X25) | ~ p601(X25) | ~ r1(X11, X25)) & ! [X26] : (~ p501(X26) | ~ p601(X26) | ~ r1(X11, X26)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & ! [X37] : (~ p202(X37) | ~ p302(X37) | ~ r1(X11, X37)) & ! [X38] : (~ p202(X38) | ~ p402(X38) | ~ r1(X11, X38)) & ! [X39] : (~ p202(X39) | ~ p502(X39) | ~ r1(X11, X39)) & ! [X40] : (~ p202(X40) | ~ p602(X40) | ~ r1(X11, X40)) & ! [X41] : (~ p302(X41) | ~ p402(X41) | ~ r1(X11, X41)) & ! [X42] : (~ p302(X42) | ~ p502(X42) | ~ r1(X11, X42)) & ! [X43] : (~ p302(X43) | ~ p602(X43) | ~ r1(X11, X43)) & ! [X44] : (~ p402(X44) | ~ p502(X44) | ~ r1(X11, X44)) & ! [X45] : (~ p402(X45) | ~ p602(X45) | ~ r1(X11, X45)) & ! [X46] : (~ p502(X46) | ~ p602(X46) | ~ r1(X11, X46)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & ! [X66] : (~ p303(X66) | ~ p403(X66) | ~ r1(X11, X66)) & ! [X67] : (~ p303(X67) | ~ p503(X67) | ~ r1(X11, X67)) & ! [X68] : (~ p303(X68) | ~ p603(X68) | ~ r1(X11, X68)) & ! [X69] : (~ p403(X69) | ~ p503(X69) | ~ r1(X11, X69)) & ! [X70] : (~ p403(X70) | ~ p603(X70) | ~ r1(X11, X70)) & ! [X71] : (~ p503(X71) | ~ p603(X71) | ~ r1(X11, X71)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & ! [X99] : (~ p404(X99) | ~ p504(X99) | ~ r1(X11, X99)) & ! [X100] : (~ p404(X100) | ~ p604(X100) | ~ r1(X11, X100)) & ! [X101] : (~ p504(X101) | ~ p604(X101) | ~ r1(X11, X101)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & ! [X136] : (~ p505(X136) | ~ p605(X136) | ~ r1(X11, X136))) | ~ sP40(X11)), inference(usedef, [], [e51])).
fof(e51, plain, ! [X11] : (sP40(X11) <=> (! [X12] : (~ p101(X12) | ~ p201(X12) | ~ r1(X11, X12)) & ! [X13] : (~ p101(X13) | ~ p301(X13) | ~ r1(X11, X13)) & ! [X14] : (~ p101(X14) | ~ p401(X14) | ~ r1(X11, X14)) & ! [X15] : (~ p101(X15) | ~ p501(X15) | ~ r1(X11, X15)) & ! [X16] : (~ p101(X16) | ~ p601(X16) | ~ r1(X11, X16)) & ! [X17] : (~ p201(X17) | ~ p301(X17) | ~ r1(X11, X17)) & ! [X18] : (~ p201(X18) | ~ p401(X18) | ~ r1(X11, X18)) & ! [X19] : (~ p201(X19) | ~ p501(X19) | ~ r1(X11, X19)) & ! [X20] : (~ p201(X20) | ~ p601(X20) | ~ r1(X11, X20)) & ! [X21] : (~ p301(X21) | ~ p401(X21) | ~ r1(X11, X21)) & ! [X22] : (~ p301(X22) | ~ p501(X22) | ~ r1(X11, X22)) & ! [X23] : (~ p301(X23) | ~ p601(X23) | ~ r1(X11, X23)) & ! [X24] : (~ p401(X24) | ~ p501(X24) | ~ r1(X11, X24)) & ! [X25] : (~ p401(X25) | ~ p601(X25) | ~ r1(X11, X25)) & ! [X26] : (~ p501(X26) | ~ p601(X26) | ~ r1(X11, X26)) & sP39(X11) & sP38(X11) & sP37(X11) & sP36(X11) & sP35(X11) & ! [X37] : (~ p202(X37) | ~ p302(X37) | ~ r1(X11, X37)) & ! [X38] : (~ p202(X38) | ~ p402(X38) | ~ r1(X11, X38)) & ! [X39] : (~ p202(X39) | ~ p502(X39) | ~ r1(X11, X39)) & ! [X40] : (~ p202(X40) | ~ p602(X40) | ~ r1(X11, X40)) & ! [X41] : (~ p302(X41) | ~ p402(X41) | ~ r1(X11, X41)) & ! [X42] : (~ p302(X42) | ~ p502(X42) | ~ r1(X11, X42)) & ! [X43] : (~ p302(X43) | ~ p602(X43) | ~ r1(X11, X43)) & ! [X44] : (~ p402(X44) | ~ p502(X44) | ~ r1(X11, X44)) & ! [X45] : (~ p402(X45) | ~ p602(X45) | ~ r1(X11, X45)) & ! [X46] : (~ p502(X46) | ~ p602(X46) | ~ r1(X11, X46)) & sP9(X11) & sP34(X11) & sP33(X11) & sP32(X11) & sP31(X11) & sP30(X11) & sP29(X11) & sP28(X11) & sP27(X11) & ! [X66] : (~ p303(X66) | ~ p403(X66) | ~ r1(X11, X66)) & ! [X67] : (~ p303(X67) | ~ p503(X67) | ~ r1(X11, X67)) & ! [X68] : (~ p303(X68) | ~ p603(X68) | ~ r1(X11, X68)) & ! [X69] : (~ p403(X69) | ~ p503(X69) | ~ r1(X11, X69)) & ! [X70] : (~ p403(X70) | ~ p603(X70) | ~ r1(X11, X70)) & ! [X71] : (~ p503(X71) | ~ p603(X71) | ~ r1(X11, X71)) & sP8(X11) & sP7(X11) & sP26(X11) & sP25(X11) & sP24(X11) & sP6(X11) & sP23(X11) & sP22(X11) & sP21(X11) & sP20(X11) & sP19(X11) & sP18(X11) & ! [X99] : (~ p404(X99) | ~ p504(X99) | ~ r1(X11, X99)) & ! [X100] : (~ p404(X100) | ~ p604(X100) | ~ r1(X11, X100)) & ! [X101] : (~ p504(X101) | ~ p604(X101) | ~ r1(X11, X101)) & sP5(X11) & sP4(X11) & sP3(X11) & sP17(X11) & sP16(X11) & sP2(X11) & sP1(X11) & sP15(X11) & sP14(X11) & sP0(X11) & sP13(X11) & sP12(X11) & sP11(X11) & sP10(X11) & ! [X136] : (~ p505(X136) | ~ p605(X136) | ~ r1(X11, X136)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f527, plain, sP40(sK91), inference(resolution, [], [f404, f227])).
fof(f227, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL684+1.005.p', reflexivity)).
fof(f404, plain, ! [X11] : (~ r1(sK91, X11) | sP40(X11)), inference(cnf_transformation, [], [f226])).
fof(f226, plain, ((p101(sK91) | ! [X1] : (p102(X1) | ~ r1(sK91, X1)) | ! [X2] : (p103(X2) | ~ r1(sK91, X2)) | ! [X3] : (p104(X3) | ~ r1(sK91, X3)) | ! [X4] : (p105(X4) | ~ r1(sK91, X4))) & (p201(sK91) | p202(sK91) | ! [X5] : (p203(X5) | ~ r1(sK91, X5)) | ! [X6] : (p204(X6) | ~ r1(sK91, X6)) | ! [X7] : (p205(X7) | ~ r1(sK91, X7))) & (p301(sK91) | p302(sK91) | p303(sK91) | ! [X8] : (p304(X8) | ~ r1(sK91, X8)) | ! [X9] : (p305(X9) | ~ r1(sK91, X9))) & (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | ! [X10] : (p405(X10) | ~ r1(sK91, X10))) & (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)) & (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)) & ! [X11] : (sP40(X11) | ~ r1(sK91, X11))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK91])], [f52, f225])).
fof(f225, plain, (? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : (sP40(X11) | ~ r1(X0, X11))) => ((p101(sK91) | ! [X1] : (p102(X1) | ~ r1(sK91, X1)) | ! [X2] : (p103(X2) | ~ r1(sK91, X2)) | ! [X3] : (p104(X3) | ~ r1(sK91, X3)) | ! [X4] : (p105(X4) | ~ r1(sK91, X4))) & (p201(sK91) | p202(sK91) | ! [X5] : (p203(X5) | ~ r1(sK91, X5)) | ! [X6] : (p204(X6) | ~ r1(sK91, X6)) | ! [X7] : (p205(X7) | ~ r1(sK91, X7))) & (p301(sK91) | p302(sK91) | p303(sK91) | ! [X8] : (p304(X8) | ~ r1(sK91, X8)) | ! [X9] : (p305(X9) | ~ r1(sK91, X9))) & (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | ! [X10] : (p405(X10) | ~ r1(sK91, X10))) & (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)) & (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)) & ! [X11] : (sP40(X11) | ~ r1(sK91, X11)))), introduced(choice_axiom, [])).
fof(f52, plain, ? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : (sP40(X11) | ~ r1(X0, X11))), inference(definition_folding, [], [f10, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11])).
fof(f11, plain, ! [X11] : (! [X125] : (? [X126] : (~ p305(X126) & r1(X125, X126)) | ? [X127] : (~ p405(X127) & r1(X125, X127)) | ~ r1(X11, X125)) | ~ sP0(X11)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X11] : (sP0(X11) <=> ! [X125] : (? [X126] : (~ p305(X126) & r1(X125, X126)) | ? [X127] : (~ p405(X127) & r1(X125, X127)) | ~ r1(X11, X125))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f12, plain, ! [X11] : (! [X118] : (? [X119] : (~ p205(X119) & r1(X118, X119)) | ? [X120] : (~ p405(X120) & r1(X118, X120)) | ~ r1(X11, X118)) | ~ sP1(X11)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X11] : (sP1(X11) <=> ! [X118] : (? [X119] : (~ p205(X119) & r1(X118, X119)) | ? [X120] : (~ p405(X120) & r1(X118, X120)) | ~ r1(X11, X118))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f13, plain, ! [X11] : (! [X115] : (? [X116] : (~ p205(X116) & r1(X115, X116)) | ? [X117] : (~ p305(X117) & r1(X115, X117)) | ~ r1(X11, X115)) | ~ sP2(X11)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X11] : (sP2(X11) <=> ! [X115] : (? [X116] : (~ p205(X116) & r1(X115, X116)) | ? [X117] : (~ p305(X117) & r1(X115, X117)) | ~ r1(X11, X115))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f14, plain, ! [X11] : (! [X108] : (? [X109] : (~ p105(X109) & r1(X108, X109)) | ? [X110] : (~ p405(X110) & r1(X108, X110)) | ~ r1(X11, X108)) | ~ sP3(X11)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X11] : (sP3(X11) <=> ! [X108] : (? [X109] : (~ p105(X109) & r1(X108, X109)) | ? [X110] : (~ p405(X110) & r1(X108, X110)) | ~ r1(X11, X108))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f15, plain, ! [X11] : (! [X105] : (? [X106] : (~ p105(X106) & r1(X105, X106)) | ? [X107] : (~ p305(X107) & r1(X105, X107)) | ~ r1(X11, X105)) | ~ sP4(X11)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X11] : (sP4(X11) <=> ! [X105] : (? [X106] : (~ p105(X106) & r1(X105, X106)) | ? [X107] : (~ p305(X107) & r1(X105, X107)) | ~ r1(X11, X105))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f16, plain, ! [X11] : (! [X102] : (? [X103] : (~ p105(X103) & r1(X102, X103)) | ? [X104] : (~ p205(X104) & r1(X102, X104)) | ~ r1(X11, X102)) | ~ sP5(X11)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X11] : (sP5(X11) <=> ! [X102] : (? [X103] : (~ p105(X103) & r1(X102, X103)) | ? [X104] : (~ p205(X104) & r1(X102, X104)) | ~ r1(X11, X102))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f17, plain, ! [X11] : (! [X84] : (? [X85] : (~ p204(X85) & r1(X84, X85)) | ? [X86] : (~ p304(X86) & r1(X84, X86)) | ~ r1(X11, X84)) | ~ sP6(X11)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X11] : (sP6(X11) <=> ! [X84] : (? [X85] : (~ p204(X85) & r1(X84, X85)) | ? [X86] : (~ p304(X86) & r1(X84, X86)) | ~ r1(X11, X84))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f18, plain, ! [X11] : (! [X75] : (? [X76] : (~ p104(X76) & r1(X75, X76)) | ? [X77] : (~ p304(X77) & r1(X75, X77)) | ~ r1(X11, X75)) | ~ sP7(X11)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X11] : (sP7(X11) <=> ! [X75] : (? [X76] : (~ p104(X76) & r1(X75, X76)) | ? [X77] : (~ p304(X77) & r1(X75, X77)) | ~ r1(X11, X75))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f19, plain, ! [X11] : (! [X72] : (? [X73] : (~ p104(X73) & r1(X72, X73)) | ? [X74] : (~ p204(X74) & r1(X72, X74)) | ~ r1(X11, X72)) | ~ sP8(X11)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X11] : (sP8(X11) <=> ! [X72] : (? [X73] : (~ p104(X73) & r1(X72, X73)) | ? [X74] : (~ p204(X74) & r1(X72, X74)) | ~ r1(X11, X72))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f20, plain, ! [X11] : (! [X47] : (? [X48] : (~ p103(X48) & r1(X47, X48)) | ? [X49] : (~ p203(X49) & r1(X47, X49)) | ~ r1(X11, X47)) | ~ sP9(X11)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X11] : (sP9(X11) <=> ! [X47] : (? [X48] : (~ p103(X48) & r1(X47, X48)) | ? [X49] : (~ p203(X49) & r1(X47, X49)) | ~ r1(X11, X47))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f21, plain, ! [X11] : (! [X134] : (? [X135] : (~ p405(X135) & r1(X134, X135)) | ~ p605(X134) | ~ r1(X11, X134)) | ~ sP10(X11)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X11] : (sP10(X11) <=> ! [X134] : (? [X135] : (~ p405(X135) & r1(X134, X135)) | ~ p605(X134) | ~ r1(X11, X134))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f22, plain, ! [X11] : (! [X132] : (? [X133] : (~ p405(X133) & r1(X132, X133)) | ~ p505(X132) | ~ r1(X11, X132)) | ~ sP11(X11)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X11] : (sP11(X11) <=> ! [X132] : (? [X133] : (~ p405(X133) & r1(X132, X133)) | ~ p505(X132) | ~ r1(X11, X132))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f23, plain, ! [X11] : (! [X130] : (? [X131] : (~ p305(X131) & r1(X130, X131)) | ~ p605(X130) | ~ r1(X11, X130)) | ~ sP12(X11)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X11] : (sP12(X11) <=> ! [X130] : (? [X131] : (~ p305(X131) & r1(X130, X131)) | ~ p605(X130) | ~ r1(X11, X130))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f24, plain, ! [X11] : (! [X128] : (? [X129] : (~ p305(X129) & r1(X128, X129)) | ~ p505(X128) | ~ r1(X11, X128)) | ~ sP13(X11)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X11] : (sP13(X11) <=> ! [X128] : (? [X129] : (~ p305(X129) & r1(X128, X129)) | ~ p505(X128) | ~ r1(X11, X128))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f25, plain, ! [X11] : (! [X123] : (? [X124] : (~ p205(X124) & r1(X123, X124)) | ~ p605(X123) | ~ r1(X11, X123)) | ~ sP14(X11)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X11] : (sP14(X11) <=> ! [X123] : (? [X124] : (~ p205(X124) & r1(X123, X124)) | ~ p605(X123) | ~ r1(X11, X123))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f26, plain, ! [X11] : (! [X121] : (? [X122] : (~ p205(X122) & r1(X121, X122)) | ~ p505(X121) | ~ r1(X11, X121)) | ~ sP15(X11)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X11] : (sP15(X11) <=> ! [X121] : (? [X122] : (~ p205(X122) & r1(X121, X122)) | ~ p505(X121) | ~ r1(X11, X121))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f27, plain, ! [X11] : (! [X113] : (? [X114] : (~ p105(X114) & r1(X113, X114)) | ~ p605(X113) | ~ r1(X11, X113)) | ~ sP16(X11)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X11] : (sP16(X11) <=> ! [X113] : (? [X114] : (~ p105(X114) & r1(X113, X114)) | ~ p605(X113) | ~ r1(X11, X113))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f28, plain, ! [X11] : (! [X111] : (? [X112] : (~ p105(X112) & r1(X111, X112)) | ~ p505(X111) | ~ r1(X11, X111)) | ~ sP17(X11)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X11] : (sP17(X11) <=> ! [X111] : (? [X112] : (~ p105(X112) & r1(X111, X112)) | ~ p505(X111) | ~ r1(X11, X111))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f29, plain, ! [X11] : (! [X97] : (? [X98] : (~ p304(X98) & r1(X97, X98)) | ~ p604(X97) | ~ r1(X11, X97)) | ~ sP18(X11)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X11] : (sP18(X11) <=> ! [X97] : (? [X98] : (~ p304(X98) & r1(X97, X98)) | ~ p604(X97) | ~ r1(X11, X97))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f30, plain, ! [X11] : (! [X95] : (? [X96] : (~ p304(X96) & r1(X95, X96)) | ~ p504(X95) | ~ r1(X11, X95)) | ~ sP19(X11)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X11] : (sP19(X11) <=> ! [X95] : (? [X96] : (~ p304(X96) & r1(X95, X96)) | ~ p504(X95) | ~ r1(X11, X95))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f31, plain, ! [X11] : (! [X93] : (? [X94] : (~ p304(X94) & r1(X93, X94)) | ~ p404(X93) | ~ r1(X11, X93)) | ~ sP20(X11)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X11] : (sP20(X11) <=> ! [X93] : (? [X94] : (~ p304(X94) & r1(X93, X94)) | ~ p404(X93) | ~ r1(X11, X93))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f32, plain, ! [X11] : (! [X91] : (? [X92] : (~ p204(X92) & r1(X91, X92)) | ~ p604(X91) | ~ r1(X11, X91)) | ~ sP21(X11)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X11] : (sP21(X11) <=> ! [X91] : (? [X92] : (~ p204(X92) & r1(X91, X92)) | ~ p604(X91) | ~ r1(X11, X91))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f33, plain, ! [X11] : (! [X89] : (? [X90] : (~ p204(X90) & r1(X89, X90)) | ~ p504(X89) | ~ r1(X11, X89)) | ~ sP22(X11)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X11] : (sP22(X11) <=> ! [X89] : (? [X90] : (~ p204(X90) & r1(X89, X90)) | ~ p504(X89) | ~ r1(X11, X89))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f34, plain, ! [X11] : (! [X87] : (? [X88] : (~ p204(X88) & r1(X87, X88)) | ~ p404(X87) | ~ r1(X11, X87)) | ~ sP23(X11)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X11] : (sP23(X11) <=> ! [X87] : (? [X88] : (~ p204(X88) & r1(X87, X88)) | ~ p404(X87) | ~ r1(X11, X87))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f35, plain, ! [X11] : (! [X82] : (? [X83] : (~ p104(X83) & r1(X82, X83)) | ~ p604(X82) | ~ r1(X11, X82)) | ~ sP24(X11)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X11] : (sP24(X11) <=> ! [X82] : (? [X83] : (~ p104(X83) & r1(X82, X83)) | ~ p604(X82) | ~ r1(X11, X82))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f36, plain, ! [X11] : (! [X80] : (? [X81] : (~ p104(X81) & r1(X80, X81)) | ~ p504(X80) | ~ r1(X11, X80)) | ~ sP25(X11)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X11] : (sP25(X11) <=> ! [X80] : (? [X81] : (~ p104(X81) & r1(X80, X81)) | ~ p504(X80) | ~ r1(X11, X80))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f37, plain, ! [X11] : (! [X78] : (? [X79] : (~ p104(X79) & r1(X78, X79)) | ~ p404(X78) | ~ r1(X11, X78)) | ~ sP26(X11)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X11] : (sP26(X11) <=> ! [X78] : (? [X79] : (~ p104(X79) & r1(X78, X79)) | ~ p404(X78) | ~ r1(X11, X78))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f38, plain, ! [X11] : (! [X64] : (? [X65] : (~ p203(X65) & r1(X64, X65)) | ~ p603(X64) | ~ r1(X11, X64)) | ~ sP27(X11)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X11] : (sP27(X11) <=> ! [X64] : (? [X65] : (~ p203(X65) & r1(X64, X65)) | ~ p603(X64) | ~ r1(X11, X64))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f39, plain, ! [X11] : (! [X62] : (? [X63] : (~ p203(X63) & r1(X62, X63)) | ~ p503(X62) | ~ r1(X11, X62)) | ~ sP28(X11)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X11] : (sP28(X11) <=> ! [X62] : (? [X63] : (~ p203(X63) & r1(X62, X63)) | ~ p503(X62) | ~ r1(X11, X62))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f40, plain, ! [X11] : (! [X60] : (? [X61] : (~ p203(X61) & r1(X60, X61)) | ~ p403(X60) | ~ r1(X11, X60)) | ~ sP29(X11)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X11] : (sP29(X11) <=> ! [X60] : (? [X61] : (~ p203(X61) & r1(X60, X61)) | ~ p403(X60) | ~ r1(X11, X60))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f41, plain, ! [X11] : (! [X58] : (? [X59] : (~ p203(X59) & r1(X58, X59)) | ~ p303(X58) | ~ r1(X11, X58)) | ~ sP30(X11)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X11] : (sP30(X11) <=> ! [X58] : (? [X59] : (~ p203(X59) & r1(X58, X59)) | ~ p303(X58) | ~ r1(X11, X58))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f42, plain, ! [X11] : (! [X56] : (? [X57] : (~ p103(X57) & r1(X56, X57)) | ~ p603(X56) | ~ r1(X11, X56)) | ~ sP31(X11)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X11] : (sP31(X11) <=> ! [X56] : (? [X57] : (~ p103(X57) & r1(X56, X57)) | ~ p603(X56) | ~ r1(X11, X56))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f43, plain, ! [X11] : (! [X54] : (? [X55] : (~ p103(X55) & r1(X54, X55)) | ~ p503(X54) | ~ r1(X11, X54)) | ~ sP32(X11)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X11] : (sP32(X11) <=> ! [X54] : (? [X55] : (~ p103(X55) & r1(X54, X55)) | ~ p503(X54) | ~ r1(X11, X54))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f44, plain, ! [X11] : (! [X52] : (? [X53] : (~ p103(X53) & r1(X52, X53)) | ~ p403(X52) | ~ r1(X11, X52)) | ~ sP33(X11)), inference(usedef, [], [e44])).
fof(e44, plain, ! [X11] : (sP33(X11) <=> ! [X52] : (? [X53] : (~ p103(X53) & r1(X52, X53)) | ~ p403(X52) | ~ r1(X11, X52))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f45, plain, ! [X11] : (! [X50] : (? [X51] : (~ p103(X51) & r1(X50, X51)) | ~ p303(X50) | ~ r1(X11, X50)) | ~ sP34(X11)), inference(usedef, [], [e45])).
fof(e45, plain, ! [X11] : (sP34(X11) <=> ! [X50] : (? [X51] : (~ p103(X51) & r1(X50, X51)) | ~ p303(X50) | ~ r1(X11, X50))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f46, plain, ! [X11] : (! [X35] : (? [X36] : (~ p102(X36) & r1(X35, X36)) | ~ p602(X35) | ~ r1(X11, X35)) | ~ sP35(X11)), inference(usedef, [], [e46])).
fof(e46, plain, ! [X11] : (sP35(X11) <=> ! [X35] : (? [X36] : (~ p102(X36) & r1(X35, X36)) | ~ p602(X35) | ~ r1(X11, X35))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f47, plain, ! [X11] : (! [X33] : (? [X34] : (~ p102(X34) & r1(X33, X34)) | ~ p502(X33) | ~ r1(X11, X33)) | ~ sP36(X11)), inference(usedef, [], [e47])).
fof(e47, plain, ! [X11] : (sP36(X11) <=> ! [X33] : (? [X34] : (~ p102(X34) & r1(X33, X34)) | ~ p502(X33) | ~ r1(X11, X33))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f48, plain, ! [X11] : (! [X31] : (? [X32] : (~ p102(X32) & r1(X31, X32)) | ~ p402(X31) | ~ r1(X11, X31)) | ~ sP37(X11)), inference(usedef, [], [e48])).
fof(e48, plain, ! [X11] : (sP37(X11) <=> ! [X31] : (? [X32] : (~ p102(X32) & r1(X31, X32)) | ~ p402(X31) | ~ r1(X11, X31))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f49, plain, ! [X11] : (! [X29] : (? [X30] : (~ p102(X30) & r1(X29, X30)) | ~ p302(X29) | ~ r1(X11, X29)) | ~ sP38(X11)), inference(usedef, [], [e49])).
fof(e49, plain, ! [X11] : (sP38(X11) <=> ! [X29] : (? [X30] : (~ p102(X30) & r1(X29, X30)) | ~ p302(X29) | ~ r1(X11, X29))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f50, plain, ! [X11] : (! [X27] : (? [X28] : (~ p102(X28) & r1(X27, X28)) | ~ p202(X27) | ~ r1(X11, X27)) | ~ sP39(X11)), inference(usedef, [], [e50])).
fof(e50, plain, ! [X11] : (sP39(X11) <=> ! [X27] : (? [X28] : (~ p102(X28) & r1(X27, X28)) | ~ p202(X27) | ~ r1(X11, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f10, plain, ? [X0] : ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0)) & ! [X11] : ((! [X12] : (~ p101(X12) | ~ p201(X12) | ~ r1(X11, X12)) & ! [X13] : (~ p101(X13) | ~ p301(X13) | ~ r1(X11, X13)) & ! [X14] : (~ p101(X14) | ~ p401(X14) | ~ r1(X11, X14)) & ! [X15] : (~ p101(X15) | ~ p501(X15) | ~ r1(X11, X15)) & ! [X16] : (~ p101(X16) | ~ p601(X16) | ~ r1(X11, X16)) & ! [X17] : (~ p201(X17) | ~ p301(X17) | ~ r1(X11, X17)) & ! [X18] : (~ p201(X18) | ~ p401(X18) | ~ r1(X11, X18)) & ! [X19] : (~ p201(X19) | ~ p501(X19) | ~ r1(X11, X19)) & ! [X20] : (~ p201(X20) | ~ p601(X20) | ~ r1(X11, X20)) & ! [X21] : (~ p301(X21) | ~ p401(X21) | ~ r1(X11, X21)) & ! [X22] : (~ p301(X22) | ~ p501(X22) | ~ r1(X11, X22)) & ! [X23] : (~ p301(X23) | ~ p601(X23) | ~ r1(X11, X23)) & ! [X24] : (~ p401(X24) | ~ p501(X24) | ~ r1(X11, X24)) & ! [X25] : (~ p401(X25) | ~ p601(X25) | ~ r1(X11, X25)) & ! [X26] : (~ p501(X26) | ~ p601(X26) | ~ r1(X11, X26)) & ! [X27] : (? [X28] : (~ p102(X28) & r1(X27, X28)) | ~ p202(X27) | ~ r1(X11, X27)) & ! [X29] : (? [X30] : (~ p102(X30) & r1(X29, X30)) | ~ p302(X29) | ~ r1(X11, X29)) & ! [X31] : (? [X32] : (~ p102(X32) & r1(X31, X32)) | ~ p402(X31) | ~ r1(X11, X31)) & ! [X33] : (? [X34] : (~ p102(X34) & r1(X33, X34)) | ~ p502(X33) | ~ r1(X11, X33)) & ! [X35] : (? [X36] : (~ p102(X36) & r1(X35, X36)) | ~ p602(X35) | ~ r1(X11, X35)) & ! [X37] : (~ p202(X37) | ~ p302(X37) | ~ r1(X11, X37)) & ! [X38] : (~ p202(X38) | ~ p402(X38) | ~ r1(X11, X38)) & ! [X39] : (~ p202(X39) | ~ p502(X39) | ~ r1(X11, X39)) & ! [X40] : (~ p202(X40) | ~ p602(X40) | ~ r1(X11, X40)) & ! [X41] : (~ p302(X41) | ~ p402(X41) | ~ r1(X11, X41)) & ! [X42] : (~ p302(X42) | ~ p502(X42) | ~ r1(X11, X42)) & ! [X43] : (~ p302(X43) | ~ p602(X43) | ~ r1(X11, X43)) & ! [X44] : (~ p402(X44) | ~ p502(X44) | ~ r1(X11, X44)) & ! [X45] : (~ p402(X45) | ~ p602(X45) | ~ r1(X11, X45)) & ! [X46] : (~ p502(X46) | ~ p602(X46) | ~ r1(X11, X46)) & ! [X47] : (? [X48] : (~ p103(X48) & r1(X47, X48)) | ? [X49] : (~ p203(X49) & r1(X47, X49)) | ~ r1(X11, X47)) & ! [X50] : (? [X51] : (~ p103(X51) & r1(X50, X51)) | ~ p303(X50) | ~ r1(X11, X50)) & ! [X52] : (? [X53] : (~ p103(X53) & r1(X52, X53)) | ~ p403(X52) | ~ r1(X11, X52)) & ! [X54] : (? [X55] : (~ p103(X55) & r1(X54, X55)) | ~ p503(X54) | ~ r1(X11, X54)) & ! [X56] : (? [X57] : (~ p103(X57) & r1(X56, X57)) | ~ p603(X56) | ~ r1(X11, X56)) & ! [X58] : (? [X59] : (~ p203(X59) & r1(X58, X59)) | ~ p303(X58) | ~ r1(X11, X58)) & ! [X60] : (? [X61] : (~ p203(X61) & r1(X60, X61)) | ~ p403(X60) | ~ r1(X11, X60)) & ! [X62] : (? [X63] : (~ p203(X63) & r1(X62, X63)) | ~ p503(X62) | ~ r1(X11, X62)) & ! [X64] : (? [X65] : (~ p203(X65) & r1(X64, X65)) | ~ p603(X64) | ~ r1(X11, X64)) & ! [X66] : (~ p303(X66) | ~ p403(X66) | ~ r1(X11, X66)) & ! [X67] : (~ p303(X67) | ~ p503(X67) | ~ r1(X11, X67)) & ! [X68] : (~ p303(X68) | ~ p603(X68) | ~ r1(X11, X68)) & ! [X69] : (~ p403(X69) | ~ p503(X69) | ~ r1(X11, X69)) & ! [X70] : (~ p403(X70) | ~ p603(X70) | ~ r1(X11, X70)) & ! [X71] : (~ p503(X71) | ~ p603(X71) | ~ r1(X11, X71)) & ! [X72] : (? [X73] : (~ p104(X73) & r1(X72, X73)) | ? [X74] : (~ p204(X74) & r1(X72, X74)) | ~ r1(X11, X72)) & ! [X75] : (? [X76] : (~ p104(X76) & r1(X75, X76)) | ? [X77] : (~ p304(X77) & r1(X75, X77)) | ~ r1(X11, X75)) & ! [X78] : (? [X79] : (~ p104(X79) & r1(X78, X79)) | ~ p404(X78) | ~ r1(X11, X78)) & ! [X80] : (? [X81] : (~ p104(X81) & r1(X80, X81)) | ~ p504(X80) | ~ r1(X11, X80)) & ! [X82] : (? [X83] : (~ p104(X83) & r1(X82, X83)) | ~ p604(X82) | ~ r1(X11, X82)) & ! [X84] : (? [X85] : (~ p204(X85) & r1(X84, X85)) | ? [X86] : (~ p304(X86) & r1(X84, X86)) | ~ r1(X11, X84)) & ! [X87] : (? [X88] : (~ p204(X88) & r1(X87, X88)) | ~ p404(X87) | ~ r1(X11, X87)) & ! [X89] : (? [X90] : (~ p204(X90) & r1(X89, X90)) | ~ p504(X89) | ~ r1(X11, X89)) & ! [X91] : (? [X92] : (~ p204(X92) & r1(X91, X92)) | ~ p604(X91) | ~ r1(X11, X91)) & ! [X93] : (? [X94] : (~ p304(X94) & r1(X93, X94)) | ~ p404(X93) | ~ r1(X11, X93)) & ! [X95] : (? [X96] : (~ p304(X96) & r1(X95, X96)) | ~ p504(X95) | ~ r1(X11, X95)) & ! [X97] : (? [X98] : (~ p304(X98) & r1(X97, X98)) | ~ p604(X97) | ~ r1(X11, X97)) & ! [X99] : (~ p404(X99) | ~ p504(X99) | ~ r1(X11, X99)) & ! [X100] : (~ p404(X100) | ~ p604(X100) | ~ r1(X11, X100)) & ! [X101] : (~ p504(X101) | ~ p604(X101) | ~ r1(X11, X101)) & ! [X102] : (? [X103] : (~ p105(X103) & r1(X102, X103)) | ? [X104] : (~ p205(X104) & r1(X102, X104)) | ~ r1(X11, X102)) & ! [X105] : (? [X106] : (~ p105(X106) & r1(X105, X106)) | ? [X107] : (~ p305(X107) & r1(X105, X107)) | ~ r1(X11, X105)) & ! [X108] : (? [X109] : (~ p105(X109) & r1(X108, X109)) | ? [X110] : (~ p405(X110) & r1(X108, X110)) | ~ r1(X11, X108)) & ! [X111] : (? [X112] : (~ p105(X112) & r1(X111, X112)) | ~ p505(X111) | ~ r1(X11, X111)) & ! [X113] : (? [X114] : (~ p105(X114) & r1(X113, X114)) | ~ p605(X113) | ~ r1(X11, X113)) & ! [X115] : (? [X116] : (~ p205(X116) & r1(X115, X116)) | ? [X117] : (~ p305(X117) & r1(X115, X117)) | ~ r1(X11, X115)) & ! [X118] : (? [X119] : (~ p205(X119) & r1(X118, X119)) | ? [X120] : (~ p405(X120) & r1(X118, X120)) | ~ r1(X11, X118)) & ! [X121] : (? [X122] : (~ p205(X122) & r1(X121, X122)) | ~ p505(X121) | ~ r1(X11, X121)) & ! [X123] : (? [X124] : (~ p205(X124) & r1(X123, X124)) | ~ p605(X123) | ~ r1(X11, X123)) & ! [X125] : (? [X126] : (~ p305(X126) & r1(X125, X126)) | ? [X127] : (~ p405(X127) & r1(X125, X127)) | ~ r1(X11, X125)) & ! [X128] : (? [X129] : (~ p305(X129) & r1(X128, X129)) | ~ p505(X128) | ~ r1(X11, X128)) & ! [X130] : (? [X131] : (~ p305(X131) & r1(X130, X131)) | ~ p605(X130) | ~ r1(X11, X130)) & ! [X132] : (? [X133] : (~ p405(X133) & r1(X132, X133)) | ~ p505(X132) | ~ r1(X11, X132)) & ! [X134] : (? [X135] : (~ p405(X135) & r1(X134, X135)) | ~ p605(X134) | ~ r1(X11, X134)) & ! [X136] : (~ p505(X136) | ~ p605(X136) | ~ r1(X11, X136))) | ~ r1(X0, X11))), inference(flattening, [], [f9])).
fof(f9, plain, ? [X0] : (((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) & ! [X11] : ((! [X12] : ((~ p101(X12) | ~ p201(X12)) | ~ r1(X11, X12)) & ! [X13] : ((~ p101(X13) | ~ p301(X13)) | ~ r1(X11, X13)) & ! [X14] : ((~ p101(X14) | ~ p401(X14)) | ~ r1(X11, X14)) & ! [X15] : ((~ p101(X15) | ~ p501(X15)) | ~ r1(X11, X15)) & ! [X16] : ((~ p101(X16) | ~ p601(X16)) | ~ r1(X11, X16)) & ! [X17] : ((~ p201(X17) | ~ p301(X17)) | ~ r1(X11, X17)) & ! [X18] : ((~ p201(X18) | ~ p401(X18)) | ~ r1(X11, X18)) & ! [X19] : ((~ p201(X19) | ~ p501(X19)) | ~ r1(X11, X19)) & ! [X20] : ((~ p201(X20) | ~ p601(X20)) | ~ r1(X11, X20)) & ! [X21] : ((~ p301(X21) | ~ p401(X21)) | ~ r1(X11, X21)) & ! [X22] : ((~ p301(X22) | ~ p501(X22)) | ~ r1(X11, X22)) & ! [X23] : ((~ p301(X23) | ~ p601(X23)) | ~ r1(X11, X23)) & ! [X24] : ((~ p401(X24) | ~ p501(X24)) | ~ r1(X11, X24)) & ! [X25] : ((~ p401(X25) | ~ p601(X25)) | ~ r1(X11, X25)) & ! [X26] : ((~ p501(X26) | ~ p601(X26)) | ~ r1(X11, X26)) & ! [X27] : ((? [X28] : (~ p102(X28) & r1(X27, X28)) | ~ p202(X27)) | ~ r1(X11, X27)) & ! [X29] : ((? [X30] : (~ p102(X30) & r1(X29, X30)) | ~ p302(X29)) | ~ r1(X11, X29)) & ! [X31] : ((? [X32] : (~ p102(X32) & r1(X31, X32)) | ~ p402(X31)) | ~ r1(X11, X31)) & ! [X33] : ((? [X34] : (~ p102(X34) & r1(X33, X34)) | ~ p502(X33)) | ~ r1(X11, X33)) & ! [X35] : ((? [X36] : (~ p102(X36) & r1(X35, X36)) | ~ p602(X35)) | ~ r1(X11, X35)) & ! [X37] : ((~ p202(X37) | ~ p302(X37)) | ~ r1(X11, X37)) & ! [X38] : ((~ p202(X38) | ~ p402(X38)) | ~ r1(X11, X38)) & ! [X39] : ((~ p202(X39) | ~ p502(X39)) | ~ r1(X11, X39)) & ! [X40] : ((~ p202(X40) | ~ p602(X40)) | ~ r1(X11, X40)) & ! [X41] : ((~ p302(X41) | ~ p402(X41)) | ~ r1(X11, X41)) & ! [X42] : ((~ p302(X42) | ~ p502(X42)) | ~ r1(X11, X42)) & ! [X43] : ((~ p302(X43) | ~ p602(X43)) | ~ r1(X11, X43)) & ! [X44] : ((~ p402(X44) | ~ p502(X44)) | ~ r1(X11, X44)) & ! [X45] : ((~ p402(X45) | ~ p602(X45)) | ~ r1(X11, X45)) & ! [X46] : ((~ p502(X46) | ~ p602(X46)) | ~ r1(X11, X46)) & ! [X47] : ((? [X48] : (~ p103(X48) & r1(X47, X48)) | ? [X49] : (~ p203(X49) & r1(X47, X49))) | ~ r1(X11, X47)) & ! [X50] : ((? [X51] : (~ p103(X51) & r1(X50, X51)) | ~ p303(X50)) | ~ r1(X11, X50)) & ! [X52] : ((? [X53] : (~ p103(X53) & r1(X52, X53)) | ~ p403(X52)) | ~ r1(X11, X52)) & ! [X54] : ((? [X55] : (~ p103(X55) & r1(X54, X55)) | ~ p503(X54)) | ~ r1(X11, X54)) & ! [X56] : ((? [X57] : (~ p103(X57) & r1(X56, X57)) | ~ p603(X56)) | ~ r1(X11, X56)) & ! [X58] : ((? [X59] : (~ p203(X59) & r1(X58, X59)) | ~ p303(X58)) | ~ r1(X11, X58)) & ! [X60] : ((? [X61] : (~ p203(X61) & r1(X60, X61)) | ~ p403(X60)) | ~ r1(X11, X60)) & ! [X62] : ((? [X63] : (~ p203(X63) & r1(X62, X63)) | ~ p503(X62)) | ~ r1(X11, X62)) & ! [X64] : ((? [X65] : (~ p203(X65) & r1(X64, X65)) | ~ p603(X64)) | ~ r1(X11, X64)) & ! [X66] : ((~ p303(X66) | ~ p403(X66)) | ~ r1(X11, X66)) & ! [X67] : ((~ p303(X67) | ~ p503(X67)) | ~ r1(X11, X67)) & ! [X68] : ((~ p303(X68) | ~ p603(X68)) | ~ r1(X11, X68)) & ! [X69] : ((~ p403(X69) | ~ p503(X69)) | ~ r1(X11, X69)) & ! [X70] : ((~ p403(X70) | ~ p603(X70)) | ~ r1(X11, X70)) & ! [X71] : ((~ p503(X71) | ~ p603(X71)) | ~ r1(X11, X71)) & ! [X72] : ((? [X73] : (~ p104(X73) & r1(X72, X73)) | ? [X74] : (~ p204(X74) & r1(X72, X74))) | ~ r1(X11, X72)) & ! [X75] : ((? [X76] : (~ p104(X76) & r1(X75, X76)) | ? [X77] : (~ p304(X77) & r1(X75, X77))) | ~ r1(X11, X75)) & ! [X78] : ((? [X79] : (~ p104(X79) & r1(X78, X79)) | ~ p404(X78)) | ~ r1(X11, X78)) & ! [X80] : ((? [X81] : (~ p104(X81) & r1(X80, X81)) | ~ p504(X80)) | ~ r1(X11, X80)) & ! [X82] : ((? [X83] : (~ p104(X83) & r1(X82, X83)) | ~ p604(X82)) | ~ r1(X11, X82)) & ! [X84] : ((? [X85] : (~ p204(X85) & r1(X84, X85)) | ? [X86] : (~ p304(X86) & r1(X84, X86))) | ~ r1(X11, X84)) & ! [X87] : ((? [X88] : (~ p204(X88) & r1(X87, X88)) | ~ p404(X87)) | ~ r1(X11, X87)) & ! [X89] : ((? [X90] : (~ p204(X90) & r1(X89, X90)) | ~ p504(X89)) | ~ r1(X11, X89)) & ! [X91] : ((? [X92] : (~ p204(X92) & r1(X91, X92)) | ~ p604(X91)) | ~ r1(X11, X91)) & ! [X93] : ((? [X94] : (~ p304(X94) & r1(X93, X94)) | ~ p404(X93)) | ~ r1(X11, X93)) & ! [X95] : ((? [X96] : (~ p304(X96) & r1(X95, X96)) | ~ p504(X95)) | ~ r1(X11, X95)) & ! [X97] : ((? [X98] : (~ p304(X98) & r1(X97, X98)) | ~ p604(X97)) | ~ r1(X11, X97)) & ! [X99] : ((~ p404(X99) | ~ p504(X99)) | ~ r1(X11, X99)) & ! [X100] : ((~ p404(X100) | ~ p604(X100)) | ~ r1(X11, X100)) & ! [X101] : ((~ p504(X101) | ~ p604(X101)) | ~ r1(X11, X101)) & ! [X102] : ((? [X103] : (~ p105(X103) & r1(X102, X103)) | ? [X104] : (~ p205(X104) & r1(X102, X104))) | ~ r1(X11, X102)) & ! [X105] : ((? [X106] : (~ p105(X106) & r1(X105, X106)) | ? [X107] : (~ p305(X107) & r1(X105, X107))) | ~ r1(X11, X105)) & ! [X108] : ((? [X109] : (~ p105(X109) & r1(X108, X109)) | ? [X110] : (~ p405(X110) & r1(X108, X110))) | ~ r1(X11, X108)) & ! [X111] : ((? [X112] : (~ p105(X112) & r1(X111, X112)) | ~ p505(X111)) | ~ r1(X11, X111)) & ! [X113] : ((? [X114] : (~ p105(X114) & r1(X113, X114)) | ~ p605(X113)) | ~ r1(X11, X113)) & ! [X115] : ((? [X116] : (~ p205(X116) & r1(X115, X116)) | ? [X117] : (~ p305(X117) & r1(X115, X117))) | ~ r1(X11, X115)) & ! [X118] : ((? [X119] : (~ p205(X119) & r1(X118, X119)) | ? [X120] : (~ p405(X120) & r1(X118, X120))) | ~ r1(X11, X118)) & ! [X121] : ((? [X122] : (~ p205(X122) & r1(X121, X122)) | ~ p505(X121)) | ~ r1(X11, X121)) & ! [X123] : ((? [X124] : (~ p205(X124) & r1(X123, X124)) | ~ p605(X123)) | ~ r1(X11, X123)) & ! [X125] : ((? [X126] : (~ p305(X126) & r1(X125, X126)) | ? [X127] : (~ p405(X127) & r1(X125, X127))) | ~ r1(X11, X125)) & ! [X128] : ((? [X129] : (~ p305(X129) & r1(X128, X129)) | ~ p505(X128)) | ~ r1(X11, X128)) & ! [X130] : ((? [X131] : (~ p305(X131) & r1(X130, X131)) | ~ p605(X130)) | ~ r1(X11, X130)) & ! [X132] : ((? [X133] : (~ p405(X133) & r1(X132, X133)) | ~ p505(X132)) | ~ r1(X11, X132)) & ! [X134] : ((? [X135] : (~ p405(X135) & r1(X134, X135)) | ~ p605(X134)) | ~ r1(X11, X134)) & ! [X136] : ((~ p505(X136) | ~ p605(X136)) | ~ r1(X11, X136))) | ~ r1(X0, X11))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X11] : (~ (~ ! [X12] : (~ (p101(X12) & p201(X12)) | ~ r1(X11, X12)) | ~ ! [X13] : (~ (p101(X13) & p301(X13)) | ~ r1(X11, X13)) | ~ ! [X14] : (~ (p101(X14) & p401(X14)) | ~ r1(X11, X14)) | ~ ! [X15] : (~ (p101(X15) & p501(X15)) | ~ r1(X11, X15)) | ~ ! [X16] : (~ (p101(X16) & p601(X16)) | ~ r1(X11, X16)) | ~ ! [X17] : (~ (p201(X17) & p301(X17)) | ~ r1(X11, X17)) | ~ ! [X18] : (~ (p201(X18) & p401(X18)) | ~ r1(X11, X18)) | ~ ! [X19] : (~ (p201(X19) & p501(X19)) | ~ r1(X11, X19)) | ~ ! [X20] : (~ (p201(X20) & p601(X20)) | ~ r1(X11, X20)) | ~ ! [X21] : (~ (p301(X21) & p401(X21)) | ~ r1(X11, X21)) | ~ ! [X22] : (~ (p301(X22) & p501(X22)) | ~ r1(X11, X22)) | ~ ! [X23] : (~ (p301(X23) & p601(X23)) | ~ r1(X11, X23)) | ~ ! [X24] : (~ (p401(X24) & p501(X24)) | ~ r1(X11, X24)) | ~ ! [X25] : (~ (p401(X25) & p601(X25)) | ~ r1(X11, X25)) | ~ ! [X26] : (~ (p501(X26) & p601(X26)) | ~ r1(X11, X26)) | ~ ! [X27] : (~ (! [X28] : (p102(X28) | ~ r1(X27, X28)) & p202(X27)) | ~ r1(X11, X27)) | ~ ! [X29] : (~ (! [X30] : (p102(X30) | ~ r1(X29, X30)) & p302(X29)) | ~ r1(X11, X29)) | ~ ! [X31] : (~ (! [X32] : (p102(X32) | ~ r1(X31, X32)) & p402(X31)) | ~ r1(X11, X31)) | ~ ! [X33] : (~ (! [X34] : (p102(X34) | ~ r1(X33, X34)) & p502(X33)) | ~ r1(X11, X33)) | ~ ! [X35] : (~ (! [X36] : (p102(X36) | ~ r1(X35, X36)) & p602(X35)) | ~ r1(X11, X35)) | ~ ! [X37] : (~ (p202(X37) & p302(X37)) | ~ r1(X11, X37)) | ~ ! [X38] : (~ (p202(X38) & p402(X38)) | ~ r1(X11, X38)) | ~ ! [X39] : (~ (p202(X39) & p502(X39)) | ~ r1(X11, X39)) | ~ ! [X40] : (~ (p202(X40) & p602(X40)) | ~ r1(X11, X40)) | ~ ! [X41] : (~ (p302(X41) & p402(X41)) | ~ r1(X11, X41)) | ~ ! [X42] : (~ (p302(X42) & p502(X42)) | ~ r1(X11, X42)) | ~ ! [X43] : (~ (p302(X43) & p602(X43)) | ~ r1(X11, X43)) | ~ ! [X44] : (~ (p402(X44) & p502(X44)) | ~ r1(X11, X44)) | ~ ! [X45] : (~ (p402(X45) & p602(X45)) | ~ r1(X11, X45)) | ~ ! [X46] : (~ (p502(X46) & p602(X46)) | ~ r1(X11, X46)) | ~ ! [X47] : (~ (! [X48] : (p103(X48) | ~ r1(X47, X48)) & ! [X49] : (p203(X49) | ~ r1(X47, X49))) | ~ r1(X11, X47)) | ~ ! [X50] : (~ (! [X51] : (p103(X51) | ~ r1(X50, X51)) & p303(X50)) | ~ r1(X11, X50)) | ~ ! [X52] : (~ (! [X53] : (p103(X53) | ~ r1(X52, X53)) & p403(X52)) | ~ r1(X11, X52)) | ~ ! [X54] : (~ (! [X55] : (p103(X55) | ~ r1(X54, X55)) & p503(X54)) | ~ r1(X11, X54)) | ~ ! [X56] : (~ (! [X57] : (p103(X57) | ~ r1(X56, X57)) & p603(X56)) | ~ r1(X11, X56)) | ~ ! [X58] : (~ (! [X59] : (p203(X59) | ~ r1(X58, X59)) & p303(X58)) | ~ r1(X11, X58)) | ~ ! [X60] : (~ (! [X61] : (p203(X61) | ~ r1(X60, X61)) & p403(X60)) | ~ r1(X11, X60)) | ~ ! [X62] : (~ (! [X63] : (p203(X63) | ~ r1(X62, X63)) & p503(X62)) | ~ r1(X11, X62)) | ~ ! [X64] : (~ (! [X65] : (p203(X65) | ~ r1(X64, X65)) & p603(X64)) | ~ r1(X11, X64)) | ~ ! [X66] : (~ (p303(X66) & p403(X66)) | ~ r1(X11, X66)) | ~ ! [X67] : (~ (p303(X67) & p503(X67)) | ~ r1(X11, X67)) | ~ ! [X68] : (~ (p303(X68) & p603(X68)) | ~ r1(X11, X68)) | ~ ! [X69] : (~ (p403(X69) & p503(X69)) | ~ r1(X11, X69)) | ~ ! [X70] : (~ (p403(X70) & p603(X70)) | ~ r1(X11, X70)) | ~ ! [X71] : (~ (p503(X71) & p603(X71)) | ~ r1(X11, X71)) | ~ ! [X72] : (~ (! [X73] : (p104(X73) | ~ r1(X72, X73)) & ! [X74] : (p204(X74) | ~ r1(X72, X74))) | ~ r1(X11, X72)) | ~ ! [X75] : (~ (! [X76] : (p104(X76) | ~ r1(X75, X76)) & ! [X77] : (p304(X77) | ~ r1(X75, X77))) | ~ r1(X11, X75)) | ~ ! [X78] : (~ (! [X79] : (p104(X79) | ~ r1(X78, X79)) & p404(X78)) | ~ r1(X11, X78)) | ~ ! [X80] : (~ (! [X81] : (p104(X81) | ~ r1(X80, X81)) & p504(X80)) | ~ r1(X11, X80)) | ~ ! [X82] : (~ (! [X83] : (p104(X83) | ~ r1(X82, X83)) & p604(X82)) | ~ r1(X11, X82)) | ~ ! [X84] : (~ (! [X85] : (p204(X85) | ~ r1(X84, X85)) & ! [X86] : (p304(X86) | ~ r1(X84, X86))) | ~ r1(X11, X84)) | ~ ! [X87] : (~ (! [X88] : (p204(X88) | ~ r1(X87, X88)) & p404(X87)) | ~ r1(X11, X87)) | ~ ! [X89] : (~ (! [X90] : (p204(X90) | ~ r1(X89, X90)) & p504(X89)) | ~ r1(X11, X89)) | ~ ! [X91] : (~ (! [X92] : (p204(X92) | ~ r1(X91, X92)) & p604(X91)) | ~ r1(X11, X91)) | ~ ! [X93] : (~ (! [X94] : (p304(X94) | ~ r1(X93, X94)) & p404(X93)) | ~ r1(X11, X93)) | ~ ! [X95] : (~ (! [X96] : (p304(X96) | ~ r1(X95, X96)) & p504(X95)) | ~ r1(X11, X95)) | ~ ! [X97] : (~ (! [X98] : (p304(X98) | ~ r1(X97, X98)) & p604(X97)) | ~ r1(X11, X97)) | ~ ! [X99] : (~ (p404(X99) & p504(X99)) | ~ r1(X11, X99)) | ~ ! [X100] : (~ (p404(X100) & p604(X100)) | ~ r1(X11, X100)) | ~ ! [X101] : (~ (p504(X101) & p604(X101)) | ~ r1(X11, X101)) | ~ ! [X102] : (~ (! [X103] : (p105(X103) | ~ r1(X102, X103)) & ! [X104] : (p205(X104) | ~ r1(X102, X104))) | ~ r1(X11, X102)) | ~ ! [X105] : (~ (! [X106] : (p105(X106) | ~ r1(X105, X106)) & ! [X107] : (p305(X107) | ~ r1(X105, X107))) | ~ r1(X11, X105)) | ~ ! [X108] : (~ (! [X109] : (p105(X109) | ~ r1(X108, X109)) & ! [X110] : (p405(X110) | ~ r1(X108, X110))) | ~ r1(X11, X108)) | ~ ! [X111] : (~ (! [X112] : (p105(X112) | ~ r1(X111, X112)) & p505(X111)) | ~ r1(X11, X111)) | ~ ! [X113] : (~ (! [X114] : (p105(X114) | ~ r1(X113, X114)) & p605(X113)) | ~ r1(X11, X113)) | ~ ! [X115] : (~ (! [X116] : (p205(X116) | ~ r1(X115, X116)) & ! [X117] : (p305(X117) | ~ r1(X115, X117))) | ~ r1(X11, X115)) | ~ ! [X118] : (~ (! [X119] : (p205(X119) | ~ r1(X118, X119)) & ! [X120] : (p405(X120) | ~ r1(X118, X120))) | ~ r1(X11, X118)) | ~ ! [X121] : (~ (! [X122] : (p205(X122) | ~ r1(X121, X122)) & p505(X121)) | ~ r1(X11, X121)) | ~ ! [X123] : (~ (! [X124] : (p205(X124) | ~ r1(X123, X124)) & p605(X123)) | ~ r1(X11, X123)) | ~ ! [X125] : (~ (! [X126] : (p305(X126) | ~ r1(X125, X126)) & ! [X127] : (p405(X127) | ~ r1(X125, X127))) | ~ r1(X11, X125)) | ~ ! [X128] : (~ (! [X129] : (p305(X129) | ~ r1(X128, X129)) & p505(X128)) | ~ r1(X11, X128)) | ~ ! [X130] : (~ (! [X131] : (p305(X131) | ~ r1(X130, X131)) & p605(X130)) | ~ r1(X11, X130)) | ~ ! [X132] : (~ (! [X133] : (p405(X133) | ~ r1(X132, X133)) & p505(X132)) | ~ r1(X11, X132)) | ~ ! [X134] : (~ (! [X135] : (p405(X135) | ~ r1(X134, X135)) & p605(X134)) | ~ r1(X11, X134)) | ~ ! [X136] : (~ (p505(X136) & p605(X136)) | ~ r1(X11, X136))) | ~ r1(X0, X11))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X2] : (p103(X2) | ~ r1(X0, X2)) | ! [X3] : (p104(X3) | ~ r1(X0, X3)) | ! [X4] : (p105(X4) | ~ r1(X0, X4))) & (p201(X0) | p202(X0) | ! [X5] : (p203(X5) | ~ r1(X0, X5)) | ! [X6] : (p204(X6) | ~ r1(X0, X6)) | ! [X7] : (p205(X7) | ~ r1(X0, X7))) & (p301(X0) | p302(X0) | p303(X0) | ! [X8] : (p304(X8) | ~ r1(X0, X8)) | ! [X9] : (p305(X9) | ~ r1(X0, X9))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X10] : (p405(X10) | ~ r1(X0, X10))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X11] : (~ (~ ! [X12] : (~ (p101(X12) & p201(X12)) | ~ r1(X11, X12)) | ~ ! [X13] : (~ (p101(X13) & p301(X13)) | ~ r1(X11, X13)) | ~ ! [X14] : (~ (p101(X14) & p401(X14)) | ~ r1(X11, X14)) | ~ ! [X15] : (~ (p101(X15) & p501(X15)) | ~ r1(X11, X15)) | ~ ! [X16] : (~ (p101(X16) & p601(X16)) | ~ r1(X11, X16)) | ~ ! [X17] : (~ (p201(X17) & p301(X17)) | ~ r1(X11, X17)) | ~ ! [X18] : (~ (p201(X18) & p401(X18)) | ~ r1(X11, X18)) | ~ ! [X19] : (~ (p201(X19) & p501(X19)) | ~ r1(X11, X19)) | ~ ! [X20] : (~ (p201(X20) & p601(X20)) | ~ r1(X11, X20)) | ~ ! [X21] : (~ (p301(X21) & p401(X21)) | ~ r1(X11, X21)) | ~ ! [X22] : (~ (p301(X22) & p501(X22)) | ~ r1(X11, X22)) | ~ ! [X23] : (~ (p301(X23) & p601(X23)) | ~ r1(X11, X23)) | ~ ! [X24] : (~ (p401(X24) & p501(X24)) | ~ r1(X11, X24)) | ~ ! [X25] : (~ (p401(X25) & p601(X25)) | ~ r1(X11, X25)) | ~ ! [X26] : (~ (p501(X26) & p601(X26)) | ~ r1(X11, X26)) | ~ ! [X27] : (~ (! [X28] : (p102(X28) | ~ r1(X27, X28)) & p202(X27)) | ~ r1(X11, X27)) | ~ ! [X29] : (~ (! [X30] : (p102(X30) | ~ r1(X29, X30)) & p302(X29)) | ~ r1(X11, X29)) | ~ ! [X31] : (~ (! [X32] : (p102(X32) | ~ r1(X31, X32)) & p402(X31)) | ~ r1(X11, X31)) | ~ ! [X33] : (~ (! [X34] : (p102(X34) | ~ r1(X33, X34)) & p502(X33)) | ~ r1(X11, X33)) | ~ ! [X35] : (~ (! [X36] : (p102(X36) | ~ r1(X35, X36)) & p602(X35)) | ~ r1(X11, X35)) | ~ ! [X37] : (~ (p202(X37) & p302(X37)) | ~ r1(X11, X37)) | ~ ! [X38] : (~ (p202(X38) & p402(X38)) | ~ r1(X11, X38)) | ~ ! [X39] : (~ (p202(X39) & p502(X39)) | ~ r1(X11, X39)) | ~ ! [X40] : (~ (p202(X40) & p602(X40)) | ~ r1(X11, X40)) | ~ ! [X41] : (~ (p302(X41) & p402(X41)) | ~ r1(X11, X41)) | ~ ! [X42] : (~ (p302(X42) & p502(X42)) | ~ r1(X11, X42)) | ~ ! [X43] : (~ (p302(X43) & p602(X43)) | ~ r1(X11, X43)) | ~ ! [X44] : (~ (p402(X44) & p502(X44)) | ~ r1(X11, X44)) | ~ ! [X45] : (~ (p402(X45) & p602(X45)) | ~ r1(X11, X45)) | ~ ! [X46] : (~ (p502(X46) & p602(X46)) | ~ r1(X11, X46)) | ~ ! [X47] : (~ (! [X48] : (p103(X48) | ~ r1(X47, X48)) & ! [X49] : (p203(X49) | ~ r1(X47, X49))) | ~ r1(X11, X47)) | ~ ! [X50] : (~ (! [X51] : (p103(X51) | ~ r1(X50, X51)) & p303(X50)) | ~ r1(X11, X50)) | ~ ! [X52] : (~ (! [X53] : (p103(X53) | ~ r1(X52, X53)) & p403(X52)) | ~ r1(X11, X52)) | ~ ! [X54] : (~ (! [X55] : (p103(X55) | ~ r1(X54, X55)) & p503(X54)) | ~ r1(X11, X54)) | ~ ! [X56] : (~ (! [X57] : (p103(X57) | ~ r1(X56, X57)) & p603(X56)) | ~ r1(X11, X56)) | ~ ! [X58] : (~ (! [X59] : (p203(X59) | ~ r1(X58, X59)) & p303(X58)) | ~ r1(X11, X58)) | ~ ! [X60] : (~ (! [X61] : (p203(X61) | ~ r1(X60, X61)) & p403(X60)) | ~ r1(X11, X60)) | ~ ! [X62] : (~ (! [X63] : (p203(X63) | ~ r1(X62, X63)) & p503(X62)) | ~ r1(X11, X62)) | ~ ! [X64] : (~ (! [X65] : (p203(X65) | ~ r1(X64, X65)) & p603(X64)) | ~ r1(X11, X64)) | ~ ! [X66] : (~ (p303(X66) & p403(X66)) | ~ r1(X11, X66)) | ~ ! [X67] : (~ (p303(X67) & p503(X67)) | ~ r1(X11, X67)) | ~ ! [X68] : (~ (p303(X68) & p603(X68)) | ~ r1(X11, X68)) | ~ ! [X69] : (~ (p403(X69) & p503(X69)) | ~ r1(X11, X69)) | ~ ! [X70] : (~ (p403(X70) & p603(X70)) | ~ r1(X11, X70)) | ~ ! [X71] : (~ (p503(X71) & p603(X71)) | ~ r1(X11, X71)) | ~ ! [X72] : (~ (! [X73] : (p104(X73) | ~ r1(X72, X73)) & ! [X74] : (p204(X74) | ~ r1(X72, X74))) | ~ r1(X11, X72)) | ~ ! [X75] : (~ (! [X76] : (p104(X76) | ~ r1(X75, X76)) & ! [X77] : (p304(X77) | ~ r1(X75, X77))) | ~ r1(X11, X75)) | ~ ! [X78] : (~ (! [X79] : (p104(X79) | ~ r1(X78, X79)) & p404(X78)) | ~ r1(X11, X78)) | ~ ! [X80] : (~ (! [X81] : (p104(X81) | ~ r1(X80, X81)) & p504(X80)) | ~ r1(X11, X80)) | ~ ! [X82] : (~ (! [X83] : (p104(X83) | ~ r1(X82, X83)) & p604(X82)) | ~ r1(X11, X82)) | ~ ! [X84] : (~ (! [X85] : (p204(X85) | ~ r1(X84, X85)) & ! [X86] : (p304(X86) | ~ r1(X84, X86))) | ~ r1(X11, X84)) | ~ ! [X87] : (~ (! [X88] : (p204(X88) | ~ r1(X87, X88)) & p404(X87)) | ~ r1(X11, X87)) | ~ ! [X89] : (~ (! [X90] : (p204(X90) | ~ r1(X89, X90)) & p504(X89)) | ~ r1(X11, X89)) | ~ ! [X91] : (~ (! [X92] : (p204(X92) | ~ r1(X91, X92)) & p604(X91)) | ~ r1(X11, X91)) | ~ ! [X93] : (~ (! [X94] : (p304(X94) | ~ r1(X93, X94)) & p404(X93)) | ~ r1(X11, X93)) | ~ ! [X95] : (~ (! [X96] : (p304(X96) | ~ r1(X95, X96)) & p504(X95)) | ~ r1(X11, X95)) | ~ ! [X97] : (~ (! [X98] : (p304(X98) | ~ r1(X97, X98)) & p604(X97)) | ~ r1(X11, X97)) | ~ ! [X99] : (~ (p404(X99) & p504(X99)) | ~ r1(X11, X99)) | ~ ! [X100] : (~ (p404(X100) & p604(X100)) | ~ r1(X11, X100)) | ~ ! [X101] : (~ (p504(X101) & p604(X101)) | ~ r1(X11, X101)) | ~ ! [X102] : (~ (! [X103] : (p105(X103) | ~ r1(X102, X103)) & ! [X104] : (p205(X104) | ~ r1(X102, X104))) | ~ r1(X11, X102)) | ~ ! [X105] : (~ (! [X106] : (p105(X106) | ~ r1(X105, X106)) & ! [X107] : (p305(X107) | ~ r1(X105, X107))) | ~ r1(X11, X105)) | ~ ! [X108] : (~ (! [X109] : (p105(X109) | ~ r1(X108, X109)) & ! [X110] : (p405(X110) | ~ r1(X108, X110))) | ~ r1(X11, X108)) | ~ ! [X111] : (~ (! [X112] : (p105(X112) | ~ r1(X111, X112)) & p505(X111)) | ~ r1(X11, X111)) | ~ ! [X113] : (~ (! [X114] : (p105(X114) | ~ r1(X113, X114)) & p605(X113)) | ~ r1(X11, X113)) | ~ ! [X115] : (~ (! [X116] : (p205(X116) | ~ r1(X115, X116)) & ! [X117] : (p305(X117) | ~ r1(X115, X117))) | ~ r1(X11, X115)) | ~ ! [X118] : (~ (! [X119] : (p205(X119) | ~ r1(X118, X119)) & ! [X120] : (p405(X120) | ~ r1(X118, X120))) | ~ r1(X11, X118)) | ~ ! [X121] : (~ (! [X122] : (p205(X122) | ~ r1(X121, X122)) & p505(X121)) | ~ r1(X11, X121)) | ~ ! [X123] : (~ (! [X124] : (p205(X124) | ~ r1(X123, X124)) & p605(X123)) | ~ r1(X11, X123)) | ~ ! [X125] : (~ (! [X126] : (p305(X126) | ~ r1(X125, X126)) & ! [X127] : (p405(X127) | ~ r1(X125, X127))) | ~ r1(X11, X125)) | ~ ! [X128] : (~ (! [X129] : (p305(X129) | ~ r1(X128, X129)) & p505(X128)) | ~ r1(X11, X128)) | ~ ! [X130] : (~ (! [X131] : (p305(X131) | ~ r1(X130, X131)) & p605(X130)) | ~ r1(X11, X130)) | ~ ! [X132] : (~ (! [X133] : (p405(X133) | ~ r1(X132, X133)) & p505(X132)) | ~ r1(X11, X132)) | ~ ! [X134] : (~ (! [X135] : (p405(X135) | ~ r1(X134, X135)) & p605(X134)) | ~ r1(X11, X134)) | ~ ! [X136] : (~ (p505(X136) & p605(X136)) | ~ r1(X11, X136))) | ~ r1(X0, X11))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X1] : (p103(X1) | ~ r1(X0, X1)) | ! [X1] : (p104(X1) | ~ r1(X0, X1)) | ! [X1] : (p105(X1) | ~ r1(X0, X1))) & (p201(X0) | p202(X0) | ! [X1] : (p203(X1) | ~ r1(X0, X1)) | ! [X1] : (p204(X1) | ~ r1(X0, X1)) | ! [X1] : (p205(X1) | ~ r1(X0, X1))) & (p301(X0) | p302(X0) | p303(X0) | ! [X1] : (p304(X1) | ~ r1(X0, X1)) | ! [X1] : (p305(X1) | ~ r1(X0, X1))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X1] : (p405(X1) | ~ r1(X0, X1))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X1] : (~ (~ ! [X0] : (~ (p101(X0) & p201(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p301(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p301(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p401(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p401(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p501(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p202(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p302(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p302(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p402(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p402(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p502(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & ! [X1] : (p203(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p303(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p303(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p403(X0) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p403(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p503(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & ! [X1] : (p204(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & ! [X1] : (p304(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & ! [X1] : (p304(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p404(X0) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p404(X0) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p504(X0) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p205(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p305(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & ! [X1] : (p305(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p405(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p405(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p505(X0) & p605(X0)) | ~ r1(X1, X0))) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ ((p101(X0) | ! [X1] : (p102(X1) | ~ r1(X0, X1)) | ! [X1] : (p103(X1) | ~ r1(X0, X1)) | ! [X1] : (p104(X1) | ~ r1(X0, X1)) | ! [X1] : (p105(X1) | ~ r1(X0, X1))) & (p201(X0) | p202(X0) | ! [X1] : (p203(X1) | ~ r1(X0, X1)) | ! [X1] : (p204(X1) | ~ r1(X0, X1)) | ! [X1] : (p205(X1) | ~ r1(X0, X1))) & (p301(X0) | p302(X0) | p303(X0) | ! [X1] : (p304(X1) | ~ r1(X0, X1)) | ! [X1] : (p305(X1) | ~ r1(X0, X1))) & (p401(X0) | p402(X0) | p403(X0) | p404(X0) | ! [X1] : (p405(X1) | ~ r1(X0, X1))) & (p501(X0) | p502(X0) | p503(X0) | p504(X0) | p505(X0)) & (p601(X0) | p602(X0) | p603(X0) | p604(X0) | p605(X0))) | ~ ! [X1] : (~ (~ ! [X0] : (~ (p101(X0) & p201(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p301(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p101(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p301(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p201(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p401(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p301(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p401(X0) & p501(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p401(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p501(X0) & p601(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p202(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p302(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p102(X1) | ~ r1(X0, X1)) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p302(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p202(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p402(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p302(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p402(X0) & p502(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p402(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p502(X0) & p602(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & ! [X1] : (p203(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p303(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p103(X1) | ~ r1(X0, X1)) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p303(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p203(X1) | ~ r1(X0, X1)) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p403(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p303(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p403(X0) & p503(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p403(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p503(X0) & p603(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & ! [X1] : (p204(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & ! [X1] : (p304(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p104(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & ! [X1] : (p304(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p204(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p404(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p304(X1) | ~ r1(X0, X1)) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p404(X0) & p504(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p404(X0) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p504(X0) & p604(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p205(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p305(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p105(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & ! [X1] : (p305(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p205(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & ! [X1] : (p405(X1) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p305(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p405(X1) | ~ r1(X0, X1)) & p505(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (! [X1] : (p405(X1) | ~ r1(X0, X1)) & p605(X0)) | ~ r1(X1, X0)) | ~ ! [X0] : (~ (p505(X0) & p605(X0)) | ~ r1(X1, X0))) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL684+1.005.p', main)).
fof(f422596, plain, (~ sP20(sK91) | (~ spl92_12 | ~ spl92_17)), inference(resolution, [], [f422589, f227])).
fof(f422589, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP20(X0)) | (~ spl92_12 | ~ spl92_17)), inference(subsumption_resolution, [], [f422588, f471])).
fof(f471, plain, (p404(sK91) | ~ spl92_17), inference(avatar_component_clause, [], [f469])).
fof(f469, plain, (spl92_17 <=> p404(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_17])])).
fof(f422588, plain, (! [X0] : (~ p404(sK91) | ~ r1(X0, sK91) | ~ sP20(X0)) | (~ spl92_12 | ~ spl92_17)), inference(resolution, [], [f422585, f343])).
fof(f343, plain, ! [X0, X1] : (~ p304(sK60(X1)) | ~ p404(X1) | ~ r1(X0, X1) | ~ sP20(X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : (! [X1] : ((~ p304(sK60(X1)) & r1(X1, sK60(X1))) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP20(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK60])], [f132, f133])).
fof(f133, plain, ! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) => (~ p304(sK60(X1)) & r1(X1, sK60(X1)))), introduced(choice_axiom, [])).
fof(f132, plain, ! [X0] : (! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP20(X0)), inference(rectify, [], [f131])).
fof(f131, plain, ! [X11] : (! [X93] : (? [X94] : (~ p304(X94) & r1(X93, X94)) | ~ p404(X93) | ~ r1(X11, X93)) | ~ sP20(X11)), inference(nnf_transformation, [], [f31])).
fof(f422585, plain, (p304(sK60(sK91)) | (~ spl92_12 | ~ spl92_17)), inference(subsumption_resolution, [], [f422584, f544])).
fof(f422584, plain, (p304(sK60(sK91)) | ~ sP20(sK91) | (~ spl92_12 | ~ spl92_17)), inference(subsumption_resolution, [], [f422156, f471])).
fof(f422156, plain, (p304(sK60(sK91)) | ~ p404(sK91) | ~ sP20(sK91) | ~ spl92_12), inference(resolution, [], [f451, f2195])).
fof(f2195, plain, ! [X0] : (r1(X0, sK60(X0)) | ~ p404(X0) | ~ sP20(X0)), inference(resolution, [], [f342, f227])).
fof(f342, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p404(X1) | r1(X1, sK60(X1)) | ~ sP20(X0)), inference(cnf_transformation, [], [f134])).
fof(f451, plain, (! [X8] : (~ r1(sK91, X8) | p304(X8)) | ~ spl92_12), inference(avatar_component_clause, [], [f450])).
fof(f450, plain, (spl92_12 <=> ! [X8] : (p304(X8) | ~ r1(sK91, X8))), introduced(avatar_definition, [new_symbols(naming, [spl92_12])])).
fof(f421611, plain, (~ spl92_7 | ~ spl92_27), inference(avatar_contradiction_clause, [], [f421610])).
fof(f421610, plain, ($false | (~ spl92_7 | ~ spl92_27)), inference(subsumption_resolution, [], [f421609, f545])).
fof(f545, plain, sP21(sK91), inference(resolution, [], [f527, f250])).
fof(f250, plain, ! [X0] : (~ sP40(X0) | sP21(X0)), inference(cnf_transformation, [], [f54])).
fof(f421609, plain, (~ sP21(sK91) | (~ spl92_7 | ~ spl92_27)), inference(resolution, [], [f421606, f227])).
fof(f421606, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP21(X0)) | (~ spl92_7 | ~ spl92_27)), inference(subsumption_resolution, [], [f421605, f513])).
fof(f513, plain, (p604(sK91) | ~ spl92_27), inference(avatar_component_clause, [], [f511])).
fof(f511, plain, (spl92_27 <=> p604(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_27])])).
fof(f421605, plain, (! [X0] : (~ p604(sK91) | ~ r1(X0, sK91) | ~ sP21(X0)) | (~ spl92_7 | ~ spl92_27)), inference(resolution, [], [f421594, f341])).
fof(f341, plain, ! [X0, X1] : (~ p204(sK59(X1)) | ~ p604(X1) | ~ r1(X0, X1) | ~ sP21(X0)), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ! [X0] : (! [X1] : ((~ p204(sK59(X1)) & r1(X1, sK59(X1))) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP21(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK59])], [f128, f129])).
fof(f129, plain, ! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) => (~ p204(sK59(X1)) & r1(X1, sK59(X1)))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X0] : (! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP21(X0)), inference(rectify, [], [f127])).
fof(f127, plain, ! [X11] : (! [X91] : (? [X92] : (~ p204(X92) & r1(X91, X92)) | ~ p604(X91) | ~ r1(X11, X91)) | ~ sP21(X11)), inference(nnf_transformation, [], [f32])).
fof(f421594, plain, (p204(sK59(sK91)) | (~ spl92_7 | ~ spl92_27)), inference(subsumption_resolution, [], [f421593, f545])).
fof(f421593, plain, (p204(sK59(sK91)) | ~ sP21(sK91) | (~ spl92_7 | ~ spl92_27)), inference(subsumption_resolution, [], [f421163, f513])).
fof(f421163, plain, (p204(sK59(sK91)) | ~ p604(sK91) | ~ sP21(sK91) | ~ spl92_7), inference(resolution, [], [f433, f2177])).
fof(f2177, plain, ! [X0] : (r1(X0, sK59(X0)) | ~ p604(X0) | ~ sP21(X0)), inference(resolution, [], [f340, f227])).
fof(f340, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p604(X1) | r1(X1, sK59(X1)) | ~ sP21(X0)), inference(cnf_transformation, [], [f130])).
fof(f433, plain, (! [X6] : (~ r1(sK91, X6) | p204(X6)) | ~ spl92_7), inference(avatar_component_clause, [], [f432])).
fof(f432, plain, (spl92_7 <=> ! [X6] : (p204(X6) | ~ r1(sK91, X6))), introduced(avatar_definition, [new_symbols(naming, [spl92_7])])).
fof(f421107, plain, (~ spl92_21 | ~ spl92_1170), inference(avatar_contradiction_clause, [], [f421106])).
fof(f421106, plain, ($false | (~ spl92_21 | ~ spl92_1170)), inference(subsumption_resolution, [], [f421105, f534])).
fof(f534, plain, sP15(sK91), inference(resolution, [], [f527, f236])).
fof(f236, plain, ! [X0] : (~ sP40(X0) | sP15(X0)), inference(cnf_transformation, [], [f54])).
fof(f421105, plain, (~ sP15(sK91) | (~ spl92_21 | ~ spl92_1170)), inference(resolution, [], [f421100, f227])).
fof(f421100, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP15(X0)) | (~ spl92_21 | ~ spl92_1170)), inference(subsumption_resolution, [], [f421099, f488])).
fof(f488, plain, (p505(sK91) | ~ spl92_21), inference(avatar_component_clause, [], [f486])).
fof(f486, plain, (spl92_21 <=> p505(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_21])])).
fof(f421099, plain, (! [X0] : (~ p505(sK91) | ~ r1(X0, sK91) | ~ sP15(X0)) | ~ spl92_1170), inference(resolution, [], [f14661, f353])).
fof(f353, plain, ! [X0, X1] : (~ p205(sK65(X1)) | ~ p505(X1) | ~ r1(X0, X1) | ~ sP15(X0)), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ! [X0] : (! [X1] : ((~ p205(sK65(X1)) & r1(X1, sK65(X1))) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP15(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK65])], [f152, f153])).
fof(f153, plain, ! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) => (~ p205(sK65(X1)) & r1(X1, sK65(X1)))), introduced(choice_axiom, [])).
fof(f152, plain, ! [X0] : (! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP15(X0)), inference(rectify, [], [f151])).
fof(f151, plain, ! [X11] : (! [X121] : (? [X122] : (~ p205(X122) & r1(X121, X122)) | ~ p505(X121) | ~ r1(X11, X121)) | ~ sP15(X11)), inference(nnf_transformation, [], [f26])).
fof(f14661, plain, (p205(sK65(sK91)) | ~ spl92_1170), inference(avatar_component_clause, [], [f14659])).
fof(f14659, plain, (spl92_1170 <=> p205(sK65(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1170])])).
fof(f420115, plain, (~ spl92_1 | ~ spl92_21), inference(avatar_contradiction_clause, [], [f420114])).
fof(f420114, plain, ($false | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f420113, f538])).
fof(f538, plain, sP17(sK91), inference(resolution, [], [f527, f240])).
fof(f240, plain, ! [X0] : (~ sP40(X0) | sP17(X0)), inference(cnf_transformation, [], [f54])).
fof(f420113, plain, (~ sP17(sK91) | (~ spl92_1 | ~ spl92_21)), inference(resolution, [], [f420106, f227])).
fof(f420106, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP17(X0)) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f420105, f488])).
fof(f420105, plain, (! [X0] : (~ p505(sK91) | ~ r1(X0, sK91) | ~ sP17(X0)) | (~ spl92_1 | ~ spl92_21)), inference(resolution, [], [f420098, f349])).
fof(f349, plain, ! [X0, X1] : (~ p105(sK63(X1)) | ~ p505(X1) | ~ r1(X0, X1) | ~ sP17(X0)), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ! [X0] : (! [X1] : ((~ p105(sK63(X1)) & r1(X1, sK63(X1))) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP17(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK63])], [f144, f145])).
fof(f145, plain, ! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) => (~ p105(sK63(X1)) & r1(X1, sK63(X1)))), introduced(choice_axiom, [])).
fof(f144, plain, ! [X0] : (! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP17(X0)), inference(rectify, [], [f143])).
fof(f143, plain, ! [X11] : (! [X111] : (? [X112] : (~ p105(X112) & r1(X111, X112)) | ~ p505(X111) | ~ r1(X11, X111)) | ~ sP17(X11)), inference(nnf_transformation, [], [f28])).
fof(f420098, plain, (p105(sK63(sK91)) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f420097, f538])).
fof(f420097, plain, (p105(sK63(sK91)) | ~ sP17(sK91) | (~ spl92_1 | ~ spl92_21)), inference(subsumption_resolution, [], [f419668, f488])).
fof(f419668, plain, (p105(sK63(sK91)) | ~ p505(sK91) | ~ sP17(sK91) | ~ spl92_1), inference(resolution, [], [f413, f2433])).
fof(f2433, plain, ! [X0] : (r1(X0, sK63(X0)) | ~ p505(X0) | ~ sP17(X0)), inference(resolution, [], [f348, f227])).
fof(f348, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p505(X1) | r1(X1, sK63(X1)) | ~ sP17(X0)), inference(cnf_transformation, [], [f146])).
fof(f413, plain, (! [X4] : (~ r1(sK91, X4) | p105(X4)) | ~ spl92_1), inference(avatar_component_clause, [], [f412])).
fof(f412, plain, (spl92_1 <=> ! [X4] : (p105(X4) | ~ r1(sK91, X4))), introduced(avatar_definition, [new_symbols(naming, [spl92_1])])).
fof(f419602, plain, ~ spl92_8842, inference(avatar_contradiction_clause, [], [f419601])).
fof(f419601, plain, ($false | ~ spl92_8842), inference(subsumption_resolution, [], [f419600, f550])).
fof(f550, plain, sP25(sK91), inference(resolution, [], [f527, f255])).
fof(f255, plain, ! [X0] : (~ sP40(X0) | sP25(X0)), inference(cnf_transformation, [], [f54])).
fof(f419600, plain, (~ sP25(sK91) | ~ spl92_8842), inference(resolution, [], [f265811, f227])).
fof(f265811, plain, (! [X51] : (~ r1(X51, sK91) | ~ sP25(X51)) | ~ spl92_8842), inference(avatar_component_clause, [], [f265810])).
fof(f265810, plain, (spl92_8842 <=> ! [X51] : (~ r1(X51, sK91) | ~ sP25(X51))), introduced(avatar_definition, [new_symbols(naming, [spl92_8842])])).
fof(f419597, plain, (spl92_8842 | ~ spl92_2 | ~ spl92_22), inference(avatar_split_clause, [], [f419596, f490, f415, f265810])).
fof(f415, plain, (spl92_2 <=> ! [X3] : (p104(X3) | ~ r1(sK91, X3))), introduced(avatar_definition, [new_symbols(naming, [spl92_2])])).
fof(f490, plain, (spl92_22 <=> p504(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_22])])).
fof(f419596, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP25(X0)) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f419595, f492])).
fof(f492, plain, (p504(sK91) | ~ spl92_22), inference(avatar_component_clause, [], [f490])).
fof(f419595, plain, (! [X0] : (~ p504(sK91) | ~ r1(X0, sK91) | ~ sP25(X0)) | (~ spl92_2 | ~ spl92_22)), inference(resolution, [], [f419582, f333])).
fof(f333, plain, ! [X0, X1] : (~ p104(sK55(X1)) | ~ p504(X1) | ~ r1(X0, X1) | ~ sP25(X0)), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : (! [X1] : ((~ p104(sK55(X1)) & r1(X1, sK55(X1))) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP25(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK55])], [f112, f113])).
fof(f113, plain, ! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) => (~ p104(sK55(X1)) & r1(X1, sK55(X1)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X0] : (! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP25(X0)), inference(rectify, [], [f111])).
fof(f111, plain, ! [X11] : (! [X80] : (? [X81] : (~ p104(X81) & r1(X80, X81)) | ~ p504(X80) | ~ r1(X11, X80)) | ~ sP25(X11)), inference(nnf_transformation, [], [f36])).
fof(f419582, plain, (p104(sK55(sK91)) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f419581, f550])).
fof(f419581, plain, (p104(sK55(sK91)) | ~ sP25(sK91) | (~ spl92_2 | ~ spl92_22)), inference(subsumption_resolution, [], [f419146, f492])).
fof(f419146, plain, (p104(sK55(sK91)) | ~ p504(sK91) | ~ sP25(sK91) | ~ spl92_2), inference(resolution, [], [f416, f2012])).
fof(f2012, plain, ! [X0] : (r1(X0, sK55(X0)) | ~ p504(X0) | ~ sP25(X0)), inference(resolution, [], [f332, f227])).
fof(f332, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p504(X1) | r1(X1, sK55(X1)) | ~ sP25(X0)), inference(cnf_transformation, [], [f114])).
fof(f416, plain, (! [X3] : (~ r1(sK91, X3) | p104(X3)) | ~ spl92_2), inference(avatar_component_clause, [], [f415])).
fof(f419096, plain, ~ spl92_8844, inference(avatar_contradiction_clause, [], [f419095])).
fof(f419095, plain, ($false | ~ spl92_8844), inference(subsumption_resolution, [], [f419094, f537])).
fof(f537, plain, sP16(sK91), inference(resolution, [], [f527, f239])).
fof(f239, plain, ! [X0] : (~ sP40(X0) | sP16(X0)), inference(cnf_transformation, [], [f54])).
fof(f419094, plain, (~ sP16(sK91) | ~ spl92_8844), inference(resolution, [], [f265820, f227])).
fof(f265820, plain, (! [X60] : (~ r1(X60, sK91) | ~ sP16(X60)) | ~ spl92_8844), inference(avatar_component_clause, [], [f265819])).
fof(f265819, plain, (spl92_8844 <=> ! [X60] : (~ r1(X60, sK91) | ~ sP16(X60))), introduced(avatar_definition, [new_symbols(naming, [spl92_8844])])).
fof(f415560, plain, (spl92_8844 | ~ spl92_1 | ~ spl92_26), inference(avatar_split_clause, [], [f415559, f507, f412, f265819])).
fof(f507, plain, (spl92_26 <=> p605(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_26])])).
fof(f415559, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP16(X0)) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f415558, f509])).
fof(f509, plain, (p605(sK91) | ~ spl92_26), inference(avatar_component_clause, [], [f507])).
fof(f415558, plain, (! [X0] : (~ p605(sK91) | ~ r1(X0, sK91) | ~ sP16(X0)) | (~ spl92_1 | ~ spl92_26)), inference(resolution, [], [f415551, f351])).
fof(f351, plain, ! [X0, X1] : (~ p105(sK64(X1)) | ~ p605(X1) | ~ r1(X0, X1) | ~ sP16(X0)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0] : (! [X1] : ((~ p105(sK64(X1)) & r1(X1, sK64(X1))) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP16(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK64])], [f148, f149])).
fof(f149, plain, ! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) => (~ p105(sK64(X1)) & r1(X1, sK64(X1)))), introduced(choice_axiom, [])).
fof(f148, plain, ! [X0] : (! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP16(X0)), inference(rectify, [], [f147])).
fof(f147, plain, ! [X11] : (! [X113] : (? [X114] : (~ p105(X114) & r1(X113, X114)) | ~ p605(X113) | ~ r1(X11, X113)) | ~ sP16(X11)), inference(nnf_transformation, [], [f27])).
fof(f415551, plain, (p105(sK64(sK91)) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f415550, f537])).
fof(f415550, plain, (p105(sK64(sK91)) | ~ sP16(sK91) | (~ spl92_1 | ~ spl92_26)), inference(subsumption_resolution, [], [f415131, f509])).
fof(f415131, plain, (p105(sK64(sK91)) | ~ p605(sK91) | ~ sP16(sK91) | ~ spl92_1), inference(resolution, [], [f413, f2449])).
fof(f2449, plain, ! [X0] : (r1(X0, sK64(X0)) | ~ p605(X0) | ~ sP16(X0)), inference(resolution, [], [f350, f227])).
fof(f350, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p605(X1) | r1(X1, sK64(X1)) | ~ sP16(X0)), inference(cnf_transformation, [], [f150])).
fof(f415064, plain, (~ spl92_3 | ~ spl92_18), inference(avatar_contradiction_clause, [], [f415063])).
fof(f415063, plain, ($false | (~ spl92_3 | ~ spl92_18)), inference(subsumption_resolution, [], [f415062, f560])).
fof(f560, plain, sP33(sK91), inference(resolution, [], [f527, f271])).
fof(f271, plain, ! [X0] : (~ sP40(X0) | sP33(X0)), inference(cnf_transformation, [], [f54])).
fof(f415062, plain, (~ sP33(sK91) | (~ spl92_3 | ~ spl92_18)), inference(resolution, [], [f415059, f227])).
fof(f415059, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP33(X0)) | (~ spl92_3 | ~ spl92_18)), inference(subsumption_resolution, [], [f415058, f475])).
fof(f475, plain, (p403(sK91) | ~ spl92_18), inference(avatar_component_clause, [], [f473])).
fof(f473, plain, (spl92_18 <=> p403(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_18])])).
fof(f415058, plain, (! [X0] : (~ p403(sK91) | ~ r1(X0, sK91) | ~ sP33(X0)) | (~ spl92_3 | ~ spl92_18)), inference(resolution, [], [f414571, f317])).
fof(f317, plain, ! [X0, X1] : (~ p103(sK47(X1)) | ~ p403(X1) | ~ r1(X0, X1) | ~ sP33(X0)), inference(cnf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : (! [X1] : ((~ p103(sK47(X1)) & r1(X1, sK47(X1))) | ~ p403(X1) | ~ r1(X0, X1)) | ~ sP33(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK47])], [f80, f81])).
fof(f81, plain, ! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) => (~ p103(sK47(X1)) & r1(X1, sK47(X1)))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X0] : (! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) | ~ p403(X1) | ~ r1(X0, X1)) | ~ sP33(X0)), inference(rectify, [], [f79])).
fof(f79, plain, ! [X11] : (! [X52] : (? [X53] : (~ p103(X53) & r1(X52, X53)) | ~ p403(X52) | ~ r1(X11, X52)) | ~ sP33(X11)), inference(nnf_transformation, [], [f44])).
fof(f414571, plain, (p103(sK47(sK91)) | (~ spl92_3 | ~ spl92_18)), inference(subsumption_resolution, [], [f414570, f560])).
fof(f414570, plain, (p103(sK47(sK91)) | ~ sP33(sK91) | (~ spl92_3 | ~ spl92_18)), inference(subsumption_resolution, [], [f414119, f475])).
fof(f414119, plain, (p103(sK47(sK91)) | ~ p403(sK91) | ~ sP33(sK91) | ~ spl92_3), inference(resolution, [], [f419, f1198])).
fof(f1198, plain, ! [X0] : (r1(X0, sK47(X0)) | ~ p403(X0) | ~ sP33(X0)), inference(resolution, [], [f316, f227])).
fof(f316, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p403(X1) | r1(X1, sK47(X1)) | ~ sP33(X0)), inference(cnf_transformation, [], [f82])).
fof(f419, plain, (! [X2] : (~ r1(sK91, X2) | p103(X2)) | ~ spl92_3), inference(avatar_component_clause, [], [f418])).
fof(f418, plain, (spl92_3 <=> ! [X2] : (p103(X2) | ~ r1(sK91, X2))), introduced(avatar_definition, [new_symbols(naming, [spl92_3])])).
fof(f414086, plain, (~ spl92_23 | ~ spl92_18), inference(avatar_split_clause, [], [f625, f473, f494])).
fof(f494, plain, (spl92_23 <=> p503(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_23])])).
fof(f625, plain, (~ p403(sK91) | ~ p503(sK91)), inference(resolution, [], [f579, f527])).
fof(f579, plain, ! [X0] : (~ sP40(X0) | ~ p403(X0) | ~ p503(X0)), inference(resolution, [], [f261, f227])).
fof(f261, plain, ! [X0, X29] : (~ r1(X0, X29) | ~ p503(X29) | ~ p403(X29) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f414085, plain, (~ spl92_14 | ~ spl92_9), inference(avatar_split_clause, [], [f643, f438, f457])).
fof(f457, plain, (spl92_14 <=> p302(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_14])])).
fof(f438, plain, (spl92_9 <=> p202(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_9])])).
fof(f643, plain, (~ p202(sK91) | ~ p302(sK91)), inference(resolution, [], [f592, f527])).
fof(f592, plain, ! [X0] : (~ sP40(X0) | ~ p202(X0) | ~ p302(X0)), inference(resolution, [], [f283, f227])).
fof(f283, plain, ! [X0, X16] : (~ r1(X0, X16) | ~ p302(X16) | ~ p202(X16) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f414084, plain, ~ spl92_8845, inference(avatar_contradiction_clause, [], [f414083])).
fof(f414083, plain, ($false | ~ spl92_8845), inference(subsumption_resolution, [], [f414082, f533])).
fof(f533, plain, sP14(sK91), inference(resolution, [], [f527, f235])).
fof(f235, plain, ! [X0] : (~ sP40(X0) | sP14(X0)), inference(cnf_transformation, [], [f54])).
fof(f414082, plain, (~ sP14(sK91) | ~ spl92_8845), inference(resolution, [], [f265824, f227])).
fof(f265824, plain, (! [X62] : (~ r1(X62, sK91) | ~ sP14(X62)) | ~ spl92_8845), inference(avatar_component_clause, [], [f265823])).
fof(f265823, plain, (spl92_8845 <=> ! [X62] : (~ r1(X62, sK91) | ~ sP14(X62))), introduced(avatar_definition, [new_symbols(naming, [spl92_8845])])).
fof(f406129, plain, (spl92_8845 | ~ spl92_26 | ~ spl92_1182), inference(avatar_split_clause, [], [f406128, f14731, f507, f265823])).
fof(f14731, plain, (spl92_1182 <=> p205(sK66(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1182])])).
fof(f406128, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP14(X0)) | (~ spl92_26 | ~ spl92_1182)), inference(subsumption_resolution, [], [f406127, f509])).
fof(f406127, plain, (! [X0] : (~ p605(sK91) | ~ r1(X0, sK91) | ~ sP14(X0)) | ~ spl92_1182), inference(resolution, [], [f14733, f355])).
fof(f355, plain, ! [X0, X1] : (~ p205(sK66(X1)) | ~ p605(X1) | ~ r1(X0, X1) | ~ sP14(X0)), inference(cnf_transformation, [], [f158])).
fof(f158, plain, ! [X0] : (! [X1] : ((~ p205(sK66(X1)) & r1(X1, sK66(X1))) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP14(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK66])], [f156, f157])).
fof(f157, plain, ! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) => (~ p205(sK66(X1)) & r1(X1, sK66(X1)))), introduced(choice_axiom, [])).
fof(f156, plain, ! [X0] : (! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP14(X0)), inference(rectify, [], [f155])).
fof(f155, plain, ! [X11] : (! [X123] : (? [X124] : (~ p205(X124) & r1(X123, X124)) | ~ p605(X123) | ~ r1(X11, X123)) | ~ sP14(X11)), inference(nnf_transformation, [], [f25])).
fof(f14733, plain, (p205(sK66(sK91)) | ~ spl92_1182), inference(avatar_component_clause, [], [f14731])).
fof(f405627, plain, (~ spl92_26 | spl92_1182 | ~ spl92_6), inference(avatar_split_clause, [], [f405626, f429, f14731, f507])).
fof(f429, plain, (spl92_6 <=> ! [X7] : (p205(X7) | ~ r1(sK91, X7))), introduced(avatar_definition, [new_symbols(naming, [spl92_6])])).
fof(f405626, plain, (p205(sK66(sK91)) | ~ p605(sK91) | ~ spl92_6), inference(subsumption_resolution, [], [f404583, f533])).
fof(f404583, plain, (p205(sK66(sK91)) | ~ p605(sK91) | ~ sP14(sK91) | ~ spl92_6), inference(resolution, [], [f430, f2481])).
fof(f2481, plain, ! [X0] : (r1(X0, sK66(X0)) | ~ p605(X0) | ~ sP14(X0)), inference(resolution, [], [f354, f227])).
fof(f354, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p605(X1) | r1(X1, sK66(X1)) | ~ sP14(X0)), inference(cnf_transformation, [], [f158])).
fof(f430, plain, (! [X7] : (~ r1(sK91, X7) | p205(X7)) | ~ spl92_6), inference(avatar_component_clause, [], [f429])).
fof(f405585, plain, (~ spl92_3 | ~ spl92_13), inference(avatar_contradiction_clause, [], [f405584])).
fof(f405584, plain, ($false | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f405583, f561])).
fof(f561, plain, sP34(sK91), inference(resolution, [], [f527, f272])).
fof(f272, plain, ! [X0] : (~ sP40(X0) | sP34(X0)), inference(cnf_transformation, [], [f54])).
fof(f405583, plain, (~ sP34(sK91) | (~ spl92_3 | ~ spl92_13)), inference(resolution, [], [f405576, f227])).
fof(f405576, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP34(X0)) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f405575, f455])).
fof(f455, plain, (p303(sK91) | ~ spl92_13), inference(avatar_component_clause, [], [f453])).
fof(f453, plain, (spl92_13 <=> p303(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_13])])).
fof(f405575, plain, (! [X0] : (~ p303(sK91) | ~ r1(X0, sK91) | ~ sP34(X0)) | (~ spl92_3 | ~ spl92_13)), inference(resolution, [], [f405566, f315])).
fof(f315, plain, ! [X0, X1] : (~ p103(sK46(X1)) | ~ p303(X1) | ~ r1(X0, X1) | ~ sP34(X0)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ! [X0] : (! [X1] : ((~ p103(sK46(X1)) & r1(X1, sK46(X1))) | ~ p303(X1) | ~ r1(X0, X1)) | ~ sP34(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46])], [f76, f77])).
fof(f77, plain, ! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) => (~ p103(sK46(X1)) & r1(X1, sK46(X1)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X0] : (! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) | ~ p303(X1) | ~ r1(X0, X1)) | ~ sP34(X0)), inference(rectify, [], [f75])).
fof(f75, plain, ! [X11] : (! [X50] : (? [X51] : (~ p103(X51) & r1(X50, X51)) | ~ p303(X50) | ~ r1(X11, X50)) | ~ sP34(X11)), inference(nnf_transformation, [], [f45])).
fof(f405566, plain, (p103(sK46(sK91)) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f405565, f561])).
fof(f405565, plain, (p103(sK46(sK91)) | ~ sP34(sK91) | (~ spl92_3 | ~ spl92_13)), inference(subsumption_resolution, [], [f405114, f455])).
fof(f405114, plain, (p103(sK46(sK91)) | ~ p303(sK91) | ~ sP34(sK91) | ~ spl92_3), inference(resolution, [], [f419, f997])).
fof(f997, plain, ! [X0] : (r1(X0, sK46(X0)) | ~ p303(X0) | ~ sP34(X0)), inference(resolution, [], [f314, f227])).
fof(f314, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p303(X1) | r1(X1, sK46(X1)) | ~ sP34(X0)), inference(cnf_transformation, [], [f78])).
fof(f405074, plain, (~ spl92_24 | ~ spl92_19), inference(avatar_split_clause, [], [f634, f477, f498])).
fof(f498, plain, (spl92_24 <=> p502(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_24])])).
fof(f477, plain, (spl92_19 <=> p402(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_19])])).
fof(f634, plain, (~ p402(sK91) | ~ p502(sK91)), inference(resolution, [], [f585, f527])).
fof(f585, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p502(X0)), inference(resolution, [], [f276, f227])).
fof(f276, plain, ! [X0, X23] : (~ r1(X0, X23) | ~ p502(X23) | ~ p402(X23) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f405073, plain, (~ spl92_15 | ~ spl92_5), inference(avatar_split_clause, [], [f662, f424, f461])).
fof(f461, plain, (spl92_15 <=> p301(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_15])])).
fof(f424, plain, (spl92_5 <=> p101(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_5])])).
fof(f662, plain, (~ p101(sK91) | ~ p301(sK91)), inference(resolution, [], [f606, f527])).
fof(f606, plain, ! [X0] : (~ sP40(X0) | ~ p101(X0) | ~ p301(X0)), inference(resolution, [], [f302, f227])).
fof(f302, plain, ! [X2, X0] : (~ r1(X0, X2) | ~ p301(X2) | ~ p101(X2) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f405002, plain, (~ spl92_24 | ~ spl92_14), inference(avatar_split_clause, [], [f637, f457, f498])).
fof(f637, plain, (~ p302(sK91) | ~ p502(sK91)), inference(resolution, [], [f587, f527])).
fof(f587, plain, ! [X0] : (~ sP40(X0) | ~ p302(X0) | ~ p502(X0)), inference(resolution, [], [f278, f227])).
fof(f278, plain, ! [X0, X21] : (~ r1(X0, X21) | ~ p502(X21) | ~ p302(X21) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f404997, plain, (spl92_1360 | ~ spl92_1 | ~ spl92_146 | ~ spl92_148), inference(avatar_split_clause, [], [f404996, f5810, f5798, f412, f17515])).
fof(f17515, plain, (spl92_1360 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP5(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1360])])).
fof(f5798, plain, (spl92_146 <=> r1(sK91, sK79(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_146])])).
fof(f5810, plain, (spl92_148 <=> p205(sK80(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_148])])).
fof(f404996, plain, (! [X1] : (~ r1(X1, sK91) | ~ sP5(X1)) | (~ spl92_1 | ~ spl92_146 | ~ spl92_148)), inference(subsumption_resolution, [], [f404995, f5812])).
fof(f5812, plain, (p205(sK80(sK91)) | ~ spl92_148), inference(avatar_component_clause, [], [f5810])).
fof(f404995, plain, (! [X1] : (~ p205(sK80(sK91)) | ~ r1(X1, sK91) | ~ sP5(X1)) | (~ spl92_1 | ~ spl92_146)), inference(resolution, [], [f404033, f383])).
fof(f383, plain, ! [X0, X1] : (~ p105(sK79(X1)) | ~ p205(sK80(X1)) | ~ r1(X0, X1) | ~ sP5(X0)), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ! [X0] : (! [X1] : ((~ p105(sK79(X1)) & r1(X1, sK79(X1))) | (~ p205(sK80(X1)) & r1(X1, sK80(X1))) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK79, sK80])], [f196, f198, f197])).
fof(f197, plain, ! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) => (~ p105(sK79(X1)) & r1(X1, sK79(X1)))), introduced(choice_axiom, [])).
fof(f198, plain, ! [X1] : (? [X3] : (~ p205(X3) & r1(X1, X3)) => (~ p205(sK80(X1)) & r1(X1, sK80(X1)))), introduced(choice_axiom, [])).
fof(f196, plain, ! [X0] : (! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) | ? [X3] : (~ p205(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP5(X0)), inference(rectify, [], [f195])).
fof(f195, plain, ! [X11] : (! [X102] : (? [X103] : (~ p105(X103) & r1(X102, X103)) | ? [X104] : (~ p205(X104) & r1(X102, X104)) | ~ r1(X11, X102)) | ~ sP5(X11)), inference(nnf_transformation, [], [f16])).
fof(f404033, plain, (p105(sK79(sK91)) | (~ spl92_1 | ~ spl92_146)), inference(resolution, [], [f413, f5800])).
fof(f5800, plain, (r1(sK91, sK79(sK91)) | ~ spl92_146), inference(avatar_component_clause, [], [f5798])).
fof(f404024, plain, (~ spl92_15 | ~ spl92_10), inference(avatar_split_clause, [], [f657, f442, f461])).
fof(f442, plain, (spl92_10 <=> p201(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_10])])).
fof(f657, plain, (~ p201(sK91) | ~ p301(sK91)), inference(resolution, [], [f602, f527])).
fof(f602, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p301(X0)), inference(resolution, [], [f298, f227])).
fof(f298, plain, ! [X6, X0] : (~ r1(X0, X6) | ~ p301(X6) | ~ p201(X6) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f403883, plain, (~ spl92_13 | ~ spl92_23), inference(avatar_split_clause, [], [f403882, f494, f453])).
fof(f403882, plain, (~ p303(sK91) | ~ spl92_23), inference(subsumption_resolution, [], [f628, f496])).
fof(f496, plain, (p503(sK91) | ~ spl92_23), inference(avatar_component_clause, [], [f494])).
fof(f628, plain, (~ p303(sK91) | ~ p503(sK91)), inference(resolution, [], [f581, f527])).
fof(f581, plain, ! [X0] : (~ sP40(X0) | ~ p303(X0) | ~ p503(X0)), inference(resolution, [], [f263, f227])).
fof(f263, plain, ! [X0, X27] : (~ r1(X0, X27) | ~ p503(X27) | ~ p303(X27) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f403879, plain, (~ spl92_20 | ~ spl92_15), inference(avatar_split_clause, [], [f652, f461, f481])).
fof(f481, plain, (spl92_20 <=> p401(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_20])])).
fof(f652, plain, (~ p301(sK91) | ~ p401(sK91)), inference(resolution, [], [f598, f527])).
fof(f598, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p401(X0)), inference(resolution, [], [f294, f227])).
fof(f294, plain, ! [X0, X10] : (~ r1(X0, X10) | ~ p401(X10) | ~ p301(X10) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f403878, plain, (~ spl92_20 | ~ spl92_10), inference(avatar_split_clause, [], [f656, f442, f481])).
fof(f656, plain, (~ p201(sK91) | ~ p401(sK91)), inference(resolution, [], [f601, f527])).
fof(f601, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p401(X0)), inference(resolution, [], [f297, f227])).
fof(f297, plain, ! [X0, X7] : (~ r1(X0, X7) | ~ p401(X7) | ~ p201(X7) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f403871, plain, (~ spl92_19 | ~ spl92_9), inference(avatar_split_clause, [], [f642, f438, f477])).
fof(f642, plain, (~ p202(sK91) | ~ p402(sK91)), inference(resolution, [], [f591, f527])).
fof(f591, plain, ! [X0] : (~ sP40(X0) | ~ p202(X0) | ~ p402(X0)), inference(resolution, [], [f282, f227])).
fof(f282, plain, ! [X0, X17] : (~ r1(X0, X17) | ~ p402(X17) | ~ p202(X17) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f403850, plain, (~ spl92_2 | ~ spl92_17), inference(avatar_contradiction_clause, [], [f403849])).
fof(f403849, plain, ($false | (~ spl92_2 | ~ spl92_17)), inference(subsumption_resolution, [], [f403848, f551])).
fof(f551, plain, sP26(sK91), inference(resolution, [], [f527, f256])).
fof(f256, plain, ! [X0] : (~ sP40(X0) | sP26(X0)), inference(cnf_transformation, [], [f54])).
fof(f403848, plain, (~ sP26(sK91) | (~ spl92_2 | ~ spl92_17)), inference(resolution, [], [f403843, f227])).
fof(f403843, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP26(X0)) | (~ spl92_2 | ~ spl92_17)), inference(subsumption_resolution, [], [f403842, f471])).
fof(f403842, plain, (! [X0] : (~ p404(sK91) | ~ r1(X0, sK91) | ~ sP26(X0)) | (~ spl92_2 | ~ spl92_17)), inference(resolution, [], [f403831, f331])).
fof(f331, plain, ! [X0, X1] : (~ p104(sK54(X1)) | ~ p404(X1) | ~ r1(X0, X1) | ~ sP26(X0)), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ! [X0] : (! [X1] : ((~ p104(sK54(X1)) & r1(X1, sK54(X1))) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP26(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f108, f109])).
fof(f109, plain, ! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) => (~ p104(sK54(X1)) & r1(X1, sK54(X1)))), introduced(choice_axiom, [])).
fof(f108, plain, ! [X0] : (! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP26(X0)), inference(rectify, [], [f107])).
fof(f107, plain, ! [X11] : (! [X78] : (? [X79] : (~ p104(X79) & r1(X78, X79)) | ~ p404(X78) | ~ r1(X11, X78)) | ~ sP26(X11)), inference(nnf_transformation, [], [f37])).
fof(f403831, plain, (p104(sK54(sK91)) | (~ spl92_2 | ~ spl92_17)), inference(subsumption_resolution, [], [f403830, f551])).
fof(f403830, plain, (p104(sK54(sK91)) | ~ sP26(sK91) | (~ spl92_2 | ~ spl92_17)), inference(subsumption_resolution, [], [f403377, f471])).
fof(f403377, plain, (p104(sK54(sK91)) | ~ p404(sK91) | ~ sP26(sK91) | ~ spl92_2), inference(resolution, [], [f416, f1907])).
fof(f1907, plain, ! [X0] : (r1(X0, sK54(X0)) | ~ p404(X0) | ~ sP26(X0)), inference(resolution, [], [f330, f227])).
fof(f330, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p404(X1) | r1(X1, sK54(X1)) | ~ sP26(X0)), inference(cnf_transformation, [], [f110])).
fof(f402713, plain, (~ spl92_4 | ~ spl92_19), inference(avatar_contradiction_clause, [], [f402712])).
fof(f402712, plain, ($false | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f402711, f565])).
fof(f565, plain, sP37(sK91), inference(resolution, [], [f527, f286])).
fof(f286, plain, ! [X0] : (~ sP40(X0) | sP37(X0)), inference(cnf_transformation, [], [f54])).
fof(f402711, plain, (~ sP37(sK91) | (~ spl92_4 | ~ spl92_19)), inference(resolution, [], [f402710, f227])).
fof(f402710, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP37(X0)) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f402709, f479])).
fof(f479, plain, (p402(sK91) | ~ spl92_19), inference(avatar_component_clause, [], [f477])).
fof(f402709, plain, (! [X0] : (~ p402(sK91) | ~ r1(X0, sK91) | ~ sP37(X0)) | (~ spl92_4 | ~ spl92_19)), inference(resolution, [], [f402708, f309])).
fof(f309, plain, ! [X0, X1] : (~ p102(sK43(X1)) | ~ p402(X1) | ~ r1(X0, X1) | ~ sP37(X0)), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ! [X0] : (! [X1] : ((~ p102(sK43(X1)) & r1(X1, sK43(X1))) | ~ p402(X1) | ~ r1(X0, X1)) | ~ sP37(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK43])], [f64, f65])).
fof(f65, plain, ! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) => (~ p102(sK43(X1)) & r1(X1, sK43(X1)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X0] : (! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) | ~ p402(X1) | ~ r1(X0, X1)) | ~ sP37(X0)), inference(rectify, [], [f63])).
fof(f63, plain, ! [X11] : (! [X31] : (? [X32] : (~ p102(X32) & r1(X31, X32)) | ~ p402(X31) | ~ r1(X11, X31)) | ~ sP37(X11)), inference(nnf_transformation, [], [f48])).
fof(f402708, plain, (p102(sK43(sK91)) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f402707, f565])).
fof(f402707, plain, (p102(sK43(sK91)) | ~ sP37(sK91) | (~ spl92_4 | ~ spl92_19)), inference(subsumption_resolution, [], [f402230, f479])).
fof(f402230, plain, (p102(sK43(sK91)) | ~ p402(sK91) | ~ sP37(sK91) | ~ spl92_4), inference(resolution, [], [f422, f745])).
fof(f745, plain, ! [X0] : (r1(X0, sK43(X0)) | ~ p402(X0) | ~ sP37(X0)), inference(resolution, [], [f308, f227])).
fof(f308, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p402(X1) | r1(X1, sK43(X1)) | ~ sP37(X0)), inference(cnf_transformation, [], [f66])).
fof(f422, plain, (! [X1] : (~ r1(sK91, X1) | p102(X1)) | ~ spl92_4), inference(avatar_component_clause, [], [f421])).
fof(f421, plain, (spl92_4 <=> ! [X1] : (p102(X1) | ~ r1(sK91, X1))), introduced(avatar_definition, [new_symbols(naming, [spl92_4])])).
fof(f401764, plain, (spl92_8843 | ~ spl92_22 | ~ spl92_10165), inference(avatar_split_clause, [], [f366124, f365030, f490, f265814])).
fof(f265814, plain, (spl92_8843 <=> ! [X54] : (~ r1(X54, sK91) | ~ sP22(X54))), introduced(avatar_definition, [new_symbols(naming, [spl92_8843])])).
fof(f365030, plain, (spl92_10165 <=> p204(sK58(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_10165])])).
fof(f366124, plain, (! [X0] : (~ p504(sK91) | ~ r1(X0, sK91) | ~ sP22(X0)) | ~ spl92_10165), inference(resolution, [], [f365032, f339])).
fof(f339, plain, ! [X0, X1] : (~ p204(sK58(X1)) | ~ p504(X1) | ~ r1(X0, X1) | ~ sP22(X0)), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ! [X0] : (! [X1] : ((~ p204(sK58(X1)) & r1(X1, sK58(X1))) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP22(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58])], [f124, f125])).
fof(f125, plain, ! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) => (~ p204(sK58(X1)) & r1(X1, sK58(X1)))), introduced(choice_axiom, [])).
fof(f124, plain, ! [X0] : (! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP22(X0)), inference(rectify, [], [f123])).
fof(f123, plain, ! [X11] : (! [X89] : (? [X90] : (~ p204(X90) & r1(X89, X90)) | ~ p504(X89) | ~ r1(X11, X89)) | ~ sP22(X11)), inference(nnf_transformation, [], [f33])).
fof(f365032, plain, (p204(sK58(sK91)) | ~ spl92_10165), inference(avatar_component_clause, [], [f365030])).
fof(f401755, plain, ~ spl92_8843, inference(avatar_contradiction_clause, [], [f401754])).
fof(f401754, plain, ($false | ~ spl92_8843), inference(subsumption_resolution, [], [f401753, f546])).
fof(f546, plain, sP22(sK91), inference(resolution, [], [f527, f251])).
fof(f251, plain, ! [X0] : (~ sP40(X0) | sP22(X0)), inference(cnf_transformation, [], [f54])).
fof(f401753, plain, (~ sP22(sK91) | ~ spl92_8843), inference(resolution, [], [f265815, f227])).
fof(f265815, plain, (! [X54] : (~ r1(X54, sK91) | ~ sP22(X54)) | ~ spl92_8843), inference(avatar_component_clause, [], [f265814])).
fof(f365033, plain, (~ spl92_22 | spl92_10165 | ~ spl92_7), inference(avatar_split_clause, [], [f365028, f432, f365030, f490])).
fof(f365028, plain, (p204(sK58(sK91)) | ~ p504(sK91) | ~ spl92_7), inference(subsumption_resolution, [], [f344422, f546])).
fof(f344422, plain, (p204(sK58(sK91)) | ~ p504(sK91) | ~ sP22(sK91) | ~ spl92_7), inference(resolution, [], [f433, f2161])).
fof(f2161, plain, ! [X0] : (r1(X0, sK58(X0)) | ~ p504(X0) | ~ sP22(X0)), inference(resolution, [], [f338, f227])).
fof(f338, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p504(X1) | r1(X1, sK58(X1)) | ~ sP22(X0)), inference(cnf_transformation, [], [f126])).
fof(f365020, plain, (~ spl92_7 | ~ spl92_1277 | spl92_1278), inference(avatar_contradiction_clause, [], [f365019])).
fof(f365019, plain, ($false | (~ spl92_7 | ~ spl92_1277 | spl92_1278)), inference(subsumption_resolution, [], [f364823, f15852])).
fof(f15852, plain, (~ p204(sK74(sK91)) | spl92_1278), inference(avatar_component_clause, [], [f15850])).
fof(f15850, plain, (spl92_1278 <=> p204(sK74(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1278])])).
fof(f364823, plain, (p204(sK74(sK91)) | (~ spl92_7 | ~ spl92_1277)), inference(resolution, [], [f15847, f433])).
fof(f15847, plain, (r1(sK91, sK74(sK91)) | ~ spl92_1277), inference(avatar_component_clause, [], [f15845])).
fof(f15845, plain, (spl92_1277 <=> r1(sK91, sK74(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1277])])).
fof(f350068, plain, (spl92_108 | spl92_1278 | ~ spl92_7), inference(avatar_split_clause, [], [f350067, f432, f15850, f5155])).
fof(f5155, plain, (spl92_108 <=> r1(sK91, sK73(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_108])])).
fof(f350067, plain, (p204(sK74(sK91)) | r1(sK91, sK73(sK91)) | ~ spl92_7), inference(subsumption_resolution, [], [f344527, f553])).
fof(f553, plain, sP8(sK91), inference(resolution, [], [f527, f258])).
fof(f258, plain, ! [X0] : (~ sP40(X0) | sP8(X0)), inference(cnf_transformation, [], [f54])).
fof(f344527, plain, (p204(sK74(sK91)) | r1(sK91, sK73(sK91)) | ~ sP8(sK91) | ~ spl92_7), inference(resolution, [], [f433, f2577])).
fof(f2577, plain, ! [X0] : (r1(X0, sK74(X0)) | r1(X0, sK73(X0)) | ~ sP8(X0)), inference(resolution, [], [f368, f227])).
fof(f368, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK74(X1)) | r1(X1, sK73(X1)) | ~ sP8(X0)), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ! [X0] : (! [X1] : ((~ p104(sK73(X1)) & r1(X1, sK73(X1))) | (~ p204(sK74(X1)) & r1(X1, sK74(X1))) | ~ r1(X0, X1)) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK73, sK74])], [f181, f183, f182])).
fof(f182, plain, ! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) => (~ p104(sK73(X1)) & r1(X1, sK73(X1)))), introduced(choice_axiom, [])).
fof(f183, plain, ! [X1] : (? [X3] : (~ p204(X3) & r1(X1, X3)) => (~ p204(sK74(X1)) & r1(X1, sK74(X1)))), introduced(choice_axiom, [])).
fof(f181, plain, ! [X0] : (! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) | ? [X3] : (~ p204(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP8(X0)), inference(rectify, [], [f180])).
fof(f180, plain, ! [X11] : (! [X72] : (? [X73] : (~ p104(X73) & r1(X72, X73)) | ? [X74] : (~ p204(X74) & r1(X72, X74)) | ~ r1(X11, X72)) | ~ sP8(X11)), inference(nnf_transformation, [], [f19])).
fof(f350064, plain, ~ spl92_1276, inference(avatar_contradiction_clause, [], [f350063])).
fof(f350063, plain, ($false | ~ spl92_1276), inference(subsumption_resolution, [], [f350062, f553])).
fof(f350062, plain, (~ sP8(sK91) | ~ spl92_1276), inference(resolution, [], [f15843, f227])).
fof(f15843, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP8(X0)) | ~ spl92_1276), inference(avatar_component_clause, [], [f15842])).
fof(f15842, plain, (spl92_1276 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP8(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1276])])).
fof(f343889, plain, (spl92_1362 | ~ spl92_6 | ~ spl92_184 | ~ spl92_185), inference(avatar_split_clause, [], [f343888, f6449, f6443, f429, f31385])).
fof(f31385, plain, (spl92_1362 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP2(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1362])])).
fof(f6443, plain, (spl92_184 <=> r1(sK91, sK85(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_184])])).
fof(f6449, plain, (spl92_185 <=> p305(sK86(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_185])])).
fof(f343888, plain, (! [X1] : (~ r1(X1, sK91) | ~ sP2(X1)) | (~ spl92_6 | ~ spl92_184 | ~ spl92_185)), inference(subsumption_resolution, [], [f343887, f6451])).
fof(f6451, plain, (p305(sK86(sK91)) | ~ spl92_185), inference(avatar_component_clause, [], [f6449])).
fof(f343887, plain, (! [X1] : (~ p305(sK86(sK91)) | ~ r1(X1, sK91) | ~ sP2(X1)) | (~ spl92_6 | ~ spl92_184)), inference(resolution, [], [f343411, f395])).
fof(f395, plain, ! [X0, X1] : (~ p205(sK85(X1)) | ~ p305(sK86(X1)) | ~ r1(X0, X1) | ~ sP2(X0)), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ! [X0] : (! [X1] : ((~ p205(sK85(X1)) & r1(X1, sK85(X1))) | (~ p305(sK86(X1)) & r1(X1, sK86(X1))) | ~ r1(X0, X1)) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK85, sK86])], [f211, f213, f212])).
fof(f212, plain, ! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) => (~ p205(sK85(X1)) & r1(X1, sK85(X1)))), introduced(choice_axiom, [])).
fof(f213, plain, ! [X1] : (? [X3] : (~ p305(X3) & r1(X1, X3)) => (~ p305(sK86(X1)) & r1(X1, sK86(X1)))), introduced(choice_axiom, [])).
fof(f211, plain, ! [X0] : (! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) | ? [X3] : (~ p305(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP2(X0)), inference(rectify, [], [f210])).
fof(f210, plain, ! [X11] : (! [X115] : (? [X116] : (~ p205(X116) & r1(X115, X116)) | ? [X117] : (~ p305(X117) & r1(X115, X117)) | ~ r1(X11, X115)) | ~ sP2(X11)), inference(nnf_transformation, [], [f13])).
fof(f343411, plain, (p205(sK85(sK91)) | (~ spl92_6 | ~ spl92_184)), inference(resolution, [], [f430, f6445])).
fof(f6445, plain, (r1(sK91, sK85(sK91)) | ~ spl92_184), inference(avatar_component_clause, [], [f6443])).
fof(f343401, plain, (~ spl92_7 | ~ spl92_17), inference(avatar_contradiction_clause, [], [f343400])).
fof(f343400, plain, ($false | (~ spl92_7 | ~ spl92_17)), inference(subsumption_resolution, [], [f343399, f547])).
fof(f547, plain, sP23(sK91), inference(resolution, [], [f527, f252])).
fof(f252, plain, ! [X0] : (~ sP40(X0) | sP23(X0)), inference(cnf_transformation, [], [f54])).
fof(f343399, plain, (~ sP23(sK91) | (~ spl92_7 | ~ spl92_17)), inference(resolution, [], [f343393, f227])).
fof(f343393, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP23(X0)) | (~ spl92_7 | ~ spl92_17)), inference(subsumption_resolution, [], [f343392, f471])).
fof(f343392, plain, (! [X0] : (~ p404(sK91) | ~ r1(X0, sK91) | ~ sP23(X0)) | (~ spl92_7 | ~ spl92_17)), inference(resolution, [], [f343389, f337])).
fof(f337, plain, ! [X0, X1] : (~ p204(sK57(X1)) | ~ p404(X1) | ~ r1(X0, X1) | ~ sP23(X0)), inference(cnf_transformation, [], [f122])).
fof(f122, plain, ! [X0] : (! [X1] : ((~ p204(sK57(X1)) & r1(X1, sK57(X1))) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP23(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57])], [f120, f121])).
fof(f121, plain, ! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) => (~ p204(sK57(X1)) & r1(X1, sK57(X1)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) | ~ p404(X1) | ~ r1(X0, X1)) | ~ sP23(X0)), inference(rectify, [], [f119])).
fof(f119, plain, ! [X11] : (! [X87] : (? [X88] : (~ p204(X88) & r1(X87, X88)) | ~ p404(X87) | ~ r1(X11, X87)) | ~ sP23(X11)), inference(nnf_transformation, [], [f34])).
fof(f343389, plain, (p204(sK57(sK91)) | (~ spl92_7 | ~ spl92_17)), inference(subsumption_resolution, [], [f343388, f547])).
fof(f343388, plain, (p204(sK57(sK91)) | ~ sP23(sK91) | (~ spl92_7 | ~ spl92_17)), inference(subsumption_resolution, [], [f342942, f471])).
fof(f342942, plain, (p204(sK57(sK91)) | ~ p404(sK91) | ~ sP23(sK91) | ~ spl92_7), inference(resolution, [], [f433, f2145])).
fof(f2145, plain, ! [X0] : (r1(X0, sK57(X0)) | ~ p404(X0) | ~ sP23(X0)), inference(resolution, [], [f336, f227])).
fof(f336, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p404(X1) | r1(X1, sK57(X1)) | ~ sP23(X0)), inference(cnf_transformation, [], [f122])).
fof(f343398, plain, (spl92_1276 | spl92_108 | ~ spl92_1278), inference(avatar_split_clause, [], [f343397, f15850, f5155, f15842])).
fof(f343397, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP8(X0)) | (spl92_108 | ~ spl92_1278)), inference(subsumption_resolution, [], [f343396, f5156])).
fof(f5156, plain, (~ r1(sK91, sK73(sK91)) | spl92_108), inference(avatar_component_clause, [], [f5155])).
fof(f343396, plain, (! [X0] : (r1(sK91, sK73(sK91)) | ~ r1(X0, sK91) | ~ sP8(X0)) | ~ spl92_1278), inference(resolution, [], [f15851, f369])).
fof(f369, plain, ! [X0, X1] : (~ p204(sK74(X1)) | r1(X1, sK73(X1)) | ~ r1(X0, X1) | ~ sP8(X0)), inference(cnf_transformation, [], [f184])).
fof(f15851, plain, (p204(sK74(sK91)) | ~ spl92_1278), inference(avatar_component_clause, [], [f15850])).
fof(f342603, plain, (spl92_8846 | ~ spl92_26 | ~ spl92_1174), inference(avatar_split_clause, [], [f342396, f14683, f507, f265827])).
fof(f265827, plain, (spl92_8846 <=> ! [X64] : (~ r1(X64, sK91) | ~ sP12(X64))), introduced(avatar_definition, [new_symbols(naming, [spl92_8846])])).
fof(f14683, plain, (spl92_1174 <=> p305(sK68(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1174])])).
fof(f342396, plain, (! [X0] : (~ p605(sK91) | ~ r1(X0, sK91) | ~ sP12(X0)) | ~ spl92_1174), inference(resolution, [], [f14685, f359])).
fof(f359, plain, ! [X0, X1] : (~ p305(sK68(X1)) | ~ p605(X1) | ~ r1(X0, X1) | ~ sP12(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : (! [X1] : ((~ p305(sK68(X1)) & r1(X1, sK68(X1))) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK68])], [f164, f165])).
fof(f165, plain, ! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) => (~ p305(sK68(X1)) & r1(X1, sK68(X1)))), introduced(choice_axiom, [])).
fof(f164, plain, ! [X0] : (! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP12(X0)), inference(rectify, [], [f163])).
fof(f163, plain, ! [X11] : (! [X130] : (? [X131] : (~ p305(X131) & r1(X130, X131)) | ~ p605(X130) | ~ r1(X11, X130)) | ~ sP12(X11)), inference(nnf_transformation, [], [f23])).
fof(f14685, plain, (p305(sK68(sK91)) | ~ spl92_1174), inference(avatar_component_clause, [], [f14683])).
fof(f342602, plain, ~ spl92_8846, inference(avatar_contradiction_clause, [], [f342601])).
fof(f342601, plain, ($false | ~ spl92_8846), inference(subsumption_resolution, [], [f342600, f530])).
fof(f530, plain, sP12(sK91), inference(resolution, [], [f527, f232])).
fof(f232, plain, ! [X0] : (~ sP40(X0) | sP12(X0)), inference(cnf_transformation, [], [f54])).
fof(f342600, plain, (~ sP12(sK91) | ~ spl92_8846), inference(resolution, [], [f265828, f227])).
fof(f265828, plain, (! [X64] : (~ r1(X64, sK91) | ~ sP12(X64)) | ~ spl92_8846), inference(avatar_component_clause, [], [f265827])).
fof(f342395, plain, (~ spl92_27 | ~ spl92_17), inference(avatar_split_clause, [], [f617, f469, f511])).
fof(f617, plain, (~ p404(sK91) | ~ p604(sK91)), inference(resolution, [], [f575, f527])).
fof(f575, plain, ! [X0] : (~ sP40(X0) | ~ p404(X0) | ~ p604(X0)), inference(resolution, [], [f245, f227])).
fof(f245, plain, ! [X33, X0] : (~ r1(X0, X33) | ~ p604(X33) | ~ p404(X33) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f342382, plain, (~ spl92_26 | spl92_1174 | ~ spl92_11), inference(avatar_split_clause, [], [f342381, f447, f14683, f507])).
fof(f447, plain, (spl92_11 <=> ! [X9] : (p305(X9) | ~ r1(sK91, X9))), introduced(avatar_definition, [new_symbols(naming, [spl92_11])])).
fof(f342381, plain, (p305(sK68(sK91)) | ~ p605(sK91) | ~ spl92_11), inference(subsumption_resolution, [], [f201532, f530])).
fof(f201532, plain, (p305(sK68(sK91)) | ~ p605(sK91) | ~ sP12(sK91) | ~ spl92_11), inference(resolution, [], [f448, f2513])).
fof(f2513, plain, ! [X0] : (r1(X0, sK68(X0)) | ~ p605(X0) | ~ sP12(X0)), inference(resolution, [], [f358, f227])).
fof(f358, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p605(X1) | r1(X1, sK68(X1)) | ~ sP12(X0)), inference(cnf_transformation, [], [f166])).
fof(f448, plain, (! [X9] : (~ r1(sK91, X9) | p305(X9)) | ~ spl92_11), inference(avatar_component_clause, [], [f447])).
fof(f340626, plain, (~ spl92_21 | ~ spl92_1162), inference(avatar_contradiction_clause, [], [f340625])).
fof(f340625, plain, ($false | (~ spl92_21 | ~ spl92_1162)), inference(subsumption_resolution, [], [f340624, f531])).
fof(f531, plain, sP13(sK91), inference(resolution, [], [f527, f233])).
fof(f233, plain, ! [X0] : (~ sP40(X0) | sP13(X0)), inference(cnf_transformation, [], [f54])).
fof(f340624, plain, (~ sP13(sK91) | (~ spl92_21 | ~ spl92_1162)), inference(resolution, [], [f340623, f227])).
fof(f340623, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP13(X0)) | (~ spl92_21 | ~ spl92_1162)), inference(subsumption_resolution, [], [f340622, f488])).
fof(f340622, plain, (! [X0] : (~ p505(sK91) | ~ r1(X0, sK91) | ~ sP13(X0)) | ~ spl92_1162), inference(resolution, [], [f14613, f357])).
fof(f357, plain, ! [X0, X1] : (~ p305(sK67(X1)) | ~ p505(X1) | ~ r1(X0, X1) | ~ sP13(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (! [X1] : ((~ p305(sK67(X1)) & r1(X1, sK67(X1))) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK67])], [f160, f161])).
fof(f161, plain, ! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) => (~ p305(sK67(X1)) & r1(X1, sK67(X1)))), introduced(choice_axiom, [])).
fof(f160, plain, ! [X0] : (! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP13(X0)), inference(rectify, [], [f159])).
fof(f159, plain, ! [X11] : (! [X128] : (? [X129] : (~ p305(X129) & r1(X128, X129)) | ~ p505(X128) | ~ r1(X11, X128)) | ~ sP13(X11)), inference(nnf_transformation, [], [f24])).
fof(f14613, plain, (p305(sK67(sK91)) | ~ spl92_1162), inference(avatar_component_clause, [], [f14611])).
fof(f14611, plain, (spl92_1162 <=> p305(sK67(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1162])])).
fof(f339793, plain, (spl92_1162 | ~ spl92_21 | ~ spl92_11), inference(avatar_split_clause, [], [f339792, f447, f486, f14611])).
fof(f339792, plain, (~ p505(sK91) | p305(sK67(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f320398, f531])).
fof(f320398, plain, (~ sP13(sK91) | ~ p505(sK91) | p305(sK67(sK91)) | ~ spl92_11), inference(resolution, [], [f201531, f227])).
fof(f201531, plain, (! [X26] : (~ r1(sK91, X26) | ~ sP13(X26) | ~ p505(X26) | p305(sK67(X26))) | ~ spl92_11), inference(resolution, [], [f448, f4243])).
fof(f4243, plain, ! [X0, X1] : (r1(X1, sK67(X0)) | ~ sP13(X0) | ~ p505(X0) | ~ r1(X1, X0)), inference(resolution, [], [f2497, f228])).
fof(f228, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f7])).
fof(f7, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL684+1.005.p', transitivity)).
fof(f2497, plain, ! [X0] : (r1(X0, sK67(X0)) | ~ p505(X0) | ~ sP13(X0)), inference(resolution, [], [f356, f227])).
fof(f356, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p505(X1) | r1(X1, sK67(X1)) | ~ sP13(X0)), inference(cnf_transformation, [], [f162])).
fof(f339733, plain, (spl92_2067 | ~ spl92_153 | ~ spl92_1 | ~ spl92_152), inference(avatar_split_clause, [], [f202841, f5951, f412, f5957, f202843])).
fof(f202843, plain, (spl92_2067 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP4(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_2067])])).
fof(f5957, plain, (spl92_153 <=> p305(sK82(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_153])])).
fof(f5951, plain, (spl92_152 <=> r1(sK91, sK81(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_152])])).
fof(f202841, plain, (! [X1] : (~ p305(sK82(sK91)) | ~ r1(X1, sK91) | ~ sP4(X1)) | (~ spl92_1 | ~ spl92_152)), inference(resolution, [], [f201013, f387])).
fof(f387, plain, ! [X0, X1] : (~ p105(sK81(X1)) | ~ p305(sK82(X1)) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0] : (! [X1] : ((~ p105(sK81(X1)) & r1(X1, sK81(X1))) | (~ p305(sK82(X1)) & r1(X1, sK82(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK81, sK82])], [f201, f203, f202])).
fof(f202, plain, ! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) => (~ p105(sK81(X1)) & r1(X1, sK81(X1)))), introduced(choice_axiom, [])).
fof(f203, plain, ! [X1] : (? [X3] : (~ p305(X3) & r1(X1, X3)) => (~ p305(sK82(X1)) & r1(X1, sK82(X1)))), introduced(choice_axiom, [])).
fof(f201, plain, ! [X0] : (! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) | ? [X3] : (~ p305(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(rectify, [], [f200])).
fof(f200, plain, ! [X11] : (! [X105] : (? [X106] : (~ p105(X106) & r1(X105, X106)) | ? [X107] : (~ p305(X107) & r1(X105, X107)) | ~ r1(X11, X105)) | ~ sP4(X11)), inference(nnf_transformation, [], [f15])).
fof(f201013, plain, (p105(sK81(sK91)) | (~ spl92_1 | ~ spl92_152)), inference(resolution, [], [f413, f5953])).
fof(f5953, plain, (r1(sK91, sK81(sK91)) | ~ spl92_152), inference(avatar_component_clause, [], [f5951])).
fof(f339732, plain, (spl92_153 | ~ spl92_11 | ~ spl92_2068), inference(avatar_split_clause, [], [f339213, f202846, f447, f5957])).
fof(f202846, plain, (spl92_2068 <=> r1(sK91, sK82(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_2068])])).
fof(f339213, plain, (p305(sK82(sK91)) | (~ spl92_11 | ~ spl92_2068)), inference(resolution, [], [f202848, f448])).
fof(f202848, plain, (r1(sK91, sK82(sK91)) | ~ spl92_2068), inference(avatar_component_clause, [], [f202846])).
fof(f338443, plain, ~ spl92_2067, inference(avatar_contradiction_clause, [], [f338442])).
fof(f338442, plain, ($false | ~ spl92_2067), inference(subsumption_resolution, [], [f338441, f540])).
fof(f540, plain, sP4(sK91), inference(resolution, [], [f527, f242])).
fof(f242, plain, ! [X0] : (~ sP40(X0) | sP4(X0)), inference(cnf_transformation, [], [f54])).
fof(f338441, plain, (~ sP4(sK91) | ~ spl92_2067), inference(resolution, [], [f202844, f227])).
fof(f202844, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP4(X0)) | ~ spl92_2067), inference(avatar_component_clause, [], [f202843])).
fof(f202849, plain, (spl92_2067 | spl92_2068 | ~ spl92_1 | ~ spl92_152), inference(avatar_split_clause, [], [f202840, f5951, f412, f202846, f202843])).
fof(f202840, plain, (! [X0] : (r1(sK91, sK82(sK91)) | ~ r1(X0, sK91) | ~ sP4(X0)) | (~ spl92_1 | ~ spl92_152)), inference(resolution, [], [f201013, f386])).
fof(f386, plain, ! [X0, X1] : (~ p105(sK81(X1)) | r1(X1, sK82(X1)) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f204])).
fof(f200981, plain, (~ spl92_2 | ~ spl92_27), inference(avatar_contradiction_clause, [], [f200980])).
fof(f200980, plain, ($false | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f200979, f549])).
fof(f549, plain, sP24(sK91), inference(resolution, [], [f527, f254])).
fof(f254, plain, ! [X0] : (~ sP40(X0) | sP24(X0)), inference(cnf_transformation, [], [f54])).
fof(f200979, plain, (~ sP24(sK91) | (~ spl92_2 | ~ spl92_27)), inference(resolution, [], [f199718, f227])).
fof(f199718, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP24(X0)) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f199717, f513])).
fof(f199717, plain, (! [X0] : (~ p604(sK91) | ~ r1(X0, sK91) | ~ sP24(X0)) | (~ spl92_2 | ~ spl92_27)), inference(resolution, [], [f199712, f335])).
fof(f335, plain, ! [X0, X1] : (~ p104(sK56(X1)) | ~ p604(X1) | ~ r1(X0, X1) | ~ sP24(X0)), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ! [X0] : (! [X1] : ((~ p104(sK56(X1)) & r1(X1, sK56(X1))) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP24(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK56])], [f116, f117])).
fof(f117, plain, ! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) => (~ p104(sK56(X1)) & r1(X1, sK56(X1)))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0] : (! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP24(X0)), inference(rectify, [], [f115])).
fof(f115, plain, ! [X11] : (! [X82] : (? [X83] : (~ p104(X83) & r1(X82, X83)) | ~ p604(X82) | ~ r1(X11, X82)) | ~ sP24(X11)), inference(nnf_transformation, [], [f35])).
fof(f199712, plain, (p104(sK56(sK91)) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f199711, f549])).
fof(f199711, plain, (p104(sK56(sK91)) | ~ sP24(sK91) | (~ spl92_2 | ~ spl92_27)), inference(subsumption_resolution, [], [f199301, f513])).
fof(f199301, plain, (p104(sK56(sK91)) | ~ p604(sK91) | ~ sP24(sK91) | ~ spl92_2), inference(resolution, [], [f416, f2129])).
fof(f2129, plain, ! [X0] : (r1(X0, sK56(X0)) | ~ p604(X0) | ~ sP24(X0)), inference(resolution, [], [f334, f227])).
fof(f334, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p604(X1) | r1(X1, sK56(X1)) | ~ sP24(X0)), inference(cnf_transformation, [], [f118])).
fof(f199215, plain, (~ spl92_27 | ~ spl92_22), inference(avatar_split_clause, [], [f614, f490, f511])).
fof(f614, plain, (~ p504(sK91) | ~ p604(sK91)), inference(resolution, [], [f574, f527])).
fof(f574, plain, ! [X0] : (~ sP40(X0) | ~ p504(X0) | ~ p604(X0)), inference(resolution, [], [f244, f227])).
fof(f244, plain, ! [X0, X34] : (~ r1(X0, X34) | ~ p604(X34) | ~ p504(X34) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f199214, plain, (spl92_1963 | ~ spl92_195 | ~ spl92_11 | ~ spl92_196), inference(avatar_split_clause, [], [f195534, f6747, f447, f6743, f195536])).
fof(f195536, plain, (spl92_1963 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP0(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1963])])).
fof(f6743, plain, (spl92_195 <=> p405(sK90(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_195])])).
fof(f6747, plain, (spl92_196 <=> r1(sK91, sK89(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_196])])).
fof(f195534, plain, (! [X1] : (~ p405(sK90(sK91)) | ~ r1(X1, sK91) | ~ sP0(X1)) | (~ spl92_11 | ~ spl92_196)), inference(resolution, [], [f192573, f403])).
fof(f403, plain, ! [X0, X1] : (~ p305(sK89(X1)) | ~ p405(sK90(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ! [X0] : (! [X1] : ((~ p305(sK89(X1)) & r1(X1, sK89(X1))) | (~ p405(sK90(X1)) & r1(X1, sK90(X1))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK89, sK90])], [f221, f223, f222])).
fof(f222, plain, ! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) => (~ p305(sK89(X1)) & r1(X1, sK89(X1)))), introduced(choice_axiom, [])).
fof(f223, plain, ! [X1] : (? [X3] : (~ p405(X3) & r1(X1, X3)) => (~ p405(sK90(X1)) & r1(X1, sK90(X1)))), introduced(choice_axiom, [])).
fof(f221, plain, ! [X0] : (! [X1] : (? [X2] : (~ p305(X2) & r1(X1, X2)) | ? [X3] : (~ p405(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f220])).
fof(f220, plain, ! [X11] : (! [X125] : (? [X126] : (~ p305(X126) & r1(X125, X126)) | ? [X127] : (~ p405(X127) & r1(X125, X127)) | ~ r1(X11, X125)) | ~ sP0(X11)), inference(nnf_transformation, [], [f11])).
fof(f192573, plain, (p305(sK89(sK91)) | (~ spl92_11 | ~ spl92_196)), inference(resolution, [], [f448, f6749])).
fof(f6749, plain, (r1(sK91, sK89(sK91)) | ~ spl92_196), inference(avatar_component_clause, [], [f6747])).
fof(f199213, plain, ~ spl92_1963, inference(avatar_contradiction_clause, [], [f199212])).
fof(f199212, plain, ($false | ~ spl92_1963), inference(subsumption_resolution, [], [f199211, f532])).
fof(f532, plain, sP0(sK91), inference(resolution, [], [f527, f234])).
fof(f234, plain, ! [X0] : (~ sP40(X0) | sP0(X0)), inference(cnf_transformation, [], [f54])).
fof(f199211, plain, (~ sP0(sK91) | ~ spl92_1963), inference(resolution, [], [f195537, f227])).
fof(f195537, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP0(X0)) | ~ spl92_1963), inference(avatar_component_clause, [], [f195536])).
fof(f198896, plain, (~ spl92_16 | spl92_195 | ~ spl92_1964), inference(avatar_contradiction_clause, [], [f198895])).
fof(f198895, plain, ($false | (~ spl92_16 | spl92_195 | ~ spl92_1964)), inference(subsumption_resolution, [], [f198816, f6744])).
fof(f6744, plain, (~ p405(sK90(sK91)) | spl92_195), inference(avatar_component_clause, [], [f6743])).
fof(f198816, plain, (p405(sK90(sK91)) | (~ spl92_16 | ~ spl92_1964)), inference(resolution, [], [f195541, f467])).
fof(f467, plain, (! [X10] : (~ r1(sK91, X10) | p405(X10)) | ~ spl92_16), inference(avatar_component_clause, [], [f466])).
fof(f466, plain, (spl92_16 <=> ! [X10] : (p405(X10) | ~ r1(sK91, X10))), introduced(avatar_definition, [new_symbols(naming, [spl92_16])])).
fof(f195541, plain, (r1(sK91, sK90(sK91)) | ~ spl92_1964), inference(avatar_component_clause, [], [f195539])).
fof(f195539, plain, (spl92_1964 <=> r1(sK91, sK90(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1964])])).
fof(f195542, plain, (spl92_1963 | spl92_1964 | ~ spl92_11 | ~ spl92_196), inference(avatar_split_clause, [], [f195533, f6747, f447, f195539, f195536])).
fof(f195533, plain, (! [X0] : (r1(sK91, sK90(sK91)) | ~ r1(X0, sK91) | ~ sP0(X0)) | (~ spl92_11 | ~ spl92_196)), inference(resolution, [], [f192573, f402])).
fof(f402, plain, ! [X0, X1] : (~ p305(sK89(X1)) | r1(X1, sK90(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f224])).
fof(f193016, plain, (~ spl92_11 | spl92_185 | ~ spl92_1363), inference(avatar_contradiction_clause, [], [f193015])).
fof(f193015, plain, ($false | (~ spl92_11 | spl92_185 | ~ spl92_1363)), inference(subsumption_resolution, [], [f192571, f6450])).
fof(f6450, plain, (~ p305(sK86(sK91)) | spl92_185), inference(avatar_component_clause, [], [f6449])).
fof(f192571, plain, (p305(sK86(sK91)) | (~ spl92_11 | ~ spl92_1363)), inference(resolution, [], [f448, f31390])).
fof(f31390, plain, (r1(sK91, sK86(sK91)) | ~ spl92_1363), inference(avatar_component_clause, [], [f31388])).
fof(f31388, plain, (spl92_1363 <=> r1(sK91, sK86(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1363])])).
fof(f192520, plain, (~ spl92_13 | ~ spl92_1910), inference(avatar_contradiction_clause, [], [f192519])).
fof(f192519, plain, ($false | (~ spl92_13 | ~ spl92_1910)), inference(subsumption_resolution, [], [f192518, f557])).
fof(f557, plain, sP30(sK91), inference(resolution, [], [f527, f268])).
fof(f268, plain, ! [X0] : (~ sP40(X0) | sP30(X0)), inference(cnf_transformation, [], [f54])).
fof(f192518, plain, (~ sP30(sK91) | (~ spl92_13 | ~ spl92_1910)), inference(resolution, [], [f184081, f227])).
fof(f184081, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP30(X0)) | (~ spl92_13 | ~ spl92_1910)), inference(subsumption_resolution, [], [f184080, f455])).
fof(f184080, plain, (! [X0] : (~ p303(sK91) | ~ r1(X0, sK91) | ~ sP30(X0)) | ~ spl92_1910), inference(resolution, [], [f183676, f323])).
fof(f323, plain, ! [X0, X1] : (~ p203(sK50(X1)) | ~ p303(X1) | ~ r1(X0, X1) | ~ sP30(X0)), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ! [X0] : (! [X1] : ((~ p203(sK50(X1)) & r1(X1, sK50(X1))) | ~ p303(X1) | ~ r1(X0, X1)) | ~ sP30(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50])], [f92, f93])).
fof(f93, plain, ! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) => (~ p203(sK50(X1)) & r1(X1, sK50(X1)))), introduced(choice_axiom, [])).
fof(f92, plain, ! [X0] : (! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) | ~ p303(X1) | ~ r1(X0, X1)) | ~ sP30(X0)), inference(rectify, [], [f91])).
fof(f91, plain, ! [X11] : (! [X58] : (? [X59] : (~ p203(X59) & r1(X58, X59)) | ~ p303(X58) | ~ r1(X11, X58)) | ~ sP30(X11)), inference(nnf_transformation, [], [f41])).
fof(f183676, plain, (p203(sK50(sK91)) | ~ spl92_1910), inference(avatar_component_clause, [], [f183674])).
fof(f183674, plain, (spl92_1910 <=> p203(sK50(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1910])])).
fof(f183677, plain, (~ spl92_13 | spl92_1910 | ~ spl92_8), inference(avatar_split_clause, [], [f183672, f435, f183674, f453])).
fof(f435, plain, (spl92_8 <=> ! [X5] : (p203(X5) | ~ r1(sK91, X5))), introduced(avatar_definition, [new_symbols(naming, [spl92_8])])).
fof(f183672, plain, (p203(sK50(sK91)) | ~ p303(sK91) | ~ spl92_8), inference(subsumption_resolution, [], [f175527, f557])).
fof(f175527, plain, (p203(sK50(sK91)) | ~ p303(sK91) | ~ sP30(sK91) | ~ spl92_8), inference(resolution, [], [f436, f1507])).
fof(f1507, plain, ! [X0] : (r1(X0, sK50(X0)) | ~ p303(X0) | ~ sP30(X0)), inference(resolution, [], [f322, f227])).
fof(f322, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p303(X1) | r1(X1, sK50(X1)) | ~ sP30(X0)), inference(cnf_transformation, [], [f94])).
fof(f436, plain, (! [X5] : (~ r1(sK91, X5) | p203(X5)) | ~ spl92_8), inference(avatar_component_clause, [], [f435])).
fof(f183659, plain, (~ spl92_14 | ~ spl92_1878), inference(avatar_contradiction_clause, [], [f183658])).
fof(f183658, plain, ($false | (~ spl92_14 | ~ spl92_1878)), inference(subsumption_resolution, [], [f183657, f566])).
fof(f566, plain, sP38(sK91), inference(resolution, [], [f527, f287])).
fof(f287, plain, ! [X0] : (~ sP40(X0) | sP38(X0)), inference(cnf_transformation, [], [f54])).
fof(f183657, plain, (~ sP38(sK91) | (~ spl92_14 | ~ spl92_1878)), inference(resolution, [], [f179400, f227])).
fof(f179400, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP38(X0)) | (~ spl92_14 | ~ spl92_1878)), inference(subsumption_resolution, [], [f179399, f459])).
fof(f459, plain, (p302(sK91) | ~ spl92_14), inference(avatar_component_clause, [], [f457])).
fof(f179399, plain, (! [X0] : (~ p302(sK91) | ~ r1(X0, sK91) | ~ sP38(X0)) | ~ spl92_1878), inference(resolution, [], [f179155, f307])).
fof(f307, plain, ! [X0, X1] : (~ p102(sK42(X1)) | ~ p302(X1) | ~ r1(X0, X1) | ~ sP38(X0)), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : (! [X1] : ((~ p102(sK42(X1)) & r1(X1, sK42(X1))) | ~ p302(X1) | ~ r1(X0, X1)) | ~ sP38(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f60, f61])).
fof(f61, plain, ! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) => (~ p102(sK42(X1)) & r1(X1, sK42(X1)))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X0] : (! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) | ~ p302(X1) | ~ r1(X0, X1)) | ~ sP38(X0)), inference(rectify, [], [f59])).
fof(f59, plain, ! [X11] : (! [X29] : (? [X30] : (~ p102(X30) & r1(X29, X30)) | ~ p302(X29) | ~ r1(X11, X29)) | ~ sP38(X11)), inference(nnf_transformation, [], [f49])).
fof(f179155, plain, (p102(sK42(sK91)) | ~ spl92_1878), inference(avatar_component_clause, [], [f179153])).
fof(f179153, plain, (spl92_1878 <=> p102(sK42(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1878])])).
fof(f179156, plain, (~ spl92_14 | spl92_1878 | ~ spl92_4), inference(avatar_split_clause, [], [f179151, f421, f179153, f457])).
fof(f179151, plain, (p102(sK42(sK91)) | ~ p302(sK91) | ~ spl92_4), inference(subsumption_resolution, [], [f177757, f566])).
fof(f177757, plain, (p102(sK42(sK91)) | ~ p302(sK91) | ~ sP38(sK91) | ~ spl92_4), inference(resolution, [], [f422, f665])).
fof(f665, plain, ! [X0] : (r1(X0, sK42(X0)) | ~ p302(X0) | ~ sP38(X0)), inference(resolution, [], [f306, f227])).
fof(f306, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p302(X1) | r1(X1, sK42(X1)) | ~ sP38(X0)), inference(cnf_transformation, [], [f62])).
fof(f179150, plain, (~ spl92_23 | ~ spl92_1875), inference(avatar_contradiction_clause, [], [f179149])).
fof(f179149, plain, ($false | (~ spl92_23 | ~ spl92_1875)), inference(subsumption_resolution, [], [f179148, f555])).
fof(f555, plain, sP28(sK91), inference(resolution, [], [f527, f266])).
fof(f266, plain, ! [X0] : (~ sP40(X0) | sP28(X0)), inference(cnf_transformation, [], [f54])).
fof(f179148, plain, (~ sP28(sK91) | (~ spl92_23 | ~ spl92_1875)), inference(resolution, [], [f179071, f227])).
fof(f179071, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP28(X0)) | (~ spl92_23 | ~ spl92_1875)), inference(subsumption_resolution, [], [f179070, f496])).
fof(f179070, plain, (! [X0] : (~ p503(sK91) | ~ r1(X0, sK91) | ~ sP28(X0)) | ~ spl92_1875), inference(resolution, [], [f179056, f327])).
fof(f327, plain, ! [X0, X1] : (~ p203(sK52(X1)) | ~ p503(X1) | ~ r1(X0, X1) | ~ sP28(X0)), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ! [X0] : (! [X1] : ((~ p203(sK52(X1)) & r1(X1, sK52(X1))) | ~ p503(X1) | ~ r1(X0, X1)) | ~ sP28(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52])], [f100, f101])).
fof(f101, plain, ! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) => (~ p203(sK52(X1)) & r1(X1, sK52(X1)))), introduced(choice_axiom, [])).
fof(f100, plain, ! [X0] : (! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) | ~ p503(X1) | ~ r1(X0, X1)) | ~ sP28(X0)), inference(rectify, [], [f99])).
fof(f99, plain, ! [X11] : (! [X62] : (? [X63] : (~ p203(X63) & r1(X62, X63)) | ~ p503(X62) | ~ r1(X11, X62)) | ~ sP28(X11)), inference(nnf_transformation, [], [f39])).
fof(f179056, plain, (p203(sK52(sK91)) | ~ spl92_1875), inference(avatar_component_clause, [], [f179054])).
fof(f179054, plain, (spl92_1875 <=> p203(sK52(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1875])])).
fof(f179057, plain, (~ spl92_23 | spl92_1875 | ~ spl92_8), inference(avatar_split_clause, [], [f179052, f435, f179054, f494])).
fof(f179052, plain, (p203(sK52(sK91)) | ~ p503(sK91) | ~ spl92_8), inference(subsumption_resolution, [], [f175531, f555])).
fof(f175531, plain, (p203(sK52(sK91)) | ~ p503(sK91) | ~ sP28(sK91) | ~ spl92_8), inference(resolution, [], [f436, f1703])).
fof(f1703, plain, ! [X0] : (r1(X0, sK52(X0)) | ~ p503(X0) | ~ sP28(X0)), inference(resolution, [], [f326, f227])).
fof(f326, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p503(X1) | r1(X1, sK52(X1)) | ~ sP28(X0)), inference(cnf_transformation, [], [f102])).
fof(f179039, plain, (~ spl92_4 | ~ spl92_29), inference(avatar_contradiction_clause, [], [f179038])).
fof(f179038, plain, ($false | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f179037, f563])).
fof(f563, plain, sP35(sK91), inference(resolution, [], [f527, f284])).
fof(f284, plain, ! [X0] : (~ sP40(X0) | sP35(X0)), inference(cnf_transformation, [], [f54])).
fof(f179037, plain, (~ sP35(sK91) | (~ spl92_4 | ~ spl92_29)), inference(resolution, [], [f178200, f227])).
fof(f178200, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP35(X0)) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f178199, f521])).
fof(f521, plain, (p602(sK91) | ~ spl92_29), inference(avatar_component_clause, [], [f519])).
fof(f519, plain, (spl92_29 <=> p602(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_29])])).
fof(f178199, plain, (! [X0] : (~ p602(sK91) | ~ r1(X0, sK91) | ~ sP35(X0)) | (~ spl92_4 | ~ spl92_29)), inference(resolution, [], [f178196, f313])).
fof(f313, plain, ! [X0, X1] : (~ p102(sK45(X1)) | ~ p602(X1) | ~ r1(X0, X1) | ~ sP35(X0)), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ! [X0] : (! [X1] : ((~ p102(sK45(X1)) & r1(X1, sK45(X1))) | ~ p602(X1) | ~ r1(X0, X1)) | ~ sP35(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45])], [f72, f73])).
fof(f73, plain, ! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) => (~ p102(sK45(X1)) & r1(X1, sK45(X1)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X0] : (! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) | ~ p602(X1) | ~ r1(X0, X1)) | ~ sP35(X0)), inference(rectify, [], [f71])).
fof(f71, plain, ! [X11] : (! [X35] : (? [X36] : (~ p102(X36) & r1(X35, X36)) | ~ p602(X35) | ~ r1(X11, X35)) | ~ sP35(X11)), inference(nnf_transformation, [], [f46])).
fof(f178196, plain, (p102(sK45(sK91)) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f178195, f563])).
fof(f178195, plain, (p102(sK45(sK91)) | ~ sP35(sK91) | (~ spl92_4 | ~ spl92_29)), inference(subsumption_resolution, [], [f177763, f521])).
fof(f177763, plain, (p102(sK45(sK91)) | ~ p602(sK91) | ~ sP35(sK91) | ~ spl92_4), inference(resolution, [], [f422, f911])).
fof(f911, plain, ! [X0] : (r1(X0, sK45(X0)) | ~ p602(X0) | ~ sP35(X0)), inference(resolution, [], [f312, f227])).
fof(f312, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p602(X1) | r1(X1, sK45(X1)) | ~ sP35(X0)), inference(cnf_transformation, [], [f74])).
fof(f177731, plain, ~ spl92_1839, inference(avatar_contradiction_clause, [], [f177730])).
fof(f177730, plain, ($false | ~ spl92_1839), inference(subsumption_resolution, [], [f177729, f562])).
fof(f562, plain, sP9(sK91), inference(resolution, [], [f527, f273])).
fof(f273, plain, ! [X0] : (~ sP40(X0) | sP9(X0)), inference(cnf_transformation, [], [f54])).
fof(f177729, plain, (~ sP9(sK91) | ~ spl92_1839), inference(resolution, [], [f175957, f227])).
fof(f175957, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP9(X0)) | ~ spl92_1839), inference(avatar_component_clause, [], [f175956])).
fof(f175956, plain, (spl92_1839 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP9(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1839])])).
fof(f177414, plain, (~ spl92_8 | spl92_1475 | ~ spl92_1840), inference(avatar_contradiction_clause, [], [f177413])).
fof(f177413, plain, ($false | (~ spl92_8 | spl92_1475 | ~ spl92_1840)), inference(subsumption_resolution, [], [f177332, f83069])).
fof(f83069, plain, (~ p203(sK72(sK91)) | spl92_1475), inference(avatar_component_clause, [], [f83068])).
fof(f83068, plain, (spl92_1475 <=> p203(sK72(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1475])])).
fof(f177332, plain, (p203(sK72(sK91)) | (~ spl92_8 | ~ spl92_1840)), inference(resolution, [], [f175961, f436])).
fof(f175961, plain, (r1(sK91, sK72(sK91)) | ~ spl92_1840), inference(avatar_component_clause, [], [f175959])).
fof(f175959, plain, (spl92_1840 <=> r1(sK91, sK72(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1840])])).
fof(f175963, plain, (spl92_1839 | ~ spl92_1475 | ~ spl92_3 | ~ spl92_102), inference(avatar_split_clause, [], [f175954, f5004, f418, f83068, f175956])).
fof(f5004, plain, (spl92_102 <=> r1(sK91, sK71(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_102])])).
fof(f175954, plain, (! [X1] : (~ p203(sK72(sK91)) | ~ r1(X1, sK91) | ~ sP9(X1)) | (~ spl92_3 | ~ spl92_102)), inference(resolution, [], [f175045, f367])).
fof(f367, plain, ! [X0, X1] : (~ p103(sK71(X1)) | ~ p203(sK72(X1)) | ~ r1(X0, X1) | ~ sP9(X0)), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ! [X0] : (! [X1] : ((~ p103(sK71(X1)) & r1(X1, sK71(X1))) | (~ p203(sK72(X1)) & r1(X1, sK72(X1))) | ~ r1(X0, X1)) | ~ sP9(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK71, sK72])], [f176, f178, f177])).
fof(f177, plain, ! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) => (~ p103(sK71(X1)) & r1(X1, sK71(X1)))), introduced(choice_axiom, [])).
fof(f178, plain, ! [X1] : (? [X3] : (~ p203(X3) & r1(X1, X3)) => (~ p203(sK72(X1)) & r1(X1, sK72(X1)))), introduced(choice_axiom, [])).
fof(f176, plain, ! [X0] : (! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) | ? [X3] : (~ p203(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP9(X0)), inference(rectify, [], [f175])).
fof(f175, plain, ! [X11] : (! [X47] : (? [X48] : (~ p103(X48) & r1(X47, X48)) | ? [X49] : (~ p203(X49) & r1(X47, X49)) | ~ r1(X11, X47)) | ~ sP9(X11)), inference(nnf_transformation, [], [f20])).
fof(f175045, plain, (p103(sK71(sK91)) | (~ spl92_3 | ~ spl92_102)), inference(resolution, [], [f419, f5006])).
fof(f5006, plain, (r1(sK91, sK71(sK91)) | ~ spl92_102), inference(avatar_component_clause, [], [f5004])).
fof(f175962, plain, (spl92_1839 | spl92_1840 | ~ spl92_3 | ~ spl92_102), inference(avatar_split_clause, [], [f175953, f5004, f418, f175959, f175956])).
fof(f175953, plain, (! [X0] : (r1(sK91, sK72(sK91)) | ~ r1(X0, sK91) | ~ sP9(X0)) | (~ spl92_3 | ~ spl92_102)), inference(resolution, [], [f175045, f366])).
fof(f366, plain, ! [X0, X1] : (~ p103(sK71(X1)) | r1(X1, sK72(X1)) | ~ r1(X0, X1) | ~ sP9(X0)), inference(cnf_transformation, [], [f179])).
fof(f175044, plain, (~ spl92_4 | ~ spl92_24), inference(avatar_contradiction_clause, [], [f175043])).
fof(f175043, plain, ($false | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f175042, f564])).
fof(f564, plain, sP36(sK91), inference(resolution, [], [f527, f285])).
fof(f285, plain, ! [X0] : (~ sP40(X0) | sP36(X0)), inference(cnf_transformation, [], [f54])).
fof(f175042, plain, (~ sP36(sK91) | (~ spl92_4 | ~ spl92_24)), inference(resolution, [], [f173899, f227])).
fof(f173899, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP36(X0)) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f173898, f500])).
fof(f500, plain, (p502(sK91) | ~ spl92_24), inference(avatar_component_clause, [], [f498])).
fof(f173898, plain, (! [X0] : (~ p502(sK91) | ~ r1(X0, sK91) | ~ sP36(X0)) | (~ spl92_4 | ~ spl92_24)), inference(resolution, [], [f173891, f311])).
fof(f311, plain, ! [X0, X1] : (~ p102(sK44(X1)) | ~ p502(X1) | ~ r1(X0, X1) | ~ sP36(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (! [X1] : ((~ p102(sK44(X1)) & r1(X1, sK44(X1))) | ~ p502(X1) | ~ r1(X0, X1)) | ~ sP36(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44])], [f68, f69])).
fof(f69, plain, ! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) => (~ p102(sK44(X1)) & r1(X1, sK44(X1)))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X0] : (! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) | ~ p502(X1) | ~ r1(X0, X1)) | ~ sP36(X0)), inference(rectify, [], [f67])).
fof(f67, plain, ! [X11] : (! [X33] : (? [X34] : (~ p102(X34) & r1(X33, X34)) | ~ p502(X33) | ~ r1(X11, X33)) | ~ sP36(X11)), inference(nnf_transformation, [], [f47])).
fof(f173891, plain, (p102(sK44(sK91)) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f173890, f564])).
fof(f173890, plain, (p102(sK44(sK91)) | ~ sP36(sK91) | (~ spl92_4 | ~ spl92_24)), inference(subsumption_resolution, [], [f173456, f500])).
fof(f173456, plain, (p102(sK44(sK91)) | ~ p502(sK91) | ~ sP36(sK91) | ~ spl92_4), inference(resolution, [], [f422, f827])).
fof(f827, plain, ! [X0] : (r1(X0, sK44(X0)) | ~ p502(X0) | ~ sP36(X0)), inference(resolution, [], [f310, f227])).
fof(f310, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p502(X1) | r1(X1, sK44(X1)) | ~ sP36(X0)), inference(cnf_transformation, [], [f70])).
fof(f173438, plain, (~ spl92_10 | ~ spl92_5), inference(avatar_split_clause, [], [f663, f424, f442])).
fof(f663, plain, (~ p101(sK91) | ~ p201(sK91)), inference(resolution, [], [f607, f527])).
fof(f607, plain, ! [X0] : (~ sP40(X0) | ~ p101(X0) | ~ p201(X0)), inference(resolution, [], [f303, f227])).
fof(f303, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p201(X1) | ~ p101(X1) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f173437, plain, (spl92_1807 | ~ spl92_189 | ~ spl92_6 | ~ spl92_190), inference(avatar_split_clause, [], [f172566, f6594, f429, f6590, f172568])).
fof(f172568, plain, (spl92_1807 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP1(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1807])])).
fof(f6590, plain, (spl92_189 <=> p405(sK88(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_189])])).
fof(f6594, plain, (spl92_190 <=> r1(sK91, sK87(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_190])])).
fof(f172566, plain, (! [X1] : (~ p405(sK88(sK91)) | ~ r1(X1, sK91) | ~ sP1(X1)) | (~ spl92_6 | ~ spl92_190)), inference(resolution, [], [f171208, f399])).
fof(f399, plain, ! [X0, X1] : (~ p205(sK87(X1)) | ~ p405(sK88(X1)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f219])).
fof(f219, plain, ! [X0] : (! [X1] : ((~ p205(sK87(X1)) & r1(X1, sK87(X1))) | (~ p405(sK88(X1)) & r1(X1, sK88(X1))) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK87, sK88])], [f216, f218, f217])).
fof(f217, plain, ! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) => (~ p205(sK87(X1)) & r1(X1, sK87(X1)))), introduced(choice_axiom, [])).
fof(f218, plain, ! [X1] : (? [X3] : (~ p405(X3) & r1(X1, X3)) => (~ p405(sK88(X1)) & r1(X1, sK88(X1)))), introduced(choice_axiom, [])).
fof(f216, plain, ! [X0] : (! [X1] : (? [X2] : (~ p205(X2) & r1(X1, X2)) | ? [X3] : (~ p405(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f215])).
fof(f215, plain, ! [X11] : (! [X118] : (? [X119] : (~ p205(X119) & r1(X118, X119)) | ? [X120] : (~ p405(X120) & r1(X118, X120)) | ~ r1(X11, X118)) | ~ sP1(X11)), inference(nnf_transformation, [], [f12])).
fof(f171208, plain, (p205(sK87(sK91)) | (~ spl92_6 | ~ spl92_190)), inference(resolution, [], [f430, f6596])).
fof(f6596, plain, (r1(sK91, sK87(sK91)) | ~ spl92_190), inference(avatar_component_clause, [], [f6594])).
fof(f173436, plain, ~ spl92_1807, inference(avatar_contradiction_clause, [], [f173435])).
fof(f173435, plain, ($false | ~ spl92_1807), inference(subsumption_resolution, [], [f173434, f535])).
fof(f535, plain, sP1(sK91), inference(resolution, [], [f527, f237])).
fof(f237, plain, ! [X0] : (~ sP40(X0) | sP1(X0)), inference(cnf_transformation, [], [f54])).
fof(f173434, plain, (~ sP1(sK91) | ~ spl92_1807), inference(resolution, [], [f172569, f227])).
fof(f172569, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP1(X0)) | ~ spl92_1807), inference(avatar_component_clause, [], [f172568])).
fof(f173118, plain, (~ spl92_16 | spl92_189 | ~ spl92_1808), inference(avatar_contradiction_clause, [], [f173117])).
fof(f173117, plain, ($false | (~ spl92_16 | spl92_189 | ~ spl92_1808)), inference(subsumption_resolution, [], [f173038, f6591])).
fof(f6591, plain, (~ p405(sK88(sK91)) | spl92_189), inference(avatar_component_clause, [], [f6590])).
fof(f173038, plain, (p405(sK88(sK91)) | (~ spl92_16 | ~ spl92_1808)), inference(resolution, [], [f172573, f467])).
fof(f172573, plain, (r1(sK91, sK88(sK91)) | ~ spl92_1808), inference(avatar_component_clause, [], [f172571])).
fof(f172571, plain, (spl92_1808 <=> r1(sK91, sK88(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1808])])).
fof(f172574, plain, (spl92_1807 | spl92_1808 | ~ spl92_6 | ~ spl92_190), inference(avatar_split_clause, [], [f172565, f6594, f429, f172571, f172568])).
fof(f172565, plain, (! [X0] : (r1(sK91, sK88(sK91)) | ~ r1(X0, sK91) | ~ sP1(X0)) | (~ spl92_6 | ~ spl92_190)), inference(resolution, [], [f171208, f398])).
fof(f398, plain, ! [X0, X1] : (~ p205(sK87(X1)) | r1(X1, sK88(X1)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f219])).
fof(f171200, plain, (~ spl92_8 | ~ spl92_28), inference(avatar_contradiction_clause, [], [f171199])).
fof(f171199, plain, ($false | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f171198, f554])).
fof(f554, plain, sP27(sK91), inference(resolution, [], [f527, f265])).
fof(f265, plain, ! [X0] : (~ sP40(X0) | sP27(X0)), inference(cnf_transformation, [], [f54])).
fof(f171198, plain, (~ sP27(sK91) | (~ spl92_8 | ~ spl92_28)), inference(resolution, [], [f166710, f227])).
fof(f166710, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP27(X0)) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f166709, f517])).
fof(f517, plain, (p603(sK91) | ~ spl92_28), inference(avatar_component_clause, [], [f515])).
fof(f515, plain, (spl92_28 <=> p603(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_28])])).
fof(f166709, plain, (! [X0] : (~ p603(sK91) | ~ r1(X0, sK91) | ~ sP27(X0)) | (~ spl92_8 | ~ spl92_28)), inference(resolution, [], [f166706, f329])).
fof(f329, plain, ! [X0, X1] : (~ p203(sK53(X1)) | ~ p603(X1) | ~ r1(X0, X1) | ~ sP27(X0)), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ! [X0] : (! [X1] : ((~ p203(sK53(X1)) & r1(X1, sK53(X1))) | ~ p603(X1) | ~ r1(X0, X1)) | ~ sP27(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK53])], [f104, f105])).
fof(f105, plain, ! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) => (~ p203(sK53(X1)) & r1(X1, sK53(X1)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X0] : (! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) | ~ p603(X1) | ~ r1(X0, X1)) | ~ sP27(X0)), inference(rectify, [], [f103])).
fof(f103, plain, ! [X11] : (! [X64] : (? [X65] : (~ p203(X65) & r1(X64, X65)) | ~ p603(X64) | ~ r1(X11, X64)) | ~ sP27(X11)), inference(nnf_transformation, [], [f38])).
fof(f166706, plain, (p203(sK53(sK91)) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f166705, f554])).
fof(f166705, plain, (p203(sK53(sK91)) | ~ sP27(sK91) | (~ spl92_8 | ~ spl92_28)), inference(subsumption_resolution, [], [f166285, f517])).
fof(f166285, plain, (p203(sK53(sK91)) | ~ p603(sK91) | ~ sP27(sK91) | ~ spl92_8), inference(resolution, [], [f436, f1804])).
fof(f1804, plain, ! [X0] : (r1(X0, sK53(X0)) | ~ p603(X0) | ~ sP27(X0)), inference(resolution, [], [f328, f227])).
fof(f328, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p603(X1) | r1(X1, sK53(X1)) | ~ sP27(X0)), inference(cnf_transformation, [], [f106])).
fof(f166247, plain, ~ spl92_1425, inference(avatar_contradiction_clause, [], [f166246])).
fof(f166246, plain, ($false | ~ spl92_1425), inference(subsumption_resolution, [], [f166245, f548])).
fof(f548, plain, sP6(sK91), inference(resolution, [], [f527, f253])).
fof(f253, plain, ! [X0] : (~ sP40(X0) | sP6(X0)), inference(cnf_transformation, [], [f54])).
fof(f166245, plain, (~ sP6(sK91) | ~ spl92_1425), inference(resolution, [], [f46403, f227])).
fof(f46403, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP6(X0)) | ~ spl92_1425), inference(avatar_component_clause, [], [f46402])).
fof(f46402, plain, (spl92_1425 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP6(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1425])])).
fof(f165098, plain, (~ spl92_24 | ~ spl92_9), inference(avatar_split_clause, [], [f641, f438, f498])).
fof(f641, plain, (~ p202(sK91) | ~ p502(sK91)), inference(resolution, [], [f590, f527])).
fof(f590, plain, ! [X0] : (~ sP40(X0) | ~ p202(X0) | ~ p502(X0)), inference(resolution, [], [f281, f227])).
fof(f281, plain, ! [X0, X18] : (~ r1(X0, X18) | ~ p502(X18) | ~ p202(X18) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f165097, plain, (~ spl92_3 | ~ spl92_28), inference(avatar_contradiction_clause, [], [f165096])).
fof(f165096, plain, ($false | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f165095, f558])).
fof(f558, plain, sP31(sK91), inference(resolution, [], [f527, f269])).
fof(f269, plain, ! [X0] : (~ sP40(X0) | sP31(X0)), inference(cnf_transformation, [], [f54])).
fof(f165095, plain, (~ sP31(sK91) | (~ spl92_3 | ~ spl92_28)), inference(resolution, [], [f160230, f227])).
fof(f160230, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP31(X0)) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f160229, f517])).
fof(f160229, plain, (! [X0] : (~ p603(sK91) | ~ r1(X0, sK91) | ~ sP31(X0)) | (~ spl92_3 | ~ spl92_28)), inference(resolution, [], [f160032, f321])).
fof(f321, plain, ! [X0, X1] : (~ p103(sK49(X1)) | ~ p603(X1) | ~ r1(X0, X1) | ~ sP31(X0)), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : (! [X1] : ((~ p103(sK49(X1)) & r1(X1, sK49(X1))) | ~ p603(X1) | ~ r1(X0, X1)) | ~ sP31(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK49])], [f88, f89])).
fof(f89, plain, ! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) => (~ p103(sK49(X1)) & r1(X1, sK49(X1)))), introduced(choice_axiom, [])).
fof(f88, plain, ! [X0] : (! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) | ~ p603(X1) | ~ r1(X0, X1)) | ~ sP31(X0)), inference(rectify, [], [f87])).
fof(f87, plain, ! [X11] : (! [X56] : (? [X57] : (~ p103(X57) & r1(X56, X57)) | ~ p603(X56) | ~ r1(X11, X56)) | ~ sP31(X11)), inference(nnf_transformation, [], [f42])).
fof(f160032, plain, (p103(sK49(sK91)) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f160031, f558])).
fof(f160031, plain, (p103(sK49(sK91)) | ~ sP31(sK91) | (~ spl92_3 | ~ spl92_28)), inference(subsumption_resolution, [], [f159605, f517])).
fof(f159605, plain, (p103(sK49(sK91)) | ~ p603(sK91) | ~ sP31(sK91) | ~ spl92_3), inference(resolution, [], [f419, f1382])).
fof(f1382, plain, ! [X0] : (r1(X0, sK49(X0)) | ~ p603(X0) | ~ sP31(X0)), inference(resolution, [], [f320, f227])).
fof(f320, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p603(X1) | r1(X1, sK49(X1)) | ~ sP31(X0)), inference(cnf_transformation, [], [f90])).
fof(f159554, plain, (~ spl92_26 | ~ spl92_1774), inference(avatar_contradiction_clause, [], [f159553])).
fof(f159553, plain, ($false | (~ spl92_26 | ~ spl92_1774)), inference(subsumption_resolution, [], [f159552, f528])).
fof(f528, plain, sP10(sK91), inference(resolution, [], [f527, f230])).
fof(f230, plain, ! [X0] : (~ sP40(X0) | sP10(X0)), inference(cnf_transformation, [], [f54])).
fof(f159552, plain, (~ sP10(sK91) | (~ spl92_26 | ~ spl92_1774)), inference(resolution, [], [f147315, f227])).
fof(f147315, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP10(X0)) | (~ spl92_26 | ~ spl92_1774)), inference(subsumption_resolution, [], [f147314, f509])).
fof(f147314, plain, (! [X0] : (~ p605(sK91) | ~ r1(X0, sK91) | ~ sP10(X0)) | ~ spl92_1774), inference(resolution, [], [f146966, f363])).
fof(f363, plain, ! [X0, X1] : (~ p405(sK70(X1)) | ~ p605(X1) | ~ r1(X0, X1) | ~ sP10(X0)), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ! [X0] : (! [X1] : ((~ p405(sK70(X1)) & r1(X1, sK70(X1))) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP10(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK70])], [f172, f173])).
fof(f173, plain, ! [X1] : (? [X2] : (~ p405(X2) & r1(X1, X2)) => (~ p405(sK70(X1)) & r1(X1, sK70(X1)))), introduced(choice_axiom, [])).
fof(f172, plain, ! [X0] : (! [X1] : (? [X2] : (~ p405(X2) & r1(X1, X2)) | ~ p605(X1) | ~ r1(X0, X1)) | ~ sP10(X0)), inference(rectify, [], [f171])).
fof(f171, plain, ! [X11] : (! [X134] : (? [X135] : (~ p405(X135) & r1(X134, X135)) | ~ p605(X134) | ~ r1(X11, X134)) | ~ sP10(X11)), inference(nnf_transformation, [], [f21])).
fof(f146966, plain, (p405(sK70(sK91)) | ~ spl92_1774), inference(avatar_component_clause, [], [f146964])).
fof(f146964, plain, (spl92_1774 <=> p405(sK70(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1774])])).
fof(f146967, plain, (~ spl92_26 | spl92_1774 | ~ spl92_16), inference(avatar_split_clause, [], [f146962, f466, f146964, f507])).
fof(f146962, plain, (p405(sK70(sK91)) | ~ p605(sK91) | ~ spl92_16), inference(subsumption_resolution, [], [f54861, f528])).
fof(f54861, plain, (p405(sK70(sK91)) | ~ p605(sK91) | ~ sP10(sK91) | ~ spl92_16), inference(resolution, [], [f467, f2545])).
fof(f2545, plain, ! [X0] : (r1(X0, sK70(X0)) | ~ p605(X0) | ~ sP10(X0)), inference(resolution, [], [f362, f227])).
fof(f362, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p605(X1) | r1(X1, sK70(X1)) | ~ sP10(X0)), inference(cnf_transformation, [], [f174])).
fof(f146961, plain, (~ spl92_27 | ~ spl92_1711), inference(avatar_contradiction_clause, [], [f146960])).
fof(f146960, plain, ($false | (~ spl92_27 | ~ spl92_1711)), inference(subsumption_resolution, [], [f146959, f542])).
fof(f542, plain, sP18(sK91), inference(resolution, [], [f527, f247])).
fof(f247, plain, ! [X0] : (~ sP40(X0) | sP18(X0)), inference(cnf_transformation, [], [f54])).
fof(f146959, plain, (~ sP18(sK91) | (~ spl92_27 | ~ spl92_1711)), inference(resolution, [], [f134264, f227])).
fof(f134264, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP18(X0)) | (~ spl92_27 | ~ spl92_1711)), inference(subsumption_resolution, [], [f134263, f513])).
fof(f134263, plain, (! [X0] : (~ p604(sK91) | ~ r1(X0, sK91) | ~ sP18(X0)) | ~ spl92_1711), inference(resolution, [], [f133679, f347])).
fof(f347, plain, ! [X0, X1] : (~ p304(sK62(X1)) | ~ p604(X1) | ~ r1(X0, X1) | ~ sP18(X0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : (! [X1] : ((~ p304(sK62(X1)) & r1(X1, sK62(X1))) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP18(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62])], [f140, f141])).
fof(f141, plain, ! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) => (~ p304(sK62(X1)) & r1(X1, sK62(X1)))), introduced(choice_axiom, [])).
fof(f140, plain, ! [X0] : (! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) | ~ p604(X1) | ~ r1(X0, X1)) | ~ sP18(X0)), inference(rectify, [], [f139])).
fof(f139, plain, ! [X11] : (! [X97] : (? [X98] : (~ p304(X98) & r1(X97, X98)) | ~ p604(X97) | ~ r1(X11, X97)) | ~ sP18(X11)), inference(nnf_transformation, [], [f29])).
fof(f133679, plain, (p304(sK62(sK91)) | ~ spl92_1711), inference(avatar_component_clause, [], [f133677])).
fof(f133677, plain, (spl92_1711 <=> p304(sK62(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1711])])).
fof(f133680, plain, (~ spl92_27 | spl92_1711 | ~ spl92_12), inference(avatar_split_clause, [], [f133675, f450, f133677, f511])).
fof(f133675, plain, (p304(sK62(sK91)) | ~ p604(sK91) | ~ spl92_12), inference(subsumption_resolution, [], [f16937, f542])).
fof(f16937, plain, (p304(sK62(sK91)) | ~ p604(sK91) | ~ sP18(sK91) | ~ spl92_12), inference(resolution, [], [f451, f2417])).
fof(f2417, plain, ! [X0] : (r1(X0, sK62(X0)) | ~ p604(X0) | ~ sP18(X0)), inference(resolution, [], [f346, f227])).
fof(f346, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p604(X1) | r1(X1, sK62(X1)) | ~ sP18(X0)), inference(cnf_transformation, [], [f142])).
fof(f133656, plain, (spl92_1467 | ~ spl92_1 | ~ spl92_157 | ~ spl92_158), inference(avatar_split_clause, [], [f133655, f6102, f6098, f412, f65786])).
fof(f65786, plain, (spl92_1467 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP3(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1467])])).
fof(f6098, plain, (spl92_157 <=> p405(sK84(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_157])])).
fof(f6102, plain, (spl92_158 <=> r1(sK91, sK83(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_158])])).
fof(f133655, plain, (! [X1] : (~ r1(X1, sK91) | ~ sP3(X1)) | (~ spl92_1 | ~ spl92_157 | ~ spl92_158)), inference(subsumption_resolution, [], [f133654, f6100])).
fof(f6100, plain, (p405(sK84(sK91)) | ~ spl92_157), inference(avatar_component_clause, [], [f6098])).
fof(f133654, plain, (! [X1] : (~ p405(sK84(sK91)) | ~ r1(X1, sK91) | ~ sP3(X1)) | (~ spl92_1 | ~ spl92_158)), inference(resolution, [], [f132519, f391])).
fof(f391, plain, ! [X0, X1] : (~ p105(sK83(X1)) | ~ p405(sK84(X1)) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f209])).
fof(f209, plain, ! [X0] : (! [X1] : ((~ p105(sK83(X1)) & r1(X1, sK83(X1))) | (~ p405(sK84(X1)) & r1(X1, sK84(X1))) | ~ r1(X0, X1)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK83, sK84])], [f206, f208, f207])).
fof(f207, plain, ! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) => (~ p105(sK83(X1)) & r1(X1, sK83(X1)))), introduced(choice_axiom, [])).
fof(f208, plain, ! [X1] : (? [X3] : (~ p405(X3) & r1(X1, X3)) => (~ p405(sK84(X1)) & r1(X1, sK84(X1)))), introduced(choice_axiom, [])).
fof(f206, plain, ! [X0] : (! [X1] : (? [X2] : (~ p105(X2) & r1(X1, X2)) | ? [X3] : (~ p405(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP3(X0)), inference(rectify, [], [f205])).
fof(f205, plain, ! [X11] : (! [X108] : (? [X109] : (~ p105(X109) & r1(X108, X109)) | ? [X110] : (~ p405(X110) & r1(X108, X110)) | ~ r1(X11, X108)) | ~ sP3(X11)), inference(nnf_transformation, [], [f14])).
fof(f132519, plain, (p105(sK83(sK91)) | (~ spl92_1 | ~ spl92_158)), inference(resolution, [], [f413, f6104])).
fof(f6104, plain, (r1(sK91, sK83(sK91)) | ~ spl92_158), inference(avatar_component_clause, [], [f6102])).
fof(f132509, plain, (~ spl92_21 | ~ spl92_1599), inference(avatar_contradiction_clause, [], [f132508])).
fof(f132508, plain, ($false | (~ spl92_21 | ~ spl92_1599)), inference(subsumption_resolution, [], [f132507, f529])).
fof(f529, plain, sP11(sK91), inference(resolution, [], [f527, f231])).
fof(f231, plain, ! [X0] : (~ sP40(X0) | sP11(X0)), inference(cnf_transformation, [], [f54])).
fof(f132507, plain, (~ sP11(sK91) | (~ spl92_21 | ~ spl92_1599)), inference(resolution, [], [f111530, f227])).
fof(f111530, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP11(X0)) | (~ spl92_21 | ~ spl92_1599)), inference(subsumption_resolution, [], [f111529, f488])).
fof(f111529, plain, (! [X0] : (~ p505(sK91) | ~ r1(X0, sK91) | ~ sP11(X0)) | ~ spl92_1599), inference(resolution, [], [f110535, f361])).
fof(f361, plain, ! [X0, X1] : (~ p405(sK69(X1)) | ~ p505(X1) | ~ r1(X0, X1) | ~ sP11(X0)), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ! [X0] : (! [X1] : ((~ p405(sK69(X1)) & r1(X1, sK69(X1))) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP11(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK69])], [f168, f169])).
fof(f169, plain, ! [X1] : (? [X2] : (~ p405(X2) & r1(X1, X2)) => (~ p405(sK69(X1)) & r1(X1, sK69(X1)))), introduced(choice_axiom, [])).
fof(f168, plain, ! [X0] : (! [X1] : (? [X2] : (~ p405(X2) & r1(X1, X2)) | ~ p505(X1) | ~ r1(X0, X1)) | ~ sP11(X0)), inference(rectify, [], [f167])).
fof(f167, plain, ! [X11] : (! [X132] : (? [X133] : (~ p405(X133) & r1(X132, X133)) | ~ p505(X132) | ~ r1(X11, X132)) | ~ sP11(X11)), inference(nnf_transformation, [], [f22])).
fof(f110535, plain, (p405(sK69(sK91)) | ~ spl92_1599), inference(avatar_component_clause, [], [f110533])).
fof(f110533, plain, (spl92_1599 <=> p405(sK69(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1599])])).
fof(f110536, plain, (~ spl92_21 | spl92_1599 | ~ spl92_16), inference(avatar_split_clause, [], [f110531, f466, f110533, f486])).
fof(f110531, plain, (p405(sK69(sK91)) | ~ p505(sK91) | ~ spl92_16), inference(subsumption_resolution, [], [f54859, f529])).
fof(f54859, plain, (p405(sK69(sK91)) | ~ p505(sK91) | ~ sP11(sK91) | ~ spl92_16), inference(resolution, [], [f467, f2529])).
fof(f2529, plain, ! [X0] : (r1(X0, sK69(X0)) | ~ p505(X0) | ~ sP11(X0)), inference(resolution, [], [f360, f227])).
fof(f360, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p505(X1) | r1(X1, sK69(X1)) | ~ sP11(X0)), inference(cnf_transformation, [], [f170])).
fof(f110474, plain, (~ spl92_3 | ~ spl92_23), inference(avatar_contradiction_clause, [], [f110473])).
fof(f110473, plain, ($false | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f110472, f559])).
fof(f559, plain, sP32(sK91), inference(resolution, [], [f527, f270])).
fof(f270, plain, ! [X0] : (~ sP40(X0) | sP32(X0)), inference(cnf_transformation, [], [f54])).
fof(f110472, plain, (~ sP32(sK91) | (~ spl92_3 | ~ spl92_23)), inference(resolution, [], [f109407, f227])).
fof(f109407, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP32(X0)) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f109406, f496])).
fof(f109406, plain, (! [X0] : (~ p503(sK91) | ~ r1(X0, sK91) | ~ sP32(X0)) | (~ spl92_3 | ~ spl92_23)), inference(resolution, [], [f109398, f319])).
fof(f319, plain, ! [X0, X1] : (~ p103(sK48(X1)) | ~ p503(X1) | ~ r1(X0, X1) | ~ sP32(X0)), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ! [X0] : (! [X1] : ((~ p103(sK48(X1)) & r1(X1, sK48(X1))) | ~ p503(X1) | ~ r1(X0, X1)) | ~ sP32(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48])], [f84, f85])).
fof(f85, plain, ! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) => (~ p103(sK48(X1)) & r1(X1, sK48(X1)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X0] : (! [X1] : (? [X2] : (~ p103(X2) & r1(X1, X2)) | ~ p503(X1) | ~ r1(X0, X1)) | ~ sP32(X0)), inference(rectify, [], [f83])).
fof(f83, plain, ! [X11] : (! [X54] : (? [X55] : (~ p103(X55) & r1(X54, X55)) | ~ p503(X54) | ~ r1(X11, X54)) | ~ sP32(X11)), inference(nnf_transformation, [], [f43])).
fof(f109398, plain, (p103(sK48(sK91)) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f109397, f559])).
fof(f109397, plain, (p103(sK48(sK91)) | ~ sP32(sK91) | (~ spl92_3 | ~ spl92_23)), inference(subsumption_resolution, [], [f108969, f496])).
fof(f108969, plain, (p103(sK48(sK91)) | ~ p503(sK91) | ~ sP32(sK91) | ~ spl92_3), inference(resolution, [], [f419, f1289])).
fof(f1289, plain, ! [X0] : (r1(X0, sK48(X0)) | ~ p503(X0) | ~ sP32(X0)), inference(resolution, [], [f318, f227])).
fof(f318, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p503(X1) | r1(X1, sK48(X1)) | ~ sP32(X0)), inference(cnf_transformation, [], [f86])).
fof(f108944, plain, (~ spl92_4 | ~ spl92_9), inference(avatar_contradiction_clause, [], [f108943])).
fof(f108943, plain, ($false | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f108942, f567])).
fof(f567, plain, sP39(sK91), inference(resolution, [], [f527, f288])).
fof(f288, plain, ! [X0] : (~ sP40(X0) | sP39(X0)), inference(cnf_transformation, [], [f54])).
fof(f108942, plain, (~ sP39(sK91) | (~ spl92_4 | ~ spl92_9)), inference(resolution, [], [f108179, f227])).
fof(f108179, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP39(X0)) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f108178, f440])).
fof(f440, plain, (p202(sK91) | ~ spl92_9), inference(avatar_component_clause, [], [f438])).
fof(f108178, plain, (! [X0] : (~ p202(sK91) | ~ r1(X0, sK91) | ~ sP39(X0)) | (~ spl92_4 | ~ spl92_9)), inference(resolution, [], [f108169, f305])).
fof(f305, plain, ! [X0, X1] : (~ p102(sK41(X1)) | ~ p202(X1) | ~ r1(X0, X1) | ~ sP39(X0)), inference(cnf_transformation, [], [f58])).
fof(f58, plain, ! [X0] : (! [X1] : ((~ p102(sK41(X1)) & r1(X1, sK41(X1))) | ~ p202(X1) | ~ r1(X0, X1)) | ~ sP39(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f56, f57])).
fof(f57, plain, ! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) => (~ p102(sK41(X1)) & r1(X1, sK41(X1)))), introduced(choice_axiom, [])).
fof(f56, plain, ! [X0] : (! [X1] : (? [X2] : (~ p102(X2) & r1(X1, X2)) | ~ p202(X1) | ~ r1(X0, X1)) | ~ sP39(X0)), inference(rectify, [], [f55])).
fof(f55, plain, ! [X11] : (! [X27] : (? [X28] : (~ p102(X28) & r1(X27, X28)) | ~ p202(X27) | ~ r1(X11, X27)) | ~ sP39(X11)), inference(nnf_transformation, [], [f50])).
fof(f108169, plain, (p102(sK41(sK91)) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f108168, f567])).
fof(f108168, plain, (p102(sK41(sK91)) | ~ sP39(sK91) | (~ spl92_4 | ~ spl92_9)), inference(subsumption_resolution, [], [f107728, f440])).
fof(f107728, plain, (p102(sK41(sK91)) | ~ p202(sK91) | ~ sP39(sK91) | ~ spl92_4), inference(resolution, [], [f422, f664])).
fof(f664, plain, ! [X0] : (r1(X0, sK41(X0)) | ~ p202(X0) | ~ sP39(X0)), inference(resolution, [], [f304, f227])).
fof(f304, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p202(X1) | r1(X1, sK41(X1)) | ~ sP39(X0)), inference(cnf_transformation, [], [f58])).
fof(f107555, plain, (~ spl92_16 | spl92_157 | ~ spl92_1468), inference(avatar_contradiction_clause, [], [f107554])).
fof(f107554, plain, ($false | (~ spl92_16 | spl92_157 | ~ spl92_1468)), inference(subsumption_resolution, [], [f107474, f6099])).
fof(f6099, plain, (~ p405(sK84(sK91)) | spl92_157), inference(avatar_component_clause, [], [f6098])).
fof(f107474, plain, (p405(sK84(sK91)) | (~ spl92_16 | ~ spl92_1468)), inference(resolution, [], [f65791, f467])).
fof(f65791, plain, (r1(sK91, sK84(sK91)) | ~ spl92_1468), inference(avatar_component_clause, [], [f65789])).
fof(f65789, plain, (spl92_1468 <=> r1(sK91, sK84(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1468])])).
fof(f107397, plain, ~ spl92_1467, inference(avatar_contradiction_clause, [], [f107396])).
fof(f107396, plain, ($false | ~ spl92_1467), inference(subsumption_resolution, [], [f107395, f539])).
fof(f539, plain, sP3(sK91), inference(resolution, [], [f527, f241])).
fof(f241, plain, ! [X0] : (~ sP40(X0) | sP3(X0)), inference(cnf_transformation, [], [f54])).
fof(f107395, plain, (~ sP3(sK91) | ~ spl92_1467), inference(resolution, [], [f65787, f227])).
fof(f65787, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP3(X0)) | ~ spl92_1467), inference(avatar_component_clause, [], [f65786])).
fof(f107318, plain, ~ spl92_1362, inference(avatar_contradiction_clause, [], [f107317])).
fof(f107317, plain, ($false | ~ spl92_1362), inference(subsumption_resolution, [], [f107316, f536])).
fof(f536, plain, sP2(sK91), inference(resolution, [], [f527, f238])).
fof(f238, plain, ! [X0] : (~ sP40(X0) | sP2(X0)), inference(cnf_transformation, [], [f54])).
fof(f107316, plain, (~ sP2(sK91) | ~ spl92_1362), inference(resolution, [], [f31386, f227])).
fof(f31386, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP2(X0)) | ~ spl92_1362), inference(avatar_component_clause, [], [f31385])).
fof(f98459, plain, (~ spl92_22 | ~ spl92_1485), inference(avatar_contradiction_clause, [], [f98458])).
fof(f98458, plain, ($false | (~ spl92_22 | ~ spl92_1485)), inference(subsumption_resolution, [], [f98457, f543])).
fof(f543, plain, sP19(sK91), inference(resolution, [], [f527, f248])).
fof(f248, plain, ! [X0] : (~ sP40(X0) | sP19(X0)), inference(cnf_transformation, [], [f54])).
fof(f98457, plain, (~ sP19(sK91) | (~ spl92_22 | ~ spl92_1485)), inference(resolution, [], [f83676, f227])).
fof(f83676, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP19(X0)) | (~ spl92_22 | ~ spl92_1485)), inference(subsumption_resolution, [], [f83675, f492])).
fof(f83675, plain, (! [X0] : (~ p504(sK91) | ~ r1(X0, sK91) | ~ sP19(X0)) | ~ spl92_1485), inference(resolution, [], [f83216, f345])).
fof(f345, plain, ! [X0, X1] : (~ p304(sK61(X1)) | ~ p504(X1) | ~ r1(X0, X1) | ~ sP19(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : (! [X1] : ((~ p304(sK61(X1)) & r1(X1, sK61(X1))) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP19(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK61])], [f136, f137])).
fof(f137, plain, ! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) => (~ p304(sK61(X1)) & r1(X1, sK61(X1)))), introduced(choice_axiom, [])).
fof(f136, plain, ! [X0] : (! [X1] : (? [X2] : (~ p304(X2) & r1(X1, X2)) | ~ p504(X1) | ~ r1(X0, X1)) | ~ sP19(X0)), inference(rectify, [], [f135])).
fof(f135, plain, ! [X11] : (! [X95] : (? [X96] : (~ p304(X96) & r1(X95, X96)) | ~ p504(X95) | ~ r1(X11, X95)) | ~ sP19(X11)), inference(nnf_transformation, [], [f30])).
fof(f83216, plain, (p304(sK61(sK91)) | ~ spl92_1485), inference(avatar_component_clause, [], [f83214])).
fof(f83214, plain, (spl92_1485 <=> p304(sK61(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1485])])).
fof(f83217, plain, (~ spl92_22 | spl92_1485 | ~ spl92_12), inference(avatar_split_clause, [], [f83212, f450, f83214, f490])).
fof(f83212, plain, (p304(sK61(sK91)) | ~ p504(sK91) | ~ spl92_12), inference(subsumption_resolution, [], [f16936, f543])).
fof(f16936, plain, (p304(sK61(sK91)) | ~ p504(sK91) | ~ sP19(sK91) | ~ spl92_12), inference(resolution, [], [f451, f2401])).
fof(f2401, plain, ! [X0] : (r1(X0, sK61(X0)) | ~ p504(X0) | ~ sP19(X0)), inference(resolution, [], [f344, f227])).
fof(f344, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p504(X1) | r1(X1, sK61(X1)) | ~ sP19(X0)), inference(cnf_transformation, [], [f138])).
fof(f83032, plain, (~ spl92_8 | spl92_102), inference(avatar_contradiction_clause, [], [f83031])).
fof(f83031, plain, ($false | (~ spl92_8 | spl92_102)), inference(subsumption_resolution, [], [f83030, f562])).
fof(f83030, plain, (~ sP9(sK91) | (~ spl92_8 | spl92_102)), inference(resolution, [], [f55087, f227])).
fof(f55087, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP9(X0)) | (~ spl92_8 | spl92_102)), inference(subsumption_resolution, [], [f55086, f5005])).
fof(f5005, plain, (~ r1(sK91, sK71(sK91)) | spl92_102), inference(avatar_component_clause, [], [f5004])).
fof(f55086, plain, (! [X0] : (r1(sK91, sK71(sK91)) | ~ r1(X0, sK91) | ~ sP9(X0)) | (~ spl92_8 | spl92_102)), inference(resolution, [], [f54793, f365])).
fof(f365, plain, ! [X0, X1] : (~ p203(sK72(X1)) | r1(X1, sK71(X1)) | ~ r1(X0, X1) | ~ sP9(X0)), inference(cnf_transformation, [], [f179])).
fof(f54793, plain, (p203(sK72(sK91)) | (~ spl92_8 | spl92_102)), inference(subsumption_resolution, [], [f54792, f562])).
fof(f54792, plain, (p203(sK72(sK91)) | ~ sP9(sK91) | (~ spl92_8 | spl92_102)), inference(subsumption_resolution, [], [f54602, f5005])).
fof(f54602, plain, (p203(sK72(sK91)) | r1(sK91, sK71(sK91)) | ~ sP9(sK91) | ~ spl92_8), inference(resolution, [], [f436, f2561])).
fof(f2561, plain, ! [X0] : (r1(X0, sK72(X0)) | r1(X0, sK71(X0)) | ~ sP9(X0)), inference(resolution, [], [f364, f227])).
fof(f364, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK72(X1)) | r1(X1, sK71(X1)) | ~ sP9(X0)), inference(cnf_transformation, [], [f179])).
fof(f65792, plain, (spl92_1467 | spl92_1468 | ~ spl92_1 | ~ spl92_158), inference(avatar_split_clause, [], [f65783, f6102, f412, f65789, f65786])).
fof(f65783, plain, (! [X0] : (r1(sK91, sK84(sK91)) | ~ r1(X0, sK91) | ~ sP3(X0)) | (~ spl92_1 | ~ spl92_158)), inference(resolution, [], [f55091, f390])).
fof(f390, plain, ! [X0, X1] : (~ p105(sK83(X1)) | r1(X1, sK84(X1)) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f209])).
fof(f55091, plain, (p105(sK83(sK91)) | (~ spl92_1 | ~ spl92_158)), inference(resolution, [], [f6104, f413])).
fof(f54279, plain, (~ spl92_12 | spl92_1283 | ~ spl92_1426), inference(avatar_contradiction_clause, [], [f54278])).
fof(f54278, plain, ($false | (~ spl92_12 | spl92_1283 | ~ spl92_1426)), inference(subsumption_resolution, [], [f54199, f16404])).
fof(f16404, plain, (~ p304(sK78(sK91)) | spl92_1283), inference(avatar_component_clause, [], [f16403])).
fof(f16403, plain, (spl92_1283 <=> p304(sK78(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1283])])).
fof(f54199, plain, (p304(sK78(sK91)) | (~ spl92_12 | ~ spl92_1426)), inference(resolution, [], [f46407, f451])).
fof(f46407, plain, (r1(sK91, sK78(sK91)) | ~ spl92_1426), inference(avatar_component_clause, [], [f46405])).
fof(f46405, plain, (spl92_1426 <=> r1(sK91, sK78(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1426])])).
fof(f46409, plain, (spl92_1425 | ~ spl92_1283 | ~ spl92_7 | ~ spl92_140), inference(avatar_split_clause, [], [f46400, f5647, f432, f16403, f46402])).
fof(f5647, plain, (spl92_140 <=> r1(sK91, sK77(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_140])])).
fof(f46400, plain, (! [X1] : (~ p304(sK78(sK91)) | ~ r1(X1, sK91) | ~ sP6(X1)) | (~ spl92_7 | ~ spl92_140)), inference(resolution, [], [f45449, f379])).
fof(f379, plain, ! [X0, X1] : (~ p204(sK77(X1)) | ~ p304(sK78(X1)) | ~ r1(X0, X1) | ~ sP6(X0)), inference(cnf_transformation, [], [f194])).
fof(f194, plain, ! [X0] : (! [X1] : ((~ p204(sK77(X1)) & r1(X1, sK77(X1))) | (~ p304(sK78(X1)) & r1(X1, sK78(X1))) | ~ r1(X0, X1)) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK77, sK78])], [f191, f193, f192])).
fof(f192, plain, ! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) => (~ p204(sK77(X1)) & r1(X1, sK77(X1)))), introduced(choice_axiom, [])).
fof(f193, plain, ! [X1] : (? [X3] : (~ p304(X3) & r1(X1, X3)) => (~ p304(sK78(X1)) & r1(X1, sK78(X1)))), introduced(choice_axiom, [])).
fof(f191, plain, ! [X0] : (! [X1] : (? [X2] : (~ p204(X2) & r1(X1, X2)) | ? [X3] : (~ p304(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP6(X0)), inference(rectify, [], [f190])).
fof(f190, plain, ! [X11] : (! [X84] : (? [X85] : (~ p204(X85) & r1(X84, X85)) | ? [X86] : (~ p304(X86) & r1(X84, X86)) | ~ r1(X11, X84)) | ~ sP6(X11)), inference(nnf_transformation, [], [f17])).
fof(f45449, plain, (p204(sK77(sK91)) | (~ spl92_7 | ~ spl92_140)), inference(resolution, [], [f433, f5649])).
fof(f5649, plain, (r1(sK91, sK77(sK91)) | ~ spl92_140), inference(avatar_component_clause, [], [f5647])).
fof(f46408, plain, (spl92_1425 | spl92_1426 | ~ spl92_7 | ~ spl92_140), inference(avatar_split_clause, [], [f46399, f5647, f432, f46405, f46402])).
fof(f46399, plain, (! [X0] : (r1(sK91, sK78(sK91)) | ~ r1(X0, sK91) | ~ sP6(X0)) | (~ spl92_7 | ~ spl92_140)), inference(resolution, [], [f45449, f378])).
fof(f378, plain, ! [X0, X1] : (~ p204(sK77(X1)) | r1(X1, sK78(X1)) | ~ r1(X0, X1) | ~ sP6(X0)), inference(cnf_transformation, [], [f194])).
fof(f45441, plain, ~ spl92_1360, inference(avatar_contradiction_clause, [], [f45440])).
fof(f45440, plain, ($false | ~ spl92_1360), inference(subsumption_resolution, [], [f45439, f541])).
fof(f541, plain, sP5(sK91), inference(resolution, [], [f527, f243])).
fof(f243, plain, ! [X0] : (~ sP40(X0) | sP5(X0)), inference(cnf_transformation, [], [f54])).
fof(f45439, plain, (~ sP5(sK91) | ~ spl92_1360), inference(resolution, [], [f17516, f227])).
fof(f17516, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP5(X0)) | ~ spl92_1360), inference(avatar_component_clause, [], [f17515])).
fof(f44415, plain, (~ spl92_6 | spl92_148 | ~ spl92_1361), inference(avatar_contradiction_clause, [], [f44414])).
fof(f44414, plain, ($false | (~ spl92_6 | spl92_148 | ~ spl92_1361)), inference(subsumption_resolution, [], [f44335, f5811])).
fof(f5811, plain, (~ p205(sK80(sK91)) | spl92_148), inference(avatar_component_clause, [], [f5810])).
fof(f44335, plain, (p205(sK80(sK91)) | (~ spl92_6 | ~ spl92_1361)), inference(resolution, [], [f17520, f430])).
fof(f17520, plain, (r1(sK91, sK80(sK91)) | ~ spl92_1361), inference(avatar_component_clause, [], [f17518])).
fof(f17518, plain, (spl92_1361 <=> r1(sK91, sK80(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1361])])).
fof(f31391, plain, (spl92_1362 | spl92_1363 | ~ spl92_6 | ~ spl92_184), inference(avatar_split_clause, [], [f31382, f6443, f429, f31388, f31385])).
fof(f31382, plain, (! [X0] : (r1(sK91, sK86(sK91)) | ~ r1(X0, sK91) | ~ sP2(X0)) | (~ spl92_6 | ~ spl92_184)), inference(resolution, [], [f17434, f394])).
fof(f394, plain, ! [X0, X1] : (~ p205(sK85(X1)) | r1(X1, sK86(X1)) | ~ r1(X0, X1) | ~ sP2(X0)), inference(cnf_transformation, [], [f214])).
fof(f17434, plain, (p205(sK85(sK91)) | (~ spl92_6 | ~ spl92_184)), inference(resolution, [], [f6445, f430])).
fof(f17521, plain, (spl92_1360 | spl92_1361 | ~ spl92_1 | ~ spl92_146), inference(avatar_split_clause, [], [f17512, f5798, f412, f17518, f17515])).
fof(f17512, plain, (! [X0] : (r1(sK91, sK80(sK91)) | ~ r1(X0, sK91) | ~ sP5(X0)) | (~ spl92_1 | ~ spl92_146)), inference(resolution, [], [f16853, f382])).
fof(f382, plain, ! [X0, X1] : (~ p105(sK79(X1)) | r1(X1, sK80(X1)) | ~ r1(X0, X1) | ~ sP5(X0)), inference(cnf_transformation, [], [f199])).
fof(f16853, plain, (p105(sK79(sK91)) | (~ spl92_1 | ~ spl92_146)), inference(resolution, [], [f413, f5800])).
fof(f16825, plain, (~ spl92_19 | ~ spl92_14), inference(avatar_split_clause, [], [f638, f457, f477])).
fof(f638, plain, (~ p302(sK91) | ~ p402(sK91)), inference(resolution, [], [f588, f527])).
fof(f588, plain, ! [X0] : (~ sP40(X0) | ~ p302(X0) | ~ p402(X0)), inference(resolution, [], [f279, f227])).
fof(f279, plain, ! [X0, X20] : (~ r1(X0, X20) | ~ p402(X20) | ~ p302(X20) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f16646, plain, (~ spl92_12 | spl92_1280 | ~ spl92_1282), inference(avatar_contradiction_clause, [], [f16645])).
fof(f16645, plain, ($false | (~ spl92_12 | spl92_1280 | ~ spl92_1282)), inference(subsumption_resolution, [], [f16566, f16229])).
fof(f16229, plain, (~ p304(sK76(sK91)) | spl92_1280), inference(avatar_component_clause, [], [f16228])).
fof(f16228, plain, (spl92_1280 <=> p304(sK76(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1280])])).
fof(f16566, plain, (p304(sK76(sK91)) | (~ spl92_12 | ~ spl92_1282)), inference(resolution, [], [f16320, f451])).
fof(f16320, plain, (r1(sK91, sK76(sK91)) | ~ spl92_1282), inference(avatar_component_clause, [], [f16318])).
fof(f16318, plain, (spl92_1282 <=> r1(sK91, sK76(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_1282])])).
fof(f16564, plain, ~ spl92_1281, inference(avatar_contradiction_clause, [], [f16563])).
fof(f16563, plain, ($false | ~ spl92_1281), inference(subsumption_resolution, [], [f16562, f552])).
fof(f552, plain, sP7(sK91), inference(resolution, [], [f527, f257])).
fof(f257, plain, ! [X0] : (~ sP40(X0) | sP7(X0)), inference(cnf_transformation, [], [f54])).
fof(f16562, plain, (~ sP7(sK91) | ~ spl92_1281), inference(resolution, [], [f16316, f227])).
fof(f16316, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP7(X0)) | ~ spl92_1281), inference(avatar_component_clause, [], [f16315])).
fof(f16315, plain, (spl92_1281 <=> ! [X0] : (~ r1(X0, sK91) | ~ sP7(X0))), introduced(avatar_definition, [new_symbols(naming, [spl92_1281])])).
fof(f16400, plain, (~ spl92_12 | spl92_140), inference(avatar_contradiction_clause, [], [f16399])).
fof(f16399, plain, ($false | (~ spl92_12 | spl92_140)), inference(subsumption_resolution, [], [f16398, f548])).
fof(f16398, plain, (~ sP6(sK91) | (~ spl92_12 | spl92_140)), inference(resolution, [], [f15997, f227])).
fof(f15997, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP6(X0)) | (~ spl92_12 | spl92_140)), inference(subsumption_resolution, [], [f15996, f5648])).
fof(f5648, plain, (~ r1(sK91, sK77(sK91)) | spl92_140), inference(avatar_component_clause, [], [f5647])).
fof(f15996, plain, (! [X0] : (r1(sK91, sK77(sK91)) | ~ r1(X0, sK91) | ~ sP6(X0)) | (~ spl92_12 | spl92_140)), inference(resolution, [], [f15977, f377])).
fof(f377, plain, ! [X0, X1] : (~ p304(sK78(X1)) | r1(X1, sK77(X1)) | ~ r1(X0, X1) | ~ sP6(X0)), inference(cnf_transformation, [], [f194])).
fof(f15977, plain, (p304(sK78(sK91)) | (~ spl92_12 | spl92_140)), inference(subsumption_resolution, [], [f15976, f548])).
fof(f15976, plain, (p304(sK78(sK91)) | ~ sP6(sK91) | (~ spl92_12 | spl92_140)), inference(subsumption_resolution, [], [f15959, f5648])).
fof(f15959, plain, (p304(sK78(sK91)) | r1(sK91, sK77(sK91)) | ~ sP6(sK91) | ~ spl92_12), inference(resolution, [], [f451, f2609])).
fof(f2609, plain, ! [X0] : (r1(X0, sK78(X0)) | r1(X0, sK77(X0)) | ~ sP6(X0)), inference(resolution, [], [f376, f227])).
fof(f376, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK78(X1)) | r1(X1, sK77(X1)) | ~ sP6(X0)), inference(cnf_transformation, [], [f194])).
fof(f16322, plain, (spl92_1281 | ~ spl92_1280 | ~ spl92_2 | ~ spl92_134), inference(avatar_split_clause, [], [f16313, f5496, f415, f16228, f16315])).
fof(f5496, plain, (spl92_134 <=> r1(sK91, sK75(sK91))), introduced(avatar_definition, [new_symbols(naming, [spl92_134])])).
fof(f16313, plain, (! [X1] : (~ p304(sK76(sK91)) | ~ r1(X1, sK91) | ~ sP7(X1)) | (~ spl92_2 | ~ spl92_134)), inference(resolution, [], [f16234, f375])).
fof(f375, plain, ! [X0, X1] : (~ p104(sK75(X1)) | ~ p304(sK76(X1)) | ~ r1(X0, X1) | ~ sP7(X0)), inference(cnf_transformation, [], [f189])).
fof(f189, plain, ! [X0] : (! [X1] : ((~ p104(sK75(X1)) & r1(X1, sK75(X1))) | (~ p304(sK76(X1)) & r1(X1, sK76(X1))) | ~ r1(X0, X1)) | ~ sP7(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK75, sK76])], [f186, f188, f187])).
fof(f187, plain, ! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) => (~ p104(sK75(X1)) & r1(X1, sK75(X1)))), introduced(choice_axiom, [])).
fof(f188, plain, ! [X1] : (? [X3] : (~ p304(X3) & r1(X1, X3)) => (~ p304(sK76(X1)) & r1(X1, sK76(X1)))), introduced(choice_axiom, [])).
fof(f186, plain, ! [X0] : (! [X1] : (? [X2] : (~ p104(X2) & r1(X1, X2)) | ? [X3] : (~ p304(X3) & r1(X1, X3)) | ~ r1(X0, X1)) | ~ sP7(X0)), inference(rectify, [], [f185])).
fof(f185, plain, ! [X11] : (! [X75] : (? [X76] : (~ p104(X76) & r1(X75, X76)) | ? [X77] : (~ p304(X77) & r1(X75, X77)) | ~ r1(X11, X75)) | ~ sP7(X11)), inference(nnf_transformation, [], [f18])).
fof(f16234, plain, (p104(sK75(sK91)) | (~ spl92_2 | ~ spl92_134)), inference(resolution, [], [f5498, f416])).
fof(f5498, plain, (r1(sK91, sK75(sK91)) | ~ spl92_134), inference(avatar_component_clause, [], [f5496])).
fof(f16321, plain, (spl92_1281 | spl92_1282 | ~ spl92_2 | ~ spl92_134), inference(avatar_split_clause, [], [f16312, f5496, f415, f16318, f16315])).
fof(f16312, plain, (! [X0] : (r1(sK91, sK76(sK91)) | ~ r1(X0, sK91) | ~ sP7(X0)) | (~ spl92_2 | ~ spl92_134)), inference(resolution, [], [f16234, f374])).
fof(f374, plain, ! [X0, X1] : (~ p104(sK75(X1)) | r1(X1, sK76(X1)) | ~ r1(X0, X1) | ~ sP7(X0)), inference(cnf_transformation, [], [f189])).
fof(f16225, plain, (~ spl92_12 | spl92_134), inference(avatar_contradiction_clause, [], [f16224])).
fof(f16224, plain, ($false | (~ spl92_12 | spl92_134)), inference(subsumption_resolution, [], [f16223, f552])).
fof(f16223, plain, (~ sP7(sK91) | (~ spl92_12 | spl92_134)), inference(resolution, [], [f15995, f227])).
fof(f15995, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP7(X0)) | (~ spl92_12 | spl92_134)), inference(subsumption_resolution, [], [f15994, f5497])).
fof(f5497, plain, (~ r1(sK91, sK75(sK91)) | spl92_134), inference(avatar_component_clause, [], [f5496])).
fof(f15994, plain, (! [X0] : (r1(sK91, sK75(sK91)) | ~ r1(X0, sK91) | ~ sP7(X0)) | (~ spl92_12 | spl92_134)), inference(resolution, [], [f15975, f373])).
fof(f373, plain, ! [X0, X1] : (~ p304(sK76(X1)) | r1(X1, sK75(X1)) | ~ r1(X0, X1) | ~ sP7(X0)), inference(cnf_transformation, [], [f189])).
fof(f15975, plain, (p304(sK76(sK91)) | (~ spl92_12 | spl92_134)), inference(subsumption_resolution, [], [f15974, f552])).
fof(f15974, plain, (p304(sK76(sK91)) | ~ sP7(sK91) | (~ spl92_12 | spl92_134)), inference(subsumption_resolution, [], [f15958, f5497])).
fof(f15958, plain, (p304(sK76(sK91)) | r1(sK91, sK75(sK91)) | ~ sP7(sK91) | ~ spl92_12), inference(resolution, [], [f451, f2593])).
fof(f2593, plain, ! [X0] : (r1(X0, sK76(X0)) | r1(X0, sK75(X0)) | ~ sP7(X0)), inference(resolution, [], [f372, f227])).
fof(f372, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK76(X1)) | r1(X1, sK75(X1)) | ~ sP7(X0)), inference(cnf_transformation, [], [f189])).
fof(f16020, plain, (~ spl92_10 | ~ spl92_25), inference(avatar_split_clause, [], [f16019, f502, f442])).
fof(f502, plain, (spl92_25 <=> p501(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_25])])).
fof(f16019, plain, (~ p201(sK91) | ~ spl92_25), inference(subsumption_resolution, [], [f655, f504])).
fof(f504, plain, (p501(sK91) | ~ spl92_25), inference(avatar_component_clause, [], [f502])).
fof(f655, plain, (~ p201(sK91) | ~ p501(sK91)), inference(resolution, [], [f600, f527])).
fof(f600, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p501(X0)), inference(resolution, [], [f296, f227])).
fof(f296, plain, ! [X0, X8] : (~ r1(X0, X8) | ~ p501(X8) | ~ p201(X8) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f16018, plain, (~ spl92_8 | ~ spl92_18), inference(avatar_contradiction_clause, [], [f16017])).
fof(f16017, plain, ($false | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f16016, f556])).
fof(f556, plain, sP29(sK91), inference(resolution, [], [f527, f267])).
fof(f267, plain, ! [X0] : (~ sP40(X0) | sP29(X0)), inference(cnf_transformation, [], [f54])).
fof(f16016, plain, (~ sP29(sK91) | (~ spl92_8 | ~ spl92_18)), inference(resolution, [], [f15991, f227])).
fof(f15991, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP29(X0)) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f15990, f475])).
fof(f15990, plain, (! [X0] : (~ p403(sK91) | ~ r1(X0, sK91) | ~ sP29(X0)) | (~ spl92_8 | ~ spl92_18)), inference(resolution, [], [f15904, f325])).
fof(f325, plain, ! [X0, X1] : (~ p203(sK51(X1)) | ~ p403(X1) | ~ r1(X0, X1) | ~ sP29(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ! [X0] : (! [X1] : ((~ p203(sK51(X1)) & r1(X1, sK51(X1))) | ~ p403(X1) | ~ r1(X0, X1)) | ~ sP29(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK51])], [f96, f97])).
fof(f97, plain, ! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) => (~ p203(sK51(X1)) & r1(X1, sK51(X1)))), introduced(choice_axiom, [])).
fof(f96, plain, ! [X0] : (! [X1] : (? [X2] : (~ p203(X2) & r1(X1, X2)) | ~ p403(X1) | ~ r1(X0, X1)) | ~ sP29(X0)), inference(rectify, [], [f95])).
fof(f95, plain, ! [X11] : (! [X60] : (? [X61] : (~ p203(X61) & r1(X60, X61)) | ~ p403(X60) | ~ r1(X11, X60)) | ~ sP29(X11)), inference(nnf_transformation, [], [f40])).
fof(f15904, plain, (p203(sK51(sK91)) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f15903, f556])).
fof(f15903, plain, (p203(sK51(sK91)) | ~ sP29(sK91) | (~ spl92_8 | ~ spl92_18)), inference(subsumption_resolution, [], [f15869, f475])).
fof(f15869, plain, (p203(sK51(sK91)) | ~ p403(sK91) | ~ sP29(sK91) | ~ spl92_8), inference(resolution, [], [f436, f1604])).
fof(f1604, plain, ! [X0] : (r1(X0, sK51(X0)) | ~ p403(X0) | ~ sP29(X0)), inference(resolution, [], [f324, f227])).
fof(f324, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p403(X1) | r1(X1, sK51(X1)) | ~ sP29(X0)), inference(cnf_transformation, [], [f98])).
fof(f16015, plain, (~ spl92_195 | spl92_196), inference(avatar_contradiction_clause, [], [f16014])).
fof(f16014, plain, ($false | (~ spl92_195 | spl92_196)), inference(subsumption_resolution, [], [f16013, f532])).
fof(f16013, plain, (~ sP0(sK91) | (~ spl92_195 | spl92_196)), inference(resolution, [], [f15184, f227])).
fof(f15184, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP0(X0)) | (~ spl92_195 | spl92_196)), inference(subsumption_resolution, [], [f15183, f6748])).
fof(f6748, plain, (~ r1(sK91, sK89(sK91)) | spl92_196), inference(avatar_component_clause, [], [f6747])).
fof(f15183, plain, (! [X0] : (r1(sK91, sK89(sK91)) | ~ r1(X0, sK91) | ~ sP0(X0)) | ~ spl92_195), inference(resolution, [], [f6745, f401])).
fof(f401, plain, ! [X0, X1] : (~ p405(sK90(X1)) | r1(X1, sK89(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f224])).
fof(f6745, plain, (p405(sK90(sK91)) | ~ spl92_195), inference(avatar_component_clause, [], [f6743])).
fof(f16012, plain, (~ spl92_189 | spl92_190), inference(avatar_contradiction_clause, [], [f16011])).
fof(f16011, plain, ($false | (~ spl92_189 | spl92_190)), inference(subsumption_resolution, [], [f16010, f535])).
fof(f16010, plain, (~ sP1(sK91) | (~ spl92_189 | spl92_190)), inference(resolution, [], [f15182, f227])).
fof(f15182, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP1(X0)) | (~ spl92_189 | spl92_190)), inference(subsumption_resolution, [], [f15181, f6595])).
fof(f6595, plain, (~ r1(sK91, sK87(sK91)) | spl92_190), inference(avatar_component_clause, [], [f6594])).
fof(f15181, plain, (! [X0] : (r1(sK91, sK87(sK91)) | ~ r1(X0, sK91) | ~ sP1(X0)) | ~ spl92_189), inference(resolution, [], [f6592, f397])).
fof(f397, plain, ! [X0, X1] : (~ p405(sK88(X1)) | r1(X1, sK87(X1)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f219])).
fof(f6592, plain, (p405(sK88(sK91)) | ~ spl92_189), inference(avatar_component_clause, [], [f6590])).
fof(f16009, plain, (~ spl92_157 | spl92_158), inference(avatar_contradiction_clause, [], [f16008])).
fof(f16008, plain, ($false | (~ spl92_157 | spl92_158)), inference(subsumption_resolution, [], [f16007, f539])).
fof(f16007, plain, (~ sP3(sK91) | (~ spl92_157 | spl92_158)), inference(resolution, [], [f15180, f227])).
fof(f15180, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP3(X0)) | (~ spl92_157 | spl92_158)), inference(subsumption_resolution, [], [f15179, f6103])).
fof(f6103, plain, (~ r1(sK91, sK83(sK91)) | spl92_158), inference(avatar_component_clause, [], [f6102])).
fof(f15179, plain, (! [X0] : (r1(sK91, sK83(sK91)) | ~ r1(X0, sK91) | ~ sP3(X0)) | ~ spl92_157), inference(resolution, [], [f6100, f389])).
fof(f389, plain, ! [X0, X1] : (~ p405(sK84(X1)) | r1(X1, sK83(X1)) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f209])).
fof(f16006, plain, (spl92_184 | ~ spl92_185), inference(avatar_contradiction_clause, [], [f16005])).
fof(f16005, plain, ($false | (spl92_184 | ~ spl92_185)), inference(subsumption_resolution, [], [f16004, f536])).
fof(f16004, plain, (~ sP2(sK91) | (spl92_184 | ~ spl92_185)), inference(resolution, [], [f12873, f227])).
fof(f12873, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP2(X0)) | (spl92_184 | ~ spl92_185)), inference(subsumption_resolution, [], [f12872, f6444])).
fof(f6444, plain, (~ r1(sK91, sK85(sK91)) | spl92_184), inference(avatar_component_clause, [], [f6443])).
fof(f12872, plain, (! [X0] : (r1(sK91, sK85(sK91)) | ~ r1(X0, sK91) | ~ sP2(X0)) | ~ spl92_185), inference(resolution, [], [f6451, f393])).
fof(f393, plain, ! [X0, X1] : (~ p305(sK86(X1)) | r1(X1, sK85(X1)) | ~ r1(X0, X1) | ~ sP2(X0)), inference(cnf_transformation, [], [f214])).
fof(f16003, plain, (spl92_152 | ~ spl92_153), inference(avatar_contradiction_clause, [], [f16002])).
fof(f16002, plain, ($false | (spl92_152 | ~ spl92_153)), inference(subsumption_resolution, [], [f16001, f540])).
fof(f16001, plain, (~ sP4(sK91) | (spl92_152 | ~ spl92_153)), inference(resolution, [], [f11390, f227])).
fof(f11390, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP4(X0)) | (spl92_152 | ~ spl92_153)), inference(subsumption_resolution, [], [f11389, f5952])).
fof(f5952, plain, (~ r1(sK91, sK81(sK91)) | spl92_152), inference(avatar_component_clause, [], [f5951])).
fof(f11389, plain, (! [X0] : (r1(sK91, sK81(sK91)) | ~ r1(X0, sK91) | ~ sP4(X0)) | ~ spl92_153), inference(resolution, [], [f5959, f385])).
fof(f385, plain, ! [X0, X1] : (~ p305(sK82(X1)) | r1(X1, sK81(X1)) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f204])).
fof(f5959, plain, (p305(sK82(sK91)) | ~ spl92_153), inference(avatar_component_clause, [], [f5957])).
fof(f16000, plain, (spl92_146 | ~ spl92_148), inference(avatar_contradiction_clause, [], [f15999])).
fof(f15999, plain, ($false | (spl92_146 | ~ spl92_148)), inference(subsumption_resolution, [], [f15998, f541])).
fof(f15998, plain, (~ sP5(sK91) | (spl92_146 | ~ spl92_148)), inference(resolution, [], [f10686, f227])).
fof(f10686, plain, (! [X0] : (~ r1(X0, sK91) | ~ sP5(X0)) | (spl92_146 | ~ spl92_148)), inference(subsumption_resolution, [], [f10685, f5799])).
fof(f5799, plain, (~ r1(sK91, sK79(sK91)) | spl92_146), inference(avatar_component_clause, [], [f5798])).
fof(f10685, plain, (! [X0] : (r1(sK91, sK79(sK91)) | ~ r1(X0, sK91) | ~ sP5(X0)) | ~ spl92_148), inference(resolution, [], [f5812, f381])).
fof(f381, plain, ! [X0, X1] : (~ p205(sK80(X1)) | r1(X1, sK79(X1)) | ~ r1(X0, X1) | ~ sP5(X0)), inference(cnf_transformation, [], [f199])).
fof(f15855, plain, (~ spl92_15 | ~ spl92_25), inference(avatar_split_clause, [], [f15854, f502, f461])).
fof(f15854, plain, (~ p301(sK91) | ~ spl92_25), inference(subsumption_resolution, [], [f651, f504])).
fof(f651, plain, (~ p301(sK91) | ~ p501(sK91)), inference(resolution, [], [f597, f527])).
fof(f597, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p501(X0)), inference(resolution, [], [f293, f227])).
fof(f293, plain, ! [X0, X11] : (~ r1(X0, X11) | ~ p501(X11) | ~ p301(X11) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f15853, plain, (spl92_1276 | ~ spl92_1278 | ~ spl92_2 | ~ spl92_108), inference(avatar_split_clause, [], [f15840, f5155, f415, f15850, f15842])).
fof(f15840, plain, (! [X1] : (~ p204(sK74(sK91)) | ~ r1(X1, sK91) | ~ sP8(X1)) | (~ spl92_2 | ~ spl92_108)), inference(resolution, [], [f15761, f371])).
fof(f371, plain, ! [X0, X1] : (~ p104(sK73(X1)) | ~ p204(sK74(X1)) | ~ r1(X0, X1) | ~ sP8(X0)), inference(cnf_transformation, [], [f184])).
fof(f15761, plain, (p104(sK73(sK91)) | (~ spl92_2 | ~ spl92_108)), inference(resolution, [], [f5157, f416])).
fof(f5157, plain, (r1(sK91, sK73(sK91)) | ~ spl92_108), inference(avatar_component_clause, [], [f5155])).
fof(f15848, plain, (spl92_1276 | spl92_1277 | ~ spl92_2 | ~ spl92_108), inference(avatar_split_clause, [], [f15839, f5155, f415, f15845, f15842])).
fof(f15839, plain, (! [X0] : (r1(sK91, sK74(sK91)) | ~ r1(X0, sK91) | ~ sP8(X0)) | (~ spl92_2 | ~ spl92_108)), inference(resolution, [], [f15761, f370])).
fof(f370, plain, ! [X0, X1] : (~ p104(sK73(X1)) | r1(X1, sK74(X1)) | ~ r1(X0, X1) | ~ sP8(X0)), inference(cnf_transformation, [], [f184])).
fof(f14817, plain, (~ spl92_18 | ~ spl92_13), inference(avatar_split_clause, [], [f629, f453, f473])).
fof(f629, plain, (~ p303(sK91) | ~ p403(sK91)), inference(resolution, [], [f582, f527])).
fof(f582, plain, ! [X0] : (~ sP40(X0) | ~ p303(X0) | ~ p403(X0)), inference(resolution, [], [f264, f227])).
fof(f264, plain, ! [X26, X0] : (~ r1(X0, X26) | ~ p403(X26) | ~ p303(X26) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f14816, plain, (~ spl92_25 | ~ spl92_20), inference(avatar_split_clause, [], [f648, f481, f502])).
fof(f648, plain, (~ p401(sK91) | ~ p501(sK91)), inference(resolution, [], [f595, f527])).
fof(f595, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p501(X0)), inference(resolution, [], [f291, f227])).
fof(f291, plain, ! [X0, X13] : (~ r1(X0, X13) | ~ p501(X13) | ~ p401(X13) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f14803, plain, (~ spl92_20 | ~ spl92_5), inference(avatar_split_clause, [], [f661, f424, f481])).
fof(f661, plain, (~ p101(sK91) | ~ p401(sK91)), inference(resolution, [], [f605, f527])).
fof(f605, plain, ! [X0] : (~ sP40(X0) | ~ p101(X0) | ~ p401(X0)), inference(resolution, [], [f301, f227])).
fof(f301, plain, ! [X0, X3] : (~ r1(X0, X3) | ~ p401(X3) | ~ p101(X3) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f14802, plain, (~ spl92_25 | ~ spl92_5), inference(avatar_split_clause, [], [f660, f424, f502])).
fof(f660, plain, (~ p101(sK91) | ~ p501(sK91)), inference(resolution, [], [f604, f527])).
fof(f604, plain, ! [X0] : (~ sP40(X0) | ~ p101(X0) | ~ p501(X0)), inference(resolution, [], [f300, f227])).
fof(f300, plain, ! [X4, X0] : (~ r1(X0, X4) | ~ p501(X4) | ~ p101(X4) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f14747, plain, (~ spl92_26 | ~ spl92_21), inference(avatar_split_clause, [], [f608, f486, f507])).
fof(f608, plain, (~ p505(sK91) | ~ p605(sK91)), inference(resolution, [], [f573, f527])).
fof(f573, plain, ! [X0] : (~ sP40(X0) | ~ p505(X0) | ~ p605(X0)), inference(resolution, [], [f229, f227])).
fof(f229, plain, ! [X35, X0] : (~ r1(X0, X35) | ~ p605(X35) | ~ p505(X35) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f14662, plain, (spl92_1170 | ~ spl92_21 | ~ spl92_6), inference(avatar_split_clause, [], [f14657, f429, f486, f14659])).
fof(f14657, plain, (~ p505(sK91) | p205(sK65(sK91)) | ~ spl92_6), inference(subsumption_resolution, [], [f3996, f534])).
fof(f3996, plain, (~ p505(sK91) | ~ sP15(sK91) | p205(sK65(sK91)) | ~ spl92_6), inference(resolution, [], [f2465, f430])).
fof(f2465, plain, ! [X0] : (r1(X0, sK65(X0)) | ~ p505(X0) | ~ sP15(X0)), inference(resolution, [], [f352, f227])).
fof(f352, plain, ! [X0, X1] : (~ r1(X0, X1) | ~ p505(X1) | r1(X1, sK65(X1)) | ~ sP15(X0)), inference(cnf_transformation, [], [f154])).
fof(f14614, plain, (~ spl92_21 | spl92_1162 | ~ spl92_11), inference(avatar_split_clause, [], [f14609, f447, f14611, f486])).
fof(f14609, plain, (p305(sK67(sK91)) | ~ p505(sK91) | ~ spl92_11), inference(subsumption_resolution, [], [f8549, f531])).
fof(f8549, plain, (p305(sK67(sK91)) | ~ p505(sK91) | ~ sP13(sK91) | ~ spl92_11), inference(resolution, [], [f448, f2497])).
fof(f14546, plain, (~ spl92_22 | ~ spl92_17), inference(avatar_split_clause, [], [f618, f469, f490])).
fof(f618, plain, (~ p404(sK91) | ~ p504(sK91)), inference(resolution, [], [f576, f527])).
fof(f576, plain, ! [X0] : (~ sP40(X0) | ~ p404(X0) | ~ p504(X0)), inference(resolution, [], [f246, f227])).
fof(f246, plain, ! [X0, X32] : (~ r1(X0, X32) | ~ p504(X32) | ~ p404(X32) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f12871, plain, (spl92_184 | spl92_185 | ~ spl92_11), inference(avatar_split_clause, [], [f12870, f447, f6449, f6443])).
fof(f12870, plain, (p305(sK86(sK91)) | r1(sK91, sK85(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f8560, f536])).
fof(f8560, plain, (p305(sK86(sK91)) | r1(sK91, sK85(sK91)) | ~ sP2(sK91) | ~ spl92_11), inference(resolution, [], [f448, f2673])).
fof(f2673, plain, ! [X0] : (r1(X0, sK86(X0)) | r1(X0, sK85(X0)) | ~ sP2(X0)), inference(resolution, [], [f392, f227])).
fof(f392, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK86(X1)) | r1(X1, sK85(X1)) | ~ sP2(X0)), inference(cnf_transformation, [], [f214])).
fof(f11388, plain, (spl92_152 | spl92_153 | ~ spl92_11), inference(avatar_split_clause, [], [f11387, f447, f5957, f5951])).
fof(f11387, plain, (p305(sK82(sK91)) | r1(sK91, sK81(sK91)) | ~ spl92_11), inference(subsumption_resolution, [], [f8558, f540])).
fof(f8558, plain, (p305(sK82(sK91)) | r1(sK91, sK81(sK91)) | ~ sP4(sK91) | ~ spl92_11), inference(resolution, [], [f448, f2641])).
fof(f2641, plain, ! [X0] : (r1(X0, sK82(X0)) | r1(X0, sK81(X0)) | ~ sP4(X0)), inference(resolution, [], [f384, f227])).
fof(f384, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK82(X1)) | r1(X1, sK81(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f204])).
fof(f6750, plain, (spl92_195 | spl92_196 | ~ spl92_16), inference(avatar_split_clause, [], [f6741, f466, f6747, f6743])).
fof(f6741, plain, (r1(sK91, sK89(sK91)) | p405(sK90(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f6659, f532])).
fof(f6659, plain, (r1(sK91, sK89(sK91)) | ~ sP0(sK91) | p405(sK90(sK91)) | ~ spl92_16), inference(resolution, [], [f2705, f467])).
fof(f2705, plain, ! [X0] : (r1(X0, sK90(X0)) | r1(X0, sK89(X0)) | ~ sP0(X0)), inference(resolution, [], [f400, f227])).
fof(f400, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK90(X1)) | r1(X1, sK89(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f224])).
fof(f6597, plain, (spl92_189 | spl92_190 | ~ spl92_16), inference(avatar_split_clause, [], [f6588, f466, f6594, f6590])).
fof(f6588, plain, (r1(sK91, sK87(sK91)) | p405(sK88(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f6506, f535])).
fof(f6506, plain, (r1(sK91, sK87(sK91)) | ~ sP1(sK91) | p405(sK88(sK91)) | ~ spl92_16), inference(resolution, [], [f2689, f467])).
fof(f2689, plain, ! [X0] : (r1(X0, sK88(X0)) | r1(X0, sK87(X0)) | ~ sP1(X0)), inference(resolution, [], [f396, f227])).
fof(f396, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK88(X1)) | r1(X1, sK87(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f219])).
fof(f6105, plain, (spl92_157 | spl92_158 | ~ spl92_16), inference(avatar_split_clause, [], [f6096, f466, f6102, f6098])).
fof(f6096, plain, (r1(sK91, sK83(sK91)) | p405(sK84(sK91)) | ~ spl92_16), inference(subsumption_resolution, [], [f6014, f539])).
fof(f6014, plain, (r1(sK91, sK83(sK91)) | ~ sP3(sK91) | p405(sK84(sK91)) | ~ spl92_16), inference(resolution, [], [f2657, f467])).
fof(f2657, plain, ! [X0] : (r1(X0, sK84(X0)) | r1(X0, sK83(X0)) | ~ sP3(X0)), inference(resolution, [], [f388, f227])).
fof(f388, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK84(X1)) | r1(X1, sK83(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f209])).
fof(f5813, plain, (spl92_148 | spl92_146 | ~ spl92_6), inference(avatar_split_clause, [], [f5808, f429, f5798, f5810])).
fof(f5808, plain, (r1(sK91, sK79(sK91)) | p205(sK80(sK91)) | ~ spl92_6), inference(subsumption_resolution, [], [f5712, f541])).
fof(f5712, plain, (r1(sK91, sK79(sK91)) | ~ sP5(sK91) | p205(sK80(sK91)) | ~ spl92_6), inference(resolution, [], [f2625, f430])).
fof(f2625, plain, ! [X0] : (r1(X0, sK80(X0)) | r1(X0, sK79(X0)) | ~ sP5(X0)), inference(resolution, [], [f380, f227])).
fof(f380, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK80(X1)) | r1(X1, sK79(X1)) | ~ sP5(X0)), inference(cnf_transformation, [], [f199])).
fof(f659, plain, (~ spl92_30 | ~ spl92_5), inference(avatar_split_clause, [], [f658, f424, f523])).
fof(f523, plain, (spl92_30 <=> p601(sK91)), introduced(avatar_definition, [new_symbols(naming, [spl92_30])])).
fof(f658, plain, (~ p101(sK91) | ~ p601(sK91)), inference(resolution, [], [f603, f527])).
fof(f603, plain, ! [X0] : (~ sP40(X0) | ~ p101(X0) | ~ p601(X0)), inference(resolution, [], [f299, f227])).
fof(f299, plain, ! [X0, X5] : (~ r1(X0, X5) | ~ p601(X5) | ~ p101(X5) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f654, plain, (~ spl92_30 | ~ spl92_10), inference(avatar_split_clause, [], [f653, f442, f523])).
fof(f653, plain, (~ p201(sK91) | ~ p601(sK91)), inference(resolution, [], [f599, f527])).
fof(f599, plain, ! [X0] : (~ sP40(X0) | ~ p201(X0) | ~ p601(X0)), inference(resolution, [], [f295, f227])).
fof(f295, plain, ! [X0, X9] : (~ r1(X0, X9) | ~ p601(X9) | ~ p201(X9) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f650, plain, (~ spl92_30 | ~ spl92_15), inference(avatar_split_clause, [], [f649, f461, f523])).
fof(f649, plain, (~ p301(sK91) | ~ p601(sK91)), inference(resolution, [], [f596, f527])).
fof(f596, plain, ! [X0] : (~ sP40(X0) | ~ p301(X0) | ~ p601(X0)), inference(resolution, [], [f292, f227])).
fof(f292, plain, ! [X0, X12] : (~ r1(X0, X12) | ~ p601(X12) | ~ p301(X12) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f647, plain, (~ spl92_30 | ~ spl92_20), inference(avatar_split_clause, [], [f646, f481, f523])).
fof(f646, plain, (~ p401(sK91) | ~ p601(sK91)), inference(resolution, [], [f594, f527])).
fof(f594, plain, ! [X0] : (~ sP40(X0) | ~ p401(X0) | ~ p601(X0)), inference(resolution, [], [f290, f227])).
fof(f290, plain, ! [X0, X14] : (~ r1(X0, X14) | ~ p601(X14) | ~ p401(X14) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f645, plain, (~ spl92_30 | ~ spl92_25), inference(avatar_split_clause, [], [f644, f502, f523])).
fof(f644, plain, (~ p501(sK91) | ~ p601(sK91)), inference(resolution, [], [f593, f527])).
fof(f593, plain, ! [X0] : (~ sP40(X0) | ~ p501(X0) | ~ p601(X0)), inference(resolution, [], [f289, f227])).
fof(f289, plain, ! [X0, X15] : (~ r1(X0, X15) | ~ p601(X15) | ~ p501(X15) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f640, plain, (~ spl92_29 | ~ spl92_9), inference(avatar_split_clause, [], [f639, f438, f519])).
fof(f639, plain, (~ p202(sK91) | ~ p602(sK91)), inference(resolution, [], [f589, f527])).
fof(f589, plain, ! [X0] : (~ sP40(X0) | ~ p202(X0) | ~ p602(X0)), inference(resolution, [], [f280, f227])).
fof(f280, plain, ! [X0, X19] : (~ r1(X0, X19) | ~ p602(X19) | ~ p202(X19) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f636, plain, (~ spl92_29 | ~ spl92_14), inference(avatar_split_clause, [], [f635, f457, f519])).
fof(f635, plain, (~ p302(sK91) | ~ p602(sK91)), inference(resolution, [], [f586, f527])).
fof(f586, plain, ! [X0] : (~ sP40(X0) | ~ p302(X0) | ~ p602(X0)), inference(resolution, [], [f277, f227])).
fof(f277, plain, ! [X0, X22] : (~ r1(X0, X22) | ~ p602(X22) | ~ p302(X22) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f633, plain, (~ spl92_29 | ~ spl92_19), inference(avatar_split_clause, [], [f632, f477, f519])).
fof(f632, plain, (~ p402(sK91) | ~ p602(sK91)), inference(resolution, [], [f584, f527])).
fof(f584, plain, ! [X0] : (~ sP40(X0) | ~ p402(X0) | ~ p602(X0)), inference(resolution, [], [f275, f227])).
fof(f275, plain, ! [X24, X0] : (~ r1(X0, X24) | ~ p602(X24) | ~ p402(X24) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f631, plain, (~ spl92_29 | ~ spl92_24), inference(avatar_split_clause, [], [f630, f498, f519])).
fof(f630, plain, (~ p502(sK91) | ~ p602(sK91)), inference(resolution, [], [f583, f527])).
fof(f583, plain, ! [X0] : (~ sP40(X0) | ~ p502(X0) | ~ p602(X0)), inference(resolution, [], [f274, f227])).
fof(f274, plain, ! [X0, X25] : (~ r1(X0, X25) | ~ p602(X25) | ~ p502(X25) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f627, plain, (~ spl92_28 | ~ spl92_13), inference(avatar_split_clause, [], [f626, f453, f515])).
fof(f626, plain, (~ p303(sK91) | ~ p603(sK91)), inference(resolution, [], [f580, f527])).
fof(f580, plain, ! [X0] : (~ sP40(X0) | ~ p303(X0) | ~ p603(X0)), inference(resolution, [], [f262, f227])).
fof(f262, plain, ! [X28, X0] : (~ r1(X0, X28) | ~ p603(X28) | ~ p303(X28) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f624, plain, (~ spl92_28 | ~ spl92_18), inference(avatar_split_clause, [], [f623, f473, f515])).
fof(f623, plain, (~ p403(sK91) | ~ p603(sK91)), inference(resolution, [], [f578, f527])).
fof(f578, plain, ! [X0] : (~ sP40(X0) | ~ p403(X0) | ~ p603(X0)), inference(resolution, [], [f260, f227])).
fof(f260, plain, ! [X30, X0] : (~ r1(X0, X30) | ~ p603(X30) | ~ p403(X30) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f622, plain, (~ spl92_28 | ~ spl92_23), inference(avatar_split_clause, [], [f621, f494, f515])).
fof(f621, plain, (~ p503(sK91) | ~ p603(sK91)), inference(resolution, [], [f577, f527])).
fof(f577, plain, ! [X0] : (~ sP40(X0) | ~ p503(X0) | ~ p603(X0)), inference(resolution, [], [f259, f227])).
fof(f259, plain, ! [X0, X31] : (~ r1(X0, X31) | ~ p603(X31) | ~ p503(X31) | ~ sP40(X0)), inference(cnf_transformation, [], [f54])).
fof(f526, plain, (spl92_26 | spl92_27 | spl92_28 | spl92_29 | spl92_30), inference(avatar_split_clause, [], [f405, f523, f519, f515, f511, f507])).
fof(f405, plain, (p601(sK91) | p602(sK91) | p603(sK91) | p604(sK91) | p605(sK91)), inference(cnf_transformation, [], [f226])).
fof(f505, plain, (spl92_21 | spl92_22 | spl92_23 | spl92_24 | spl92_25), inference(avatar_split_clause, [], [f406, f502, f498, f494, f490, f486])).
fof(f406, plain, (p501(sK91) | p502(sK91) | p503(sK91) | p504(sK91) | p505(sK91)), inference(cnf_transformation, [], [f226])).
fof(f484, plain, (spl92_16 | spl92_17 | spl92_18 | spl92_19 | spl92_20), inference(avatar_split_clause, [], [f407, f481, f477, f473, f469, f466])).
fof(f407, plain, ! [X10] : (p401(sK91) | p402(sK91) | p403(sK91) | p404(sK91) | p405(X10) | ~ r1(sK91, X10)), inference(cnf_transformation, [], [f226])).
fof(f464, plain, (spl92_11 | spl92_12 | spl92_13 | spl92_14 | spl92_15), inference(avatar_split_clause, [], [f408, f461, f457, f453, f450, f447])).
fof(f408, plain, ! [X8, X9] : (p301(sK91) | p302(sK91) | p303(sK91) | p304(X8) | ~ r1(sK91, X8) | p305(X9) | ~ r1(sK91, X9)), inference(cnf_transformation, [], [f226])).
fof(f445, plain, (spl92_6 | spl92_7 | spl92_8 | spl92_9 | spl92_10), inference(avatar_split_clause, [], [f409, f442, f438, f435, f432, f429])).
fof(f409, plain, ! [X6, X7, X5] : (p201(sK91) | p202(sK91) | p203(X5) | ~ r1(sK91, X5) | p204(X6) | ~ r1(sK91, X6) | p205(X7) | ~ r1(sK91, X7)), inference(cnf_transformation, [], [f226])).
fof(f427, plain, (spl92_1 | spl92_2 | spl92_3 | spl92_4 | spl92_5), inference(avatar_split_clause, [], [f410, f424, f421, f418, f415, f412])).
fof(f410, plain, ! [X4, X2, X3, X1] : (p101(sK91) | p102(X1) | ~ r1(sK91, X1) | p103(X2) | ~ r1(sK91, X2) | p104(X3) | ~ r1(sK91, X3) | p105(X4) | ~ r1(sK91, X4)), inference(cnf_transformation, [], [f226])).