fof(f9665, plain, $false, inference(avatar_sat_refutation, [], [f135, f139, f147, f152, f157, f974, f1008, f2108, f2177, f3680, f5476, f6173, f7081, f7135, f7232, f7235, f7258, f7374, f8070, f8076, f8079, f9664])).
fof(f9664, plain, (spl36_109 | ~ spl36_108 | ~ spl36_153 | ~ spl36_246 | spl36_287), inference(avatar_split_clause, [], [f9663, f6164, f5473, f3677, f2102, f2106])).
fof(f2106, plain, (spl36_109 <=> ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2))), introduced(avatar_definition, [new_symbols(naming, [spl36_109])])).
fof(f2102, plain, (spl36_108 <=> p2(sK22(sK25(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl36_108])])).
fof(f3677, plain, (spl36_153 <=> r1(sK25(sK32), sK22(sK25(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl36_153])])).
fof(f5473, plain, (spl36_246 <=> r1(sK22(sK25(sK32)), sK23(sK25(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl36_246])])).
fof(f6164, plain, (spl36_287 <=> p2(sK25(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl36_287])])).
fof(f9663, plain, (! [X0] : (~ r1(X0, sK32) | ~ sP0(X0)) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246 | spl36_287)), inference(subsumption_resolution, [], [f9662, f108])).
fof(f108, plain, r1(sK24, sK32), inference(cnf_transformation, [], [f63])).
fof(f63, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK25(X1), X3)) & ~ p2(sK25(X1)) & r1(X1, sK25(X1))) | p2(X1) | ~ r1(sK24, X1)) & ((! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(sK26, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(sK26, X9)) & ~ p2(sK26)) | sP1(sK26)) & r1(sK24, sK26)) | sP5(sK24)) & ! [X11] : ((p1(sK27(X11)) & (~ p1(sK28(X11)) & r1(sK27(X11), sK28(X11))) & r1(X11, sK27(X11))) | p1(X11) | ~ r1(sK24, X11)) & (~ p1(sK29) & r1(sK24, sK29)) & ! [X15] : ((p2(sK30(X15)) & (~ p2(sK31(X15)) & r1(sK30(X15), sK31(X15))) & r1(X15, sK30(X15))) | p2(X15) | ~ r1(sK24, X15)) & (~ p2(sK32) & r1(sK24, sK32)) & ! [X19] : ((p3(sK33(X19)) & (~ p3(sK34(X19)) & r1(sK33(X19), sK34(X19))) & r1(X19, sK33(X19))) | p3(X19) | ~ r1(sK24, X19)) & (~ p3(sK35) & r1(sK24, sK35))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25, sK26, sK27, sK28, sK29, sK30, sK31, sK32, sK33, sK34, sK35])], [f50, f62, f61, f60, f59, f58, f57, f56, f55, f54, f53, f52, f51])).
fof(f51, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP1(X5)) & r1(X0, X5)) | sP5(X0)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(X0, X11)) & ? [X14] : (~ p1(X14) & r1(X0, X14)) & ! [X15] : (? [X16] : (p2(X16) & ? [X17] : (~ p2(X17) & r1(X16, X17)) & r1(X15, X16)) | p2(X15) | ~ r1(X0, X15)) & ? [X18] : (~ p2(X18) & r1(X0, X18)) & ! [X19] : (? [X20] : (p3(X20) & ? [X21] : (~ p3(X21) & r1(X20, X21)) & r1(X19, X20)) | p3(X19) | ~ r1(X0, X19)) & ? [X22] : (~ p3(X22) & r1(X0, X22))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK24, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP1(X5)) & r1(sK24, X5)) | sP5(sK24)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(sK24, X11)) & ? [X14] : (~ p1(X14) & r1(sK24, X14)) & ! [X15] : (? [X16] : (p2(X16) & ? [X17] : (~ p2(X17) & r1(X16, X17)) & r1(X15, X16)) | p2(X15) | ~ r1(sK24, X15)) & ? [X18] : (~ p2(X18) & r1(sK24, X18)) & ! [X19] : (? [X20] : (p3(X20) & ? [X21] : (~ p3(X21) & r1(X20, X21)) & r1(X19, X20)) | p3(X19) | ~ r1(sK24, X19)) & ? [X22] : (~ p3(X22) & r1(sK24, X22)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK25(X1), X3)) & ~ p2(sK25(X1)) & r1(X1, sK25(X1)))), introduced(choice_axiom, [])).
fof(f53, plain, (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP1(X5)) & r1(sK24, X5)) => (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(sK26, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(sK26, X9)) & ~ p2(sK26)) | sP1(sK26)) & r1(sK24, sK26))), introduced(choice_axiom, [])).
fof(f54, plain, ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) => (p1(sK27(X11)) & ? [X13] : (~ p1(X13) & r1(sK27(X11), X13)) & r1(X11, sK27(X11)))), introduced(choice_axiom, [])).
fof(f55, plain, ! [X11] : (? [X13] : (~ p1(X13) & r1(sK27(X11), X13)) => (~ p1(sK28(X11)) & r1(sK27(X11), sK28(X11)))), introduced(choice_axiom, [])).
fof(f56, plain, (? [X14] : (~ p1(X14) & r1(sK24, X14)) => (~ p1(sK29) & r1(sK24, sK29))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X15] : (? [X16] : (p2(X16) & ? [X17] : (~ p2(X17) & r1(X16, X17)) & r1(X15, X16)) => (p2(sK30(X15)) & ? [X17] : (~ p2(X17) & r1(sK30(X15), X17)) & r1(X15, sK30(X15)))), introduced(choice_axiom, [])).
fof(f58, plain, ! [X15] : (? [X17] : (~ p2(X17) & r1(sK30(X15), X17)) => (~ p2(sK31(X15)) & r1(sK30(X15), sK31(X15)))), introduced(choice_axiom, [])).
fof(f59, plain, (? [X18] : (~ p2(X18) & r1(sK24, X18)) => (~ p2(sK32) & r1(sK24, sK32))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X19] : (? [X20] : (p3(X20) & ? [X21] : (~ p3(X21) & r1(X20, X21)) & r1(X19, X20)) => (p3(sK33(X19)) & ? [X21] : (~ p3(X21) & r1(sK33(X19), X21)) & r1(X19, sK33(X19)))), introduced(choice_axiom, [])).
fof(f61, plain, ! [X19] : (? [X21] : (~ p3(X21) & r1(sK33(X19), X21)) => (~ p3(sK34(X19)) & r1(sK33(X19), sK34(X19)))), introduced(choice_axiom, [])).
fof(f62, plain, (? [X22] : (~ p3(X22) & r1(sK24, X22)) => (~ p3(sK35) & r1(sK24, sK35))), introduced(choice_axiom, [])).
fof(f50, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP1(X5)) & r1(X0, X5)) | sP5(X0)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(X0, X11)) & ? [X14] : (~ p1(X14) & r1(X0, X14)) & ! [X15] : (? [X16] : (p2(X16) & ? [X17] : (~ p2(X17) & r1(X16, X17)) & r1(X15, X16)) | p2(X15) | ~ r1(X0, X15)) & ? [X18] : (~ p2(X18) & r1(X0, X18)) & ! [X19] : (? [X20] : (p3(X20) & ? [X21] : (~ p3(X21) & r1(X20, X21)) & r1(X19, X20)) | p3(X19) | ~ r1(X0, X19)) & ? [X22] : (~ p3(X22) & r1(X0, X22))), inference(rectify, [], [f13])).
fof(f13, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP3(X6) | sP4(X6) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | sP1(X5)) & r1(X0, X5)) | sP5(X0)) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & ! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51)) & ! [X52] : (? [X53] : (p3(X53) & ? [X54] : (~ p3(X54) & r1(X53, X54)) & r1(X52, X53)) | p3(X52) | ~ r1(X0, X52)) & ? [X55] : (~ p3(X55) & r1(X0, X55))), inference(definition_folding, [], [f6, e12, e11, e10, e9, e8, e7])).
fof(f7, plain, ! [X0] : (! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)) | ~ sP0(X0)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X0] : (sP0(X0) <=> ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X5] : ((! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31))) | ~ sP1(X5)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X5] : (sP1(X5) <=> (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ! [X16] : (! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)) | ~ sP2(X16)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X16] : (sP2(X16) <=> ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ! [X6] : ((! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12))) | ~ sP3(X6)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X6] : (sP3(X6) <=> (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f11, plain, ! [X6] : (! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | sP2(X16))) | ~ r1(X6, X16)) | ~ sP4(X6)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X6] : (sP4(X6) <=> ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | sP2(X16))) | ~ r1(X6, X16))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f12, plain, ! [X0] : (((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | sP0(X0))) | ~ sP5(X0)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X0] : (sP5(X0) <=> ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | sP0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f6, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12))) | ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))) & r1(X0, X5)) | ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40))))) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & ! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51)) & ! [X52] : (? [X53] : (p3(X53) & ? [X54] : (~ p3(X54) & r1(X53, X54)) & r1(X52, X53)) | p3(X52) | ~ r1(X0, X52)) & ? [X55] : (~ p3(X55) & r1(X0, X55))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : (((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12)))) | ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))) & r1(X0, X5)) | ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & ! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51)) & ! [X52] : (? [X53] : (p3(X53) & ? [X54] : (~ p3(X54) & r1(X53, X54)) & r1(X52, X53)) | p3(X52) | ~ r1(X0, X52)) & ? [X55] : (~ p3(X55) & r1(X0, X55))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ ((~ ! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p2(X6)) & (~ ! [X9] : (~ ! [X10] : (~ p2(X10) | ! [X11] : (p2(X11) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) | ! [X12] : (! [X13] : (~ ! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | p2(X13) | ~ r1(X12, X13)) | ~ r1(X6, X12)))) | ! [X16] : (((~ ! [X17] : (~ p2(X17) | ! [X18] : (p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | p2(X16)) & (~ ! [X19] : (~ ! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | p2(X19) | ~ r1(X16, X19)) | ! [X22] : (! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p2(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) | ((~ ! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) | p2(X5)) & (~ ! [X28] : (~ ! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) | ! [X31] : (! [X32] : (~ ! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) | p2(X32) | ~ r1(X31, X32)) | ~ r1(X5, X31)))) | ~ r1(X0, X5)) | ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X0, X35)) | p2(X0)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X0, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) | ~ ! [X44] : (~ ! [X45] : (~ p1(X45) | ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) | ! [X47] : (p1(X47) | ~ r1(X0, X47)) | ~ ! [X48] : (~ ! [X49] : (~ p2(X49) | ! [X50] : (p2(X50) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) | ! [X51] : (p2(X51) | ~ r1(X0, X51)) | ~ ! [X52] : (~ ! [X53] : (~ p3(X53) | ! [X54] : (p3(X54) | ~ r1(X53, X54)) | ~ r1(X52, X53)) | p3(X52) | ~ r1(X0, X52)) | ! [X55] : (p3(X55) | ~ r1(X0, X55))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ ((~ ! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p2(X6)) & (~ ! [X9] : (~ ! [X10] : (~ p2(X10) | ! [X11] : (p2(X11) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) | ! [X12] : (! [X13] : (~ ! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | p2(X13) | ~ r1(X12, X13)) | ~ r1(X6, X12)))) | ! [X16] : (((~ ! [X17] : (~ p2(X17) | ! [X18] : (p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | p2(X16)) & (~ ! [X19] : (~ ! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | p2(X19) | ~ r1(X16, X19)) | ! [X22] : (! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p2(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) | ((~ ! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) | p2(X5)) & (~ ! [X28] : (~ ! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) | ! [X31] : (! [X32] : (~ ! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) | p2(X32) | ~ r1(X31, X32)) | ~ r1(X5, X31)))) | ~ r1(X0, X5)) | ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X0, X35)) | p2(X0)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X0, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) | ~ ! [X44] : (~ ! [X45] : (~ p1(X45) | ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) | ! [X47] : (p1(X47) | ~ r1(X0, X47)) | ~ ! [X48] : (~ ! [X49] : (~ p2(X49) | ! [X50] : (p2(X50) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) | ! [X51] : (p2(X51) | ~ r1(X0, X51)) | ~ ! [X52] : (~ ! [X53] : (~ p3(X53) | ! [X54] : (p3(X54) | ~ r1(X53, X54)) | ~ r1(X52, X53)) | p3(X52) | ~ r1(X0, X52)) | ! [X55] : (p3(X55) | ~ r1(X0, X55))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL642+1.001.p', main)).
fof(f9662, plain, (! [X0] : (~ r1(X0, sK32) | ~ sP0(X0) | ~ r1(sK24, sK32)) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246 | spl36_287)), inference(subsumption_resolution, [], [f9661, f109])).
fof(f109, plain, ~ p2(sK32), inference(cnf_transformation, [], [f63])).
fof(f9661, plain, (! [X0] : (~ r1(X0, sK32) | ~ sP0(X0) | p2(sK32) | ~ r1(sK24, sK32)) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246 | spl36_287)), inference(resolution, [], [f6525, f125])).
fof(f125, plain, ! [X1] : (r1(X1, sK25(X1)) | p2(X1) | ~ r1(sK24, X1)), inference(cnf_transformation, [], [f63])).
fof(f6525, plain, (! [X0, X1] : (~ r1(X0, sK25(sK32)) | ~ r1(X1, X0) | ~ sP0(X1)) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246 | spl36_287)), inference(subsumption_resolution, [], [f6524, f6165])).
fof(f6165, plain, (~ p2(sK25(sK32)) | spl36_287), inference(avatar_component_clause, [], [f6164])).
fof(f6524, plain, (! [X0, X1] : (p2(sK25(sK32)) | ~ r1(X0, sK25(sK32)) | ~ r1(X1, X0) | ~ sP0(X1)) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246)), inference(resolution, [], [f6222, f100])).
fof(f100, plain, ! [X2, X0, X1] : (~ p2(sK23(X2)) | p2(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f49])).
fof(f49, plain, ! [X0] : (! [X1] : (! [X2] : ((p2(sK22(X2)) & (~ p2(sK23(X2)) & r1(sK22(X2), sK23(X2))) & r1(X2, sK22(X2))) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22, sK23])], [f46, f48, f47])).
fof(f47, plain, ! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) => (p2(sK22(X2)) & ? [X4] : (~ p2(X4) & r1(sK22(X2), X4)) & r1(X2, sK22(X2)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X2] : (? [X4] : (~ p2(X4) & r1(sK22(X2), X4)) => (~ p2(sK23(X2)) & r1(sK22(X2), sK23(X2)))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f45])).
fof(f45, plain, ! [X0] : (! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)) | ~ sP0(X0)), inference(nnf_transformation, [], [f7])).
fof(f6222, plain, (p2(sK23(sK25(sK32))) | (~ spl36_108 | ~ spl36_153 | ~ spl36_246)), inference(resolution, [], [f5475, f4765])).
fof(f4765, plain, (! [X0] : (~ r1(sK22(sK25(sK32)), X0) | p2(X0)) | (~ spl36_108 | ~ spl36_153)), inference(subsumption_resolution, [], [f4764, f108])).
fof(f4764, plain, (! [X0] : (p2(X0) | ~ r1(sK22(sK25(sK32)), X0) | ~ r1(sK24, sK32)) | (~ spl36_108 | ~ spl36_153)), inference(subsumption_resolution, [], [f4763, f109])).
fof(f4763, plain, (! [X0] : (p2(X0) | ~ r1(sK22(sK25(sK32)), X0) | p2(sK32) | ~ r1(sK24, sK32)) | (~ spl36_108 | ~ spl36_153)), inference(subsumption_resolution, [], [f4732, f2104])).
fof(f2104, plain, (p2(sK22(sK25(sK32))) | ~ spl36_108), inference(avatar_component_clause, [], [f2102])).
fof(f4732, plain, (! [X0] : (p2(X0) | ~ r1(sK22(sK25(sK32)), X0) | ~ p2(sK22(sK25(sK32))) | p2(sK32) | ~ r1(sK24, sK32)) | ~ spl36_153), inference(resolution, [], [f3679, f127])).
fof(f127, plain, ! [X4, X3, X1] : (~ r1(sK25(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK24, X1)), inference(cnf_transformation, [], [f63])).
fof(f3679, plain, (r1(sK25(sK32), sK22(sK25(sK32))) | ~ spl36_153), inference(avatar_component_clause, [], [f3677])).
fof(f5475, plain, (r1(sK22(sK25(sK32)), sK23(sK25(sK32))) | ~ spl36_246), inference(avatar_component_clause, [], [f5473])).
fof(f8079, plain, (~ spl36_4 | ~ spl36_53), inference(avatar_contradiction_clause, [], [f8078])).
fof(f8078, plain, ($false | (~ spl36_4 | ~ spl36_53)), inference(subsumption_resolution, [], [f8077, f143])).
fof(f143, plain, (sP1(sK26) | ~ spl36_4), inference(avatar_component_clause, [], [f141])).
fof(f141, plain, (spl36_4 <=> sP1(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl36_4])])).
fof(f8077, plain, (~ sP1(sK26) | ~ spl36_53), inference(resolution, [], [f423, f1724])).
fof(f1724, plain, ! [X0] : (~ sP3(sK20(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f1723])).
fof(f1723, plain, ! [X0] : (~ sP3(sK20(X0)) | ~ sP1(X0) | ~ sP3(sK20(X0)) | ~ sP1(X0)), inference(resolution, [], [f1706, f91])).
fof(f91, plain, ! [X0] : (r1(sK20(X0), sK21(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f44, plain, ! [X0] : ((! [X1] : ((p2(sK18(X1)) & (~ p2(sK19(X1)) & r1(sK18(X1), sK19(X1))) & r1(X1, sK18(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK21(X0), X6)) & ~ p2(sK21(X0)) & r1(sK20(X0), sK21(X0))) & r1(X0, sK20(X0)))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18, sK19, sK20, sK21])], [f39, f43, f42, f41, f40])).
fof(f40, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK18(X1)) & ? [X3] : (~ p2(X3) & r1(sK18(X1), X3)) & r1(X1, sK18(X1)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK18(X1), X3)) => (~ p2(sK19(X1)) & r1(sK18(X1), sK19(X1)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK20(X0), X5)) & r1(X0, sK20(X0)))), introduced(choice_axiom, [])).
fof(f43, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK20(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK21(X0), X6)) & ~ p2(sK21(X0)) & r1(sK20(X0), sK21(X0)))), introduced(choice_axiom, [])).
fof(f39, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP1(X0)), inference(rectify, [], [f38])).
fof(f38, plain, ! [X5] : ((! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31))) | ~ sP1(X5)), inference(nnf_transformation, [], [f8])).
fof(f1706, plain, ! [X0, X1] : (~ r1(X1, sK21(X0)) | ~ sP3(sK20(X0)) | ~ sP1(X0) | ~ sP3(X1)), inference(subsumption_resolution, [], [f1705, f92])).
fof(f92, plain, ! [X0] : (~ p2(sK21(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f1705, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP3(sK20(X0)) | p2(sK21(X0)) | ~ r1(X1, sK21(X0)) | ~ sP3(X1)), inference(resolution, [], [f1687, f84])).
fof(f84, plain, ! [X0, X1] : (~ p2(sK13(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((! [X1] : ((p2(sK12(X1)) & (~ p2(sK13(X1)) & r1(sK12(X1), sK13(X1))) & r1(X1, sK12(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK15(X0), X6)) & ~ p2(sK15(X0)) & r1(sK14(X0), sK15(X0))) & r1(X0, sK14(X0)))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12, sK13, sK14, sK15])], [f27, f31, f30, f29, f28])).
fof(f28, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK12(X1)) & ? [X3] : (~ p2(X3) & r1(sK12(X1), X3)) & r1(X1, sK12(X1)))), introduced(choice_axiom, [])).
fof(f29, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK12(X1), X3)) => (~ p2(sK13(X1)) & r1(sK12(X1), sK13(X1)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK14(X0), X5)) & r1(X0, sK14(X0)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK14(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK15(X0), X6)) & ~ p2(sK15(X0)) & r1(sK14(X0), sK15(X0)))), introduced(choice_axiom, [])).
fof(f27, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP3(X0)), inference(rectify, [], [f26])).
fof(f26, plain, ! [X6] : ((! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12))) | ~ sP3(X6)), inference(nnf_transformation, [], [f10])).
fof(f1687, plain, ! [X0] : (p2(sK13(sK21(X0))) | ~ sP1(X0) | ~ sP3(sK20(X0))), inference(duplicate_literal_removal, [], [f1678])).
fof(f1678, plain, ! [X0] : (~ sP1(X0) | p2(sK13(sK21(X0))) | ~ sP3(sK20(X0)) | ~ sP3(sK20(X0)) | ~ sP1(X0)), inference(resolution, [], [f1120, f694])).
fof(f694, plain, ! [X10] : (r1(sK12(sK21(X10)), sK13(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f681, f92])).
fof(f681, plain, ! [X10] : (p2(sK21(X10)) | r1(sK12(sK21(X10)), sK13(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f83, f91])).
fof(f83, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK12(X1), sK13(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f32])).
fof(f1120, plain, ! [X0, X1] : (~ r1(sK12(sK21(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP3(sK20(X0))), inference(subsumption_resolution, [], [f1119, f528])).
fof(f528, plain, ! [X10] : (p2(sK12(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f515, f92])).
fof(f515, plain, ! [X10] : (p2(sK21(X10)) | p2(sK12(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f85, f91])).
fof(f85, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK12(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f32])).
fof(f1119, plain, ! [X0, X1] : (~ sP3(sK20(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK12(sK21(X0)), X1) | ~ p2(sK12(sK21(X0)))), inference(duplicate_literal_removal, [], [f1102])).
fof(f1102, plain, ! [X0, X1] : (~ sP3(sK20(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK12(sK21(X0)), X1) | ~ p2(sK12(sK21(X0))) | ~ sP1(X0)), inference(resolution, [], [f620, f93])).
fof(f93, plain, ! [X6, X0, X7] : (~ r1(sK21(X0), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f620, plain, ! [X10] : (r1(sK21(X10), sK12(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f607, f92])).
fof(f607, plain, ! [X10] : (p2(sK21(X10)) | r1(sK21(X10), sK12(sK21(X10))) | ~ sP3(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f82, f91])).
fof(f82, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK12(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f32])).
fof(f423, plain, (sP3(sK20(sK26)) | ~ spl36_53), inference(avatar_component_clause, [], [f421])).
fof(f421, plain, (spl36_53 <=> sP3(sK20(sK26))), introduced(avatar_definition, [new_symbols(naming, [spl36_53])])).
fof(f8076, plain, (spl36_53 | ~ spl36_55 | ~ spl36_3 | ~ spl36_4 | spl36_54), inference(avatar_split_clause, [], [f8075, f425, f141, f137, f429, f421])).
fof(f429, plain, (spl36_55 <=> p2(sK20(sK26))), introduced(avatar_definition, [new_symbols(naming, [spl36_55])])).
fof(f137, plain, (spl36_3 <=> ! [X6] : (~ p2(X6) | ~ r1(sK26, X6) | sP4(X6) | sP3(X6))), introduced(avatar_definition, [new_symbols(naming, [spl36_3])])).
fof(f425, plain, (spl36_54 <=> sP4(sK20(sK26))), introduced(avatar_definition, [new_symbols(naming, [spl36_54])])).
fof(f8075, plain, (~ p2(sK20(sK26)) | sP3(sK20(sK26)) | (~ spl36_3 | ~ spl36_4 | spl36_54)), inference(subsumption_resolution, [], [f8074, f143])).
fof(f8074, plain, (~ p2(sK20(sK26)) | sP3(sK20(sK26)) | ~ sP1(sK26) | (~ spl36_3 | spl36_54)), inference(subsumption_resolution, [], [f397, f426])).
fof(f426, plain, (~ sP4(sK20(sK26)) | spl36_54), inference(avatar_component_clause, [], [f425])).
fof(f397, plain, (~ p2(sK20(sK26)) | sP4(sK20(sK26)) | sP3(sK20(sK26)) | ~ sP1(sK26) | ~ spl36_3), inference(resolution, [], [f138, f90])).
fof(f90, plain, ! [X0] : (r1(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f138, plain, (! [X6] : (~ r1(sK26, X6) | ~ p2(X6) | sP4(X6) | sP3(X6)) | ~ spl36_3), inference(avatar_component_clause, [], [f137])).
fof(f8070, plain, (~ spl36_4 | spl36_55 | ~ spl36_403), inference(avatar_contradiction_clause, [], [f8069])).
fof(f8069, plain, ($false | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(subsumption_resolution, [], [f8068, f143])).
fof(f8068, plain, (~ sP1(sK26) | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(duplicate_literal_removal, [], [f8067])).
fof(f8067, plain, (~ sP1(sK26) | ~ sP1(sK26) | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(resolution, [], [f8066, f90])).
fof(f8066, plain, (! [X0] : (~ r1(X0, sK20(sK26)) | ~ sP1(X0)) | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(subsumption_resolution, [], [f8065, f431])).
fof(f431, plain, (~ p2(sK20(sK26)) | spl36_55), inference(avatar_component_clause, [], [f429])).
fof(f8065, plain, (! [X0] : (p2(sK20(sK26)) | ~ r1(X0, sK20(sK26)) | ~ sP1(X0)) | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(resolution, [], [f7996, f96])).
fof(f96, plain, ! [X0, X1] : (~ p2(sK19(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f7996, plain, (p2(sK19(sK20(sK26))) | (~ spl36_4 | spl36_55 | ~ spl36_403)), inference(subsumption_resolution, [], [f7995, f143])).
fof(f7995, plain, (p2(sK19(sK20(sK26))) | ~ sP1(sK26) | (spl36_55 | ~ spl36_403)), inference(subsumption_resolution, [], [f7986, f431])).
fof(f7986, plain, (p2(sK19(sK20(sK26))) | p2(sK20(sK26)) | ~ sP1(sK26) | ~ spl36_403), inference(resolution, [], [f7257, f714])).
fof(f714, plain, ! [X3] : (r1(sK18(sK20(X3)), sK19(sK20(X3))) | p2(sK20(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f699])).
fof(f699, plain, ! [X3] : (p2(sK20(X3)) | r1(sK18(sK20(X3)), sK19(sK20(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f95, f90])).
fof(f95, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK18(X1), sK19(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f7257, plain, (! [X3] : (~ r1(sK18(sK20(sK26)), X3) | p2(X3)) | ~ spl36_403), inference(avatar_component_clause, [], [f7256])).
fof(f7256, plain, (spl36_403 <=> ! [X3] : (p2(X3) | ~ r1(sK18(sK20(sK26)), X3))), introduced(avatar_definition, [new_symbols(naming, [spl36_403])])).
fof(f7374, plain, (~ spl36_4 | spl36_55 | spl36_402), inference(avatar_contradiction_clause, [], [f7373])).
fof(f7373, plain, ($false | (~ spl36_4 | spl36_55 | spl36_402)), inference(subsumption_resolution, [], [f7372, f143])).
fof(f7372, plain, (~ sP1(sK26) | (spl36_55 | spl36_402)), inference(subsumption_resolution, [], [f7371, f431])).
fof(f7371, plain, (p2(sK20(sK26)) | ~ sP1(sK26) | spl36_402), inference(resolution, [], [f7254, f548])).
fof(f548, plain, ! [X3] : (p2(sK18(sK20(X3))) | p2(sK20(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f533])).
fof(f533, plain, ! [X3] : (p2(sK20(X3)) | p2(sK18(sK20(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f97, f90])).
fof(f97, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK18(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f7254, plain, (~ p2(sK18(sK20(sK26))) | spl36_402), inference(avatar_component_clause, [], [f7252])).
fof(f7252, plain, (spl36_402 <=> p2(sK18(sK20(sK26)))), introduced(avatar_definition, [new_symbols(naming, [spl36_402])])).
fof(f7258, plain, (~ spl36_402 | spl36_403 | ~ spl36_4 | spl36_55 | ~ spl36_74), inference(avatar_split_clause, [], [f7250, f961, f429, f141, f7256, f7252])).
fof(f961, plain, (spl36_74 <=> ! [X7, X6] : (~ p2(X6) | p2(X7) | ~ r1(X6, X7) | ~ r1(sK20(sK26), X6))), introduced(avatar_definition, [new_symbols(naming, [spl36_74])])).
fof(f7250, plain, (! [X3] : (p2(X3) | ~ r1(sK18(sK20(sK26)), X3) | ~ p2(sK18(sK20(sK26)))) | (~ spl36_4 | spl36_55 | ~ spl36_74)), inference(subsumption_resolution, [], [f7249, f143])).
fof(f7249, plain, (! [X3] : (p2(X3) | ~ r1(sK18(sK20(sK26)), X3) | ~ p2(sK18(sK20(sK26))) | ~ sP1(sK26)) | (spl36_55 | ~ spl36_74)), inference(subsumption_resolution, [], [f7239, f431])).
fof(f7239, plain, (! [X3] : (p2(X3) | ~ r1(sK18(sK20(sK26)), X3) | ~ p2(sK18(sK20(sK26))) | p2(sK20(sK26)) | ~ sP1(sK26)) | ~ spl36_74), inference(resolution, [], [f962, f640])).
fof(f640, plain, ! [X3] : (r1(sK20(X3), sK18(sK20(X3))) | p2(sK20(X3)) | ~ sP1(X3)), inference(duplicate_literal_removal, [], [f625])).
fof(f625, plain, ! [X3] : (p2(sK20(X3)) | r1(sK20(X3), sK18(sK20(X3))) | ~ sP1(X3) | ~ sP1(X3)), inference(resolution, [], [f94, f90])).
fof(f94, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK18(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f44])).
fof(f962, plain, (! [X6, X7] : (~ r1(sK20(sK26), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6)) | ~ spl36_74), inference(avatar_component_clause, [], [f961])).
fof(f7235, plain, (~ spl36_4 | ~ spl36_54), inference(avatar_contradiction_clause, [], [f7234])).
fof(f7234, plain, ($false | (~ spl36_4 | ~ spl36_54)), inference(subsumption_resolution, [], [f7233, f143])).
fof(f7233, plain, (~ sP1(sK26) | ~ spl36_54), inference(resolution, [], [f427, f1566])).
fof(f1566, plain, ! [X0] : (~ sP4(sK20(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f1565])).
fof(f1565, plain, ! [X0] : (~ sP4(sK20(X0)) | ~ sP1(X0) | ~ sP4(sK20(X0)) | ~ sP1(X0)), inference(resolution, [], [f1548, f91])).
fof(f1548, plain, ! [X0, X1] : (~ r1(X1, sK21(X0)) | ~ sP4(sK20(X0)) | ~ sP1(X0) | ~ sP4(X1)), inference(subsumption_resolution, [], [f1547, f92])).
fof(f1547, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP4(sK20(X0)) | p2(sK21(X0)) | ~ r1(X1, sK21(X0)) | ~ sP4(X1)), inference(resolution, [], [f1529, f76])).
fof(f76, plain, ! [X0, X1] : (~ p2(sK10(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (! [X1] : ((((p2(sK9(X1)) & (~ p2(sK10(X1)) & r1(sK9(X1), sK10(X1))) & r1(X1, sK9(X1))) | p2(X1)) & ((! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK11(X1), X5)) & ~ p2(sK11(X1)) & r1(X1, sK11(X1))) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f21, f24, f23, f22])).
fof(f22, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK9(X1)) & ? [X3] : (~ p2(X3) & r1(sK9(X1), X3)) & r1(X1, sK9(X1)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK9(X1), X3)) => (~ p2(sK10(X1)) & r1(sK9(X1), sK10(X1)))), introduced(choice_axiom, [])).
fof(f24, plain, ! [X1] : (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) => (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK11(X1), X5)) & ~ p2(sK11(X1)) & r1(X1, sK11(X1)))), introduced(choice_axiom, [])).
fof(f21, plain, ! [X0] : (! [X1] : (((? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1)) & (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(rectify, [], [f20])).
fof(f20, plain, ! [X6] : (! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | sP2(X16))) | ~ r1(X6, X16)) | ~ sP4(X6)), inference(nnf_transformation, [], [f11])).
fof(f1529, plain, ! [X0] : (p2(sK10(sK21(X0))) | ~ sP1(X0) | ~ sP4(sK20(X0))), inference(duplicate_literal_removal, [], [f1520])).
fof(f1520, plain, ! [X0] : (~ sP1(X0) | p2(sK10(sK21(X0))) | ~ sP4(sK20(X0)) | ~ sP4(sK20(X0)) | ~ sP1(X0)), inference(resolution, [], [f1066, f669])).
fof(f669, plain, ! [X10] : (r1(sK9(sK21(X10)), sK10(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f657, f92])).
fof(f657, plain, ! [X10] : (p2(sK21(X10)) | r1(sK9(sK21(X10)), sK10(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f75, f91])).
fof(f75, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK9(X1), sK10(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f25])).
fof(f1066, plain, ! [X0, X1] : (~ r1(sK9(sK21(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP4(sK20(X0))), inference(subsumption_resolution, [], [f1065, f483])).
fof(f483, plain, ! [X10] : (p2(sK9(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f471, f92])).
fof(f471, plain, ! [X10] : (p2(sK21(X10)) | p2(sK9(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f77, f91])).
fof(f77, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK9(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f25])).
fof(f1065, plain, ! [X0, X1] : (~ sP4(sK20(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK9(sK21(X0)), X1) | ~ p2(sK9(sK21(X0)))), inference(duplicate_literal_removal, [], [f1048])).
fof(f1048, plain, ! [X0, X1] : (~ sP4(sK20(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK9(sK21(X0)), X1) | ~ p2(sK9(sK21(X0))) | ~ sP1(X0)), inference(resolution, [], [f595, f93])).
fof(f595, plain, ! [X10] : (r1(sK21(X10), sK9(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(subsumption_resolution, [], [f583, f92])).
fof(f583, plain, ! [X10] : (p2(sK21(X10)) | r1(sK21(X10), sK9(sK21(X10))) | ~ sP4(sK20(X10)) | ~ sP1(X10)), inference(resolution, [], [f74, f91])).
fof(f74, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK9(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f25])).
fof(f427, plain, (sP4(sK20(sK26)) | ~ spl36_54), inference(avatar_component_clause, [], [f425])).
fof(f7232, plain, (spl36_6 | ~ spl36_7 | ~ spl36_384), inference(avatar_contradiction_clause, [], [f7231])).
fof(f7231, plain, ($false | (spl36_6 | ~ spl36_7 | ~ spl36_384)), inference(subsumption_resolution, [], [f7230, f156])).
fof(f156, plain, (r1(sK24, sK26) | ~ spl36_7), inference(avatar_component_clause, [], [f154])).
fof(f154, plain, (spl36_7 <=> r1(sK24, sK26)), introduced(avatar_definition, [new_symbols(naming, [spl36_7])])).
fof(f7230, plain, (~ r1(sK24, sK26) | (spl36_6 | ~ spl36_7 | ~ spl36_384)), inference(subsumption_resolution, [], [f7229, f151])).
fof(f151, plain, (~ p2(sK26) | spl36_6), inference(avatar_component_clause, [], [f149])).
fof(f149, plain, (spl36_6 <=> p2(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl36_6])])).
fof(f7229, plain, (p2(sK26) | ~ r1(sK24, sK26) | (spl36_6 | ~ spl36_7 | ~ spl36_384)), inference(resolution, [], [f7167, f112])).
fof(f112, plain, ! [X15] : (~ p2(sK31(X15)) | p2(X15) | ~ r1(sK24, X15)), inference(cnf_transformation, [], [f63])).
fof(f7167, plain, (p2(sK31(sK26)) | (spl36_6 | ~ spl36_7 | ~ spl36_384)), inference(subsumption_resolution, [], [f7166, f156])).
fof(f7166, plain, (p2(sK31(sK26)) | ~ r1(sK24, sK26) | (spl36_6 | ~ spl36_384)), inference(subsumption_resolution, [], [f7156, f151])).
fof(f7156, plain, (p2(sK31(sK26)) | p2(sK26) | ~ r1(sK24, sK26) | ~ spl36_384), inference(resolution, [], [f7134, f111])).
fof(f111, plain, ! [X15] : (r1(sK30(X15), sK31(X15)) | p2(X15) | ~ r1(sK24, X15)), inference(cnf_transformation, [], [f63])).
fof(f7134, plain, (! [X6] : (~ r1(sK30(sK26), X6) | p2(X6)) | ~ spl36_384), inference(avatar_component_clause, [], [f7133])).
fof(f7133, plain, (spl36_384 <=> ! [X6] : (~ r1(sK30(sK26), X6) | p2(X6))), introduced(avatar_definition, [new_symbols(naming, [spl36_384])])).
fof(f7135, plain, (spl36_6 | spl36_384 | ~ spl36_5 | ~ spl36_7 | ~ spl36_37), inference(avatar_split_clause, [], [f7131, f320, f154, f145, f7133, f149])).
fof(f145, plain, (spl36_5 <=> ! [X9, X10] : (~ p2(X9) | ~ r1(sK26, X9) | ~ r1(X9, X10) | p2(X10))), introduced(avatar_definition, [new_symbols(naming, [spl36_5])])).
fof(f320, plain, (spl36_37 <=> p2(sK30(sK26))), introduced(avatar_definition, [new_symbols(naming, [spl36_37])])).
fof(f7131, plain, (! [X6] : (~ r1(sK30(sK26), X6) | p2(X6) | p2(sK26)) | (~ spl36_5 | ~ spl36_7 | ~ spl36_37)), inference(subsumption_resolution, [], [f7130, f156])).
fof(f7130, plain, (! [X6] : (~ r1(sK30(sK26), X6) | p2(X6) | p2(sK26) | ~ r1(sK24, sK26)) | (~ spl36_5 | ~ spl36_37)), inference(subsumption_resolution, [], [f1015, f322])).
fof(f322, plain, (p2(sK30(sK26)) | ~ spl36_37), inference(avatar_component_clause, [], [f320])).
fof(f1015, plain, (! [X6] : (~ p2(sK30(sK26)) | ~ r1(sK30(sK26), X6) | p2(X6) | p2(sK26) | ~ r1(sK24, sK26)) | ~ spl36_5), inference(resolution, [], [f146, f110])).
fof(f110, plain, ! [X15] : (r1(X15, sK30(X15)) | p2(X15) | ~ r1(sK24, X15)), inference(cnf_transformation, [], [f63])).
fof(f146, plain, (! [X10, X9] : (~ r1(sK26, X9) | ~ p2(X9) | ~ r1(X9, X10) | p2(X10)) | ~ spl36_5), inference(avatar_component_clause, [], [f145])).
fof(f7081, plain, (spl36_30 | ~ spl36_1), inference(avatar_split_clause, [], [f2913, f129, f279])).
fof(f279, plain, (spl36_30 <=> sP0(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl36_30])])).
fof(f129, plain, (spl36_1 <=> sP5(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl36_1])])).
fof(f2913, plain, (~ sP5(sK24) | sP0(sK24)), inference(duplicate_literal_removal, [], [f2912])).
fof(f2912, plain, (~ sP5(sK24) | sP0(sK24) | sP0(sK24) | ~ sP5(sK24)), inference(resolution, [], [f2911, f64])).
fof(f64, plain, ! [X0] : (r1(X0, sK8(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : ((((p2(sK6(X0)) & (~ p2(sK7(X0)) & r1(sK6(X0), sK7(X0))) & r1(X0, sK6(X0))) | p2(X0)) & ((! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK8(X0), X4)) & ~ p2(sK8(X0)) & r1(X0, sK8(X0))) | sP0(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7, sK8])], [f15, f18, f17, f16])).
fof(f16, plain, ! [X0] : (? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) => (p2(sK6(X0)) & ? [X2] : (~ p2(X2) & r1(sK6(X0), X2)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f17, plain, ! [X0] : (? [X2] : (~ p2(X2) & r1(sK6(X0), X2)) => (~ p2(sK7(X0)) & r1(sK6(X0), sK7(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) => (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK8(X0), X4)) & ~ p2(sK8(X0)) & r1(X0, sK8(X0)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X0] : (((? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) | p2(X0)) & (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) | sP0(X0))) | ~ sP5(X0)), inference(rectify, [], [f14])).
fof(f14, plain, ! [X0] : (((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | sP0(X0))) | ~ sP5(X0)), inference(nnf_transformation, [], [f12])).
fof(f2911, plain, ! [X0] : (~ r1(sK24, sK8(X0)) | ~ sP5(X0) | sP0(X0)), inference(subsumption_resolution, [], [f2910, f65])).
fof(f65, plain, ! [X0] : (~ p2(sK8(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f19])).
fof(f2910, plain, ! [X0] : (sP0(X0) | ~ sP5(X0) | ~ r1(sK24, sK8(X0)) | p2(sK8(X0))), inference(subsumption_resolution, [], [f2909, f112])).
fof(f2909, plain, ! [X0] : (p2(sK31(sK8(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK24, sK8(X0)) | p2(sK8(X0))), inference(duplicate_literal_removal, [], [f2900])).
fof(f2900, plain, ! [X0] : (p2(sK31(sK8(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK24, sK8(X0)) | p2(sK8(X0)) | ~ r1(sK24, sK8(X0))), inference(resolution, [], [f890, f111])).
fof(f890, plain, ! [X12, X13] : (~ r1(sK30(sK8(X13)), X12) | p2(X12) | sP0(X13) | ~ sP5(X13) | ~ r1(sK24, sK8(X13))), inference(subsumption_resolution, [], [f889, f65])).
fof(f889, plain, ! [X12, X13] : (p2(X12) | ~ r1(sK30(sK8(X13)), X12) | sP0(X13) | ~ sP5(X13) | p2(sK8(X13)) | ~ r1(sK24, sK8(X13))), inference(subsumption_resolution, [], [f885, f113])).
fof(f113, plain, ! [X15] : (~ r1(sK24, X15) | p2(X15) | p2(sK30(X15))), inference(cnf_transformation, [], [f63])).
fof(f885, plain, ! [X12, X13] : (p2(X12) | ~ r1(sK30(sK8(X13)), X12) | ~ p2(sK30(sK8(X13))) | sP0(X13) | ~ sP5(X13) | p2(sK8(X13)) | ~ r1(sK24, sK8(X13))), inference(resolution, [], [f66, f110])).
fof(f66, plain, ! [X4, X0, X5] : (~ r1(sK8(X0), X4) | p2(X5) | ~ r1(X4, X5) | ~ p2(X4) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f19])).
fof(f6173, plain, ~ spl36_287, inference(avatar_contradiction_clause, [], [f6172])).
fof(f6172, plain, ($false | ~ spl36_287), inference(subsumption_resolution, [], [f6171, f108])).
fof(f6171, plain, (~ r1(sK24, sK32) | ~ spl36_287), inference(subsumption_resolution, [], [f6170, f109])).
fof(f6170, plain, (p2(sK32) | ~ r1(sK24, sK32) | ~ spl36_287), inference(resolution, [], [f6166, f126])).
fof(f126, plain, ! [X1] : (~ p2(sK25(X1)) | p2(X1) | ~ r1(sK24, X1)), inference(cnf_transformation, [], [f63])).
fof(f6166, plain, (p2(sK25(sK32)) | ~ spl36_287), inference(avatar_component_clause, [], [f6164])).
fof(f5476, plain, (spl36_246 | spl36_109), inference(avatar_split_clause, [], [f5471, f2106, f5473])).
fof(f5471, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | r1(sK22(sK25(sK32)), sK23(sK25(sK32)))), inference(subsumption_resolution, [], [f5461, f109])).
fof(f5461, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | p2(sK32) | r1(sK22(sK25(sK32)), sK23(sK25(sK32)))), inference(resolution, [], [f934, f108])).
fof(f934, plain, ! [X8, X9] : (~ r1(sK24, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | r1(sK22(sK25(X8)), sK23(sK25(X8)))), inference(subsumption_resolution, [], [f919, f126])).
fof(f919, plain, ! [X8, X9] : (p2(sK25(X8)) | r1(sK22(sK25(X8)), sK23(sK25(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK24, X8)), inference(resolution, [], [f99, f125])).
fof(f99, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(sK22(X2), sK23(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f49])).
fof(f3680, plain, (spl36_153 | spl36_109), inference(avatar_split_clause, [], [f3675, f2106, f3677])).
fof(f3675, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | r1(sK25(sK32), sK22(sK25(sK32)))), inference(subsumption_resolution, [], [f3665, f109])).
fof(f3665, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | p2(sK32) | r1(sK25(sK32), sK22(sK25(sK32)))), inference(resolution, [], [f874, f108])).
fof(f874, plain, ! [X8, X9] : (~ r1(sK24, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | r1(sK25(X8), sK22(sK25(X8)))), inference(subsumption_resolution, [], [f859, f126])).
fof(f859, plain, ! [X8, X9] : (p2(sK25(X8)) | r1(sK25(X8), sK22(sK25(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK24, X8)), inference(resolution, [], [f98, f125])).
fof(f98, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(X2, sK22(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f49])).
fof(f2177, plain, (~ spl36_30 | ~ spl36_109), inference(avatar_contradiction_clause, [], [f2176])).
fof(f2176, plain, ($false | (~ spl36_30 | ~ spl36_109)), inference(subsumption_resolution, [], [f2175, f281])).
fof(f281, plain, (sP0(sK24) | ~ spl36_30), inference(avatar_component_clause, [], [f279])).
fof(f2175, plain, (~ sP0(sK24) | ~ spl36_109), inference(resolution, [], [f2107, f108])).
fof(f2107, plain, (! [X2] : (~ r1(X2, sK32) | ~ sP0(X2)) | ~ spl36_109), inference(avatar_component_clause, [], [f2106])).
fof(f2108, plain, (spl36_108 | spl36_109), inference(avatar_split_clause, [], [f2100, f2106, f2102])).
fof(f2100, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | p2(sK22(sK25(sK32)))), inference(subsumption_resolution, [], [f2082, f109])).
fof(f2082, plain, ! [X2] : (~ r1(X2, sK32) | ~ sP0(X2) | p2(sK32) | p2(sK22(sK25(sK32)))), inference(resolution, [], [f807, f108])).
fof(f807, plain, ! [X8, X9] : (~ r1(sK24, X8) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | p2(sK22(sK25(X8)))), inference(subsumption_resolution, [], [f792, f126])).
fof(f792, plain, ! [X8, X9] : (p2(sK25(X8)) | p2(sK22(sK25(X8))) | ~ r1(X9, X8) | ~ sP0(X9) | p2(X8) | ~ r1(sK24, X8)), inference(resolution, [], [f101, f125])).
fof(f101, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | p2(sK22(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f49])).
fof(f1008, plain, (spl36_37 | spl36_6 | ~ spl36_7), inference(avatar_split_clause, [], [f354, f154, f149, f320])).
fof(f354, plain, (p2(sK26) | p2(sK30(sK26)) | ~ spl36_7), inference(resolution, [], [f156, f113])).
fof(f974, plain, (~ spl36_4 | spl36_53 | spl36_54 | spl36_74 | ~ spl36_2), inference(avatar_split_clause, [], [f954, f133, f961, f425, f421, f141])).
fof(f133, plain, (spl36_2 <=> ! [X7, X8, X6] : (~ p2(X7) | ~ r1(sK26, X6) | sP4(X6) | sP3(X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | p2(X8))), introduced(avatar_definition, [new_symbols(naming, [spl36_2])])).
fof(f954, plain, (! [X6, X7] : (~ p2(X6) | sP4(sK20(sK26)) | sP3(sK20(sK26)) | ~ r1(sK20(sK26), X6) | ~ r1(X6, X7) | p2(X7) | ~ sP1(sK26)) | ~ spl36_2), inference(resolution, [], [f134, f90])).
fof(f134, plain, (! [X6, X8, X7] : (~ r1(sK26, X6) | ~ p2(X7) | sP4(X6) | sP3(X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | p2(X8)) | ~ spl36_2), inference(avatar_component_clause, [], [f133])).
fof(f157, plain, (spl36_1 | spl36_7), inference(avatar_split_clause, [], [f120, f154, f129])).
fof(f120, plain, (r1(sK24, sK26) | sP5(sK24)), inference(cnf_transformation, [], [f63])).
fof(f152, plain, (spl36_1 | spl36_4 | ~ spl36_6), inference(avatar_split_clause, [], [f121, f149, f141, f129])).
fof(f121, plain, (~ p2(sK26) | sP1(sK26) | sP5(sK24)), inference(cnf_transformation, [], [f63])).
fof(f147, plain, (spl36_1 | spl36_4 | spl36_5), inference(avatar_split_clause, [], [f122, f145, f141, f129])).
fof(f122, plain, ! [X10, X9] : (~ p2(X9) | p2(X10) | ~ r1(X9, X10) | ~ r1(sK26, X9) | sP1(sK26) | sP5(sK24)), inference(cnf_transformation, [], [f63])).
fof(f139, plain, (spl36_1 | spl36_3), inference(avatar_split_clause, [], [f123, f137, f129])).
fof(f123, plain, ! [X6] : (~ p2(X6) | sP3(X6) | sP4(X6) | ~ r1(sK26, X6) | sP5(sK24)), inference(cnf_transformation, [], [f63])).
fof(f135, plain, (spl36_1 | spl36_2), inference(avatar_split_clause, [], [f124, f133, f129])).
fof(f124, plain, ! [X6, X8, X7] : (~ p2(X7) | p2(X8) | ~ r1(X7, X8) | ~ r1(X6, X7) | sP3(X6) | sP4(X6) | ~ r1(sK26, X6) | sP5(sK24)), inference(cnf_transformation, [], [f63])).