fof(f411990, plain, $false, inference(avatar_sat_refutation, [], [f678, f8272, f8277, f8282, f8678, f8713, f10300, f13949, f89476, f90295, f90328, f233939, f395170, f395199, f401332, f403213, f403217, f403275, f405702, f411646])).
fof(f411646, plain, (spl85_75 | ~ spl85_15589), inference(avatar_contradiction_clause, [], [f411645])).
fof(f411645, plain, ($false | (spl85_75 | ~ spl85_15589)), inference(subsumption_resolution, [], [f406231, f212])).
fof(f212, plain, r1(sK54, sK57), inference(cnf_transformation, [], [f106])).
fof(f106, plain, (! [X1] : ((! [X2] : ((~ p1(sK14(X2)) & r1(X2, sK14(X2))) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & ((! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1))) | ! [X9] : ((~ p1(sK16(X9)) & r1(X9, sK16(X9))) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : ((~ p1(sK17(X13)) & r1(X13, sK17(X13))) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ((! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12))) | ! [X20] : ((~ p1(sK19(X20)) & r1(X20, sK19(X20))) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : ((~ p1(sK20(X25)) & r1(X25, sK20(X25))) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & ((! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24))) | ! [X32] : ((~ p1(sK22(X32)) & r1(X32, sK22(X32))) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : ((~ p1(sK23(X38)) & r1(X38, sK23(X38))) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & ((! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37))) | ! [X45] : ((~ p1(sK25(X45)) & r1(X45, sK25(X45))) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : ((~ p1(sK26(X52)) & r1(X52, sK26(X52))) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & ((! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51))) | ! [X59] : ((~ p1(sK28(X59)) & r1(X59, sK28(X59))) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : ((~ p1(sK29(X67)) & r1(X67, sK29(X67))) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & ((! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(sK30(X66), X72)) & ~ p1(sK30(X66)) & r1(X66, sK30(X66))) | ! [X74] : ((~ p1(sK31(X74)) & r1(X74, sK31(X74))) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(sK13, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : ((~ p1(sK32(X83)) & r1(X83, sK32(X83))) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & ((! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(sK33(X82), X88)) & ~ p1(sK33(X82)) & r1(X82, sK33(X82))) | ! [X90] : ((~ p1(sK34(X90)) & r1(X90, sK34(X90))) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(sK13, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : ((~ p1(sK35(X100)) & r1(X100, sK35(X100))) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & ((! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(sK36(X99), X105)) & ~ p1(sK36(X99)) & r1(X99, sK36(X99))) | ! [X107] : ((~ p1(sK37(X107)) & r1(X107, sK37(X107))) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(sK13, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : ((~ p1(sK38(X118)) & r1(X118, sK38(X118))) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & ((! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(sK39(X117), X123)) & ~ p1(sK39(X117)) & r1(X117, sK39(X117))) | ! [X125] : ((~ p1(sK40(X125)) & r1(X125, sK40(X125))) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(sK13, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : ((~ p1(sK41(X137)) & r1(X137, sK41(X137))) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & ((! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(sK42(X136), X142)) & ~ p1(sK42(X136)) & r1(X136, sK42(X136))) | ! [X144] : ((~ p1(sK43(X144)) & r1(X144, sK43(X144))) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(sK13, X127)) & (((((((((((! [X157] : ((p1(sK55(X157)) & (~ p1(sK56(X157)) & r1(sK55(X157), sK56(X157))) & r1(X157, sK55(X157))) | p1(X157) | ~ r1(sK54, X157)) & (! [X161] : (p1(X161) | ~ r1(sK57, X161)) & r1(sK54, sK57)) & (~ p1(sK58) & r1(sK54, sK58)) & r1(sK53, sK54)) & r1(sK52, sK53)) & r1(sK51, sK52)) & r1(sK50, sK51)) & r1(sK49, sK50)) & r1(sK48, sK49)) & r1(sK47, sK48)) & r1(sK46, sK47)) & r1(sK45, sK46)) & r1(sK44, sK45)) & r1(sK13, sK44)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : ((((! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(sK59(X173), X175)) & ~ p1(sK59(X173)) & r1(X173, sK59(X173))) | sP1(X173)) & ! [X177] : ((~ p1(sK60(X177)) & r1(X177, sK60(X177))) | ! [X179] : (! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X177, X179)) | ~ r1(X173, X177)) & sP2(X173) & sP3(X173)) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(sK13, X163)) & ! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : (! [X190] : (! [X191] : (! [X192] : ((! [X193] : ((~ p1(sK61(X193)) & r1(X193, sK61(X193))) | ! [X195] : (! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X193, X195)) | ~ r1(X192, X193)) & ((! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(sK62(X192), X198)) & ~ p1(sK62(X192)) & r1(X192, sK62(X192))) | ! [X200] : ((~ p1(sK63(X200)) & r1(X200, sK63(X200))) | ~ r1(X192, X200)) | p1(X192))) | ~ r1(X191, X192)) | ~ r1(X190, X191)) | ~ r1(X189, X190)) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(sK13, X181)) & ! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : (! [X212] : (! [X213] : (! [X214] : ((! [X215] : ((~ p1(sK64(X215)) & r1(X215, sK64(X215))) | ! [X217] : (! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X215, X217)) | ~ r1(X214, X215)) & ((! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(sK65(X214), X220)) & ~ p1(sK65(X214)) & r1(X214, sK65(X214))) | ! [X222] : ((~ p1(sK66(X222)) & r1(X222, sK66(X222))) | ~ r1(X214, X222)) | p1(X214))) | ~ r1(X213, X214)) | ~ r1(X212, X213)) | ~ r1(X211, X212)) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(sK13, X202)) & ! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : (! [X235] : (! [X236] : (! [X237] : ((! [X238] : ((~ p1(sK67(X238)) & r1(X238, sK67(X238))) | ! [X240] : (! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X238, X240)) | ~ r1(X237, X238)) & ((! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(sK68(X237), X243)) & ~ p1(sK68(X237)) & r1(X237, sK68(X237))) | ! [X245] : ((~ p1(sK69(X245)) & r1(X245, sK69(X245))) | ~ r1(X237, X245)) | p1(X237))) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | ~ r1(X234, X235)) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(sK13, X224)) & ! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : (! [X257] : (! [X258] : (! [X259] : (! [X260] : (! [X261] : ((! [X262] : ((~ p1(sK70(X262)) & r1(X262, sK70(X262))) | ! [X264] : (! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X262, X264)) | ~ r1(X261, X262)) & ((! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(sK71(X261), X267)) & ~ p1(sK71(X261)) & r1(X261, sK71(X261))) | ! [X269] : ((~ p1(sK72(X269)) & r1(X269, sK72(X269))) | ~ r1(X261, X269)) | p1(X261))) | ~ r1(X260, X261)) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | ~ r1(X257, X258)) | ~ r1(X256, X257)) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(sK13, X247)) & ! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : (! [X281] : (! [X282] : (! [X283] : (! [X284] : (! [X285] : (! [X286] : ((! [X287] : ((~ p1(sK73(X287)) & r1(X287, sK73(X287))) | ! [X289] : (! [X290] : (p1(X290) | ~ r1(X289, X290)) | ~ r1(X287, X289)) | ~ r1(X286, X287)) & ((! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(sK74(X286), X292)) & ~ p1(sK74(X286)) & r1(X286, sK74(X286))) | ! [X294] : ((~ p1(sK75(X294)) & r1(X294, sK75(X294))) | ~ r1(X286, X294)) | p1(X286))) | ~ r1(X285, X286)) | ~ r1(X284, X285)) | ~ r1(X283, X284)) | ~ r1(X282, X283)) | ~ r1(X281, X282)) | ~ r1(X280, X281)) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(sK13, X271)) & ! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : (! [X306] : (! [X307] : (! [X308] : (! [X309] : (! [X310] : (! [X311] : (! [X312] : ((! [X313] : ((~ p1(sK76(X313)) & r1(X313, sK76(X313))) | ! [X315] : (! [X316] : (p1(X316) | ~ r1(X315, X316)) | ~ r1(X313, X315)) | ~ r1(X312, X313)) & ((! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(sK77(X312), X318)) & ~ p1(sK77(X312)) & r1(X312, sK77(X312))) | ! [X320] : ((~ p1(sK78(X320)) & r1(X320, sK78(X320))) | ~ r1(X312, X320)) | p1(X312))) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | ~ r1(X309, X310)) | ~ r1(X308, X309)) | ~ r1(X307, X308)) | ~ r1(X306, X307)) | ~ r1(X305, X306)) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(sK13, X296)) & ! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : (! [X332] : (! [X333] : (! [X334] : (! [X335] : (! [X336] : (! [X337] : (! [X338] : (! [X339] : ((! [X340] : ((~ p1(sK79(X340)) & r1(X340, sK79(X340))) | ! [X342] : (! [X343] : (p1(X343) | ~ r1(X342, X343)) | ~ r1(X340, X342)) | ~ r1(X339, X340)) & ((! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(sK80(X339), X345)) & ~ p1(sK80(X339)) & r1(X339, sK80(X339))) | ! [X347] : ((~ p1(sK81(X347)) & r1(X347, sK81(X347))) | ~ r1(X339, X347)) | p1(X339))) | ~ r1(X338, X339)) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | ~ r1(X335, X336)) | ~ r1(X334, X335)) | ~ r1(X333, X334)) | ~ r1(X332, X333)) | ~ r1(X331, X332)) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(sK13, X322)) & ! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : (! [X359] : (! [X360] : (! [X361] : (! [X362] : (! [X363] : (! [X364] : (! [X365] : (! [X366] : (! [X367] : ((! [X368] : ((~ p1(sK82(X368)) & r1(X368, sK82(X368))) | ! [X370] : (! [X371] : (p1(X371) | ~ r1(X370, X371)) | ~ r1(X368, X370)) | ~ r1(X367, X368)) & ((! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(sK83(X367), X373)) & ~ p1(sK83(X367)) & r1(X367, sK83(X367))) | ! [X375] : ((~ p1(sK84(X375)) & r1(X375, sK84(X375))) | ~ r1(X367, X375)) | p1(X367))) | ~ r1(X366, X367)) | ~ r1(X365, X366)) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | ~ r1(X362, X363)) | ~ r1(X361, X362)) | ~ r1(X360, X361)) | ~ r1(X359, X360)) | ~ r1(X358, X359)) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(sK13, X349))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15, sK16, sK17, sK18, sK19, sK20, sK21, sK22, sK23, sK24, sK25, sK26, sK27, sK28, sK29, sK30, sK31, sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49, sK50, sK51, sK52, sK53, sK54, sK55, sK56, sK57, sK58, sK59, sK60, sK61, sK62, sK63, sK64, sK65, sK66, sK67, sK68, sK69, sK70, sK71, sK72, sK73, sK74, sK75, sK76, sK77, sK78, sK79, sK80, sK81, sK82, sK83, sK84])], [f33, f105, f104, f103, f102, f101, f100, f99, f98, f97, f96, f95, f94, f93, f92, f91, f90, f89, f88, f87, f86, f85, f84, f83, f82, f81, f80, f79, f78, f77, f76, f75, f74, f73, f72, f71, f70, f69, f68, f67, f66, f65, f64, f63, f62, f61, f60, f59, f58, f57, f56, f55, f54, f53, f52, f51, f50, f49, f48, f47, f46, f45, f44, f43, f42, f41, f40, f39, f38, f37, f36, f35, f34])).
fof(f34, plain, (? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(X0, X146)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | sP1(X173)) & ! [X177] : (? [X178] : (~ p1(X178) & r1(X177, X178)) | ! [X179] : (! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X177, X179)) | ~ r1(X173, X177)) & sP2(X173) & sP3(X173)) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) & ! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : (! [X190] : (! [X191] : (! [X192] : ((! [X193] : (? [X194] : (~ p1(X194) & r1(X193, X194)) | ! [X195] : (! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X193, X195)) | ~ r1(X192, X193)) & (? [X197] : (! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(X197, X198)) & ~ p1(X197) & r1(X192, X197)) | ! [X200] : (? [X201] : (~ p1(X201) & r1(X200, X201)) | ~ r1(X192, X200)) | p1(X192))) | ~ r1(X191, X192)) | ~ r1(X190, X191)) | ~ r1(X189, X190)) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X0, X181)) & ! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : (! [X212] : (! [X213] : (! [X214] : ((! [X215] : (? [X216] : (~ p1(X216) & r1(X215, X216)) | ! [X217] : (! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X215, X217)) | ~ r1(X214, X215)) & (? [X219] : (! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(X219, X220)) & ~ p1(X219) & r1(X214, X219)) | ! [X222] : (? [X223] : (~ p1(X223) & r1(X222, X223)) | ~ r1(X214, X222)) | p1(X214))) | ~ r1(X213, X214)) | ~ r1(X212, X213)) | ~ r1(X211, X212)) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X0, X202)) & ! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : (! [X235] : (! [X236] : (! [X237] : ((! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ! [X240] : (! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X238, X240)) | ~ r1(X237, X238)) & (? [X242] : (! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(X242, X243)) & ~ p1(X242) & r1(X237, X242)) | ! [X245] : (? [X246] : (~ p1(X246) & r1(X245, X246)) | ~ r1(X237, X245)) | p1(X237))) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | ~ r1(X234, X235)) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X0, X224)) & ! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : (! [X257] : (! [X258] : (! [X259] : (! [X260] : (! [X261] : ((! [X262] : (? [X263] : (~ p1(X263) & r1(X262, X263)) | ! [X264] : (! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X262, X264)) | ~ r1(X261, X262)) & (? [X266] : (! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(X266, X267)) & ~ p1(X266) & r1(X261, X266)) | ! [X269] : (? [X270] : (~ p1(X270) & r1(X269, X270)) | ~ r1(X261, X269)) | p1(X261))) | ~ r1(X260, X261)) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | ~ r1(X257, X258)) | ~ r1(X256, X257)) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X0, X247)) & ! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : (! [X281] : (! [X282] : (! [X283] : (! [X284] : (! [X285] : (! [X286] : ((! [X287] : (? [X288] : (~ p1(X288) & r1(X287, X288)) | ! [X289] : (! [X290] : (p1(X290) | ~ r1(X289, X290)) | ~ r1(X287, X289)) | ~ r1(X286, X287)) & (? [X291] : (! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(X291, X292)) & ~ p1(X291) & r1(X286, X291)) | ! [X294] : (? [X295] : (~ p1(X295) & r1(X294, X295)) | ~ r1(X286, X294)) | p1(X286))) | ~ r1(X285, X286)) | ~ r1(X284, X285)) | ~ r1(X283, X284)) | ~ r1(X282, X283)) | ~ r1(X281, X282)) | ~ r1(X280, X281)) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X0, X271)) & ! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : (! [X306] : (! [X307] : (! [X308] : (! [X309] : (! [X310] : (! [X311] : (! [X312] : ((! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ! [X315] : (! [X316] : (p1(X316) | ~ r1(X315, X316)) | ~ r1(X313, X315)) | ~ r1(X312, X313)) & (? [X317] : (! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(X317, X318)) & ~ p1(X317) & r1(X312, X317)) | ! [X320] : (? [X321] : (~ p1(X321) & r1(X320, X321)) | ~ r1(X312, X320)) | p1(X312))) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | ~ r1(X309, X310)) | ~ r1(X308, X309)) | ~ r1(X307, X308)) | ~ r1(X306, X307)) | ~ r1(X305, X306)) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X0, X296)) & ! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : (! [X332] : (! [X333] : (! [X334] : (! [X335] : (! [X336] : (! [X337] : (! [X338] : (! [X339] : ((! [X340] : (? [X341] : (~ p1(X341) & r1(X340, X341)) | ! [X342] : (! [X343] : (p1(X343) | ~ r1(X342, X343)) | ~ r1(X340, X342)) | ~ r1(X339, X340)) & (? [X344] : (! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(X344, X345)) & ~ p1(X344) & r1(X339, X344)) | ! [X347] : (? [X348] : (~ p1(X348) & r1(X347, X348)) | ~ r1(X339, X347)) | p1(X339))) | ~ r1(X338, X339)) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | ~ r1(X335, X336)) | ~ r1(X334, X335)) | ~ r1(X333, X334)) | ~ r1(X332, X333)) | ~ r1(X331, X332)) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X0, X322)) & ! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : (! [X359] : (! [X360] : (! [X361] : (! [X362] : (! [X363] : (! [X364] : (! [X365] : (! [X366] : (! [X367] : ((! [X368] : (? [X369] : (~ p1(X369) & r1(X368, X369)) | ! [X370] : (! [X371] : (p1(X371) | ~ r1(X370, X371)) | ~ r1(X368, X370)) | ~ r1(X367, X368)) & (? [X372] : (! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(X372, X373)) & ~ p1(X372) & r1(X367, X372)) | ! [X375] : (? [X376] : (~ p1(X376) & r1(X375, X376)) | ~ r1(X367, X375)) | p1(X367))) | ~ r1(X366, X367)) | ~ r1(X365, X366)) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | ~ r1(X362, X363)) | ~ r1(X361, X362)) | ~ r1(X360, X361)) | ~ r1(X359, X360)) | ~ r1(X358, X359)) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X0, X349))) => (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(sK13, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(sK13, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(sK13, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(sK13, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(sK13, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(sK13, X146)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | sP1(X173)) & ! [X177] : (? [X178] : (~ p1(X178) & r1(X177, X178)) | ! [X179] : (! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X177, X179)) | ~ r1(X173, X177)) & sP2(X173) & sP3(X173)) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(sK13, X163)) & ! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : (! [X190] : (! [X191] : (! [X192] : ((! [X193] : (? [X194] : (~ p1(X194) & r1(X193, X194)) | ! [X195] : (! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X193, X195)) | ~ r1(X192, X193)) & (? [X197] : (! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(X197, X198)) & ~ p1(X197) & r1(X192, X197)) | ! [X200] : (? [X201] : (~ p1(X201) & r1(X200, X201)) | ~ r1(X192, X200)) | p1(X192))) | ~ r1(X191, X192)) | ~ r1(X190, X191)) | ~ r1(X189, X190)) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(sK13, X181)) & ! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : (! [X212] : (! [X213] : (! [X214] : ((! [X215] : (? [X216] : (~ p1(X216) & r1(X215, X216)) | ! [X217] : (! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X215, X217)) | ~ r1(X214, X215)) & (? [X219] : (! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(X219, X220)) & ~ p1(X219) & r1(X214, X219)) | ! [X222] : (? [X223] : (~ p1(X223) & r1(X222, X223)) | ~ r1(X214, X222)) | p1(X214))) | ~ r1(X213, X214)) | ~ r1(X212, X213)) | ~ r1(X211, X212)) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(sK13, X202)) & ! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : (! [X235] : (! [X236] : (! [X237] : ((! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ! [X240] : (! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X238, X240)) | ~ r1(X237, X238)) & (? [X242] : (! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(X242, X243)) & ~ p1(X242) & r1(X237, X242)) | ! [X245] : (? [X246] : (~ p1(X246) & r1(X245, X246)) | ~ r1(X237, X245)) | p1(X237))) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | ~ r1(X234, X235)) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(sK13, X224)) & ! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : (! [X257] : (! [X258] : (! [X259] : (! [X260] : (! [X261] : ((! [X262] : (? [X263] : (~ p1(X263) & r1(X262, X263)) | ! [X264] : (! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X262, X264)) | ~ r1(X261, X262)) & (? [X266] : (! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(X266, X267)) & ~ p1(X266) & r1(X261, X266)) | ! [X269] : (? [X270] : (~ p1(X270) & r1(X269, X270)) | ~ r1(X261, X269)) | p1(X261))) | ~ r1(X260, X261)) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | ~ r1(X257, X258)) | ~ r1(X256, X257)) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(sK13, X247)) & ! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : (! [X281] : (! [X282] : (! [X283] : (! [X284] : (! [X285] : (! [X286] : ((! [X287] : (? [X288] : (~ p1(X288) & r1(X287, X288)) | ! [X289] : (! [X290] : (p1(X290) | ~ r1(X289, X290)) | ~ r1(X287, X289)) | ~ r1(X286, X287)) & (? [X291] : (! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(X291, X292)) & ~ p1(X291) & r1(X286, X291)) | ! [X294] : (? [X295] : (~ p1(X295) & r1(X294, X295)) | ~ r1(X286, X294)) | p1(X286))) | ~ r1(X285, X286)) | ~ r1(X284, X285)) | ~ r1(X283, X284)) | ~ r1(X282, X283)) | ~ r1(X281, X282)) | ~ r1(X280, X281)) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(sK13, X271)) & ! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : (! [X306] : (! [X307] : (! [X308] : (! [X309] : (! [X310] : (! [X311] : (! [X312] : ((! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ! [X315] : (! [X316] : (p1(X316) | ~ r1(X315, X316)) | ~ r1(X313, X315)) | ~ r1(X312, X313)) & (? [X317] : (! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(X317, X318)) & ~ p1(X317) & r1(X312, X317)) | ! [X320] : (? [X321] : (~ p1(X321) & r1(X320, X321)) | ~ r1(X312, X320)) | p1(X312))) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | ~ r1(X309, X310)) | ~ r1(X308, X309)) | ~ r1(X307, X308)) | ~ r1(X306, X307)) | ~ r1(X305, X306)) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(sK13, X296)) & ! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : (! [X332] : (! [X333] : (! [X334] : (! [X335] : (! [X336] : (! [X337] : (! [X338] : (! [X339] : ((! [X340] : (? [X341] : (~ p1(X341) & r1(X340, X341)) | ! [X342] : (! [X343] : (p1(X343) | ~ r1(X342, X343)) | ~ r1(X340, X342)) | ~ r1(X339, X340)) & (? [X344] : (! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(X344, X345)) & ~ p1(X344) & r1(X339, X344)) | ! [X347] : (? [X348] : (~ p1(X348) & r1(X347, X348)) | ~ r1(X339, X347)) | p1(X339))) | ~ r1(X338, X339)) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | ~ r1(X335, X336)) | ~ r1(X334, X335)) | ~ r1(X333, X334)) | ~ r1(X332, X333)) | ~ r1(X331, X332)) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(sK13, X322)) & ! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : (! [X359] : (! [X360] : (! [X361] : (! [X362] : (! [X363] : (! [X364] : (! [X365] : (! [X366] : (! [X367] : ((! [X368] : (? [X369] : (~ p1(X369) & r1(X368, X369)) | ! [X370] : (! [X371] : (p1(X371) | ~ r1(X370, X371)) | ~ r1(X368, X370)) | ~ r1(X367, X368)) & (? [X372] : (! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(X372, X373)) & ~ p1(X372) & r1(X367, X372)) | ! [X375] : (? [X376] : (~ p1(X376) & r1(X375, X376)) | ~ r1(X367, X375)) | p1(X367))) | ~ r1(X366, X367)) | ~ r1(X365, X366)) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | ~ r1(X362, X363)) | ~ r1(X361, X362)) | ~ r1(X360, X361)) | ~ r1(X359, X360)) | ~ r1(X358, X359)) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(sK13, X349)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK14(X2)) & r1(X2, sK14(X2)))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) => (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) => (~ p1(sK16(X9)) & r1(X9, sK16(X9)))), introduced(choice_axiom, [])).
fof(f38, plain, ! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) => (~ p1(sK17(X13)) & r1(X13, sK17(X13)))), introduced(choice_axiom, [])).
fof(f39, plain, ! [X12] : (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) => (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) => (~ p1(sK19(X20)) & r1(X20, sK19(X20)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) => (~ p1(sK20(X25)) & r1(X25, sK20(X25)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X24] : (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) => (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24)))), introduced(choice_axiom, [])).
fof(f43, plain, ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) => (~ p1(sK22(X32)) & r1(X32, sK22(X32)))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) => (~ p1(sK23(X38)) & r1(X38, sK23(X38)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X37] : (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) => (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37)))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) => (~ p1(sK25(X45)) & r1(X45, sK25(X45)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) => (~ p1(sK26(X52)) & r1(X52, sK26(X52)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X51] : (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) => (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51)))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) => (~ p1(sK28(X59)) & r1(X59, sK28(X59)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) => (~ p1(sK29(X67)) & r1(X67, sK29(X67)))), introduced(choice_axiom, [])).
fof(f51, plain, ! [X66] : (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) => (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(sK30(X66), X72)) & ~ p1(sK30(X66)) & r1(X66, sK30(X66)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) => (~ p1(sK31(X74)) & r1(X74, sK31(X74)))), introduced(choice_axiom, [])).
fof(f53, plain, ! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) => (~ p1(sK32(X83)) & r1(X83, sK32(X83)))), introduced(choice_axiom, [])).
fof(f54, plain, ! [X82] : (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) => (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(sK33(X82), X88)) & ~ p1(sK33(X82)) & r1(X82, sK33(X82)))), introduced(choice_axiom, [])).
fof(f55, plain, ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) => (~ p1(sK34(X90)) & r1(X90, sK34(X90)))), introduced(choice_axiom, [])).
fof(f56, plain, ! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) => (~ p1(sK35(X100)) & r1(X100, sK35(X100)))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X99] : (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) => (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(sK36(X99), X105)) & ~ p1(sK36(X99)) & r1(X99, sK36(X99)))), introduced(choice_axiom, [])).
fof(f58, plain, ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) => (~ p1(sK37(X107)) & r1(X107, sK37(X107)))), introduced(choice_axiom, [])).
fof(f59, plain, ! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) => (~ p1(sK38(X118)) & r1(X118, sK38(X118)))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X117] : (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) => (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(sK39(X117), X123)) & ~ p1(sK39(X117)) & r1(X117, sK39(X117)))), introduced(choice_axiom, [])).
fof(f61, plain, ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) => (~ p1(sK40(X125)) & r1(X125, sK40(X125)))), introduced(choice_axiom, [])).
fof(f62, plain, ! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) => (~ p1(sK41(X137)) & r1(X137, sK41(X137)))), introduced(choice_axiom, [])).
fof(f63, plain, ! [X136] : (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) => (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(sK42(X136), X142)) & ~ p1(sK42(X136)) & r1(X136, sK42(X136)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) => (~ p1(sK43(X144)) & r1(X144, sK43(X144)))), introduced(choice_axiom, [])).
fof(f65, plain, (? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(sK13, X146)) => (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(sK44, X147)) & r1(sK13, sK44))), introduced(choice_axiom, [])).
fof(f66, plain, (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(sK44, X147)) => (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(sK45, X148)) & r1(sK44, sK45))), introduced(choice_axiom, [])).
fof(f67, plain, (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(sK45, X148)) => (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(sK46, X149)) & r1(sK45, sK46))), introduced(choice_axiom, [])).
fof(f68, plain, (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(sK46, X149)) => (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(sK47, X150)) & r1(sK46, sK47))), introduced(choice_axiom, [])).
fof(f69, plain, (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(sK47, X150)) => (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(sK48, X151)) & r1(sK47, sK48))), introduced(choice_axiom, [])).
fof(f70, plain, (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(sK48, X151)) => (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(sK49, X152)) & r1(sK48, sK49))), introduced(choice_axiom, [])).
fof(f71, plain, (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(sK49, X152)) => (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(sK50, X153)) & r1(sK49, sK50))), introduced(choice_axiom, [])).
fof(f72, plain, (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(sK50, X153)) => (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(sK51, X154)) & r1(sK50, sK51))), introduced(choice_axiom, [])).
fof(f73, plain, (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(sK51, X154)) => (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(sK52, X155)) & r1(sK51, sK52))), introduced(choice_axiom, [])).
fof(f74, plain, (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(sK52, X155)) => (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(sK53, X156)) & r1(sK52, sK53))), introduced(choice_axiom, [])).
fof(f75, plain, (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(sK53, X156)) => (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(sK54, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(sK54, X160)) & ? [X162] : (~ p1(X162) & r1(sK54, X162)) & r1(sK53, sK54))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) => (p1(sK55(X157)) & ? [X159] : (~ p1(X159) & r1(sK55(X157), X159)) & r1(X157, sK55(X157)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X157] : (? [X159] : (~ p1(X159) & r1(sK55(X157), X159)) => (~ p1(sK56(X157)) & r1(sK55(X157), sK56(X157)))), introduced(choice_axiom, [])).
fof(f78, plain, (? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(sK54, X160)) => (! [X161] : (p1(X161) | ~ r1(sK57, X161)) & r1(sK54, sK57))), introduced(choice_axiom, [])).
fof(f79, plain, (? [X162] : (~ p1(X162) & r1(sK54, X162)) => (~ p1(sK58) & r1(sK54, sK58))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X173] : (? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) => (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(sK59(X173), X175)) & ~ p1(sK59(X173)) & r1(X173, sK59(X173)))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X177] : (? [X178] : (~ p1(X178) & r1(X177, X178)) => (~ p1(sK60(X177)) & r1(X177, sK60(X177)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X193] : (? [X194] : (~ p1(X194) & r1(X193, X194)) => (~ p1(sK61(X193)) & r1(X193, sK61(X193)))), introduced(choice_axiom, [])).
fof(f83, plain, ! [X192] : (? [X197] : (! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(X197, X198)) & ~ p1(X197) & r1(X192, X197)) => (! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(sK62(X192), X198)) & ~ p1(sK62(X192)) & r1(X192, sK62(X192)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X200] : (? [X201] : (~ p1(X201) & r1(X200, X201)) => (~ p1(sK63(X200)) & r1(X200, sK63(X200)))), introduced(choice_axiom, [])).
fof(f85, plain, ! [X215] : (? [X216] : (~ p1(X216) & r1(X215, X216)) => (~ p1(sK64(X215)) & r1(X215, sK64(X215)))), introduced(choice_axiom, [])).
fof(f86, plain, ! [X214] : (? [X219] : (! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(X219, X220)) & ~ p1(X219) & r1(X214, X219)) => (! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(sK65(X214), X220)) & ~ p1(sK65(X214)) & r1(X214, sK65(X214)))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X222] : (? [X223] : (~ p1(X223) & r1(X222, X223)) => (~ p1(sK66(X222)) & r1(X222, sK66(X222)))), introduced(choice_axiom, [])).
fof(f88, plain, ! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) => (~ p1(sK67(X238)) & r1(X238, sK67(X238)))), introduced(choice_axiom, [])).
fof(f89, plain, ! [X237] : (? [X242] : (! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(X242, X243)) & ~ p1(X242) & r1(X237, X242)) => (! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(sK68(X237), X243)) & ~ p1(sK68(X237)) & r1(X237, sK68(X237)))), introduced(choice_axiom, [])).
fof(f90, plain, ! [X245] : (? [X246] : (~ p1(X246) & r1(X245, X246)) => (~ p1(sK69(X245)) & r1(X245, sK69(X245)))), introduced(choice_axiom, [])).
fof(f91, plain, ! [X262] : (? [X263] : (~ p1(X263) & r1(X262, X263)) => (~ p1(sK70(X262)) & r1(X262, sK70(X262)))), introduced(choice_axiom, [])).
fof(f92, plain, ! [X261] : (? [X266] : (! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(X266, X267)) & ~ p1(X266) & r1(X261, X266)) => (! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(sK71(X261), X267)) & ~ p1(sK71(X261)) & r1(X261, sK71(X261)))), introduced(choice_axiom, [])).
fof(f93, plain, ! [X269] : (? [X270] : (~ p1(X270) & r1(X269, X270)) => (~ p1(sK72(X269)) & r1(X269, sK72(X269)))), introduced(choice_axiom, [])).
fof(f94, plain, ! [X287] : (? [X288] : (~ p1(X288) & r1(X287, X288)) => (~ p1(sK73(X287)) & r1(X287, sK73(X287)))), introduced(choice_axiom, [])).
fof(f95, plain, ! [X286] : (? [X291] : (! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(X291, X292)) & ~ p1(X291) & r1(X286, X291)) => (! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(sK74(X286), X292)) & ~ p1(sK74(X286)) & r1(X286, sK74(X286)))), introduced(choice_axiom, [])).
fof(f96, plain, ! [X294] : (? [X295] : (~ p1(X295) & r1(X294, X295)) => (~ p1(sK75(X294)) & r1(X294, sK75(X294)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) => (~ p1(sK76(X313)) & r1(X313, sK76(X313)))), introduced(choice_axiom, [])).
fof(f98, plain, ! [X312] : (? [X317] : (! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(X317, X318)) & ~ p1(X317) & r1(X312, X317)) => (! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(sK77(X312), X318)) & ~ p1(sK77(X312)) & r1(X312, sK77(X312)))), introduced(choice_axiom, [])).
fof(f99, plain, ! [X320] : (? [X321] : (~ p1(X321) & r1(X320, X321)) => (~ p1(sK78(X320)) & r1(X320, sK78(X320)))), introduced(choice_axiom, [])).
fof(f100, plain, ! [X340] : (? [X341] : (~ p1(X341) & r1(X340, X341)) => (~ p1(sK79(X340)) & r1(X340, sK79(X340)))), introduced(choice_axiom, [])).
fof(f101, plain, ! [X339] : (? [X344] : (! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(X344, X345)) & ~ p1(X344) & r1(X339, X344)) => (! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(sK80(X339), X345)) & ~ p1(sK80(X339)) & r1(X339, sK80(X339)))), introduced(choice_axiom, [])).
fof(f102, plain, ! [X347] : (? [X348] : (~ p1(X348) & r1(X347, X348)) => (~ p1(sK81(X347)) & r1(X347, sK81(X347)))), introduced(choice_axiom, [])).
fof(f103, plain, ! [X368] : (? [X369] : (~ p1(X369) & r1(X368, X369)) => (~ p1(sK82(X368)) & r1(X368, sK82(X368)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X367] : (? [X372] : (! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(X372, X373)) & ~ p1(X372) & r1(X367, X372)) => (! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(sK83(X367), X373)) & ~ p1(sK83(X367)) & r1(X367, sK83(X367)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X375] : (? [X376] : (~ p1(X376) & r1(X375, X376)) => (~ p1(sK84(X375)) & r1(X375, sK84(X375)))), introduced(choice_axiom, [])).
fof(f33, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(X0, X146)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | sP1(X173)) & ! [X177] : (? [X178] : (~ p1(X178) & r1(X177, X178)) | ! [X179] : (! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X177, X179)) | ~ r1(X173, X177)) & sP2(X173) & sP3(X173)) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) & ! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : (! [X190] : (! [X191] : (! [X192] : ((! [X193] : (? [X194] : (~ p1(X194) & r1(X193, X194)) | ! [X195] : (! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X193, X195)) | ~ r1(X192, X193)) & (? [X197] : (! [X198] : (~ p1(X198) | ! [X199] : (p1(X199) | ~ r1(X198, X199)) | ~ r1(X197, X198)) & ~ p1(X197) & r1(X192, X197)) | ! [X200] : (? [X201] : (~ p1(X201) & r1(X200, X201)) | ~ r1(X192, X200)) | p1(X192))) | ~ r1(X191, X192)) | ~ r1(X190, X191)) | ~ r1(X189, X190)) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X0, X181)) & ! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : (! [X212] : (! [X213] : (! [X214] : ((! [X215] : (? [X216] : (~ p1(X216) & r1(X215, X216)) | ! [X217] : (! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X215, X217)) | ~ r1(X214, X215)) & (? [X219] : (! [X220] : (~ p1(X220) | ! [X221] : (p1(X221) | ~ r1(X220, X221)) | ~ r1(X219, X220)) & ~ p1(X219) & r1(X214, X219)) | ! [X222] : (? [X223] : (~ p1(X223) & r1(X222, X223)) | ~ r1(X214, X222)) | p1(X214))) | ~ r1(X213, X214)) | ~ r1(X212, X213)) | ~ r1(X211, X212)) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X0, X202)) & ! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : (! [X235] : (! [X236] : (! [X237] : ((! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ! [X240] : (! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X238, X240)) | ~ r1(X237, X238)) & (? [X242] : (! [X243] : (~ p1(X243) | ! [X244] : (p1(X244) | ~ r1(X243, X244)) | ~ r1(X242, X243)) & ~ p1(X242) & r1(X237, X242)) | ! [X245] : (? [X246] : (~ p1(X246) & r1(X245, X246)) | ~ r1(X237, X245)) | p1(X237))) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | ~ r1(X234, X235)) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X0, X224)) & ! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : (! [X257] : (! [X258] : (! [X259] : (! [X260] : (! [X261] : ((! [X262] : (? [X263] : (~ p1(X263) & r1(X262, X263)) | ! [X264] : (! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X262, X264)) | ~ r1(X261, X262)) & (? [X266] : (! [X267] : (~ p1(X267) | ! [X268] : (p1(X268) | ~ r1(X267, X268)) | ~ r1(X266, X267)) & ~ p1(X266) & r1(X261, X266)) | ! [X269] : (? [X270] : (~ p1(X270) & r1(X269, X270)) | ~ r1(X261, X269)) | p1(X261))) | ~ r1(X260, X261)) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | ~ r1(X257, X258)) | ~ r1(X256, X257)) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X0, X247)) & ! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : (! [X281] : (! [X282] : (! [X283] : (! [X284] : (! [X285] : (! [X286] : ((! [X287] : (? [X288] : (~ p1(X288) & r1(X287, X288)) | ! [X289] : (! [X290] : (p1(X290) | ~ r1(X289, X290)) | ~ r1(X287, X289)) | ~ r1(X286, X287)) & (? [X291] : (! [X292] : (~ p1(X292) | ! [X293] : (p1(X293) | ~ r1(X292, X293)) | ~ r1(X291, X292)) & ~ p1(X291) & r1(X286, X291)) | ! [X294] : (? [X295] : (~ p1(X295) & r1(X294, X295)) | ~ r1(X286, X294)) | p1(X286))) | ~ r1(X285, X286)) | ~ r1(X284, X285)) | ~ r1(X283, X284)) | ~ r1(X282, X283)) | ~ r1(X281, X282)) | ~ r1(X280, X281)) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X0, X271)) & ! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : (! [X306] : (! [X307] : (! [X308] : (! [X309] : (! [X310] : (! [X311] : (! [X312] : ((! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ! [X315] : (! [X316] : (p1(X316) | ~ r1(X315, X316)) | ~ r1(X313, X315)) | ~ r1(X312, X313)) & (? [X317] : (! [X318] : (~ p1(X318) | ! [X319] : (p1(X319) | ~ r1(X318, X319)) | ~ r1(X317, X318)) & ~ p1(X317) & r1(X312, X317)) | ! [X320] : (? [X321] : (~ p1(X321) & r1(X320, X321)) | ~ r1(X312, X320)) | p1(X312))) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | ~ r1(X309, X310)) | ~ r1(X308, X309)) | ~ r1(X307, X308)) | ~ r1(X306, X307)) | ~ r1(X305, X306)) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X0, X296)) & ! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : (! [X332] : (! [X333] : (! [X334] : (! [X335] : (! [X336] : (! [X337] : (! [X338] : (! [X339] : ((! [X340] : (? [X341] : (~ p1(X341) & r1(X340, X341)) | ! [X342] : (! [X343] : (p1(X343) | ~ r1(X342, X343)) | ~ r1(X340, X342)) | ~ r1(X339, X340)) & (? [X344] : (! [X345] : (~ p1(X345) | ! [X346] : (p1(X346) | ~ r1(X345, X346)) | ~ r1(X344, X345)) & ~ p1(X344) & r1(X339, X344)) | ! [X347] : (? [X348] : (~ p1(X348) & r1(X347, X348)) | ~ r1(X339, X347)) | p1(X339))) | ~ r1(X338, X339)) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | ~ r1(X335, X336)) | ~ r1(X334, X335)) | ~ r1(X333, X334)) | ~ r1(X332, X333)) | ~ r1(X331, X332)) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X0, X322)) & ! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : (! [X359] : (! [X360] : (! [X361] : (! [X362] : (! [X363] : (! [X364] : (! [X365] : (! [X366] : (! [X367] : ((! [X368] : (? [X369] : (~ p1(X369) & r1(X368, X369)) | ! [X370] : (! [X371] : (p1(X371) | ~ r1(X370, X371)) | ~ r1(X368, X370)) | ~ r1(X367, X368)) & (? [X372] : (! [X373] : (~ p1(X373) | ! [X374] : (p1(X374) | ~ r1(X373, X374)) | ~ r1(X372, X373)) & ~ p1(X372) & r1(X367, X372)) | ! [X375] : (? [X376] : (~ p1(X376) & r1(X375, X376)) | ~ r1(X367, X375)) | p1(X367))) | ~ r1(X366, X367)) | ~ r1(X365, X366)) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | ~ r1(X362, X363)) | ~ r1(X361, X362)) | ~ r1(X360, X361)) | ~ r1(X359, X360)) | ~ r1(X358, X359)) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X0, X349))), inference(rectify, [], [f11])).
fof(f11, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(X0, X146)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | sP1(X173)) & ! [X181] : (? [X182] : (~ p1(X182) & r1(X181, X182)) | ! [X183] : (! [X184] : (p1(X184) | ~ r1(X183, X184)) | ~ r1(X181, X183)) | ~ r1(X173, X181)) & sP2(X173) & sP3(X173)) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) & ! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X0, X200)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : ((! [X234] : (? [X235] : (~ p1(X235) & r1(X234, X235)) | ! [X236] : (! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X234, X236)) | ~ r1(X233, X234)) & (? [X238] : (! [X239] : (~ p1(X239) | ! [X240] : (p1(X240) | ~ r1(X239, X240)) | ~ r1(X238, X239)) & ~ p1(X238) & r1(X233, X238)) | ! [X241] : (? [X242] : (~ p1(X242) & r1(X241, X242)) | ~ r1(X233, X241)) | p1(X233))) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221)) & ! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : ((! [X257] : (? [X258] : (~ p1(X258) & r1(X257, X258)) | ! [X259] : (! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X257, X259)) | ~ r1(X256, X257)) & (? [X261] : (! [X262] : (~ p1(X262) | ! [X263] : (p1(X263) | ~ r1(X262, X263)) | ~ r1(X261, X262)) & ~ p1(X261) & r1(X256, X261)) | ! [X264] : (? [X265] : (~ p1(X265) & r1(X264, X265)) | ~ r1(X256, X264)) | p1(X256))) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X0, X243)) & ! [X266] : (! [X267] : (! [X268] : (! [X269] : (! [X270] : (! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : ((! [X281] : (? [X282] : (~ p1(X282) & r1(X281, X282)) | ! [X283] : (! [X284] : (p1(X284) | ~ r1(X283, X284)) | ~ r1(X281, X283)) | ~ r1(X280, X281)) & (? [X285] : (! [X286] : (~ p1(X286) | ! [X287] : (p1(X287) | ~ r1(X286, X287)) | ~ r1(X285, X286)) & ~ p1(X285) & r1(X280, X285)) | ! [X288] : (? [X289] : (~ p1(X289) & r1(X288, X289)) | ~ r1(X280, X288)) | p1(X280))) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X270, X271)) | ~ r1(X269, X270)) | ~ r1(X268, X269)) | ~ r1(X267, X268)) | ~ r1(X266, X267)) | ~ r1(X0, X266)) & ! [X290] : (! [X291] : (! [X292] : (! [X293] : (! [X294] : (! [X295] : (! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : ((! [X306] : (? [X307] : (~ p1(X307) & r1(X306, X307)) | ! [X308] : (! [X309] : (p1(X309) | ~ r1(X308, X309)) | ~ r1(X306, X308)) | ~ r1(X305, X306)) & (? [X310] : (! [X311] : (~ p1(X311) | ! [X312] : (p1(X312) | ~ r1(X311, X312)) | ~ r1(X310, X311)) & ~ p1(X310) & r1(X305, X310)) | ! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ~ r1(X305, X313)) | p1(X305))) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X295, X296)) | ~ r1(X294, X295)) | ~ r1(X293, X294)) | ~ r1(X292, X293)) | ~ r1(X291, X292)) | ~ r1(X290, X291)) | ~ r1(X0, X290)) & ! [X315] : (! [X316] : (! [X317] : (! [X318] : (! [X319] : (! [X320] : (! [X321] : (! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : ((! [X332] : (? [X333] : (~ p1(X333) & r1(X332, X333)) | ! [X334] : (! [X335] : (p1(X335) | ~ r1(X334, X335)) | ~ r1(X332, X334)) | ~ r1(X331, X332)) & (? [X336] : (! [X337] : (~ p1(X337) | ! [X338] : (p1(X338) | ~ r1(X337, X338)) | ~ r1(X336, X337)) & ~ p1(X336) & r1(X331, X336)) | ! [X339] : (? [X340] : (~ p1(X340) & r1(X339, X340)) | ~ r1(X331, X339)) | p1(X331))) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X321, X322)) | ~ r1(X320, X321)) | ~ r1(X319, X320)) | ~ r1(X318, X319)) | ~ r1(X317, X318)) | ~ r1(X316, X317)) | ~ r1(X315, X316)) | ~ r1(X0, X315)) & ! [X341] : (! [X342] : (! [X343] : (! [X344] : (! [X345] : (! [X346] : (! [X347] : (! [X348] : (! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : ((! [X359] : (? [X360] : (~ p1(X360) & r1(X359, X360)) | ! [X361] : (! [X362] : (p1(X362) | ~ r1(X361, X362)) | ~ r1(X359, X361)) | ~ r1(X358, X359)) & (? [X363] : (! [X364] : (~ p1(X364) | ! [X365] : (p1(X365) | ~ r1(X364, X365)) | ~ r1(X363, X364)) & ~ p1(X363) & r1(X358, X363)) | ! [X366] : (? [X367] : (~ p1(X367) & r1(X366, X367)) | ~ r1(X358, X366)) | p1(X358))) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X348, X349)) | ~ r1(X347, X348)) | ~ r1(X346, X347)) | ~ r1(X345, X346)) | ~ r1(X344, X345)) | ~ r1(X343, X344)) | ~ r1(X342, X343)) | ~ r1(X341, X342)) | ~ r1(X0, X341)) & ! [X368] : (! [X369] : (! [X370] : (! [X371] : (! [X372] : (! [X373] : (! [X374] : (! [X375] : (! [X376] : (! [X377] : (! [X378] : (! [X379] : (! [X380] : (! [X381] : (! [X382] : (! [X383] : (! [X384] : (! [X385] : (! [X386] : ((! [X387] : (? [X388] : (~ p1(X388) & r1(X387, X388)) | ! [X389] : (! [X390] : (p1(X390) | ~ r1(X389, X390)) | ~ r1(X387, X389)) | ~ r1(X386, X387)) & (? [X391] : (! [X392] : (~ p1(X392) | ! [X393] : (p1(X393) | ~ r1(X392, X393)) | ~ r1(X391, X392)) & ~ p1(X391) & r1(X386, X391)) | ! [X394] : (? [X395] : (~ p1(X395) & r1(X394, X395)) | ~ r1(X386, X394)) | p1(X386))) | ~ r1(X385, X386)) | ~ r1(X384, X385)) | ~ r1(X383, X384)) | ~ r1(X382, X383)) | ~ r1(X381, X382)) | ~ r1(X380, X381)) | ~ r1(X379, X380)) | ~ r1(X378, X379)) | ~ r1(X377, X378)) | ~ r1(X376, X377)) | ~ r1(X375, X376)) | ~ r1(X374, X375)) | ~ r1(X373, X374)) | ~ r1(X372, X373)) | ~ r1(X371, X372)) | ~ r1(X370, X371)) | ~ r1(X369, X370)) | ~ r1(X368, X369)) | ~ r1(X0, X368))), inference(definition_folding, [], [f6, e10, e9, e8, e7])).
fof(f7, plain, ! [X173] : (? [X190] : (! [X191] : ((p1(X191) & ? [X192] : (~ p1(X192) & r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) & p1(X190) & ? [X195] : (~ p1(X195) & r1(X190, X195)) & r1(X173, X190)) | ~ sP0(X173)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X173] : (sP0(X173) <=> ? [X190] : (! [X191] : ((p1(X191) & ? [X192] : (~ p1(X192) & r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) & p1(X190) & ? [X195] : (~ p1(X195) & r1(X190, X195)) & r1(X173, X190))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X173] : (! [X177] : (! [X178] : (? [X179] : (p1(X179) & ? [X180] : (~ p1(X180) & r1(X179, X180)) & r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177)) | ~ sP1(X173)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X173] : (sP1(X173) <=> ! [X177] : (! [X178] : (? [X179] : (p1(X179) & ? [X180] : (~ p1(X180) & r1(X179, X180)) & r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ! [X173] : (? [X185] : (! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) & ~ p1(X185) & r1(X173, X185)) | ! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173) | ~ sP2(X173)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X173] : (sP2(X173) <=> (? [X185] : (! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) & ~ p1(X185) & r1(X173, X185)) | ! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ! [X173] : (sP0(X173) | ! [X196] : (? [X197] : (p1(X197) & ? [X198] : (~ p1(X198) & r1(X197, X198)) & r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)) | ~ sP3(X173)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X173] : (sP3(X173) <=> (sP0(X173) | ! [X196] : (? [X197] : (p1(X197) & ? [X198] : (~ p1(X198) & r1(X197, X198)) & r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f6, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(X0, X146)) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | ! [X177] : (! [X178] : (? [X179] : (p1(X179) & ? [X180] : (~ p1(X180) & r1(X179, X180)) & r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177))) & ! [X181] : (? [X182] : (~ p1(X182) & r1(X181, X182)) | ! [X183] : (! [X184] : (p1(X184) | ~ r1(X183, X184)) | ~ r1(X181, X183)) | ~ r1(X173, X181)) & (? [X185] : (! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) & ~ p1(X185) & r1(X173, X185)) | ! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173)) & (? [X190] : (! [X191] : ((p1(X191) & ? [X192] : (~ p1(X192) & r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) & p1(X190) & ? [X195] : (~ p1(X195) & r1(X190, X195)) & r1(X173, X190)) | ! [X196] : (? [X197] : (p1(X197) & ? [X198] : (~ p1(X198) & r1(X197, X198)) & r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)))) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) & ! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X0, X200)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : ((! [X234] : (? [X235] : (~ p1(X235) & r1(X234, X235)) | ! [X236] : (! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X234, X236)) | ~ r1(X233, X234)) & (? [X238] : (! [X239] : (~ p1(X239) | ! [X240] : (p1(X240) | ~ r1(X239, X240)) | ~ r1(X238, X239)) & ~ p1(X238) & r1(X233, X238)) | ! [X241] : (? [X242] : (~ p1(X242) & r1(X241, X242)) | ~ r1(X233, X241)) | p1(X233))) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221)) & ! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : ((! [X257] : (? [X258] : (~ p1(X258) & r1(X257, X258)) | ! [X259] : (! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X257, X259)) | ~ r1(X256, X257)) & (? [X261] : (! [X262] : (~ p1(X262) | ! [X263] : (p1(X263) | ~ r1(X262, X263)) | ~ r1(X261, X262)) & ~ p1(X261) & r1(X256, X261)) | ! [X264] : (? [X265] : (~ p1(X265) & r1(X264, X265)) | ~ r1(X256, X264)) | p1(X256))) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X0, X243)) & ! [X266] : (! [X267] : (! [X268] : (! [X269] : (! [X270] : (! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : ((! [X281] : (? [X282] : (~ p1(X282) & r1(X281, X282)) | ! [X283] : (! [X284] : (p1(X284) | ~ r1(X283, X284)) | ~ r1(X281, X283)) | ~ r1(X280, X281)) & (? [X285] : (! [X286] : (~ p1(X286) | ! [X287] : (p1(X287) | ~ r1(X286, X287)) | ~ r1(X285, X286)) & ~ p1(X285) & r1(X280, X285)) | ! [X288] : (? [X289] : (~ p1(X289) & r1(X288, X289)) | ~ r1(X280, X288)) | p1(X280))) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X270, X271)) | ~ r1(X269, X270)) | ~ r1(X268, X269)) | ~ r1(X267, X268)) | ~ r1(X266, X267)) | ~ r1(X0, X266)) & ! [X290] : (! [X291] : (! [X292] : (! [X293] : (! [X294] : (! [X295] : (! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : ((! [X306] : (? [X307] : (~ p1(X307) & r1(X306, X307)) | ! [X308] : (! [X309] : (p1(X309) | ~ r1(X308, X309)) | ~ r1(X306, X308)) | ~ r1(X305, X306)) & (? [X310] : (! [X311] : (~ p1(X311) | ! [X312] : (p1(X312) | ~ r1(X311, X312)) | ~ r1(X310, X311)) & ~ p1(X310) & r1(X305, X310)) | ! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ~ r1(X305, X313)) | p1(X305))) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X295, X296)) | ~ r1(X294, X295)) | ~ r1(X293, X294)) | ~ r1(X292, X293)) | ~ r1(X291, X292)) | ~ r1(X290, X291)) | ~ r1(X0, X290)) & ! [X315] : (! [X316] : (! [X317] : (! [X318] : (! [X319] : (! [X320] : (! [X321] : (! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : ((! [X332] : (? [X333] : (~ p1(X333) & r1(X332, X333)) | ! [X334] : (! [X335] : (p1(X335) | ~ r1(X334, X335)) | ~ r1(X332, X334)) | ~ r1(X331, X332)) & (? [X336] : (! [X337] : (~ p1(X337) | ! [X338] : (p1(X338) | ~ r1(X337, X338)) | ~ r1(X336, X337)) & ~ p1(X336) & r1(X331, X336)) | ! [X339] : (? [X340] : (~ p1(X340) & r1(X339, X340)) | ~ r1(X331, X339)) | p1(X331))) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X321, X322)) | ~ r1(X320, X321)) | ~ r1(X319, X320)) | ~ r1(X318, X319)) | ~ r1(X317, X318)) | ~ r1(X316, X317)) | ~ r1(X315, X316)) | ~ r1(X0, X315)) & ! [X341] : (! [X342] : (! [X343] : (! [X344] : (! [X345] : (! [X346] : (! [X347] : (! [X348] : (! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : ((! [X359] : (? [X360] : (~ p1(X360) & r1(X359, X360)) | ! [X361] : (! [X362] : (p1(X362) | ~ r1(X361, X362)) | ~ r1(X359, X361)) | ~ r1(X358, X359)) & (? [X363] : (! [X364] : (~ p1(X364) | ! [X365] : (p1(X365) | ~ r1(X364, X365)) | ~ r1(X363, X364)) & ~ p1(X363) & r1(X358, X363)) | ! [X366] : (? [X367] : (~ p1(X367) & r1(X366, X367)) | ~ r1(X358, X366)) | p1(X358))) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X348, X349)) | ~ r1(X347, X348)) | ~ r1(X346, X347)) | ~ r1(X345, X346)) | ~ r1(X344, X345)) | ~ r1(X343, X344)) | ~ r1(X342, X343)) | ~ r1(X341, X342)) | ~ r1(X0, X341)) & ! [X368] : (! [X369] : (! [X370] : (! [X371] : (! [X372] : (! [X373] : (! [X374] : (! [X375] : (! [X376] : (! [X377] : (! [X378] : (! [X379] : (! [X380] : (! [X381] : (! [X382] : (! [X383] : (! [X384] : (! [X385] : (! [X386] : ((! [X387] : (? [X388] : (~ p1(X388) & r1(X387, X388)) | ! [X389] : (! [X390] : (p1(X390) | ~ r1(X389, X390)) | ~ r1(X387, X389)) | ~ r1(X386, X387)) & (? [X391] : (! [X392] : (~ p1(X392) | ! [X393] : (p1(X393) | ~ r1(X392, X393)) | ~ r1(X391, X392)) & ~ p1(X391) & r1(X386, X391)) | ! [X394] : (? [X395] : (~ p1(X395) & r1(X394, X395)) | ~ r1(X386, X394)) | p1(X386))) | ~ r1(X385, X386)) | ~ r1(X384, X385)) | ~ r1(X383, X384)) | ~ r1(X382, X383)) | ~ r1(X381, X382)) | ~ r1(X380, X381)) | ~ r1(X379, X380)) | ~ r1(X378, X379)) | ~ r1(X377, X378)) | ~ r1(X376, X377)) | ~ r1(X375, X376)) | ~ r1(X374, X375)) | ~ r1(X373, X374)) | ~ r1(X372, X373)) | ~ r1(X371, X372)) | ~ r1(X370, X371)) | ~ r1(X369, X370)) | ~ r1(X368, X369)) | ~ r1(X0, X368))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : ((! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (? [X101] : (~ p1(X101) & r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (? [X104] : (! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) & ~ p1(X104) & r1(X99, X104)) | ! [X107] : (? [X108] : (~ p1(X108) & r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (? [X119] : (~ p1(X119) & r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (? [X122] : (! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) & ~ p1(X122) & r1(X117, X122)) | ! [X125] : (? [X126] : (~ p1(X126) & r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (? [X141] : (! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) & ~ p1(X141) & r1(X136, X141)) | ! [X144] : (? [X145] : (~ p1(X145) & r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ? [X146] : (? [X147] : (? [X148] : (? [X149] : (? [X150] : (? [X151] : (? [X152] : (? [X153] : (? [X154] : (? [X155] : (? [X156] : (! [X157] : (? [X158] : (p1(X158) & ? [X159] : (~ p1(X159) & r1(X158, X159)) & r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) & ? [X160] : (! [X161] : (p1(X161) | ~ r1(X160, X161)) & r1(X156, X160)) & ? [X162] : (~ p1(X162) & r1(X156, X162)) & r1(X155, X156)) & r1(X154, X155)) & r1(X153, X154)) & r1(X152, X153)) & r1(X151, X152)) & r1(X150, X151)) & r1(X149, X150)) & r1(X148, X149)) & r1(X147, X148)) & r1(X146, X147)) & r1(X0, X146))) & ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((? [X174] : (! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) & ~ p1(X174) & r1(X173, X174)) | ! [X177] : (! [X178] : (? [X179] : (p1(X179) & ? [X180] : (~ p1(X180) & r1(X179, X180)) & r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177))) & ! [X181] : (? [X182] : (~ p1(X182) & r1(X181, X182)) | ! [X183] : (! [X184] : (p1(X184) | ~ r1(X183, X184)) | ~ r1(X181, X183)) | ~ r1(X173, X181)) & (? [X185] : (! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) & ~ p1(X185) & r1(X173, X185)) | ! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173)) & (? [X190] : (! [X191] : ((p1(X191) & ? [X192] : (~ p1(X192) & r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) & p1(X190) & ? [X195] : (~ p1(X195) & r1(X190, X195)) & r1(X173, X190)) | ! [X196] : (? [X197] : (p1(X197) & ? [X198] : (~ p1(X198) & r1(X197, X198)) & r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)))) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) & ! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X0, X200)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : ((! [X234] : (? [X235] : (~ p1(X235) & r1(X234, X235)) | ! [X236] : (! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X234, X236)) | ~ r1(X233, X234)) & (? [X238] : (! [X239] : (~ p1(X239) | ! [X240] : (p1(X240) | ~ r1(X239, X240)) | ~ r1(X238, X239)) & ~ p1(X238) & r1(X233, X238)) | ! [X241] : (? [X242] : (~ p1(X242) & r1(X241, X242)) | ~ r1(X233, X241)) | p1(X233))) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221)) & ! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : ((! [X257] : (? [X258] : (~ p1(X258) & r1(X257, X258)) | ! [X259] : (! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X257, X259)) | ~ r1(X256, X257)) & (? [X261] : (! [X262] : (~ p1(X262) | ! [X263] : (p1(X263) | ~ r1(X262, X263)) | ~ r1(X261, X262)) & ~ p1(X261) & r1(X256, X261)) | ! [X264] : (? [X265] : (~ p1(X265) & r1(X264, X265)) | ~ r1(X256, X264)) | p1(X256))) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X0, X243)) & ! [X266] : (! [X267] : (! [X268] : (! [X269] : (! [X270] : (! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : ((! [X281] : (? [X282] : (~ p1(X282) & r1(X281, X282)) | ! [X283] : (! [X284] : (p1(X284) | ~ r1(X283, X284)) | ~ r1(X281, X283)) | ~ r1(X280, X281)) & (? [X285] : (! [X286] : (~ p1(X286) | ! [X287] : (p1(X287) | ~ r1(X286, X287)) | ~ r1(X285, X286)) & ~ p1(X285) & r1(X280, X285)) | ! [X288] : (? [X289] : (~ p1(X289) & r1(X288, X289)) | ~ r1(X280, X288)) | p1(X280))) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X270, X271)) | ~ r1(X269, X270)) | ~ r1(X268, X269)) | ~ r1(X267, X268)) | ~ r1(X266, X267)) | ~ r1(X0, X266)) & ! [X290] : (! [X291] : (! [X292] : (! [X293] : (! [X294] : (! [X295] : (! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : ((! [X306] : (? [X307] : (~ p1(X307) & r1(X306, X307)) | ! [X308] : (! [X309] : (p1(X309) | ~ r1(X308, X309)) | ~ r1(X306, X308)) | ~ r1(X305, X306)) & (? [X310] : (! [X311] : (~ p1(X311) | ! [X312] : (p1(X312) | ~ r1(X311, X312)) | ~ r1(X310, X311)) & ~ p1(X310) & r1(X305, X310)) | ! [X313] : (? [X314] : (~ p1(X314) & r1(X313, X314)) | ~ r1(X305, X313)) | p1(X305))) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X295, X296)) | ~ r1(X294, X295)) | ~ r1(X293, X294)) | ~ r1(X292, X293)) | ~ r1(X291, X292)) | ~ r1(X290, X291)) | ~ r1(X0, X290)) & ! [X315] : (! [X316] : (! [X317] : (! [X318] : (! [X319] : (! [X320] : (! [X321] : (! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : ((! [X332] : (? [X333] : (~ p1(X333) & r1(X332, X333)) | ! [X334] : (! [X335] : (p1(X335) | ~ r1(X334, X335)) | ~ r1(X332, X334)) | ~ r1(X331, X332)) & (? [X336] : (! [X337] : (~ p1(X337) | ! [X338] : (p1(X338) | ~ r1(X337, X338)) | ~ r1(X336, X337)) & ~ p1(X336) & r1(X331, X336)) | ! [X339] : (? [X340] : (~ p1(X340) & r1(X339, X340)) | ~ r1(X331, X339)) | p1(X331))) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X321, X322)) | ~ r1(X320, X321)) | ~ r1(X319, X320)) | ~ r1(X318, X319)) | ~ r1(X317, X318)) | ~ r1(X316, X317)) | ~ r1(X315, X316)) | ~ r1(X0, X315)) & ! [X341] : (! [X342] : (! [X343] : (! [X344] : (! [X345] : (! [X346] : (! [X347] : (! [X348] : (! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : ((! [X359] : (? [X360] : (~ p1(X360) & r1(X359, X360)) | ! [X361] : (! [X362] : (p1(X362) | ~ r1(X361, X362)) | ~ r1(X359, X361)) | ~ r1(X358, X359)) & (? [X363] : (! [X364] : (~ p1(X364) | ! [X365] : (p1(X365) | ~ r1(X364, X365)) | ~ r1(X363, X364)) & ~ p1(X363) & r1(X358, X363)) | ! [X366] : (? [X367] : (~ p1(X367) & r1(X366, X367)) | ~ r1(X358, X366)) | p1(X358))) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X348, X349)) | ~ r1(X347, X348)) | ~ r1(X346, X347)) | ~ r1(X345, X346)) | ~ r1(X344, X345)) | ~ r1(X343, X344)) | ~ r1(X342, X343)) | ~ r1(X341, X342)) | ~ r1(X0, X341)) & ! [X368] : (! [X369] : (! [X370] : (! [X371] : (! [X372] : (! [X373] : (! [X374] : (! [X375] : (! [X376] : (! [X377] : (! [X378] : (! [X379] : (! [X380] : (! [X381] : (! [X382] : (! [X383] : (! [X384] : (! [X385] : (! [X386] : ((! [X387] : (? [X388] : (~ p1(X388) & r1(X387, X388)) | ! [X389] : (! [X390] : (p1(X390) | ~ r1(X389, X390)) | ~ r1(X387, X389)) | ~ r1(X386, X387)) & (? [X391] : (! [X392] : (~ p1(X392) | ! [X393] : (p1(X393) | ~ r1(X392, X393)) | ~ r1(X391, X392)) & ~ p1(X391) & r1(X386, X391)) | ! [X394] : (? [X395] : (~ p1(X395) & r1(X394, X395)) | ~ r1(X386, X394)) | p1(X386))) | ~ r1(X385, X386)) | ~ r1(X384, X385)) | ~ r1(X383, X384)) | ~ r1(X382, X383)) | ~ r1(X381, X382)) | ~ r1(X380, X381)) | ~ r1(X379, X380)) | ~ r1(X378, X379)) | ~ r1(X377, X378)) | ~ r1(X376, X377)) | ~ r1(X375, X376)) | ~ r1(X374, X375)) | ~ r1(X373, X374)) | ~ r1(X372, X373)) | ~ r1(X371, X372)) | ~ r1(X370, X371)) | ~ r1(X369, X370)) | ~ r1(X368, X369)) | ~ r1(X0, X368))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (~ ! [X68] : (p1(X68) | ~ r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (~ ! [X71] : (~ ! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) | p1(X71) | ~ r1(X66, X71)) | ! [X74] : (~ ! [X75] : (p1(X75) | ~ r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (~ ! [X84] : (p1(X84) | ~ r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (~ ! [X87] : (~ ! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | p1(X87) | ~ r1(X82, X87)) | ! [X90] : (~ ! [X91] : (p1(X91) | ~ r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (~ ! [X101] : (p1(X101) | ~ r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (~ ! [X104] : (~ ! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | p1(X104) | ~ r1(X99, X104)) | ! [X107] : (~ ! [X108] : (p1(X108) | ~ r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (~ ! [X119] : (p1(X119) | ~ r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (~ ! [X122] : (~ ! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | p1(X122) | ~ r1(X117, X122)) | ! [X125] : (~ ! [X126] : (p1(X126) | ~ r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (~ ! [X138] : (p1(X138) | ~ r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (~ ! [X141] : (~ ! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | p1(X141) | ~ r1(X136, X141)) | ! [X144] : (~ ! [X145] : (p1(X145) | ~ r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ~ ! [X146] : (! [X147] : (! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (! [X153] : (! [X154] : (! [X155] : (! [X156] : (~ ! [X157] : (~ ! [X158] : (~ p1(X158) | ! [X159] : (p1(X159) | ~ r1(X158, X159)) | ~ r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) | ! [X160] : (~ ! [X161] : (p1(X161) | ~ r1(X160, X161)) | ~ r1(X156, X160)) | ! [X162] : (p1(X162) | ~ r1(X156, X162)) | ~ r1(X155, X156)) | ~ r1(X154, X155)) | ~ r1(X153, X154)) | ~ r1(X152, X153)) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X0, X146))) | ~ ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((~ ! [X174] : (~ ! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) | p1(X174) | ~ r1(X173, X174)) | ! [X177] : (! [X178] : (~ ! [X179] : (~ p1(X179) | ! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177))) & ! [X181] : (~ ! [X182] : (p1(X182) | ~ r1(X181, X182)) | ! [X183] : (! [X184] : (p1(X184) | ~ r1(X183, X184)) | ~ r1(X181, X183)) | ~ r1(X173, X181)) & (~ ! [X185] : (~ ! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | p1(X185) | ~ r1(X173, X185)) | ! [X188] : (~ ! [X189] : (p1(X189) | ~ r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173)) & (~ ! [X190] : (~ ! [X191] : (~ (~ p1(X191) | ! [X192] : (p1(X192) | ~ r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) | ~ p1(X190) | ! [X195] : (p1(X195) | ~ r1(X190, X195)) | ~ r1(X173, X190)) | ! [X196] : (~ ! [X197] : (~ p1(X197) | ! [X198] : (p1(X198) | ~ r1(X197, X198)) | ~ r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)))) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) | ~ ! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (~ ! [X213] : (p1(X213) | ~ r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (~ ! [X216] : (~ ! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) | p1(X216) | ~ r1(X211, X216)) | ! [X219] : (~ ! [X220] : (p1(X220) | ~ r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X0, X200)) | ~ ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : ((! [X234] : (~ ! [X235] : (p1(X235) | ~ r1(X234, X235)) | ! [X236] : (! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X234, X236)) | ~ r1(X233, X234)) & (~ ! [X238] : (~ ! [X239] : (~ p1(X239) | ! [X240] : (p1(X240) | ~ r1(X239, X240)) | ~ r1(X238, X239)) | p1(X238) | ~ r1(X233, X238)) | ! [X241] : (~ ! [X242] : (p1(X242) | ~ r1(X241, X242)) | ~ r1(X233, X241)) | p1(X233))) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221)) | ~ ! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : ((! [X257] : (~ ! [X258] : (p1(X258) | ~ r1(X257, X258)) | ! [X259] : (! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X257, X259)) | ~ r1(X256, X257)) & (~ ! [X261] : (~ ! [X262] : (~ p1(X262) | ! [X263] : (p1(X263) | ~ r1(X262, X263)) | ~ r1(X261, X262)) | p1(X261) | ~ r1(X256, X261)) | ! [X264] : (~ ! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X256, X264)) | p1(X256))) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X0, X243)) | ~ ! [X266] : (! [X267] : (! [X268] : (! [X269] : (! [X270] : (! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : ((! [X281] : (~ ! [X282] : (p1(X282) | ~ r1(X281, X282)) | ! [X283] : (! [X284] : (p1(X284) | ~ r1(X283, X284)) | ~ r1(X281, X283)) | ~ r1(X280, X281)) & (~ ! [X285] : (~ ! [X286] : (~ p1(X286) | ! [X287] : (p1(X287) | ~ r1(X286, X287)) | ~ r1(X285, X286)) | p1(X285) | ~ r1(X280, X285)) | ! [X288] : (~ ! [X289] : (p1(X289) | ~ r1(X288, X289)) | ~ r1(X280, X288)) | p1(X280))) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X270, X271)) | ~ r1(X269, X270)) | ~ r1(X268, X269)) | ~ r1(X267, X268)) | ~ r1(X266, X267)) | ~ r1(X0, X266)) | ~ ! [X290] : (! [X291] : (! [X292] : (! [X293] : (! [X294] : (! [X295] : (! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : ((! [X306] : (~ ! [X307] : (p1(X307) | ~ r1(X306, X307)) | ! [X308] : (! [X309] : (p1(X309) | ~ r1(X308, X309)) | ~ r1(X306, X308)) | ~ r1(X305, X306)) & (~ ! [X310] : (~ ! [X311] : (~ p1(X311) | ! [X312] : (p1(X312) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | p1(X310) | ~ r1(X305, X310)) | ! [X313] : (~ ! [X314] : (p1(X314) | ~ r1(X313, X314)) | ~ r1(X305, X313)) | p1(X305))) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X295, X296)) | ~ r1(X294, X295)) | ~ r1(X293, X294)) | ~ r1(X292, X293)) | ~ r1(X291, X292)) | ~ r1(X290, X291)) | ~ r1(X0, X290)) | ~ ! [X315] : (! [X316] : (! [X317] : (! [X318] : (! [X319] : (! [X320] : (! [X321] : (! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : ((! [X332] : (~ ! [X333] : (p1(X333) | ~ r1(X332, X333)) | ! [X334] : (! [X335] : (p1(X335) | ~ r1(X334, X335)) | ~ r1(X332, X334)) | ~ r1(X331, X332)) & (~ ! [X336] : (~ ! [X337] : (~ p1(X337) | ! [X338] : (p1(X338) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | p1(X336) | ~ r1(X331, X336)) | ! [X339] : (~ ! [X340] : (p1(X340) | ~ r1(X339, X340)) | ~ r1(X331, X339)) | p1(X331))) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X321, X322)) | ~ r1(X320, X321)) | ~ r1(X319, X320)) | ~ r1(X318, X319)) | ~ r1(X317, X318)) | ~ r1(X316, X317)) | ~ r1(X315, X316)) | ~ r1(X0, X315)) | ~ ! [X341] : (! [X342] : (! [X343] : (! [X344] : (! [X345] : (! [X346] : (! [X347] : (! [X348] : (! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : ((! [X359] : (~ ! [X360] : (p1(X360) | ~ r1(X359, X360)) | ! [X361] : (! [X362] : (p1(X362) | ~ r1(X361, X362)) | ~ r1(X359, X361)) | ~ r1(X358, X359)) & (~ ! [X363] : (~ ! [X364] : (~ p1(X364) | ! [X365] : (p1(X365) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | p1(X363) | ~ r1(X358, X363)) | ! [X366] : (~ ! [X367] : (p1(X367) | ~ r1(X366, X367)) | ~ r1(X358, X366)) | p1(X358))) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X348, X349)) | ~ r1(X347, X348)) | ~ r1(X346, X347)) | ~ r1(X345, X346)) | ~ r1(X344, X345)) | ~ r1(X343, X344)) | ~ r1(X342, X343)) | ~ r1(X341, X342)) | ~ r1(X0, X341)) | ~ ! [X368] : (! [X369] : (! [X370] : (! [X371] : (! [X372] : (! [X373] : (! [X374] : (! [X375] : (! [X376] : (! [X377] : (! [X378] : (! [X379] : (! [X380] : (! [X381] : (! [X382] : (! [X383] : (! [X384] : (! [X385] : (! [X386] : ((! [X387] : (~ ! [X388] : (p1(X388) | ~ r1(X387, X388)) | ! [X389] : (! [X390] : (p1(X390) | ~ r1(X389, X390)) | ~ r1(X387, X389)) | ~ r1(X386, X387)) & (~ ! [X391] : (~ ! [X392] : (~ p1(X392) | ! [X393] : (p1(X393) | ~ r1(X392, X393)) | ~ r1(X391, X392)) | p1(X391) | ~ r1(X386, X391)) | ! [X394] : (~ ! [X395] : (p1(X395) | ~ r1(X394, X395)) | ~ r1(X386, X394)) | p1(X386))) | ~ r1(X385, X386)) | ~ r1(X384, X385)) | ~ r1(X383, X384)) | ~ r1(X382, X383)) | ~ r1(X381, X382)) | ~ r1(X380, X381)) | ~ r1(X379, X380)) | ~ r1(X378, X379)) | ~ r1(X377, X378)) | ~ r1(X376, X377)) | ~ r1(X375, X376)) | ~ r1(X374, X375)) | ~ r1(X373, X374)) | ~ r1(X372, X373)) | ~ r1(X371, X372)) | ~ r1(X370, X371)) | ~ r1(X369, X370)) | ~ r1(X368, X369)) | ~ r1(X0, X368))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (~ ! [X68] : (p1(X68) | ~ r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (~ ! [X71] : (~ ! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) | p1(X71) | ~ r1(X66, X71)) | ! [X74] : (~ ! [X75] : (p1(X75) | ~ r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (~ ! [X84] : (p1(X84) | ~ r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (~ ! [X87] : (~ ! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | p1(X87) | ~ r1(X82, X87)) | ! [X90] : (~ ! [X91] : (p1(X91) | ~ r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : ((! [X100] : (~ ! [X101] : (p1(X101) | ~ r1(X100, X101)) | ! [X102] : (! [X103] : (p1(X103) | ~ r1(X102, X103)) | ~ r1(X100, X102)) | ~ r1(X99, X100)) & (~ ! [X104] : (~ ! [X105] : (~ p1(X105) | ! [X106] : (p1(X106) | ~ r1(X105, X106)) | ~ r1(X104, X105)) | p1(X104) | ~ r1(X99, X104)) | ! [X107] : (~ ! [X108] : (p1(X108) | ~ r1(X107, X108)) | ~ r1(X99, X107)) | p1(X99))) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92)) & ! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (! [X114] : (! [X115] : (! [X116] : (! [X117] : ((! [X118] : (~ ! [X119] : (p1(X119) | ~ r1(X118, X119)) | ! [X120] : (! [X121] : (p1(X121) | ~ r1(X120, X121)) | ~ r1(X118, X120)) | ~ r1(X117, X118)) & (~ ! [X122] : (~ ! [X123] : (~ p1(X123) | ! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | p1(X122) | ~ r1(X117, X122)) | ! [X125] : (~ ! [X126] : (p1(X126) | ~ r1(X125, X126)) | ~ r1(X117, X125)) | p1(X117))) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | ~ r1(X113, X114)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X0, X109)) & ! [X127] : (! [X128] : (! [X129] : (! [X130] : (! [X131] : (! [X132] : (! [X133] : (! [X134] : (! [X135] : (! [X136] : ((! [X137] : (~ ! [X138] : (p1(X138) | ~ r1(X137, X138)) | ! [X139] : (! [X140] : (p1(X140) | ~ r1(X139, X140)) | ~ r1(X137, X139)) | ~ r1(X136, X137)) & (~ ! [X141] : (~ ! [X142] : (~ p1(X142) | ! [X143] : (p1(X143) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | p1(X141) | ~ r1(X136, X141)) | ! [X144] : (~ ! [X145] : (p1(X145) | ~ r1(X144, X145)) | ~ r1(X136, X144)) | p1(X136))) | ~ r1(X135, X136)) | ~ r1(X134, X135)) | ~ r1(X133, X134)) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | ~ r1(X130, X131)) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X0, X127)) & ~ ! [X146] : (! [X147] : (! [X148] : (! [X149] : (! [X150] : (! [X151] : (! [X152] : (! [X153] : (! [X154] : (! [X155] : (! [X156] : (~ ! [X157] : (~ ! [X158] : (~ p1(X158) | ! [X159] : (p1(X159) | ~ r1(X158, X159)) | ~ r1(X157, X158)) | p1(X157) | ~ r1(X156, X157)) | ! [X160] : (~ ! [X161] : (p1(X161) | ~ r1(X160, X161)) | ~ r1(X156, X160)) | ! [X162] : (p1(X162) | ~ r1(X156, X162)) | ~ r1(X155, X156)) | ~ r1(X154, X155)) | ~ r1(X153, X154)) | ~ r1(X152, X153)) | ~ r1(X151, X152)) | ~ r1(X150, X151)) | ~ r1(X149, X150)) | ~ r1(X148, X149)) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X0, X146))) | ~ ! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : (! [X169] : (! [X170] : (! [X171] : (! [X172] : (! [X173] : (((~ ! [X174] : (~ ! [X175] : (~ p1(X175) | ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X174, X175)) | p1(X174) | ~ r1(X173, X174)) | ! [X177] : (! [X178] : (~ ! [X179] : (~ p1(X179) | ! [X180] : (p1(X180) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177))) & ! [X181] : (~ ! [X182] : (p1(X182) | ~ r1(X181, X182)) | ! [X183] : (! [X184] : (p1(X184) | ~ r1(X183, X184)) | ~ r1(X181, X183)) | ~ r1(X173, X181)) & (~ ! [X185] : (~ ! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | p1(X185) | ~ r1(X173, X185)) | ! [X188] : (~ ! [X189] : (p1(X189) | ~ r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173)) & (~ ! [X190] : (~ ! [X191] : (~ (~ p1(X191) | ! [X192] : (p1(X192) | ~ r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) | ~ p1(X190) | ! [X195] : (p1(X195) | ~ r1(X190, X195)) | ~ r1(X173, X190)) | ! [X196] : (~ ! [X197] : (~ p1(X197) | ! [X198] : (p1(X198) | ~ r1(X197, X198)) | ~ r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)))) | ~ r1(X172, X173)) | ~ r1(X171, X172)) | ~ r1(X170, X171)) | ~ r1(X169, X170)) | ~ r1(X168, X169)) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X0, X163)) | ~ ! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (~ ! [X213] : (p1(X213) | ~ r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (~ ! [X216] : (~ ! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) | p1(X216) | ~ r1(X211, X216)) | ! [X219] : (~ ! [X220] : (p1(X220) | ~ r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X0, X200)) | ~ ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : ((! [X234] : (~ ! [X235] : (p1(X235) | ~ r1(X234, X235)) | ! [X236] : (! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X234, X236)) | ~ r1(X233, X234)) & (~ ! [X238] : (~ ! [X239] : (~ p1(X239) | ! [X240] : (p1(X240) | ~ r1(X239, X240)) | ~ r1(X238, X239)) | p1(X238) | ~ r1(X233, X238)) | ! [X241] : (~ ! [X242] : (p1(X242) | ~ r1(X241, X242)) | ~ r1(X233, X241)) | p1(X233))) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221)) | ~ ! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : (! [X254] : (! [X255] : (! [X256] : ((! [X257] : (~ ! [X258] : (p1(X258) | ~ r1(X257, X258)) | ! [X259] : (! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X257, X259)) | ~ r1(X256, X257)) & (~ ! [X261] : (~ ! [X262] : (~ p1(X262) | ! [X263] : (p1(X263) | ~ r1(X262, X263)) | ~ r1(X261, X262)) | p1(X261) | ~ r1(X256, X261)) | ! [X264] : (~ ! [X265] : (p1(X265) | ~ r1(X264, X265)) | ~ r1(X256, X264)) | p1(X256))) | ~ r1(X255, X256)) | ~ r1(X254, X255)) | ~ r1(X253, X254)) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X0, X243)) | ~ ! [X266] : (! [X267] : (! [X268] : (! [X269] : (! [X270] : (! [X271] : (! [X272] : (! [X273] : (! [X274] : (! [X275] : (! [X276] : (! [X277] : (! [X278] : (! [X279] : (! [X280] : ((! [X281] : (~ ! [X282] : (p1(X282) | ~ r1(X281, X282)) | ! [X283] : (! [X284] : (p1(X284) | ~ r1(X283, X284)) | ~ r1(X281, X283)) | ~ r1(X280, X281)) & (~ ! [X285] : (~ ! [X286] : (~ p1(X286) | ! [X287] : (p1(X287) | ~ r1(X286, X287)) | ~ r1(X285, X286)) | p1(X285) | ~ r1(X280, X285)) | ! [X288] : (~ ! [X289] : (p1(X289) | ~ r1(X288, X289)) | ~ r1(X280, X288)) | p1(X280))) | ~ r1(X279, X280)) | ~ r1(X278, X279)) | ~ r1(X277, X278)) | ~ r1(X276, X277)) | ~ r1(X275, X276)) | ~ r1(X274, X275)) | ~ r1(X273, X274)) | ~ r1(X272, X273)) | ~ r1(X271, X272)) | ~ r1(X270, X271)) | ~ r1(X269, X270)) | ~ r1(X268, X269)) | ~ r1(X267, X268)) | ~ r1(X266, X267)) | ~ r1(X0, X266)) | ~ ! [X290] : (! [X291] : (! [X292] : (! [X293] : (! [X294] : (! [X295] : (! [X296] : (! [X297] : (! [X298] : (! [X299] : (! [X300] : (! [X301] : (! [X302] : (! [X303] : (! [X304] : (! [X305] : ((! [X306] : (~ ! [X307] : (p1(X307) | ~ r1(X306, X307)) | ! [X308] : (! [X309] : (p1(X309) | ~ r1(X308, X309)) | ~ r1(X306, X308)) | ~ r1(X305, X306)) & (~ ! [X310] : (~ ! [X311] : (~ p1(X311) | ! [X312] : (p1(X312) | ~ r1(X311, X312)) | ~ r1(X310, X311)) | p1(X310) | ~ r1(X305, X310)) | ! [X313] : (~ ! [X314] : (p1(X314) | ~ r1(X313, X314)) | ~ r1(X305, X313)) | p1(X305))) | ~ r1(X304, X305)) | ~ r1(X303, X304)) | ~ r1(X302, X303)) | ~ r1(X301, X302)) | ~ r1(X300, X301)) | ~ r1(X299, X300)) | ~ r1(X298, X299)) | ~ r1(X297, X298)) | ~ r1(X296, X297)) | ~ r1(X295, X296)) | ~ r1(X294, X295)) | ~ r1(X293, X294)) | ~ r1(X292, X293)) | ~ r1(X291, X292)) | ~ r1(X290, X291)) | ~ r1(X0, X290)) | ~ ! [X315] : (! [X316] : (! [X317] : (! [X318] : (! [X319] : (! [X320] : (! [X321] : (! [X322] : (! [X323] : (! [X324] : (! [X325] : (! [X326] : (! [X327] : (! [X328] : (! [X329] : (! [X330] : (! [X331] : ((! [X332] : (~ ! [X333] : (p1(X333) | ~ r1(X332, X333)) | ! [X334] : (! [X335] : (p1(X335) | ~ r1(X334, X335)) | ~ r1(X332, X334)) | ~ r1(X331, X332)) & (~ ! [X336] : (~ ! [X337] : (~ p1(X337) | ! [X338] : (p1(X338) | ~ r1(X337, X338)) | ~ r1(X336, X337)) | p1(X336) | ~ r1(X331, X336)) | ! [X339] : (~ ! [X340] : (p1(X340) | ~ r1(X339, X340)) | ~ r1(X331, X339)) | p1(X331))) | ~ r1(X330, X331)) | ~ r1(X329, X330)) | ~ r1(X328, X329)) | ~ r1(X327, X328)) | ~ r1(X326, X327)) | ~ r1(X325, X326)) | ~ r1(X324, X325)) | ~ r1(X323, X324)) | ~ r1(X322, X323)) | ~ r1(X321, X322)) | ~ r1(X320, X321)) | ~ r1(X319, X320)) | ~ r1(X318, X319)) | ~ r1(X317, X318)) | ~ r1(X316, X317)) | ~ r1(X315, X316)) | ~ r1(X0, X315)) | ~ ! [X341] : (! [X342] : (! [X343] : (! [X344] : (! [X345] : (! [X346] : (! [X347] : (! [X348] : (! [X349] : (! [X350] : (! [X351] : (! [X352] : (! [X353] : (! [X354] : (! [X355] : (! [X356] : (! [X357] : (! [X358] : ((! [X359] : (~ ! [X360] : (p1(X360) | ~ r1(X359, X360)) | ! [X361] : (! [X362] : (p1(X362) | ~ r1(X361, X362)) | ~ r1(X359, X361)) | ~ r1(X358, X359)) & (~ ! [X363] : (~ ! [X364] : (~ p1(X364) | ! [X365] : (p1(X365) | ~ r1(X364, X365)) | ~ r1(X363, X364)) | p1(X363) | ~ r1(X358, X363)) | ! [X366] : (~ ! [X367] : (p1(X367) | ~ r1(X366, X367)) | ~ r1(X358, X366)) | p1(X358))) | ~ r1(X357, X358)) | ~ r1(X356, X357)) | ~ r1(X355, X356)) | ~ r1(X354, X355)) | ~ r1(X353, X354)) | ~ r1(X352, X353)) | ~ r1(X351, X352)) | ~ r1(X350, X351)) | ~ r1(X349, X350)) | ~ r1(X348, X349)) | ~ r1(X347, X348)) | ~ r1(X346, X347)) | ~ r1(X345, X346)) | ~ r1(X344, X345)) | ~ r1(X343, X344)) | ~ r1(X342, X343)) | ~ r1(X341, X342)) | ~ r1(X0, X341)) | ~ ! [X368] : (! [X369] : (! [X370] : (! [X371] : (! [X372] : (! [X373] : (! [X374] : (! [X375] : (! [X376] : (! [X377] : (! [X378] : (! [X379] : (! [X380] : (! [X381] : (! [X382] : (! [X383] : (! [X384] : (! [X385] : (! [X386] : ((! [X387] : (~ ! [X388] : (p1(X388) | ~ r1(X387, X388)) | ! [X389] : (! [X390] : (p1(X390) | ~ r1(X389, X390)) | ~ r1(X387, X389)) | ~ r1(X386, X387)) & (~ ! [X391] : (~ ! [X392] : (~ p1(X392) | ! [X393] : (p1(X393) | ~ r1(X392, X393)) | ~ r1(X391, X392)) | p1(X391) | ~ r1(X386, X391)) | ! [X394] : (~ ! [X395] : (p1(X395) | ~ r1(X394, X395)) | ~ r1(X386, X394)) | p1(X386))) | ~ r1(X385, X386)) | ~ r1(X384, X385)) | ~ r1(X383, X384)) | ~ r1(X382, X383)) | ~ r1(X381, X382)) | ~ r1(X380, X381)) | ~ r1(X379, X380)) | ~ r1(X378, X379)) | ~ r1(X377, X378)) | ~ r1(X376, X377)) | ~ r1(X375, X376)) | ~ r1(X374, X375)) | ~ r1(X373, X374)) | ~ r1(X372, X373)) | ~ r1(X371, X372)) | ~ r1(X370, X371)) | ~ r1(X369, X370)) | ~ r1(X368, X369)) | ~ r1(X0, X368))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) & ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1)) & (~ ! [X0] : (~ ! [X1] : (~ (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0))) | ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (((~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) & ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1)) & (~ ! [X0] : (~ ! [X1] : (~ (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0))) | ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL640+1.020.p', main)).
fof(f406231, plain, (~ r1(sK54, sK57) | (spl85_75 | ~ spl85_15589)), inference(resolution, [], [f401346, f676])).
fof(f676, plain, (~ r1(sK57, sK7(sK57)) | spl85_75), inference(avatar_component_clause, [], [f675])).
fof(f675, plain, (spl85_75 <=> r1(sK57, sK7(sK57))), introduced(avatar_definition, [new_symbols(naming, [spl85_75])])).
fof(f401346, plain, (! [X1] : (r1(X1, sK7(X1)) | ~ r1(sK54, X1)) | ~ spl85_15589), inference(avatar_component_clause, [], [f401345])).
fof(f401345, plain, (spl85_15589 <=> ! [X1] : (r1(X1, sK7(X1)) | ~ r1(sK54, X1))), introduced(avatar_definition, [new_symbols(naming, [spl85_15589])])).
fof(f405702, plain, (~ spl85_75 | ~ spl85_15591), inference(avatar_contradiction_clause, [], [f405701])).
fof(f405701, plain, ($false | (~ spl85_75 | ~ spl85_15591)), inference(subsumption_resolution, [], [f405699, f212])).
fof(f405699, plain, (~ r1(sK54, sK57) | (~ spl85_75 | ~ spl85_15591)), inference(resolution, [], [f401357, f395355])).
fof(f395355, plain, (p1(sK7(sK57)) | ~ spl85_75), inference(resolution, [], [f677, f213])).
fof(f213, plain, ! [X161] : (~ r1(sK57, X161) | p1(X161)), inference(cnf_transformation, [], [f106])).
fof(f677, plain, (r1(sK57, sK7(sK57)) | ~ spl85_75), inference(avatar_component_clause, [], [f675])).
fof(f401357, plain, (! [X2] : (~ p1(sK7(X2)) | ~ r1(sK54, X2)) | ~ spl85_15591), inference(avatar_component_clause, [], [f401356])).
fof(f401356, plain, (spl85_15591 <=> ! [X2] : (~ p1(sK7(X2)) | ~ r1(sK54, X2))), introduced(avatar_definition, [new_symbols(naming, [spl85_15591])])).
fof(f403275, plain, (~ spl85_73 | spl85_555 | ~ spl85_15938), inference(avatar_contradiction_clause, [], [f403274])).
fof(f403274, plain, ($false | (~ spl85_73 | spl85_555 | ~ spl85_15938)), inference(subsumption_resolution, [], [f403273, f668])).
fof(f668, plain, (r1(sK54, sK6(sK54)) | ~ spl85_73), inference(avatar_component_clause, [], [f666])).
fof(f666, plain, (spl85_73 <=> r1(sK54, sK6(sK54))), introduced(avatar_definition, [new_symbols(naming, [spl85_73])])).
fof(f403273, plain, (~ r1(sK54, sK6(sK54)) | (spl85_555 | ~ spl85_15938)), inference(subsumption_resolution, [], [f403272, f8686])).
fof(f8686, plain, (~ p1(sK6(sK54)) | spl85_555), inference(avatar_component_clause, [], [f8685])).
fof(f8685, plain, (spl85_555 <=> p1(sK6(sK54))), introduced(avatar_definition, [new_symbols(naming, [spl85_555])])).
fof(f403272, plain, (p1(sK6(sK54)) | ~ r1(sK54, sK6(sK54)) | ~ spl85_15938), inference(resolution, [], [f403212, f216])).
fof(f216, plain, ! [X157] : (~ p1(sK56(X157)) | p1(X157) | ~ r1(sK54, X157)), inference(cnf_transformation, [], [f106])).
fof(f403212, plain, (p1(sK56(sK6(sK54))) | ~ spl85_15938), inference(avatar_component_clause, [], [f403210])).
fof(f403210, plain, (spl85_15938 <=> p1(sK56(sK6(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl85_15938])])).
fof(f403217, plain, (spl85_15591 | spl85_15938 | spl85_35 | ~ spl85_72 | ~ spl85_73 | ~ spl85_554), inference(avatar_split_clause, [], [f403216, f8681, f666, f662, f460, f403210, f401356])).
fof(f460, plain, (spl85_35 <=> p1(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl85_35])])).
fof(f662, plain, (spl85_72 <=> sP2(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl85_72])])).
fof(f8681, plain, (spl85_554 <=> p1(sK55(sK6(sK54)))), introduced(avatar_definition, [new_symbols(naming, [spl85_554])])).
fof(f403216, plain, (! [X3] : (p1(sK56(sK6(sK54))) | ~ p1(sK7(X3)) | ~ r1(sK54, X3)) | (spl85_35 | ~ spl85_72 | ~ spl85_73 | ~ spl85_554)), inference(subsumption_resolution, [], [f403215, f668])).
fof(f403215, plain, (! [X3] : (p1(sK56(sK6(sK54))) | ~ p1(sK7(X3)) | ~ r1(sK54, X3) | ~ r1(sK54, sK6(sK54))) | (spl85_35 | ~ spl85_72 | ~ spl85_554)), inference(subsumption_resolution, [], [f403214, f663])).
fof(f663, plain, (sP2(sK54) | ~ spl85_72), inference(avatar_component_clause, [], [f662])).
fof(f403214, plain, (! [X3] : (p1(sK56(sK6(sK54))) | ~ p1(sK7(X3)) | ~ r1(sK54, X3) | ~ sP2(sK54) | ~ r1(sK54, sK6(sK54))) | (spl85_35 | ~ spl85_554)), inference(subsumption_resolution, [], [f403205, f461])).
fof(f461, plain, (~ p1(sK54) | spl85_35), inference(avatar_component_clause, [], [f460])).
fof(f403205, plain, (! [X3] : (p1(sK56(sK6(sK54))) | ~ p1(sK7(X3)) | ~ r1(sK54, X3) | p1(sK54) | ~ sP2(sK54) | ~ r1(sK54, sK6(sK54))) | ~ spl85_554), inference(resolution, [], [f8683, f3108])).
fof(f3108, plain, ! [X0, X1] : (~ p1(sK55(sK6(X0))) | p1(sK56(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0))), inference(subsumption_resolution, [], [f3107, f114])).
fof(f114, plain, ! [X4, X0] : (~ p1(sK7(X4)) | ~ p1(sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : ((! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0))) | ! [X4] : ((~ p1(sK7(X4)) & r1(X4, sK7(X4))) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7])], [f18, f20, f19])).
fof(f19, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) => (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) => (~ p1(sK7(X4)) & r1(X4, sK7(X4)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) | ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X173] : (? [X185] : (! [X186] : (~ p1(X186) | ! [X187] : (p1(X187) | ~ r1(X186, X187)) | ~ r1(X185, X186)) & ~ p1(X185) & r1(X173, X185)) | ! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ~ r1(X173, X188)) | p1(X173) | ~ sP2(X173)), inference(nnf_transformation, [], [f9])).
fof(f3107, plain, ! [X0, X1] : (p1(sK56(sK6(X0))) | ~ p1(sK55(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f3098])).
fof(f3098, plain, ! [X0, X1] : (p1(sK56(sK6(X0))) | ~ p1(sK55(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK54, sK6(X0))), inference(resolution, [], [f947, f215])).
fof(f215, plain, ! [X157] : (r1(sK55(X157), sK56(X157)) | p1(X157) | ~ r1(sK54, X157)), inference(cnf_transformation, [], [f106])).
fof(f947, plain, ! [X4, X5, X3] : (~ r1(sK55(sK6(X4)), X3) | p1(X3) | ~ p1(sK55(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK54, sK6(X4))), inference(subsumption_resolution, [], [f945, f114])).
fof(f945, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK55(sK6(X4)), X3) | ~ p1(sK55(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK54, sK6(X4))), inference(resolution, [], [f116, f214])).
fof(f214, plain, ! [X157] : (r1(X157, sK55(X157)) | p1(X157) | ~ r1(sK54, X157)), inference(cnf_transformation, [], [f106])).
fof(f116, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | ~ p1(sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f8683, plain, (p1(sK55(sK6(sK54))) | ~ spl85_554), inference(avatar_component_clause, [], [f8681])).
fof(f403213, plain, (spl85_15589 | spl85_15938 | spl85_35 | ~ spl85_72 | ~ spl85_73 | ~ spl85_554), inference(avatar_split_clause, [], [f403208, f8681, f666, f662, f460, f403210, f401345])).
fof(f403208, plain, (! [X2] : (p1(sK56(sK6(sK54))) | r1(X2, sK7(X2)) | ~ r1(sK54, X2)) | (spl85_35 | ~ spl85_72 | ~ spl85_73 | ~ spl85_554)), inference(subsumption_resolution, [], [f403207, f668])).
fof(f403207, plain, (! [X2] : (p1(sK56(sK6(sK54))) | r1(X2, sK7(X2)) | ~ r1(sK54, X2) | ~ r1(sK54, sK6(sK54))) | (spl85_35 | ~ spl85_72 | ~ spl85_554)), inference(subsumption_resolution, [], [f403206, f663])).
fof(f403206, plain, (! [X2] : (p1(sK56(sK6(sK54))) | r1(X2, sK7(X2)) | ~ r1(sK54, X2) | ~ sP2(sK54) | ~ r1(sK54, sK6(sK54))) | (spl85_35 | ~ spl85_554)), inference(subsumption_resolution, [], [f403204, f461])).
fof(f403204, plain, (! [X2] : (p1(sK56(sK6(sK54))) | r1(X2, sK7(X2)) | ~ r1(sK54, X2) | p1(sK54) | ~ sP2(sK54) | ~ r1(sK54, sK6(sK54))) | ~ spl85_554), inference(resolution, [], [f8683, f3194])).
fof(f3194, plain, ! [X0, X1] : (~ p1(sK55(sK6(X0))) | p1(sK56(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0))), inference(subsumption_resolution, [], [f3193, f113])).
fof(f113, plain, ! [X4, X0] : (~ p1(sK6(X0)) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f3193, plain, ! [X0, X1] : (p1(sK56(sK6(X0))) | ~ p1(sK55(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f3184])).
fof(f3184, plain, ! [X0, X1] : (p1(sK56(sK6(X0))) | ~ p1(sK55(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK54, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK54, sK6(X0))), inference(resolution, [], [f958, f215])).
fof(f958, plain, ! [X4, X5, X3] : (~ r1(sK55(sK6(X4)), X3) | p1(X3) | ~ p1(sK55(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK54, sK6(X4))), inference(subsumption_resolution, [], [f956, f113])).
fof(f956, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK55(sK6(X4)), X3) | ~ p1(sK55(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK54, sK6(X4))), inference(resolution, [], [f115, f214])).
fof(f115, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f401332, plain, (spl85_35 | ~ spl85_72 | spl85_73 | ~ spl85_75), inference(avatar_contradiction_clause, [], [f401331])).
fof(f401331, plain, ($false | (spl85_35 | ~ spl85_72 | spl85_73 | ~ spl85_75)), inference(subsumption_resolution, [], [f401330, f663])).
fof(f401330, plain, (~ sP2(sK54) | (spl85_35 | spl85_73 | ~ spl85_75)), inference(subsumption_resolution, [], [f401329, f461])).
fof(f401329, plain, (p1(sK54) | ~ sP2(sK54) | (spl85_73 | ~ spl85_75)), inference(subsumption_resolution, [], [f395383, f212])).
fof(f395383, plain, (~ r1(sK54, sK57) | p1(sK54) | ~ sP2(sK54) | (spl85_73 | ~ spl85_75)), inference(resolution, [], [f395377, f667])).
fof(f667, plain, (~ r1(sK54, sK6(sK54)) | spl85_73), inference(avatar_component_clause, [], [f666])).
fof(f395377, plain, (! [X0] : (r1(X0, sK6(X0)) | ~ r1(X0, sK57) | p1(X0) | ~ sP2(X0)) | ~ spl85_75), inference(resolution, [], [f395355, f112])).
fof(f112, plain, ! [X4, X0] : (~ p1(sK7(X4)) | r1(X0, sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f395199, plain, (~ spl85_109 | ~ spl85_806), inference(avatar_split_clause, [], [f89924, f13946, f855])).
fof(f855, plain, (spl85_109 <=> sP0(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl85_109])])).
fof(f13946, plain, (spl85_806 <=> sP1(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl85_806])])).
fof(f89924, plain, (~ sP0(sK54) | ~ spl85_806), inference(resolution, [], [f13948, f1110])).
fof(f1110, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f1109])).
fof(f1109, plain, ! [X0] : (~ sP0(X0) | ~ sP1(X0) | ~ sP1(X0) | ~ sP0(X0)), inference(resolution, [], [f1104, f121])).
fof(f121, plain, ! [X0] : (r1(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((! [X2] : ((p1(X2) & (~ p1(sK11(X2)) & r1(X2, sK11(X2)))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0))) & r1(X0, sK10(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12])], [f28, f31, f30, f29])).
fof(f29, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) => (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & ? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) & r1(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK11(X2)) & r1(X2, sK11(X2)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) => (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f27])).
fof(f27, plain, ! [X173] : (? [X190] : (! [X191] : ((p1(X191) & ? [X192] : (~ p1(X192) & r1(X191, X192))) | ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X191, X193)) | ~ r1(X190, X191)) & p1(X190) & ? [X195] : (~ p1(X195) & r1(X190, X195)) & r1(X173, X190)) | ~ sP0(X173)), inference(nnf_transformation, [], [f7])).
fof(f1104, plain, ! [X0, X1] : (~ r1(X1, sK10(X0)) | ~ sP0(X0) | ~ sP1(X0) | ~ sP1(X1)), inference(duplicate_literal_removal, [], [f1103])).
fof(f1103, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | ~ r1(X1, sK10(X0)) | ~ sP1(X1) | ~ sP0(X0)), inference(resolution, [], [f1102, f122])).
fof(f122, plain, ! [X0] : (r1(sK10(X0), sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f1102, plain, ! [X2, X0, X1] : (~ r1(X1, sK12(X0)) | ~ sP1(X0) | ~ sP0(X0) | ~ r1(X2, X1) | ~ sP1(X2)), inference(subsumption_resolution, [], [f1101, f123])).
fof(f123, plain, ! [X0] : (~ p1(sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f1101, plain, ! [X2, X0, X1] : (~ sP0(X0) | ~ sP1(X0) | p1(sK12(X0)) | ~ r1(X1, sK12(X0)) | ~ r1(X2, X1) | ~ sP1(X2)), inference(resolution, [], [f1100, f119])).
fof(f119, plain, ! [X2, X0, X1] : (~ p1(sK9(X2)) | p1(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (! [X1] : (! [X2] : ((p1(sK8(X2)) & (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2))) & r1(X2, sK8(X2))) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9])], [f23, f25, f24])).
fof(f24, plain, ! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) => (p1(sK8(X2)) & ? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) & r1(X2, sK8(X2)))), introduced(choice_axiom, [])).
fof(f25, plain, ! [X2] : (? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) => (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f22])).
fof(f22, plain, ! [X173] : (! [X177] : (! [X178] : (? [X179] : (p1(X179) & ? [X180] : (~ p1(X180) & r1(X179, X180)) & r1(X178, X179)) | p1(X178) | ~ r1(X177, X178)) | ~ r1(X173, X177)) | ~ sP1(X173)), inference(nnf_transformation, [], [f8])).
fof(f1100, plain, ! [X0] : (p1(sK9(sK12(X0))) | ~ sP0(X0) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f1092])).
fof(f1092, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0) | p1(sK9(sK12(X0))) | ~ sP1(X0)), inference(resolution, [], [f1088, f1038])).
fof(f1038, plain, ! [X0, X1] : (~ r1(sK8(sK12(X0)), X1) | ~ sP0(X0) | p1(X1) | ~ sP1(X0)), inference(subsumption_resolution, [], [f1037, f996])).
fof(f996, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f995])).
fof(f995, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f331, f121])).
fof(f331, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | p1(sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f316, f123])).
fof(f316, plain, ! [X4, X5] : (p1(sK12(X4)) | p1(sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f120, f122])).
fof(f120, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | p1(sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f1037, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0)))), inference(duplicate_literal_removal, [], [f1029])).
fof(f1029, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0))) | ~ sP0(X0)), inference(resolution, [], [f1028, f908])).
fof(f908, plain, ! [X2, X0, X1] : (~ r1(sK12(X2), X0) | p1(X1) | ~ r1(X0, X1) | ~ p1(X0) | ~ sP0(X2)), inference(subsumption_resolution, [], [f907, f123])).
fof(f907, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2)), inference(duplicate_literal_removal, [], [f904])).
fof(f904, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2) | ~ sP0(X2)), inference(resolution, [], [f127, f122])).
fof(f127, plain, ! [X4, X2, X0, X5] : (~ r1(sK10(X0), X2) | ~ p1(X4) | p1(X5) | ~ r1(X4, X5) | ~ r1(X2, X4) | p1(X2) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f1028, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f1027])).
fof(f1027, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f493, f121])).
fof(f493, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f478, f123])).
fof(f478, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f117, f122])).
fof(f117, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(X2, sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f1088, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f1087])).
fof(f1087, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f696, f121])).
fof(f696, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f681, f123])).
fof(f681, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f118, f122])).
fof(f118, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(sK8(X2), sK9(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f13948, plain, (sP1(sK54) | ~ spl85_806), inference(avatar_component_clause, [], [f13946])).
fof(f395170, plain, ~ spl85_4139, inference(avatar_contradiction_clause, [], [f395169])).
fof(f395169, plain, ($false | ~ spl85_4139), inference(subsumption_resolution, [], [f395140, f212])).
fof(f395140, plain, (~ r1(sK54, sK57) | ~ spl85_4139), inference(resolution, [], [f395084, f209])).
fof(f209, plain, r1(sK53, sK54), inference(cnf_transformation, [], [f106])).
fof(f395084, plain, (! [X0] : (~ r1(sK53, X0) | ~ r1(X0, sK57)) | ~ spl85_4139), inference(resolution, [], [f395028, f208])).
fof(f208, plain, r1(sK52, sK53), inference(cnf_transformation, [], [f106])).
fof(f395028, plain, (! [X0, X1] : (~ r1(sK52, X1) | ~ r1(X0, sK57) | ~ r1(X1, X0)) | ~ spl85_4139), inference(resolution, [], [f394972, f207])).
fof(f207, plain, r1(sK51, sK52), inference(cnf_transformation, [], [f106])).
fof(f394972, plain, (! [X2, X0, X1] : (~ r1(sK51, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl85_4139), inference(resolution, [], [f394916, f206])).
fof(f206, plain, r1(sK50, sK51), inference(cnf_transformation, [], [f106])).
fof(f394916, plain, (! [X2, X0, X3, X1] : (~ r1(sK50, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl85_4139), inference(resolution, [], [f394860, f205])).
fof(f205, plain, r1(sK49, sK50), inference(cnf_transformation, [], [f106])).
fof(f394860, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK49, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X0)) | ~ spl85_4139), inference(resolution, [], [f394803, f204])).
fof(f204, plain, r1(sK48, sK49), inference(cnf_transformation, [], [f106])).
fof(f394803, plain, (! [X4, X2, X0, X5, X3, X1] : (~ r1(sK48, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X0)) | ~ spl85_4139), inference(resolution, [], [f394748, f203])).
fof(f203, plain, r1(sK47, sK48), inference(cnf_transformation, [], [f106])).
fof(f394748, plain, (! [X26, X24, X23, X21, X27, X25, X22] : (~ r1(sK47, X22) | ~ r1(X21, sK57) | ~ r1(X22, X23) | ~ r1(X23, X24) | ~ r1(X24, X25) | ~ r1(X25, X26) | ~ r1(X26, X27) | ~ r1(X27, X21)) | ~ spl85_4139), inference(resolution, [], [f394686, f202])).
fof(f202, plain, r1(sK46, sK47), inference(cnf_transformation, [], [f106])).
fof(f394686, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK46, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X0)) | ~ spl85_4139), inference(resolution, [], [f394658, f201])).
fof(f201, plain, r1(sK45, sK46), inference(cnf_transformation, [], [f106])).
fof(f394658, plain, (! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK45, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | ~ r1(X8, X0)) | ~ spl85_4139), inference(resolution, [], [f394592, f200])).
fof(f200, plain, r1(sK44, sK45), inference(cnf_transformation, [], [f106])).
fof(f394592, plain, (! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK44, X1) | ~ r1(X0, sK57) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | ~ r1(X8, X9) | ~ r1(X9, X0)) | ~ spl85_4139), inference(resolution, [], [f90327, f199])).
fof(f199, plain, r1(sK13, sK44), inference(cnf_transformation, [], [f106])).
fof(f90327, plain, (! [X6, X4, X2, X12, X10, X8, X7, X5, X3, X11, X9] : (~ r1(sK13, X12) | ~ r1(X2, sK57) | ~ r1(X12, X11) | ~ r1(X11, X10) | ~ r1(X10, X9) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2)) | ~ spl85_4139), inference(avatar_component_clause, [], [f90326])).
fof(f90326, plain, (spl85_4139 <=> ! [X9, X11, X3, X5, X7, X8, X10, X12, X2, X4, X6] : (~ r1(X2, sK57) | ~ r1(sK13, X12) | ~ r1(X12, X11) | ~ r1(X11, X10) | ~ r1(X10, X9) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2))), introduced(avatar_definition, [new_symbols(naming, [spl85_4139])])).
fof(f233939, plain, ~ spl85_805, inference(avatar_contradiction_clause, [], [f233938])).
fof(f233938, plain, ($false | ~ spl85_805), inference(subsumption_resolution, [], [f233913, f209])).
fof(f233913, plain, (~ r1(sK53, sK54) | ~ spl85_805), inference(resolution, [], [f233888, f208])).
fof(f233888, plain, (! [X0] : (~ r1(sK52, X0) | ~ r1(X0, sK54)) | ~ spl85_805), inference(resolution, [], [f233861, f207])).
fof(f233861, plain, (! [X0, X1] : (~ r1(sK51, X1) | ~ r1(X0, sK54) | ~ r1(X1, X0)) | ~ spl85_805), inference(resolution, [], [f233833, f206])).
fof(f233833, plain, (! [X2, X0, X1] : (~ r1(sK50, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl85_805), inference(resolution, [], [f233780, f205])).
fof(f233780, plain, (! [X2, X0, X3, X1] : (~ r1(sK49, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl85_805), inference(resolution, [], [f233723, f204])).
fof(f233723, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK48, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X0)) | ~ spl85_805), inference(resolution, [], [f233699, f203])).
fof(f233699, plain, (! [X23, X21, X19, X22, X20, X18] : (~ r1(sK47, X19) | ~ r1(X18, sK54) | ~ r1(X19, X20) | ~ r1(X20, X21) | ~ r1(X21, X22) | ~ r1(X22, X23) | ~ r1(X23, X18)) | ~ spl85_805), inference(resolution, [], [f233671, f202])).
fof(f233671, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK46, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X0)) | ~ spl85_805), inference(resolution, [], [f233647, f201])).
fof(f233647, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK45, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X0)) | ~ spl85_805), inference(resolution, [], [f233623, f200])).
fof(f233623, plain, (! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK44, X1) | ~ r1(X0, sK54) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | ~ r1(X8, X0)) | ~ spl85_805), inference(resolution, [], [f13944, f199])).
fof(f13944, plain, (! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK13, X9) | ~ r1(X0, sK54) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0)) | ~ spl85_805), inference(avatar_component_clause, [], [f13943])).
fof(f13943, plain, (spl85_805 <=> ! [X9, X1, X3, X5, X7, X8, X0, X2, X4, X6] : (~ r1(X0, sK54) | ~ r1(sK13, X9) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl85_805])])).
fof(f90328, plain, (spl85_4139 | spl85_4080 | ~ spl85_4079), inference(avatar_split_clause, [], [f90324, f89470, f89474, f90326])).
fof(f89474, plain, (spl85_4080 <=> ! [X3, X2] : (~ r1(sK57, X2) | p1(X3) | ~ r1(X2, X3))), introduced(avatar_definition, [new_symbols(naming, [spl85_4080])])).
fof(f89470, plain, (spl85_4079 <=> r1(sK57, sK60(sK57))), introduced(avatar_definition, [new_symbols(naming, [spl85_4079])])).
fof(f90324, plain, (! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X11, X9] : (p1(X0) | ~ r1(X1, X0) | ~ r1(sK57, X1) | ~ r1(X2, sK57) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(sK13, X12)) | ~ spl85_4079), inference(resolution, [], [f90316, f195])).
fof(f195, plain, ! [X177, X167, X165, X163, X173, X171, X169, X180, X166, X164, X172, X170, X168, X179] : (~ p1(sK60(X177)) | p1(X180) | ~ r1(X179, X180) | ~ r1(X177, X179) | ~ r1(X173, X177) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | ~ r1(sK13, X163)), inference(cnf_transformation, [], [f106])).
fof(f90316, plain, (p1(sK60(sK57)) | ~ spl85_4079), inference(resolution, [], [f89472, f213])).
fof(f89472, plain, (r1(sK57, sK60(sK57)) | ~ spl85_4079), inference(avatar_component_clause, [], [f89470])).
fof(f90295, plain, (~ spl85_35 | ~ spl85_108 | spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080), inference(avatar_contradiction_clause, [], [f90294])).
fof(f90294, plain, ($false | (~ spl85_35 | ~ spl85_108 | spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f90271, f211])).
fof(f211, plain, ~ p1(sK58), inference(cnf_transformation, [], [f106])).
fof(f90271, plain, (p1(sK58) | (~ spl85_35 | ~ spl85_108 | spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(resolution, [], [f90250, f210])).
fof(f210, plain, r1(sK54, sK58), inference(cnf_transformation, [], [f106])).
fof(f90250, plain, (! [X0] : (~ r1(sK54, X0) | p1(X0)) | (~ spl85_35 | ~ spl85_108 | spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f90249, f852])).
fof(f852, plain, (sP3(sK54) | ~ spl85_108), inference(avatar_component_clause, [], [f851])).
fof(f851, plain, (spl85_108 <=> sP3(sK54)), introduced(avatar_definition, [new_symbols(naming, [spl85_108])])).
fof(f90249, plain, (! [X0] : (p1(X0) | ~ r1(sK54, X0) | ~ sP3(sK54)) | (~ spl85_35 | spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f90248, f462])).
fof(f462, plain, (p1(sK54) | ~ spl85_35), inference(avatar_component_clause, [], [f460])).
fof(f90248, plain, (! [X0] : (~ p1(sK54) | p1(X0) | ~ r1(sK54, X0) | ~ sP3(sK54)) | (spl85_109 | ~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f90247, f856])).
fof(f856, plain, (~ sP0(sK54) | spl85_109), inference(avatar_component_clause, [], [f855])).
fof(f90247, plain, (! [X0] : (sP0(sK54) | ~ p1(sK54) | p1(X0) | ~ r1(sK54, X0) | ~ sP3(sK54)) | (~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(resolution, [], [f89842, f212])).
fof(f89842, plain, (! [X0, X1] : (~ r1(X0, sK57) | sP0(X0) | ~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ sP3(X0)) | (~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(resolution, [], [f89810, f109])).
fof(f109, plain, ! [X4, X0, X1] : (~ p1(sK5(X1)) | sP0(X0) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | ~ r1(X0, X4) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (sP0(X0) | ! [X1] : ((p1(sK4(X1)) & (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1))) & r1(X1, sK4(X1))) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f13, f15, f14])).
fof(f14, plain, ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) => (p1(sK4(X1)) & ? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) & r1(X1, sK4(X1)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X1] : (? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) => (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1)))), introduced(choice_axiom, [])).
fof(f13, plain, ! [X0] : (sP0(X0) | ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(rectify, [], [f12])).
fof(f12, plain, ! [X173] : (sP0(X173) | ! [X196] : (? [X197] : (p1(X197) & ? [X198] : (~ p1(X198) & r1(X197, X198)) & r1(X196, X197)) | ~ r1(X173, X196)) | ~ p1(X173) | ! [X199] : (p1(X199) | ~ r1(X173, X199)) | ~ sP3(X173)), inference(nnf_transformation, [], [f10])).
fof(f89810, plain, (p1(sK5(sK57)) | (~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f89788, f212])).
fof(f89788, plain, (p1(sK5(sK57)) | ~ r1(sK54, sK57) | (~ spl85_546 | ~ spl85_547 | ~ spl85_4080)), inference(resolution, [], [f89771, f8276])).
fof(f8276, plain, (! [X17] : (r1(sK4(X17), sK5(X17)) | ~ r1(sK54, X17)) | ~ spl85_546), inference(avatar_component_clause, [], [f8275])).
fof(f8275, plain, (spl85_546 <=> ! [X17] : (r1(sK4(X17), sK5(X17)) | ~ r1(sK54, X17))), introduced(avatar_definition, [new_symbols(naming, [spl85_546])])).
fof(f89771, plain, (! [X0] : (~ r1(sK4(sK57), X0) | p1(X0)) | (~ spl85_547 | ~ spl85_4080)), inference(subsumption_resolution, [], [f89751, f212])).
fof(f89751, plain, (! [X0] : (p1(X0) | ~ r1(sK4(sK57), X0) | ~ r1(sK54, sK57)) | (~ spl85_547 | ~ spl85_4080)), inference(resolution, [], [f89475, f8281])).
fof(f8281, plain, (! [X17] : (r1(X17, sK4(X17)) | ~ r1(sK54, X17)) | ~ spl85_547), inference(avatar_component_clause, [], [f8280])).
fof(f8280, plain, (spl85_547 <=> ! [X17] : (r1(X17, sK4(X17)) | ~ r1(sK54, X17))), introduced(avatar_definition, [new_symbols(naming, [spl85_547])])).
fof(f89475, plain, (! [X2, X3] : (~ r1(sK57, X2) | p1(X3) | ~ r1(X2, X3)) | ~ spl85_4080), inference(avatar_component_clause, [], [f89474])).
fof(f89476, plain, (spl85_4079 | spl85_4080), inference(avatar_split_clause, [], [f89440, f89474, f89470])).
fof(f89440, plain, ! [X2, X3] : (~ r1(sK57, X2) | ~ r1(X2, X3) | p1(X3) | r1(sK57, sK60(sK57))), inference(resolution, [], [f89382, f212])).
fof(f89382, plain, ! [X2, X0, X1] : (~ r1(sK54, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | p1(X2) | r1(X0, sK60(X0))), inference(resolution, [], [f89360, f209])).
fof(f89360, plain, ! [X2, X0, X3, X1] : (~ r1(sK53, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | p1(X3) | r1(X1, sK60(X1))), inference(resolution, [], [f89338, f208])).
fof(f89338, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK52, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | p1(X4) | r1(X2, sK60(X2))), inference(resolution, [], [f89316, f207])).
fof(f89316, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK51, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | p1(X5) | r1(X3, sK60(X3))), inference(resolution, [], [f89294, f206])).
fof(f89294, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK50, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | p1(X6) | r1(X4, sK60(X4))), inference(resolution, [], [f48059, f205])).
fof(f48059, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK49, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X7) | p1(X7) | r1(X5, sK60(X5))), inference(resolution, [], [f17618, f204])).
fof(f17618, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK48, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ r1(X1, X6) | ~ r1(X6, X7) | ~ r1(X7, X8) | p1(X8) | r1(X6, sK60(X6))), inference(resolution, [], [f10237, f203])).
fof(f10237, plain, ! [X30, X39, X37, X35, X33, X31, X38, X36, X34, X32] : (~ r1(sK47, X37) | ~ r1(X32, X30) | ~ r1(X33, X32) | ~ r1(X34, X33) | ~ r1(X35, X34) | ~ r1(X36, X35) | ~ r1(X37, X36) | ~ r1(X30, X31) | ~ r1(X31, X38) | ~ r1(X38, X39) | p1(X39) | r1(X31, sK60(X31))), inference(resolution, [], [f6555, f202])).
fof(f6555, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (~ r1(sK46, X9) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X0, X1) | ~ r1(X1, X10) | p1(X10) | r1(X0, sK60(X0))), inference(resolution, [], [f3915, f201])).
fof(f3915, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK45, X11) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X0, X1) | p1(X1) | r1(X2, sK60(X2))), inference(resolution, [], [f1583, f200])).
fof(f1583, plain, ! [X6, X4, X2, X0, X12, X10, X8, X7, X5, X3, X1, X11, X9] : (~ r1(sK44, X12) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(X11, X10) | ~ r1(X12, X11) | p1(X0) | r1(X2, sK60(X2))), inference(resolution, [], [f194, f199])).
fof(f194, plain, ! [X177, X167, X165, X163, X173, X171, X169, X180, X166, X164, X172, X170, X168, X179] : (~ r1(sK13, X163) | p1(X180) | ~ r1(X179, X180) | ~ r1(X177, X179) | ~ r1(X173, X177) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | r1(X177, sK60(X177))), inference(cnf_transformation, [], [f106])).
fof(f13949, plain, (spl85_805 | spl85_806), inference(avatar_split_clause, [], [f13941, f13946, f13943])).
fof(f13941, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (sP1(sK54) | ~ r1(X0, sK54) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(sK13, X9)), inference(subsumption_resolution, [], [f13940, f209])).
fof(f13940, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (sP1(sK54) | ~ r1(sK53, sK54) | ~ r1(X0, sK54) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(sK13, X9)), inference(duplicate_literal_removal, [], [f13302])).
fof(f13302, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (sP1(sK54) | ~ r1(sK53, sK54) | ~ r1(X0, sK54) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(sK13, X9) | sP1(sK54)), inference(resolution, [], [f13291, f3940])).
fof(f3940, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (~ r1(sK54, sK59(X0)) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(sK13, X10) | sP1(X0)), inference(subsumption_resolution, [], [f3939, f197])).
fof(f197, plain, ! [X167, X165, X163, X173, X171, X169, X166, X164, X172, X170, X168] : (~ p1(sK59(X173)) | sP1(X173) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | ~ r1(sK13, X163)), inference(cnf_transformation, [], [f106])).
fof(f3939, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(sK13, X10) | ~ r1(sK54, sK59(X0)) | p1(sK59(X0))), inference(subsumption_resolution, [], [f3938, f216])).
fof(f3938, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (p1(sK56(sK59(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(sK13, X10) | ~ r1(sK54, sK59(X0)) | p1(sK59(X0))), inference(duplicate_literal_removal, [], [f3929])).
fof(f3929, plain, ! [X6, X4, X2, X0, X10, X8, X7, X5, X3, X1, X9] : (p1(sK56(sK59(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X10, X9) | ~ r1(sK13, X10) | ~ r1(sK54, sK59(X0)) | p1(sK59(X0)) | ~ r1(sK54, sK59(X0))), inference(resolution, [], [f1335, f215])).
fof(f1335, plain, ! [X14, X12, X23, X21, X19, X17, X15, X13, X22, X20, X18, X16] : (~ r1(sK55(sK59(X13)), X12) | p1(X12) | sP1(X13) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(X18, X17) | ~ r1(X19, X18) | ~ r1(X20, X19) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK13, X23) | ~ r1(sK54, sK59(X13))), inference(subsumption_resolution, [], [f1334, f197])).
fof(f1334, plain, ! [X14, X12, X23, X21, X19, X17, X15, X13, X22, X20, X18, X16] : (p1(X12) | ~ r1(sK55(sK59(X13)), X12) | sP1(X13) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(X18, X17) | ~ r1(X19, X18) | ~ r1(X20, X19) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK13, X23) | p1(sK59(X13)) | ~ r1(sK54, sK59(X13))), inference(subsumption_resolution, [], [f1332, f217])).
fof(f217, plain, ! [X157] : (~ r1(sK54, X157) | p1(X157) | p1(sK55(X157))), inference(cnf_transformation, [], [f106])).
fof(f1332, plain, ! [X14, X12, X23, X21, X19, X17, X15, X13, X22, X20, X18, X16] : (p1(X12) | ~ r1(sK55(sK59(X13)), X12) | ~ p1(sK55(sK59(X13))) | sP1(X13) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(X18, X17) | ~ r1(X19, X18) | ~ r1(X20, X19) | ~ r1(X21, X20) | ~ r1(X22, X21) | ~ r1(X23, X22) | ~ r1(sK13, X23) | p1(sK59(X13)) | ~ r1(sK54, sK59(X13))), inference(resolution, [], [f198, f214])).
fof(f198, plain, ! [X167, X165, X163, X175, X173, X171, X169, X176, X166, X164, X172, X170, X168] : (~ r1(sK59(X173), X175) | p1(X176) | ~ r1(X175, X176) | ~ p1(X175) | sP1(X173) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | ~ r1(sK13, X163)), inference(cnf_transformation, [], [f106])).
fof(f13291, plain, ! [X0] : (r1(X0, sK59(X0)) | sP1(X0) | ~ r1(sK53, X0)), inference(resolution, [], [f13260, f208])).
fof(f13260, plain, ! [X0, X1] : (~ r1(sK52, X0) | ~ r1(X0, X1) | sP1(X1) | r1(X1, sK59(X1))), inference(resolution, [], [f13247, f207])).
fof(f13247, plain, ! [X2, X0, X1] : (~ r1(sK51, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP1(X2) | r1(X2, sK59(X2))), inference(resolution, [], [f13225, f206])).
fof(f13225, plain, ! [X2, X0, X3, X1] : (~ r1(sK50, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP1(X3) | r1(X3, sK59(X3))), inference(resolution, [], [f13200, f205])).
fof(f13200, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK49, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | sP1(X4) | r1(X4, sK59(X4))), inference(resolution, [], [f13187, f204])).
fof(f13187, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK48, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | sP1(X5) | r1(X5, sK59(X5))), inference(resolution, [], [f8514, f203])).
fof(f8514, plain, ! [X26, X24, X23, X21, X27, X25, X22] : (~ r1(sK47, X25) | ~ r1(X23, X21) | ~ r1(X24, X23) | ~ r1(X25, X24) | ~ r1(X21, X22) | ~ r1(X22, X26) | ~ r1(X26, X27) | sP1(X27) | r1(X27, sK59(X27))), inference(resolution, [], [f6261, f202])).
fof(f6261, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK46, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | ~ r1(X1, X7) | sP1(X7) | r1(X7, sK59(X7))), inference(resolution, [], [f3603, f201])).
fof(f3603, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK45, X8) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X0, X1) | sP1(X1) | r1(X1, sK59(X1))), inference(resolution, [], [f1184, f200])).
fof(f1184, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK44, X9) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | sP1(X0) | r1(X0, sK59(X0))), inference(resolution, [], [f196, f199])).
fof(f196, plain, ! [X167, X165, X163, X173, X171, X169, X166, X164, X172, X170, X168] : (~ r1(sK13, X163) | sP1(X173) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | r1(X173, sK59(X173))), inference(cnf_transformation, [], [f106])).
fof(f10300, plain, (spl85_35 | ~ spl85_72 | ~ spl85_555), inference(avatar_contradiction_clause, [], [f10299])).
fof(f10299, plain, ($false | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(subsumption_resolution, [], [f10298, f663])).
fof(f10298, plain, (~ sP2(sK54) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(subsumption_resolution, [], [f10297, f461])).
fof(f10297, plain, (p1(sK54) | ~ sP2(sK54) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(subsumption_resolution, [], [f10296, f212])).
fof(f10296, plain, (~ r1(sK54, sK57) | p1(sK54) | ~ sP2(sK54) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(resolution, [], [f10065, f8687])).
fof(f8687, plain, (p1(sK6(sK54)) | ~ spl85_555), inference(avatar_component_clause, [], [f8685])).
fof(f10065, plain, (! [X1] : (~ p1(sK6(X1)) | ~ r1(X1, sK57) | p1(X1) | ~ sP2(X1)) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(resolution, [], [f10048, f114])).
fof(f10048, plain, (p1(sK7(sK57)) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(subsumption_resolution, [], [f8998, f212])).
fof(f8998, plain, (~ r1(sK54, sK57) | p1(sK7(sK57)) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(resolution, [], [f8716, f213])).
fof(f8716, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK54, X0)) | (spl85_35 | ~ spl85_72 | ~ spl85_555)), inference(subsumption_resolution, [], [f8715, f663])).
fof(f8715, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK54, X0) | ~ sP2(sK54)) | (spl85_35 | ~ spl85_555)), inference(subsumption_resolution, [], [f8714, f461])).
fof(f8714, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK54, X0) | p1(sK54) | ~ sP2(sK54)) | ~ spl85_555), inference(resolution, [], [f8687, f113])).
fof(f8713, plain, (spl85_554 | spl85_555 | ~ spl85_73), inference(avatar_split_clause, [], [f8705, f666, f8685, f8681])).
fof(f8705, plain, (p1(sK6(sK54)) | p1(sK55(sK6(sK54))) | ~ spl85_73), inference(resolution, [], [f668, f217])).
fof(f8678, plain, spl85_72, inference(avatar_split_clause, [], [f8666, f662])).
fof(f8666, plain, sP2(sK54), inference(resolution, [], [f8656, f209])).
fof(f8656, plain, ! [X0] : (~ r1(sK53, X0) | sP2(X0)), inference(resolution, [], [f8633, f208])).
fof(f8633, plain, ! [X0, X1] : (~ r1(sK52, X0) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f8609, f207])).
fof(f8609, plain, ! [X2, X0, X1] : (~ r1(sK51, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP2(X2)), inference(resolution, [], [f8585, f206])).
fof(f8585, plain, ! [X2, X0, X3, X1] : (~ r1(sK50, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP2(X3)), inference(resolution, [], [f8572, f205])).
fof(f8572, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK49, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | sP2(X4)), inference(resolution, [], [f8546, f204])).
fof(f8546, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK48, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | sP2(X5)), inference(resolution, [], [f8068, f203])).
fof(f8068, plain, ! [X26, X24, X23, X21, X27, X25, X22] : (~ r1(sK47, X24) | ~ r1(X23, X21) | ~ r1(X24, X23) | ~ r1(X21, X22) | ~ r1(X22, X25) | ~ r1(X25, X26) | ~ r1(X26, X27) | sP2(X27)), inference(resolution, [], [f6048, f202])).
fof(f6048, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK46, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ r1(X1, X6) | ~ r1(X6, X7) | sP2(X7)), inference(resolution, [], [f3448, f201])).
fof(f3448, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK45, X7) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X0, X1) | ~ r1(X1, X8) | sP2(X8)), inference(resolution, [], [f1089, f200])).
fof(f1089, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK44, X9) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f193, f199])).
fof(f193, plain, ! [X167, X165, X163, X173, X171, X169, X166, X164, X172, X170, X168] : (~ r1(sK13, X163) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | sP2(X173)), inference(cnf_transformation, [], [f106])).
fof(f8282, plain, (~ spl85_108 | spl85_109 | ~ spl85_35 | spl85_547), inference(avatar_split_clause, [], [f8278, f8280, f460, f855, f851])).
fof(f8278, plain, ! [X17] : (r1(X17, sK4(X17)) | ~ r1(sK54, X17) | ~ p1(sK54) | sP0(sK54) | ~ sP3(sK54)), inference(subsumption_resolution, [], [f897, f211])).
fof(f897, plain, ! [X17] : (r1(X17, sK4(X17)) | ~ r1(sK54, X17) | ~ p1(sK54) | p1(sK58) | sP0(sK54) | ~ sP3(sK54)), inference(resolution, [], [f107, f210])).
fof(f107, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(X1, sK4(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f8277, plain, (~ spl85_108 | spl85_109 | ~ spl85_35 | spl85_546), inference(avatar_split_clause, [], [f8273, f8275, f460, f855, f851])).
fof(f8273, plain, ! [X17] : (r1(sK4(X17), sK5(X17)) | ~ r1(sK54, X17) | ~ p1(sK54) | sP0(sK54) | ~ sP3(sK54)), inference(subsumption_resolution, [], [f926, f211])).
fof(f926, plain, ! [X17] : (r1(sK4(X17), sK5(X17)) | ~ r1(sK54, X17) | ~ p1(sK54) | p1(sK58) | sP0(sK54) | ~ sP3(sK54)), inference(resolution, [], [f108, f210])).
fof(f108, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(sK4(X1), sK5(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f8272, plain, spl85_108, inference(avatar_split_clause, [], [f8206, f851])).
fof(f8206, plain, sP3(sK54), inference(resolution, [], [f8150, f209])).
fof(f8150, plain, ! [X0] : (~ r1(sK53, X0) | sP3(X0)), inference(resolution, [], [f8136, f208])).
fof(f8136, plain, ! [X0, X1] : (~ r1(sK52, X0) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f8124, f207])).
fof(f8124, plain, ! [X2, X0, X1] : (~ r1(sK51, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP3(X2)), inference(resolution, [], [f8099, f206])).
fof(f8099, plain, ! [X2, X0, X3, X1] : (~ r1(sK50, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP3(X3)), inference(resolution, [], [f8078, f205])).
fof(f8078, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK49, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | sP3(X4)), inference(resolution, [], [f8053, f204])).
fof(f8053, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK48, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | sP3(X5)), inference(resolution, [], [f8043, f203])).
fof(f8043, plain, ! [X26, X24, X23, X21, X27, X25, X22] : (~ r1(sK47, X24) | ~ r1(X23, X21) | ~ r1(X24, X23) | ~ r1(X21, X22) | ~ r1(X22, X25) | ~ r1(X25, X26) | ~ r1(X26, X27) | sP3(X27)), inference(resolution, [], [f6039, f202])).
fof(f6039, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK46, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ r1(X1, X6) | ~ r1(X6, X7) | sP3(X7)), inference(resolution, [], [f3424, f201])).
fof(f3424, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK45, X7) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X0, X1) | ~ r1(X1, X8) | sP3(X8)), inference(resolution, [], [f1084, f200])).
fof(f1084, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK44, X9) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f192, f199])).
fof(f192, plain, ! [X167, X165, X163, X173, X171, X169, X166, X164, X172, X170, X168] : (~ r1(sK13, X163) | ~ r1(X172, X173) | ~ r1(X171, X172) | ~ r1(X170, X171) | ~ r1(X169, X170) | ~ r1(X168, X169) | ~ r1(X167, X168) | ~ r1(X166, X167) | ~ r1(X165, X166) | ~ r1(X164, X165) | ~ r1(X163, X164) | sP3(X173)), inference(cnf_transformation, [], [f106])).
fof(f678, plain, (~ spl85_72 | spl85_35 | spl85_73 | spl85_75), inference(avatar_split_clause, [], [f511, f675, f666, f460, f662])).
fof(f511, plain, (r1(sK57, sK7(sK57)) | r1(sK54, sK6(sK54)) | p1(sK54) | ~ sP2(sK54)), inference(resolution, [], [f111, f212])).
fof(f111, plain, ! [X4, X0] : (~ r1(X0, X4) | r1(X4, sK7(X4)) | r1(X0, sK6(X0)) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).