fof(f159136, plain, $false, inference(avatar_sat_refutation, [], [f657, f661, f669, f674, f679, f958, f2221, f24250, f42968, f42972, f156458, f156461, f156830, f156842, f156862, f157996, f158000, f158074, f158075, f158097, f158197, f158199, f159135])).
fof(f159135, plain, (~ spl91_62 | spl91_132 | ~ spl91_354), inference(avatar_contradiction_clause, [], [f159134])).
fof(f159134, plain, ($false | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(subsumption_resolution, [], [f159133, f665])).
fof(f665, plain, (sP1(sK81) | ~ spl91_62), inference(avatar_component_clause, [], [f663])).
fof(f663, plain, (spl91_62 <=> sP1(sK81)), introduced(avatar_definition, [new_symbols(naming, [spl91_62])])).
fof(f159133, plain, (~ sP1(sK81) | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(duplicate_literal_removal, [], [f159132])).
fof(f159132, plain, (~ sP1(sK81) | ~ sP1(sK81) | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(resolution, [], [f158284, f286])).
fof(f286, plain, ! [X0] : (r1(X0, sK60(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0] : ((! [X1] : ((p2(sK58(X1)) & (~ p2(sK59(X1)) & r1(sK58(X1), sK59(X1))) & r1(X1, sK58(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK61(X0), X6)) & ~ p2(sK61(X0)) & r1(sK60(X0), sK61(X0))) & r1(X0, sK60(X0)))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58, sK59, sK60, sK61])], [f126, f130, f129, f128, f127])).
fof(f127, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK58(X1)) & ? [X3] : (~ p2(X3) & r1(sK58(X1), X3)) & r1(X1, sK58(X1)))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK58(X1), X3)) => (~ p2(sK59(X1)) & r1(sK58(X1), sK59(X1)))), introduced(choice_axiom, [])).
fof(f129, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK60(X0), X5)) & r1(X0, sK60(X0)))), introduced(choice_axiom, [])).
fof(f130, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK60(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK61(X0), X6)) & ~ p2(sK61(X0)) & r1(sK60(X0), sK61(X0)))), introduced(choice_axiom, [])).
fof(f126, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP1(X0)), inference(rectify, [], [f125])).
fof(f125, plain, ! [X92] : ((! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118))) | ~ sP1(X92)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X92] : ((! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118))) | ~ sP1(X92)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X92] : (sP1(X92) <=> (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f158284, plain, (! [X0] : (~ r1(X0, sK60(sK81)) | ~ sP1(X0)) | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(subsumption_resolution, [], [f158283, f1085])).
fof(f1085, plain, (~ p2(sK60(sK81)) | spl91_132), inference(avatar_component_clause, [], [f1083])).
fof(f1083, plain, (spl91_132 <=> p2(sK60(sK81))), introduced(avatar_definition, [new_symbols(naming, [spl91_132])])).
fof(f158283, plain, (! [X0] : (p2(sK60(sK81)) | ~ r1(X0, sK60(sK81)) | ~ sP1(X0)) | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(resolution, [], [f158223, f292])).
fof(f292, plain, ! [X0, X1] : (~ p2(sK59(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f158223, plain, (p2(sK59(sK60(sK81))) | (~ spl91_62 | spl91_132 | ~ spl91_354)), inference(subsumption_resolution, [], [f158222, f665])).
fof(f158222, plain, (p2(sK59(sK60(sK81))) | ~ sP1(sK81) | (spl91_132 | ~ spl91_354)), inference(subsumption_resolution, [], [f158204, f1085])).
fof(f158204, plain, (p2(sK59(sK60(sK81))) | p2(sK60(sK81)) | ~ sP1(sK81) | ~ spl91_354), inference(resolution, [], [f156519, f1572])).
fof(f1572, plain, ! [X12] : (r1(sK58(sK60(X12)), sK59(sK60(X12))) | p2(sK60(X12)) | ~ sP1(X12)), inference(duplicate_literal_removal, [], [f1550])).
fof(f1550, plain, ! [X12] : (p2(sK60(X12)) | r1(sK58(sK60(X12)), sK59(sK60(X12))) | ~ sP1(X12) | ~ sP1(X12)), inference(resolution, [], [f291, f286])).
fof(f291, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK58(X1), sK59(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f156519, plain, (! [X3] : (~ r1(sK58(sK60(sK81)), X3) | p2(X3)) | ~ spl91_354), inference(avatar_component_clause, [], [f156518])).
fof(f156518, plain, (spl91_354 <=> ! [X3] : (p2(X3) | ~ r1(sK58(sK60(sK81)), X3))), introduced(avatar_definition, [new_symbols(naming, [spl91_354])])).
fof(f158199, plain, (~ spl91_62 | spl91_132 | spl91_354 | ~ spl91_188), inference(avatar_split_clause, [], [f158198, f2218, f156518, f1083, f663])).
fof(f2218, plain, (spl91_188 <=> ! [X25, X24] : (~ p2(X24) | p2(X25) | ~ r1(X24, X25) | ~ r1(sK60(sK81), X24))), introduced(avatar_definition, [new_symbols(naming, [spl91_188])])).
fof(f158198, plain, (! [X3] : (p2(X3) | ~ r1(sK58(sK60(sK81)), X3) | p2(sK60(sK81)) | ~ sP1(sK81)) | ~ spl91_188), inference(subsumption_resolution, [], [f156465, f1240])).
fof(f1240, plain, ! [X12] : (p2(sK58(sK60(X12))) | p2(sK60(X12)) | ~ sP1(X12)), inference(duplicate_literal_removal, [], [f1218])).
fof(f1218, plain, ! [X12] : (p2(sK60(X12)) | p2(sK58(sK60(X12))) | ~ sP1(X12) | ~ sP1(X12)), inference(resolution, [], [f293, f286])).
fof(f293, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK58(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f156465, plain, (! [X3] : (p2(X3) | ~ r1(sK58(sK60(sK81)), X3) | ~ p2(sK58(sK60(sK81))) | p2(sK60(sK81)) | ~ sP1(sK81)) | ~ spl91_188), inference(resolution, [], [f2219, f1423])).
fof(f1423, plain, ! [X12] : (r1(sK60(X12), sK58(sK60(X12))) | p2(sK60(X12)) | ~ sP1(X12)), inference(duplicate_literal_removal, [], [f1401])).
fof(f1401, plain, ! [X12] : (p2(sK60(X12)) | r1(sK60(X12), sK58(sK60(X12))) | ~ sP1(X12) | ~ sP1(X12)), inference(resolution, [], [f290, f286])).
fof(f290, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK58(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f2219, plain, (! [X24, X25] : (~ r1(sK60(sK81), X24) | p2(X25) | ~ r1(X24, X25) | ~ p2(X24)) | ~ spl91_188), inference(avatar_component_clause, [], [f2218])).
fof(f158197, plain, (spl91_64 | ~ spl91_65 | ~ spl91_539), inference(avatar_contradiction_clause, [], [f158196])).
fof(f158196, plain, ($false | (spl91_64 | ~ spl91_65 | ~ spl91_539)), inference(subsumption_resolution, [], [f158195, f678])).
fof(f678, plain, (r1(sK64, sK81) | ~ spl91_65), inference(avatar_component_clause, [], [f676])).
fof(f676, plain, (spl91_65 <=> r1(sK64, sK81)), introduced(avatar_definition, [new_symbols(naming, [spl91_65])])).
fof(f158195, plain, (~ r1(sK64, sK81) | (spl91_64 | ~ spl91_65 | ~ spl91_539)), inference(subsumption_resolution, [], [f158194, f673])).
fof(f673, plain, (~ p2(sK81) | spl91_64), inference(avatar_component_clause, [], [f671])).
fof(f671, plain, (spl91_64 <=> p2(sK81)), introduced(avatar_definition, [new_symbols(naming, [spl91_64])])).
fof(f158194, plain, (p2(sK81) | ~ r1(sK64, sK81) | (spl91_64 | ~ spl91_65 | ~ spl91_539)), inference(resolution, [], [f158138, f308])).
fof(f308, plain, ! [X51] : (~ p2(sK86(X51)) | p2(X51) | ~ r1(sK64, X51)), inference(cnf_transformation, [], [f165])).
fof(f165, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK65(X1), X3)) & ~ p2(sK65(X1)) & r1(X1, sK65(X1))) | p2(X1) | ~ r1(sK64, X1)) & ((! [X6] : ((r1(X6, sK67(X6)) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK66, X6)) & r1(sK66, sK68) & ~ p1(sK66) & r1(sK64, sK66)) | ! [X11] : ~ r1(sK64, X11) | p1(sK64)) & ((! [X13] : ((r1(X13, sK70(X13)) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK69, X13)) & r1(sK69, sK71) & ~ p1(sK69) & ~ p2(sK69) & r1(sK64, sK69)) | ! [X18] : ~ r1(sK64, X18) | p1(sK64) | p2(sK64)) & ((sP21(sK72) & r1(sK72, sK73) & ~ p1(sK72) & ~ p2(sK72) & ~ p3(sK72) & r1(sK64, sK72)) | ! [X21] : ~ r1(sK64, X21) | p1(sK64) | p2(sK64) | p3(sK64)) & ((sP20(sK74) & r1(sK74, sK75) & ~ p1(sK74) & ~ p2(sK74) & ~ p3(sK74) & ~ p4(sK74) & r1(sK64, sK74)) | ! [X24] : ~ r1(sK64, X24) | p1(sK64) | p2(sK64) | p3(sK64) | p4(sK64)) & ((sP18(sK76) & sP19(sK76) & ~ p1(sK76) & r1(sK64, sK76)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(sK64, X26)) | p1(sK64)) & ((sP16(sK77) & sP17(sK77) & ~ p1(sK77) & ~ p2(sK77) & r1(sK64, sK77)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(sK64, X29)) | p1(sK64) | p2(sK64)) & ((sP15(sK78) & sP14(sK78) & ~ p1(sK78) & ~ p2(sK78) & ~ p3(sK78) & r1(sK64, sK78)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(sK64, X32)) | p1(sK64) | p2(sK64) | p3(sK64)) & ((sP12(sK79) & sP11(sK79) & ~ p1(sK79) & ~ p2(sK79) & ~ p3(sK79) & ~ p4(sK79) & r1(sK64, sK79)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(sK64, X35)) | p1(sK64) | p2(sK64) | p3(sK64) | p4(sK64)) & ((sP8(sK80) & sP9(sK80) & ~ p1(sK80) & r1(sK64, sK80)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(sK64, X38)) | p1(sK64)) & ((! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(sK81, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(sK81, X45)) & ~ p2(sK81)) | sP1(sK81)) & r1(sK64, sK81)) | sP5(sK64)) & ! [X47] : ((p1(sK82(X47)) & (~ p1(sK83(X47)) & r1(sK82(X47), sK83(X47))) & r1(X47, sK82(X47))) | p1(X47) | ~ r1(sK64, X47)) & (~ p1(sK84) & r1(sK64, sK84)) & ! [X51] : ((p2(sK85(X51)) & (~ p2(sK86(X51)) & r1(sK85(X51), sK86(X51))) & r1(X51, sK85(X51))) | p2(X51) | ~ r1(sK64, X51)) & (~ p2(sK87) & r1(sK64, sK87)) & ! [X55] : ((p3(sK88(X55)) & (~ p3(sK89(X55)) & r1(sK88(X55), sK89(X55))) & r1(X55, sK88(X55))) | p3(X55) | ~ r1(sK64, X55)) & (~ p3(sK90) & r1(sK64, sK90))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK64, sK65, sK66, sK67, sK68, sK69, sK70, sK71, sK72, sK73, sK74, sK75, sK76, sK77, sK78, sK79, sK80, sK81, sK82, sK83, sK84, sK85, sK86, sK87, sK88, sK89, sK90])], [f137, f164, f163, f162, f161, f160, f159, f158, f157, f156, f155, f154, f153, f152, f151, f150, f149, f148, f147, f146, f145, f144, f143, f142, f141, f140, f139, f138])).
fof(f138, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP21(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP20(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (sP18(X25) & sP19(X25) & ~ p1(X25) & r1(X0, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | p1(X0)) & (? [X28] : (sP16(X28) & sP17(X28) & ~ p1(X28) & ~ p2(X28) & r1(X0, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X0, X29)) | p1(X0) | p2(X0)) & (? [X31] : (sP15(X31) & sP14(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(X0, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0)) & (? [X34] : (sP12(X34) & sP11(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(X0, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X0, X35)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X37] : (sP8(X37) & sP9(X37) & ~ p1(X37) & r1(X0, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X0, X38)) | p1(X0)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP1(X41)) & r1(X0, X41)) | sP5(X0)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(X0, X47)) & ? [X50] : (~ p1(X50) & r1(X0, X50)) & ! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X0, X51)) & ? [X54] : (~ p2(X54) & r1(X0, X54)) & ! [X55] : (? [X56] : (p3(X56) & ? [X57] : (~ p3(X57) & r1(X56, X57)) & r1(X55, X56)) | p3(X55) | ~ r1(X0, X55)) & ? [X58] : (~ p3(X58) & r1(X0, X58))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK64, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK64, X5)) | ! [X11] : ~ r1(sK64, X11) | p1(sK64)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK64, X12)) | ! [X18] : ~ r1(sK64, X18) | p1(sK64) | p2(sK64)) & (? [X19] : (sP21(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK64, X19)) | ! [X21] : ~ r1(sK64, X21) | p1(sK64) | p2(sK64) | p3(sK64)) & (? [X22] : (sP20(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK64, X22)) | ! [X24] : ~ r1(sK64, X24) | p1(sK64) | p2(sK64) | p3(sK64) | p4(sK64)) & (? [X25] : (sP18(X25) & sP19(X25) & ~ p1(X25) & r1(sK64, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(sK64, X26)) | p1(sK64)) & (? [X28] : (sP16(X28) & sP17(X28) & ~ p1(X28) & ~ p2(X28) & r1(sK64, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(sK64, X29)) | p1(sK64) | p2(sK64)) & (? [X31] : (sP15(X31) & sP14(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(sK64, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(sK64, X32)) | p1(sK64) | p2(sK64) | p3(sK64)) & (? [X34] : (sP12(X34) & sP11(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(sK64, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(sK64, X35)) | p1(sK64) | p2(sK64) | p3(sK64) | p4(sK64)) & (? [X37] : (sP8(X37) & sP9(X37) & ~ p1(X37) & r1(sK64, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(sK64, X38)) | p1(sK64)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP1(X41)) & r1(sK64, X41)) | sP5(sK64)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(sK64, X47)) & ? [X50] : (~ p1(X50) & r1(sK64, X50)) & ! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(sK64, X51)) & ? [X54] : (~ p2(X54) & r1(sK64, X54)) & ! [X55] : (? [X56] : (p3(X56) & ? [X57] : (~ p3(X57) & r1(X56, X57)) & r1(X55, X56)) | p3(X55) | ~ r1(sK64, X55)) & ? [X58] : (~ p3(X58) & r1(sK64, X58)))), introduced(choice_axiom, [])).
fof(f139, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK65(X1), X3)) & ~ p2(sK65(X1)) & r1(X1, sK65(X1)))), introduced(choice_axiom, [])).
fof(f140, plain, (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(sK64, X5)) => (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(sK66, X6)) & ? [X10] : r1(sK66, X10) & ~ p1(sK66) & r1(sK64, sK66))), introduced(choice_axiom, [])).
fof(f141, plain, ! [X6] : (? [X7] : r1(X6, X7) => r1(X6, sK67(X6))), introduced(choice_axiom, [])).
fof(f142, plain, (? [X10] : r1(sK66, X10) => r1(sK66, sK68)), introduced(choice_axiom, [])).
fof(f143, plain, (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(sK64, X12)) => (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(sK69, X13)) & ? [X17] : r1(sK69, X17) & ~ p1(sK69) & ~ p2(sK69) & r1(sK64, sK69))), introduced(choice_axiom, [])).
fof(f144, plain, ! [X13] : (? [X14] : r1(X13, X14) => r1(X13, sK70(X13))), introduced(choice_axiom, [])).
fof(f145, plain, (? [X17] : r1(sK69, X17) => r1(sK69, sK71)), introduced(choice_axiom, [])).
fof(f146, plain, (? [X19] : (sP21(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(sK64, X19)) => (sP21(sK72) & ? [X20] : r1(sK72, X20) & ~ p1(sK72) & ~ p2(sK72) & ~ p3(sK72) & r1(sK64, sK72))), introduced(choice_axiom, [])).
fof(f147, plain, (? [X20] : r1(sK72, X20) => r1(sK72, sK73)), introduced(choice_axiom, [])).
fof(f148, plain, (? [X22] : (sP20(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(sK64, X22)) => (sP20(sK74) & ? [X23] : r1(sK74, X23) & ~ p1(sK74) & ~ p2(sK74) & ~ p3(sK74) & ~ p4(sK74) & r1(sK64, sK74))), introduced(choice_axiom, [])).
fof(f149, plain, (? [X23] : r1(sK74, X23) => r1(sK74, sK75)), introduced(choice_axiom, [])).
fof(f150, plain, (? [X25] : (sP18(X25) & sP19(X25) & ~ p1(X25) & r1(sK64, X25)) => (sP18(sK76) & sP19(sK76) & ~ p1(sK76) & r1(sK64, sK76))), introduced(choice_axiom, [])).
fof(f151, plain, (? [X28] : (sP16(X28) & sP17(X28) & ~ p1(X28) & ~ p2(X28) & r1(sK64, X28)) => (sP16(sK77) & sP17(sK77) & ~ p1(sK77) & ~ p2(sK77) & r1(sK64, sK77))), introduced(choice_axiom, [])).
fof(f152, plain, (? [X31] : (sP15(X31) & sP14(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(sK64, X31)) => (sP15(sK78) & sP14(sK78) & ~ p1(sK78) & ~ p2(sK78) & ~ p3(sK78) & r1(sK64, sK78))), introduced(choice_axiom, [])).
fof(f153, plain, (? [X34] : (sP12(X34) & sP11(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(sK64, X34)) => (sP12(sK79) & sP11(sK79) & ~ p1(sK79) & ~ p2(sK79) & ~ p3(sK79) & ~ p4(sK79) & r1(sK64, sK79))), introduced(choice_axiom, [])).
fof(f154, plain, (? [X37] : (sP8(X37) & sP9(X37) & ~ p1(X37) & r1(sK64, X37)) => (sP8(sK80) & sP9(sK80) & ~ p1(sK80) & r1(sK64, sK80))), introduced(choice_axiom, [])).
fof(f155, plain, (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP1(X41)) & r1(sK64, X41)) => (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(sK81, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(sK81, X45)) & ~ p2(sK81)) | sP1(sK81)) & r1(sK64, sK81))), introduced(choice_axiom, [])).
fof(f156, plain, ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) => (p1(sK82(X47)) & ? [X49] : (~ p1(X49) & r1(sK82(X47), X49)) & r1(X47, sK82(X47)))), introduced(choice_axiom, [])).
fof(f157, plain, ! [X47] : (? [X49] : (~ p1(X49) & r1(sK82(X47), X49)) => (~ p1(sK83(X47)) & r1(sK82(X47), sK83(X47)))), introduced(choice_axiom, [])).
fof(f158, plain, (? [X50] : (~ p1(X50) & r1(sK64, X50)) => (~ p1(sK84) & r1(sK64, sK84))), introduced(choice_axiom, [])).
fof(f159, plain, ! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) => (p2(sK85(X51)) & ? [X53] : (~ p2(X53) & r1(sK85(X51), X53)) & r1(X51, sK85(X51)))), introduced(choice_axiom, [])).
fof(f160, plain, ! [X51] : (? [X53] : (~ p2(X53) & r1(sK85(X51), X53)) => (~ p2(sK86(X51)) & r1(sK85(X51), sK86(X51)))), introduced(choice_axiom, [])).
fof(f161, plain, (? [X54] : (~ p2(X54) & r1(sK64, X54)) => (~ p2(sK87) & r1(sK64, sK87))), introduced(choice_axiom, [])).
fof(f162, plain, ! [X55] : (? [X56] : (p3(X56) & ? [X57] : (~ p3(X57) & r1(X56, X57)) & r1(X55, X56)) => (p3(sK88(X55)) & ? [X57] : (~ p3(X57) & r1(sK88(X55), X57)) & r1(X55, sK88(X55)))), introduced(choice_axiom, [])).
fof(f163, plain, ! [X55] : (? [X57] : (~ p3(X57) & r1(sK88(X55), X57)) => (~ p3(sK89(X55)) & r1(sK88(X55), sK89(X55)))), introduced(choice_axiom, [])).
fof(f164, plain, (? [X58] : (~ p3(X58) & r1(sK64, X58)) => (~ p3(sK90) & r1(sK64, sK90))), introduced(choice_axiom, [])).
fof(f137, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP21(X19) & ? [X20] : r1(X19, X20) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X21] : ~ r1(X0, X21) | p1(X0) | p2(X0) | p3(X0)) & (? [X22] : (sP20(X22) & ? [X23] : r1(X22, X23) & ~ p1(X22) & ~ p2(X22) & ~ p3(X22) & ~ p4(X22) & r1(X0, X22)) | ! [X24] : ~ r1(X0, X24) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X25] : (sP18(X25) & sP19(X25) & ~ p1(X25) & r1(X0, X25)) | ! [X26] : (! [X27] : ~ r1(X26, X27) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | p1(X0)) & (? [X28] : (sP16(X28) & sP17(X28) & ~ p1(X28) & ~ p2(X28) & r1(X0, X28)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X0, X29)) | p1(X0) | p2(X0)) & (? [X31] : (sP15(X31) & sP14(X31) & ~ p1(X31) & ~ p2(X31) & ~ p3(X31) & r1(X0, X31)) | ! [X32] : (! [X33] : ~ r1(X32, X33) | p1(X32) | p2(X32) | p3(X32) | p4(X32) | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0)) & (? [X34] : (sP12(X34) & sP11(X34) & ~ p1(X34) & ~ p2(X34) & ~ p3(X34) & ~ p4(X34) & r1(X0, X34)) | ! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X0, X35)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X37] : (sP8(X37) & sP9(X37) & ~ p1(X37) & r1(X0, X37)) | ! [X38] : (! [X39] : (! [X40] : ~ r1(X39, X40) | p1(X39) | p2(X39) | p3(X39) | p4(X39) | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X0, X38)) | p1(X0)) & (? [X41] : (! [X42] : ((! [X43] : (~ p2(X43) | ! [X44] : (p2(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p2(X42)) | sP3(X42) | sP4(X42) | ~ r1(X41, X42)) & ((! [X45] : (~ p2(X45) | ! [X46] : (p2(X46) | ~ r1(X45, X46)) | ~ r1(X41, X45)) & ~ p2(X41)) | sP1(X41)) & r1(X0, X41)) | sP5(X0)) & ! [X47] : (? [X48] : (p1(X48) & ? [X49] : (~ p1(X49) & r1(X48, X49)) & r1(X47, X48)) | p1(X47) | ~ r1(X0, X47)) & ? [X50] : (~ p1(X50) & r1(X0, X50)) & ! [X51] : (? [X52] : (p2(X52) & ? [X53] : (~ p2(X53) & r1(X52, X53)) & r1(X51, X52)) | p2(X51) | ~ r1(X0, X51)) & ? [X54] : (~ p2(X54) & r1(X0, X54)) & ! [X55] : (? [X56] : (p3(X56) & ? [X57] : (~ p3(X57) & r1(X56, X57)) & r1(X55, X56)) | p3(X55) | ~ r1(X0, X55)) & ? [X58] : (~ p3(X58) & r1(X0, X58))), inference(rectify, [], [f30])).
fof(f30, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (sP21(X19) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (sP20(X26) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (sP18(X33) & sP19(X33) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (sP16(X44) & sP17(X44) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (sP15(X55) & sP14(X55) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (sP12(X66) & sP11(X66) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (sP8(X77) & sP9(X77) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : ((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | sP3(X93) | sP4(X93) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | sP1(X92)) & r1(X0, X92)) | sP5(X0)) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & ! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138)) & ! [X139] : (? [X140] : (p3(X140) & ? [X141] : (~ p3(X141) & r1(X140, X141)) & r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) & ? [X142] : (~ p3(X142) & r1(X0, X142))), inference(definition_folding, [], [f7, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8])).
fof(f8, plain, ! [X0] : (! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)) | ~ sP0(X0)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X0] : (sP0(X0) <=> ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X103] : (! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)) | ~ sP2(X103)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X103] : (sP2(X103) <=> ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f11, plain, ! [X93] : ((! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ~ sP3(X93)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X93] : (sP3(X93) <=> (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f12, plain, ! [X93] : (! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP2(X103))) | ~ r1(X93, X103)) | ~ sP4(X93)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X93] : (sP4(X93) <=> ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP2(X103))) | ~ r1(X93, X103))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f13, plain, ! [X0] : (((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP0(X0))) | ~ sP5(X0)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X0] : (sP5(X0) <=> ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f14, plain, ! [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) | ~ sP6(X86)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X86] : (sP6(X86) <=> ? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f15, plain, ! [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) | ~ sP7(X79)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X79] : (sP7(X79) <=> ? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f16, plain, ! [X77] : (! [X78] : ((? [X79] : (sP7(X79) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ~ sP8(X77)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X77] : (sP8(X77) <=> ! [X78] : ((? [X79] : (sP7(X79) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f17, plain, ! [X77] : (? [X86] : (sP6(X86) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) | ~ sP9(X77)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X77] : (sP9(X77) <=> ? [X86] : (sP6(X86) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f18, plain, ! [X67] : (? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) | ~ sP10(X67)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X67] : (sP10(X67) <=> ? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f19, plain, ! [X66] : (? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) | ~ sP11(X66)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X66] : (sP11(X66) <=> ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f20, plain, ! [X66] : (! [X67] : ((sP10(X67) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ~ sP12(X66)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X66] : (sP12(X66) <=> ! [X67] : ((sP10(X67) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f21, plain, ! [X56] : (? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) | ~ sP13(X56)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X56] : (sP13(X56) <=> ? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f22, plain, ! [X55] : (? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) | ~ sP14(X55)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X55] : (sP14(X55) <=> ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f23, plain, ! [X55] : (! [X56] : ((sP13(X56) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ~ sP15(X55)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X55] : (sP15(X55) <=> ! [X56] : ((sP13(X56) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f24, plain, ! [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ~ sP16(X44)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X44] : (sP16(X44) <=> ! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f25, plain, ! [X44] : (? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) | ~ sP17(X44)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X44] : (sP17(X44) <=> ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f26, plain, ! [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ~ sP18(X33)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X33] : (sP18(X33) <=> ! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f27, plain, ! [X33] : (? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) | ~ sP19(X33)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X33] : (sP19(X33) <=> ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f28, plain, ! [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ~ sP20(X26)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X26] : (sP20(X26) <=> ! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f29, plain, ! [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ~ sP21(X19)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X19] : (sP21(X19) <=> ! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f7, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) & ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) & ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (! [X56] : ((? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) & ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (! [X67] : ((? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) & ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (! [X78] : ((? [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) & ? [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : ((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))) & r1(X0, X92)) | ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127))))) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & ! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138)) & ! [X139] : (? [X140] : (p3(X140) & ? [X141] : (~ p3(X141) & r1(X140, X141)) & r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) & ? [X142] : (~ p3(X142) & r1(X0, X142))), inference(flattening, [], [f6])).
fof(f6, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((? [X7] : r1(X6, X7) & ~ p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) & ? [X10] : r1(X5, X10) & ~ p1(X5) & r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (? [X12] : (! [X13] : ((? [X14] : r1(X13, X14) & ~ p1(X13) & ~ p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ? [X17] : r1(X12, X17) & ~ p1(X12) & ~ p2(X12) & r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (? [X19] : (! [X20] : ((? [X21] : r1(X20, X21) & ~ p1(X20) & ~ p2(X20) & ~ p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) & ? [X24] : r1(X19, X24) & ~ p1(X19) & ~ p2(X19) & ~ p3(X19) & r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (? [X26] : (! [X27] : ((? [X28] : r1(X27, X28) & ~ p1(X27) & ~ p2(X27) & ~ p3(X27) & ~ p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) & ? [X31] : r1(X26, X31) & ~ p1(X26) & ~ p2(X26) & ~ p3(X26) & ~ p4(X26) & r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X33] : (! [X34] : ((? [X35] : (? [X36] : r1(X35, X36) & ~ p1(X35) & ~ p2(X35) & ~ p3(X35) & ~ p4(X35) & r1(X34, X35)) & ~ p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) & ? [X40] : (? [X41] : r1(X40, X41) & ~ p1(X40) & ~ p2(X40) & ~ p3(X40) & ~ p4(X40) & r1(X33, X40)) & ~ p1(X33) & r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (? [X44] : (! [X45] : ((? [X46] : (? [X47] : r1(X46, X47) & ~ p1(X46) & ~ p2(X46) & ~ p3(X46) & ~ p4(X46) & r1(X45, X46)) & ~ p1(X45) & ~ p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) & ? [X51] : (? [X52] : r1(X51, X52) & ~ p1(X51) & ~ p2(X51) & ~ p3(X51) & ~ p4(X51) & r1(X44, X51)) & ~ p1(X44) & ~ p2(X44) & r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (? [X55] : (! [X56] : ((? [X57] : (? [X58] : r1(X57, X58) & ~ p1(X57) & ~ p2(X57) & ~ p3(X57) & ~ p4(X57) & r1(X56, X57)) & ~ p1(X56) & ~ p2(X56) & ~ p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) & ? [X62] : (? [X63] : r1(X62, X63) & ~ p1(X62) & ~ p2(X62) & ~ p3(X62) & ~ p4(X62) & r1(X55, X62)) & ~ p1(X55) & ~ p2(X55) & ~ p3(X55) & r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (? [X66] : (! [X67] : ((? [X68] : (? [X69] : r1(X68, X69) & ~ p1(X68) & ~ p2(X68) & ~ p3(X68) & ~ p4(X68) & r1(X67, X68)) & ~ p1(X67) & ~ p2(X67) & ~ p3(X67) & ~ p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) & ? [X73] : (? [X74] : r1(X73, X74) & ~ p1(X73) & ~ p2(X73) & ~ p3(X73) & ~ p4(X73) & r1(X66, X73)) & ~ p1(X66) & ~ p2(X66) & ~ p3(X66) & ~ p4(X66) & r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (? [X77] : (! [X78] : ((? [X79] : (? [X80] : (? [X81] : r1(X80, X81) & ~ p1(X80) & ~ p2(X80) & ~ p3(X80) & ~ p4(X80) & r1(X79, X80)) & ~ p1(X79) & ~ p2(X79) & ~ p3(X79) & ~ p4(X79) & r1(X78, X79)) & ~ p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) & ? [X86] : (? [X87] : (? [X88] : r1(X87, X88) & ~ p1(X87) & ~ p2(X87) & ~ p3(X87) & ~ p4(X87) & r1(X86, X87)) & ~ p1(X86) & ~ p2(X86) & ~ p3(X86) & ~ p4(X86) & r1(X77, X86)) & ~ p1(X77) & r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (? [X92] : (! [X93] : (((! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) & ~ p2(X93)) | (! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99)))) | ! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | ! [X109] : (! [X110] : (? [X111] : (p2(X111) & ? [X112] : (~ p2(X112) & r1(X111, X112)) & r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) & ((! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) & ~ p2(X92)) | (! [X115] : (? [X116] : (p2(X116) & ? [X117] : (~ p2(X117) & r1(X116, X117)) & r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) & ? [X118] : (? [X119] : (! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) & ~ p2(X119) & r1(X118, X119)) & r1(X92, X118)))) & r1(X0, X92)) | ((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | ! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) & ! [X131] : (? [X132] : (p1(X132) & ? [X133] : (~ p1(X133) & r1(X132, X133)) & r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) & ? [X134] : (~ p1(X134) & r1(X0, X134)) & ! [X135] : (? [X136] : (p2(X136) & ? [X137] : (~ p2(X137) & r1(X136, X137)) & r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) & ? [X138] : (~ p2(X138) & r1(X0, X138)) & ! [X139] : (? [X140] : (p3(X140) & ? [X141] : (~ p3(X141) & r1(X140, X141)) & r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) & ? [X142] : (~ p3(X142) & r1(X0, X142))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ~ r1(X40, X41) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ~ r1(X46, X47) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ~ r1(X51, X52) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ~ r1(X57, X58) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ~ r1(X62, X63) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ~ r1(X68, X69) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ~ r1(X73, X74) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ~ r1(X80, X81) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ~ r1(X87, X88) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138)) | ~ ! [X139] : (~ ! [X140] : (~ p3(X140) | ! [X141] : (p3(X141) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) | ! [X142] : (p3(X142) | ~ r1(X0, X142))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ~ r1(X6, X7) | p1(X6)) | ! [X8] : (! [X9] : ~ r1(X8, X9) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ~ r1(X5, X10) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ~ r1(X0, X11) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ~ r1(X13, X14) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ~ r1(X15, X16) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ~ r1(X12, X17) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ~ r1(X0, X18) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ~ r1(X20, X21) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ~ r1(X22, X23) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ~ r1(X19, X24) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ~ r1(X0, X25) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ~ r1(X27, X28) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ~ r1(X29, X30) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ~ r1(X26, X31) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ~ r1(X0, X32) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ~ r1(X35, X36) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ~ r1(X38, X39) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ~ r1(X40, X41) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ~ r1(X42, X43) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ~ r1(X46, X47) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ~ r1(X49, X50) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ~ r1(X51, X52) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ~ r1(X53, X54) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ~ r1(X57, X58) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ~ r1(X60, X61) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ~ r1(X62, X63) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ~ r1(X64, X65) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ~ r1(X68, X69) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ~ r1(X71, X72) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ~ r1(X73, X74) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ~ r1(X75, X76) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ~ r1(X80, X81) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ~ r1(X84, X85) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ~ r1(X87, X88) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ~ r1(X90, X91) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138)) | ~ ! [X139] : (~ ! [X140] : (~ p3(X140) | ! [X141] : (p3(X141) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) | ! [X142] : (p3(X142) | ~ r1(X0, X142))), inference(true_and_false_elimination, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ (! [X7] : ($false | ~ r1(X6, X7)) | p1(X6)) | ! [X8] : (! [X9] : ($false | ~ r1(X8, X9)) | p1(X8) | ~ r1(X6, X8)) | ~ r1(X5, X6)) | ! [X10] : ($false | ~ r1(X5, X10)) | p1(X5) | ~ r1(X0, X5)) | ! [X11] : ($false | ~ r1(X0, X11)) | p1(X0)) & (~ ! [X12] : (~ ! [X13] : (~ (! [X14] : ($false | ~ r1(X13, X14)) | p1(X13) | p2(X13)) | ! [X15] : (! [X16] : ($false | ~ r1(X15, X16)) | p1(X15) | p2(X15) | ~ r1(X13, X15)) | ~ r1(X12, X13)) | ! [X17] : ($false | ~ r1(X12, X17)) | p1(X12) | p2(X12) | ~ r1(X0, X12)) | ! [X18] : ($false | ~ r1(X0, X18)) | p1(X0) | p2(X0)) & (~ ! [X19] : (~ ! [X20] : (~ (! [X21] : ($false | ~ r1(X20, X21)) | p1(X20) | p2(X20) | p3(X20)) | ! [X22] : (! [X23] : ($false | ~ r1(X22, X23)) | p1(X22) | p2(X22) | p3(X22) | ~ r1(X20, X22)) | ~ r1(X19, X20)) | ! [X24] : ($false | ~ r1(X19, X24)) | p1(X19) | p2(X19) | p3(X19) | ~ r1(X0, X19)) | ! [X25] : ($false | ~ r1(X0, X25)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X26] : (~ ! [X27] : (~ (! [X28] : ($false | ~ r1(X27, X28)) | p1(X27) | p2(X27) | p3(X27) | p4(X27)) | ! [X29] : (! [X30] : ($false | ~ r1(X29, X30)) | p1(X29) | p2(X29) | p3(X29) | p4(X29) | ~ r1(X27, X29)) | ~ r1(X26, X27)) | ! [X31] : ($false | ~ r1(X26, X31)) | p1(X26) | p2(X26) | p3(X26) | p4(X26) | ~ r1(X0, X26)) | ! [X32] : ($false | ~ r1(X0, X32)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X33] : (~ ! [X34] : (~ (! [X35] : (! [X36] : ($false | ~ r1(X35, X36)) | p1(X35) | p2(X35) | p3(X35) | p4(X35) | ~ r1(X34, X35)) | p1(X34)) | ! [X37] : (! [X38] : (! [X39] : ($false | ~ r1(X38, X39)) | p1(X38) | p2(X38) | p3(X38) | p4(X38) | ~ r1(X37, X38)) | p1(X37) | ~ r1(X34, X37)) | ~ r1(X33, X34)) | ! [X40] : (! [X41] : ($false | ~ r1(X40, X41)) | p1(X40) | p2(X40) | p3(X40) | p4(X40) | ~ r1(X33, X40)) | p1(X33) | ~ r1(X0, X33)) | ! [X42] : (! [X43] : ($false | ~ r1(X42, X43)) | p1(X42) | p2(X42) | p3(X42) | p4(X42) | ~ r1(X0, X42)) | p1(X0)) & (~ ! [X44] : (~ ! [X45] : (~ (! [X46] : (! [X47] : ($false | ~ r1(X46, X47)) | p1(X46) | p2(X46) | p3(X46) | p4(X46) | ~ r1(X45, X46)) | p1(X45) | p2(X45)) | ! [X48] : (! [X49] : (! [X50] : ($false | ~ r1(X49, X50)) | p1(X49) | p2(X49) | p3(X49) | p4(X49) | ~ r1(X48, X49)) | p1(X48) | p2(X48) | ~ r1(X45, X48)) | ~ r1(X44, X45)) | ! [X51] : (! [X52] : ($false | ~ r1(X51, X52)) | p1(X51) | p2(X51) | p3(X51) | p4(X51) | ~ r1(X44, X51)) | p1(X44) | p2(X44) | ~ r1(X0, X44)) | ! [X53] : (! [X54] : ($false | ~ r1(X53, X54)) | p1(X53) | p2(X53) | p3(X53) | p4(X53) | ~ r1(X0, X53)) | p1(X0) | p2(X0)) & (~ ! [X55] : (~ ! [X56] : (~ (! [X57] : (! [X58] : ($false | ~ r1(X57, X58)) | p1(X57) | p2(X57) | p3(X57) | p4(X57) | ~ r1(X56, X57)) | p1(X56) | p2(X56) | p3(X56)) | ! [X59] : (! [X60] : (! [X61] : ($false | ~ r1(X60, X61)) | p1(X60) | p2(X60) | p3(X60) | p4(X60) | ~ r1(X59, X60)) | p1(X59) | p2(X59) | p3(X59) | ~ r1(X56, X59)) | ~ r1(X55, X56)) | ! [X62] : (! [X63] : ($false | ~ r1(X62, X63)) | p1(X62) | p2(X62) | p3(X62) | p4(X62) | ~ r1(X55, X62)) | p1(X55) | p2(X55) | p3(X55) | ~ r1(X0, X55)) | ! [X64] : (! [X65] : ($false | ~ r1(X64, X65)) | p1(X64) | p2(X64) | p3(X64) | p4(X64) | ~ r1(X0, X64)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X66] : (~ ! [X67] : (~ (! [X68] : (! [X69] : ($false | ~ r1(X68, X69)) | p1(X68) | p2(X68) | p3(X68) | p4(X68) | ~ r1(X67, X68)) | p1(X67) | p2(X67) | p3(X67) | p4(X67)) | ! [X70] : (! [X71] : (! [X72] : ($false | ~ r1(X71, X72)) | p1(X71) | p2(X71) | p3(X71) | p4(X71) | ~ r1(X70, X71)) | p1(X70) | p2(X70) | p3(X70) | p4(X70) | ~ r1(X67, X70)) | ~ r1(X66, X67)) | ! [X73] : (! [X74] : ($false | ~ r1(X73, X74)) | p1(X73) | p2(X73) | p3(X73) | p4(X73) | ~ r1(X66, X73)) | p1(X66) | p2(X66) | p3(X66) | p4(X66) | ~ r1(X0, X66)) | ! [X75] : (! [X76] : ($false | ~ r1(X75, X76)) | p1(X75) | p2(X75) | p3(X75) | p4(X75) | ~ r1(X0, X75)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X77] : (~ ! [X78] : (~ (! [X79] : (! [X80] : (! [X81] : ($false | ~ r1(X80, X81)) | p1(X80) | p2(X80) | p3(X80) | p4(X80) | ~ r1(X79, X80)) | p1(X79) | p2(X79) | p3(X79) | p4(X79) | ~ r1(X78, X79)) | p1(X78)) | ! [X82] : (! [X83] : (! [X84] : (! [X85] : ($false | ~ r1(X84, X85)) | p1(X84) | p2(X84) | p3(X84) | p4(X84) | ~ r1(X83, X84)) | p1(X83) | p2(X83) | p3(X83) | p4(X83) | ~ r1(X82, X83)) | p1(X82) | ~ r1(X78, X82)) | ~ r1(X77, X78)) | ! [X86] : (! [X87] : (! [X88] : ($false | ~ r1(X87, X88)) | p1(X87) | p2(X87) | p3(X87) | p4(X87) | ~ r1(X86, X87)) | p1(X86) | p2(X86) | p3(X86) | p4(X86) | ~ r1(X77, X86)) | p1(X77) | ~ r1(X0, X77)) | ! [X89] : (! [X90] : (! [X91] : ($false | ~ r1(X90, X91)) | p1(X90) | p2(X90) | p3(X90) | p4(X90) | ~ r1(X89, X90)) | p1(X89) | p2(X89) | p3(X89) | p4(X89) | ~ r1(X0, X89)) | p1(X0)) & (~ ! [X92] : (~ ! [X93] : (~ ((~ ! [X94] : (~ p2(X94) | ! [X95] : (p2(X95) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | p2(X93)) & (~ ! [X96] : (~ ! [X97] : (~ p2(X97) | ! [X98] : (p2(X98) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) | ! [X99] : (! [X100] : (~ ! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p2(X100) | ~ r1(X99, X100)) | ~ r1(X93, X99)))) | ! [X103] : (((~ ! [X104] : (~ p2(X104) | ! [X105] : (p2(X105) | ~ r1(X104, X105)) | ~ r1(X103, X104)) | p2(X103)) & (~ ! [X106] : (~ ! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | p2(X106) | ~ r1(X103, X106)) | ! [X109] : (! [X110] : (~ ! [X111] : (~ p2(X111) | ! [X112] : (p2(X112) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | p2(X110) | ~ r1(X109, X110)) | ~ r1(X103, X109)))) | ~ r1(X93, X103)) | ~ r1(X92, X93)) | ((~ ! [X113] : (~ p2(X113) | ! [X114] : (p2(X114) | ~ r1(X113, X114)) | ~ r1(X92, X113)) | p2(X92)) & (~ ! [X115] : (~ ! [X116] : (~ p2(X116) | ! [X117] : (p2(X117) | ~ r1(X116, X117)) | ~ r1(X115, X116)) | p2(X115) | ~ r1(X92, X115)) | ! [X118] : (! [X119] : (~ ! [X120] : (~ p2(X120) | ! [X121] : (p2(X121) | ~ r1(X120, X121)) | ~ r1(X119, X120)) | p2(X119) | ~ r1(X118, X119)) | ~ r1(X92, X118)))) | ~ r1(X0, X92)) | ((~ ! [X122] : (~ p2(X122) | ! [X123] : (p2(X123) | ~ r1(X122, X123)) | ~ r1(X0, X122)) | p2(X0)) & (~ ! [X124] : (~ ! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | p2(X124) | ~ r1(X0, X124)) | ! [X127] : (! [X128] : (~ ! [X129] : (~ p2(X129) | ! [X130] : (p2(X130) | ~ r1(X129, X130)) | ~ r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)))))) | ~ ! [X131] : (~ ! [X132] : (~ p1(X132) | ! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X131, X132)) | p1(X131) | ~ r1(X0, X131)) | ! [X134] : (p1(X134) | ~ r1(X0, X134)) | ~ ! [X135] : (~ ! [X136] : (~ p2(X136) | ! [X137] : (p2(X137) | ~ r1(X136, X137)) | ~ r1(X135, X136)) | p2(X135) | ~ r1(X0, X135)) | ! [X138] : (p2(X138) | ~ r1(X0, X138)) | ~ ! [X139] : (~ ! [X140] : (~ p3(X140) | ! [X141] : (p3(X141) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | p3(X139) | ~ r1(X0, X139)) | ! [X142] : (p3(X142) | ~ r1(X0, X142))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (! [X0] : ($false | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : ($false | ~ r1(X0, X1)) | p1(X0) | p2(X0) | p3(X0) | p4(X0) | ~ r1(X1, X0)) | p1(X1) | p2(X1) | p3(X1) | p4(X1) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p3(X1) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL642+1.010.p', main)).
fof(f158138, plain, (p2(sK86(sK81)) | (spl91_64 | ~ spl91_65 | ~ spl91_539)), inference(subsumption_resolution, [], [f158137, f678])).
fof(f158137, plain, (p2(sK86(sK81)) | ~ r1(sK64, sK81) | (spl91_64 | ~ spl91_539)), inference(subsumption_resolution, [], [f158118, f673])).
fof(f158118, plain, (p2(sK86(sK81)) | p2(sK81) | ~ r1(sK64, sK81) | ~ spl91_539), inference(resolution, [], [f158096, f307])).
fof(f307, plain, ! [X51] : (r1(sK85(X51), sK86(X51)) | p2(X51) | ~ r1(sK64, X51)), inference(cnf_transformation, [], [f165])).
fof(f158096, plain, (! [X15] : (~ r1(sK85(sK81), X15) | p2(X15)) | ~ spl91_539), inference(avatar_component_clause, [], [f158095])).
fof(f158095, plain, (spl91_539 <=> ! [X15] : (~ r1(sK85(sK81), X15) | p2(X15))), introduced(avatar_definition, [new_symbols(naming, [spl91_539])])).
fof(f158097, plain, (spl91_64 | spl91_539 | ~ spl91_63 | ~ spl91_65 | ~ spl91_111), inference(avatar_split_clause, [], [f158093, f955, f676, f667, f158095, f671])).
fof(f667, plain, (spl91_63 <=> ! [X46, X45] : (~ p2(X45) | ~ r1(sK81, X45) | ~ r1(X45, X46) | p2(X46))), introduced(avatar_definition, [new_symbols(naming, [spl91_63])])).
fof(f955, plain, (spl91_111 <=> p2(sK85(sK81))), introduced(avatar_definition, [new_symbols(naming, [spl91_111])])).
fof(f158093, plain, (! [X15] : (~ r1(sK85(sK81), X15) | p2(X15) | p2(sK81)) | (~ spl91_63 | ~ spl91_65 | ~ spl91_111)), inference(subsumption_resolution, [], [f158092, f678])).
fof(f158092, plain, (! [X15] : (~ r1(sK85(sK81), X15) | p2(X15) | p2(sK81) | ~ r1(sK64, sK81)) | (~ spl91_63 | ~ spl91_111)), inference(subsumption_resolution, [], [f43247, f957])).
fof(f957, plain, (p2(sK85(sK81)) | ~ spl91_111), inference(avatar_component_clause, [], [f955])).
fof(f43247, plain, (! [X15] : (~ p2(sK85(sK81)) | ~ r1(sK85(sK81), X15) | p2(X15) | p2(sK81) | ~ r1(sK64, sK81)) | ~ spl91_63), inference(resolution, [], [f668, f306])).
fof(f306, plain, ! [X51] : (r1(X51, sK85(X51)) | p2(X51) | ~ r1(sK64, X51)), inference(cnf_transformation, [], [f165])).
fof(f668, plain, (! [X45, X46] : (~ r1(sK81, X45) | ~ p2(X45) | ~ r1(X45, X46) | p2(X46)) | ~ spl91_63), inference(avatar_component_clause, [], [f667])).
fof(f158075, plain, (spl91_478 | ~ spl91_265 | ~ spl91_286 | ~ spl91_422), inference(avatar_split_clause, [], [f158001, f156859, f42965, f24244, f157680])).
fof(f157680, plain, (spl91_478 <=> p2(sK63(sK65(sK87)))), introduced(avatar_definition, [new_symbols(naming, [spl91_478])])).
fof(f24244, plain, (spl91_265 <=> p2(sK62(sK65(sK87)))), introduced(avatar_definition, [new_symbols(naming, [spl91_265])])).
fof(f42965, plain, (spl91_286 <=> r1(sK65(sK87), sK62(sK65(sK87)))), introduced(avatar_definition, [new_symbols(naming, [spl91_286])])).
fof(f156859, plain, (spl91_422 <=> r1(sK62(sK65(sK87)), sK63(sK65(sK87)))), introduced(avatar_definition, [new_symbols(naming, [spl91_422])])).
fof(f158001, plain, (p2(sK63(sK65(sK87))) | (~ spl91_265 | ~ spl91_286 | ~ spl91_422)), inference(resolution, [], [f157099, f156861])).
fof(f156861, plain, (r1(sK62(sK65(sK87)), sK63(sK65(sK87))) | ~ spl91_422), inference(avatar_component_clause, [], [f156859])).
fof(f157099, plain, (! [X0] : (~ r1(sK62(sK65(sK87)), X0) | p2(X0)) | (~ spl91_265 | ~ spl91_286)), inference(subsumption_resolution, [], [f157098, f304])).
fof(f304, plain, r1(sK64, sK87), inference(cnf_transformation, [], [f165])).
fof(f157098, plain, (! [X0] : (p2(X0) | ~ r1(sK62(sK65(sK87)), X0) | ~ r1(sK64, sK87)) | (~ spl91_265 | ~ spl91_286)), inference(subsumption_resolution, [], [f157097, f305])).
fof(f305, plain, ~ p2(sK87), inference(cnf_transformation, [], [f165])).
fof(f157097, plain, (! [X0] : (p2(X0) | ~ r1(sK62(sK65(sK87)), X0) | p2(sK87) | ~ r1(sK64, sK87)) | (~ spl91_265 | ~ spl91_286)), inference(subsumption_resolution, [], [f156868, f24246])).
fof(f24246, plain, (p2(sK62(sK65(sK87))) | ~ spl91_265), inference(avatar_component_clause, [], [f24244])).
fof(f156868, plain, (! [X0] : (p2(X0) | ~ r1(sK62(sK65(sK87)), X0) | ~ p2(sK62(sK65(sK87))) | p2(sK87) | ~ r1(sK64, sK87)) | ~ spl91_286), inference(resolution, [], [f42967, f374])).
fof(f374, plain, ! [X4, X3, X1] : (~ r1(sK65(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK64, X1)), inference(cnf_transformation, [], [f165])).
fof(f42967, plain, (r1(sK65(sK87), sK62(sK65(sK87))) | ~ spl91_286), inference(avatar_component_clause, [], [f42965])).
fof(f158074, plain, (spl91_266 | ~ spl91_526), inference(avatar_split_clause, [], [f158073, f157990, f24248])).
fof(f24248, plain, (spl91_266 <=> ! [X1] : (~ r1(X1, sK87) | ~ sP0(X1))), introduced(avatar_definition, [new_symbols(naming, [spl91_266])])).
fof(f157990, plain, (spl91_526 <=> ! [X1, X0] : (~ r1(X0, sK65(sK87)) | ~ sP0(X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl91_526])])).
fof(f158073, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK87)) | ~ spl91_526), inference(subsumption_resolution, [], [f158072, f304])).
fof(f158072, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK87) | ~ r1(sK64, sK87)) | ~ spl91_526), inference(subsumption_resolution, [], [f158071, f305])).
fof(f158071, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK87) | p2(sK87) | ~ r1(sK64, sK87)) | ~ spl91_526), inference(resolution, [], [f157991, f372])).
fof(f372, plain, ! [X1] : (r1(X1, sK65(X1)) | p2(X1) | ~ r1(sK64, X1)), inference(cnf_transformation, [], [f165])).
fof(f157991, plain, (! [X0, X1] : (~ r1(X0, sK65(sK87)) | ~ sP0(X1) | ~ r1(X1, X0)) | ~ spl91_526), inference(avatar_component_clause, [], [f157990])).
fof(f158000, plain, ~ spl91_527, inference(avatar_contradiction_clause, [], [f157999])).
fof(f157999, plain, ($false | ~ spl91_527), inference(subsumption_resolution, [], [f157998, f304])).
fof(f157998, plain, (~ r1(sK64, sK87) | ~ spl91_527), inference(subsumption_resolution, [], [f157997, f305])).
fof(f157997, plain, (p2(sK87) | ~ r1(sK64, sK87) | ~ spl91_527), inference(resolution, [], [f157995, f373])).
fof(f373, plain, ! [X1] : (~ p2(sK65(X1)) | p2(X1) | ~ r1(sK64, X1)), inference(cnf_transformation, [], [f165])).
fof(f157995, plain, (p2(sK65(sK87)) | ~ spl91_527), inference(avatar_component_clause, [], [f157993])).
fof(f157993, plain, (spl91_527 <=> p2(sK65(sK87))), introduced(avatar_definition, [new_symbols(naming, [spl91_527])])).
fof(f157996, plain, (spl91_526 | spl91_527 | ~ spl91_478), inference(avatar_split_clause, [], [f157988, f157680, f157993, f157990])).
fof(f157988, plain, (! [X0, X1] : (p2(sK65(sK87)) | ~ r1(X0, sK65(sK87)) | ~ r1(X1, X0) | ~ sP0(X1)) | ~ spl91_478), inference(resolution, [], [f157682, f296])).
fof(f296, plain, ! [X2, X0, X1] : (~ p2(sK63(X2)) | p2(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ! [X0] : (! [X1] : (! [X2] : ((p2(sK62(X2)) & (~ p2(sK63(X2)) & r1(sK62(X2), sK63(X2))) & r1(X2, sK62(X2))) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK62, sK63])], [f133, f135, f134])).
fof(f134, plain, ! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) => (p2(sK62(X2)) & ? [X4] : (~ p2(X4) & r1(sK62(X2), X4)) & r1(X2, sK62(X2)))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X2] : (? [X4] : (~ p2(X4) & r1(sK62(X2), X4)) => (~ p2(sK63(X2)) & r1(sK62(X2), sK63(X2)))), introduced(choice_axiom, [])).
fof(f133, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p2(X3) & ? [X4] : (~ p2(X4) & r1(X3, X4)) & r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f132])).
fof(f132, plain, ! [X0] : (! [X127] : (! [X128] : (? [X129] : (p2(X129) & ? [X130] : (~ p2(X130) & r1(X129, X130)) & r1(X128, X129)) | p2(X128) | ~ r1(X127, X128)) | ~ r1(X0, X127)) | ~ sP0(X0)), inference(nnf_transformation, [], [f8])).
fof(f157682, plain, (p2(sK63(sK65(sK87))) | ~ spl91_478), inference(avatar_component_clause, [], [f157680])).
fof(f156862, plain, (spl91_422 | spl91_266), inference(avatar_split_clause, [], [f156857, f24248, f156859])).
fof(f156857, plain, ! [X2] : (~ r1(X2, sK87) | ~ sP0(X2) | r1(sK62(sK65(sK87)), sK63(sK65(sK87)))), inference(subsumption_resolution, [], [f48075, f305])).
fof(f48075, plain, ! [X2] : (~ r1(X2, sK87) | ~ sP0(X2) | p2(sK87) | r1(sK62(sK65(sK87)), sK63(sK65(sK87)))), inference(resolution, [], [f2025, f304])).
fof(f2025, plain, ! [X26, X27] : (~ r1(sK64, X26) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | r1(sK62(sK65(X26)), sK63(sK65(X26)))), inference(subsumption_resolution, [], [f1994, f373])).
fof(f1994, plain, ! [X26, X27] : (p2(sK65(X26)) | r1(sK62(sK65(X26)), sK63(sK65(X26))) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | ~ r1(sK64, X26)), inference(resolution, [], [f295, f372])).
fof(f295, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(sK62(X2), sK63(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f136])).
fof(f156842, plain, (~ spl91_62 | ~ spl91_132 | ~ spl91_61 | spl91_130 | spl91_131), inference(avatar_split_clause, [], [f156841, f1079, f1075, f659, f1083, f663])).
fof(f659, plain, (spl91_61 <=> ! [X42] : (~ p2(X42) | ~ r1(sK81, X42) | sP4(X42) | sP3(X42))), introduced(avatar_definition, [new_symbols(naming, [spl91_61])])).
fof(f1075, plain, (spl91_130 <=> sP3(sK60(sK81))), introduced(avatar_definition, [new_symbols(naming, [spl91_130])])).
fof(f1079, plain, (spl91_131 <=> sP4(sK60(sK81))), introduced(avatar_definition, [new_symbols(naming, [spl91_131])])).
fof(f156841, plain, (~ p2(sK60(sK81)) | ~ sP1(sK81) | (~ spl91_61 | spl91_130 | spl91_131)), inference(subsumption_resolution, [], [f156840, f1076])).
fof(f1076, plain, (~ sP3(sK60(sK81)) | spl91_130), inference(avatar_component_clause, [], [f1075])).
fof(f156840, plain, (~ p2(sK60(sK81)) | sP3(sK60(sK81)) | ~ sP1(sK81) | (~ spl91_61 | spl91_131)), inference(subsumption_resolution, [], [f1037, f1080])).
fof(f1080, plain, (~ sP4(sK60(sK81)) | spl91_131), inference(avatar_component_clause, [], [f1079])).
fof(f1037, plain, (~ p2(sK60(sK81)) | sP4(sK60(sK81)) | sP3(sK60(sK81)) | ~ sP1(sK81) | ~ spl91_61), inference(resolution, [], [f660, f286])).
fof(f660, plain, (! [X42] : (~ r1(sK81, X42) | ~ p2(X42) | sP4(X42) | sP3(X42)) | ~ spl91_61), inference(avatar_component_clause, [], [f659])).
fof(f156830, plain, (spl91_104 | ~ spl91_59), inference(avatar_split_clause, [], [f36338, f651, f914])).
fof(f914, plain, (spl91_104 <=> sP0(sK64)), introduced(avatar_definition, [new_symbols(naming, [spl91_104])])).
fof(f651, plain, (spl91_59 <=> sP5(sK64)), introduced(avatar_definition, [new_symbols(naming, [spl91_59])])).
fof(f36338, plain, (~ sP5(sK64) | sP0(sK64)), inference(duplicate_literal_removal, [], [f36337])).
fof(f36337, plain, (~ sP5(sK64) | sP0(sK64) | sP0(sK64) | ~ sP5(sK64)), inference(resolution, [], [f36336, f260])).
fof(f260, plain, ! [X0] : (r1(X0, sK48(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ! [X0] : ((((p2(sK46(X0)) & (~ p2(sK47(X0)) & r1(sK46(X0), sK47(X0))) & r1(X0, sK46(X0))) | p2(X0)) & ((! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK48(X0), X4)) & ~ p2(sK48(X0)) & r1(X0, sK48(X0))) | sP0(X0))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46, sK47, sK48])], [f102, f105, f104, f103])).
fof(f103, plain, ! [X0] : (? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) => (p2(sK46(X0)) & ? [X2] : (~ p2(X2) & r1(sK46(X0), X2)) & r1(X0, sK46(X0)))), introduced(choice_axiom, [])).
fof(f104, plain, ! [X0] : (? [X2] : (~ p2(X2) & r1(sK46(X0), X2)) => (~ p2(sK47(X0)) & r1(sK46(X0), sK47(X0)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X0] : (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) => (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(sK48(X0), X4)) & ~ p2(sK48(X0)) & r1(X0, sK48(X0)))), introduced(choice_axiom, [])).
fof(f102, plain, ! [X0] : (((? [X1] : (p2(X1) & ? [X2] : (~ p2(X2) & r1(X1, X2)) & r1(X0, X1)) | p2(X0)) & (? [X3] : (! [X4] : (~ p2(X4) | ! [X5] : (p2(X5) | ~ r1(X4, X5)) | ~ r1(X3, X4)) & ~ p2(X3) & r1(X0, X3)) | sP0(X0))) | ~ sP5(X0)), inference(rectify, [], [f101])).
fof(f101, plain, ! [X0] : (((? [X122] : (p2(X122) & ? [X123] : (~ p2(X123) & r1(X122, X123)) & r1(X0, X122)) | p2(X0)) & (? [X124] : (! [X125] : (~ p2(X125) | ! [X126] : (p2(X126) | ~ r1(X125, X126)) | ~ r1(X124, X125)) & ~ p2(X124) & r1(X0, X124)) | sP0(X0))) | ~ sP5(X0)), inference(nnf_transformation, [], [f13])).
fof(f36336, plain, ! [X0] : (~ r1(sK64, sK48(X0)) | ~ sP5(X0) | sP0(X0)), inference(subsumption_resolution, [], [f36335, f261])).
fof(f261, plain, ! [X0] : (~ p2(sK48(X0)) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f106])).
fof(f36335, plain, ! [X0] : (sP0(X0) | ~ sP5(X0) | ~ r1(sK64, sK48(X0)) | p2(sK48(X0))), inference(subsumption_resolution, [], [f36334, f308])).
fof(f36334, plain, ! [X0] : (p2(sK86(sK48(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK64, sK48(X0)) | p2(sK48(X0))), inference(duplicate_literal_removal, [], [f36316])).
fof(f36316, plain, ! [X0] : (p2(sK86(sK48(X0))) | sP0(X0) | ~ sP5(X0) | ~ r1(sK64, sK48(X0)) | p2(sK48(X0)) | ~ r1(sK64, sK48(X0))), inference(resolution, [], [f1931, f307])).
fof(f1931, plain, ! [X30, X31] : (~ r1(sK85(sK48(X31)), X30) | p2(X30) | sP0(X31) | ~ sP5(X31) | ~ r1(sK64, sK48(X31))), inference(subsumption_resolution, [], [f1930, f261])).
fof(f1930, plain, ! [X30, X31] : (p2(X30) | ~ r1(sK85(sK48(X31)), X30) | sP0(X31) | ~ sP5(X31) | p2(sK48(X31)) | ~ r1(sK64, sK48(X31))), inference(subsumption_resolution, [], [f1926, f309])).
fof(f309, plain, ! [X51] : (~ r1(sK64, X51) | p2(X51) | p2(sK85(X51))), inference(cnf_transformation, [], [f165])).
fof(f1926, plain, ! [X30, X31] : (p2(X30) | ~ r1(sK85(sK48(X31)), X30) | ~ p2(sK85(sK48(X31))) | sP0(X31) | ~ sP5(X31) | p2(sK48(X31)) | ~ r1(sK64, sK48(X31))), inference(resolution, [], [f262, f306])).
fof(f262, plain, ! [X4, X0, X5] : (~ r1(sK48(X0), X4) | p2(X5) | ~ r1(X4, X5) | ~ p2(X4) | sP0(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f106])).
fof(f156461, plain, (~ spl91_62 | ~ spl91_131), inference(avatar_contradiction_clause, [], [f156460])).
fof(f156460, plain, ($false | (~ spl91_62 | ~ spl91_131)), inference(subsumption_resolution, [], [f156459, f665])).
fof(f156459, plain, (~ sP1(sK81) | ~ spl91_131), inference(resolution, [], [f1081, f12387])).
fof(f12387, plain, ! [X0] : (~ sP4(sK60(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f12386])).
fof(f12386, plain, ! [X0] : (~ sP4(sK60(X0)) | ~ sP1(X0) | ~ sP4(sK60(X0)) | ~ sP1(X0)), inference(resolution, [], [f11528, f287])).
fof(f287, plain, ! [X0] : (r1(sK60(X0), sK61(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f11528, plain, ! [X0, X1] : (~ r1(X1, sK61(X0)) | ~ sP4(sK60(X0)) | ~ sP1(X0) | ~ sP4(X1)), inference(subsumption_resolution, [], [f11527, f288])).
fof(f288, plain, ! [X0] : (~ p2(sK61(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f11527, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP4(sK60(X0)) | p2(sK61(X0)) | ~ r1(X1, sK61(X0)) | ~ sP4(X1)), inference(resolution, [], [f11426, f272])).
fof(f272, plain, ! [X0, X1] : (~ p2(sK50(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP4(X0)), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ! [X0] : (! [X1] : ((((p2(sK49(X1)) & (~ p2(sK50(X1)) & r1(sK49(X1), sK50(X1))) & r1(X1, sK49(X1))) | p2(X1)) & ((! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK51(X1), X5)) & ~ p2(sK51(X1)) & r1(X1, sK51(X1))) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK49, sK50, sK51])], [f108, f111, f110, f109])).
fof(f109, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK49(X1)) & ? [X3] : (~ p2(X3) & r1(sK49(X1), X3)) & r1(X1, sK49(X1)))), introduced(choice_axiom, [])).
fof(f110, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK49(X1), X3)) => (~ p2(sK50(X1)) & r1(sK49(X1), sK50(X1)))), introduced(choice_axiom, [])).
fof(f111, plain, ! [X1] : (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) => (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(sK51(X1), X5)) & ~ p2(sK51(X1)) & r1(X1, sK51(X1)))), introduced(choice_axiom, [])).
fof(f108, plain, ! [X0] : (! [X1] : (((? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1)) & (? [X4] : (! [X5] : (~ p2(X5) | ! [X6] : (p2(X6) | ~ r1(X5, X6)) | ~ r1(X4, X5)) & ~ p2(X4) & r1(X1, X4)) | sP2(X1))) | ~ r1(X0, X1)) | ~ sP4(X0)), inference(rectify, [], [f107])).
fof(f107, plain, ! [X93] : (! [X103] : (((? [X104] : (p2(X104) & ? [X105] : (~ p2(X105) & r1(X104, X105)) & r1(X103, X104)) | p2(X103)) & (? [X106] : (! [X107] : (~ p2(X107) | ! [X108] : (p2(X108) | ~ r1(X107, X108)) | ~ r1(X106, X107)) & ~ p2(X106) & r1(X103, X106)) | sP2(X103))) | ~ r1(X93, X103)) | ~ sP4(X93)), inference(nnf_transformation, [], [f12])).
fof(f11426, plain, ! [X0] : (p2(sK50(sK61(X0))) | ~ sP1(X0) | ~ sP4(sK60(X0))), inference(duplicate_literal_removal, [], [f11408])).
fof(f11408, plain, ! [X0] : (~ sP1(X0) | p2(sK50(sK61(X0))) | ~ sP4(sK60(X0)) | ~ sP4(sK60(X0)) | ~ sP1(X0)), inference(resolution, [], [f5867, f1486])).
fof(f1486, plain, ! [X27] : (r1(sK49(sK61(X27)), sK50(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1466, f288])).
fof(f1466, plain, ! [X27] : (p2(sK61(X27)) | r1(sK49(sK61(X27)), sK50(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f271, f287])).
fof(f271, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK49(X1), sK50(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f112])).
fof(f5867, plain, ! [X0, X1] : (~ r1(sK49(sK61(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP4(sK60(X0))), inference(subsumption_resolution, [], [f5866, f1134])).
fof(f1134, plain, ! [X27] : (p2(sK49(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1114, f288])).
fof(f1114, plain, ! [X27] : (p2(sK61(X27)) | p2(sK49(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f273, f287])).
fof(f273, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK49(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f112])).
fof(f5866, plain, ! [X0, X1] : (~ sP4(sK60(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK49(sK61(X0)), X1) | ~ p2(sK49(sK61(X0)))), inference(duplicate_literal_removal, [], [f5821])).
fof(f5821, plain, ! [X0, X1] : (~ sP4(sK60(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK49(sK61(X0)), X1) | ~ p2(sK49(sK61(X0))) | ~ sP1(X0)), inference(resolution, [], [f1337, f289])).
fof(f289, plain, ! [X6, X0, X7] : (~ r1(sK61(X0), X6) | p2(X7) | ~ r1(X6, X7) | ~ p2(X6) | ~ sP1(X0)), inference(cnf_transformation, [], [f131])).
fof(f1337, plain, ! [X27] : (r1(sK61(X27), sK49(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1317, f288])).
fof(f1317, plain, ! [X27] : (p2(sK61(X27)) | r1(sK61(X27), sK49(sK61(X27))) | ~ sP4(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f270, f287])).
fof(f270, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK49(X1)) | ~ sP4(X0)), inference(cnf_transformation, [], [f112])).
fof(f1081, plain, (sP4(sK60(sK81)) | ~ spl91_131), inference(avatar_component_clause, [], [f1079])).
fof(f156458, plain, (~ spl91_62 | ~ spl91_130), inference(avatar_contradiction_clause, [], [f156457])).
fof(f156457, plain, ($false | (~ spl91_62 | ~ spl91_130)), inference(subsumption_resolution, [], [f156456, f665])).
fof(f156456, plain, (~ sP1(sK81) | ~ spl91_130), inference(resolution, [], [f156455, f1077])).
fof(f1077, plain, (sP3(sK60(sK81)) | ~ spl91_130), inference(avatar_component_clause, [], [f1075])).
fof(f156455, plain, ! [X0] : (~ sP3(sK60(X0)) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f156454])).
fof(f156454, plain, ! [X0] : (~ sP3(sK60(X0)) | ~ sP1(X0) | ~ sP3(sK60(X0)) | ~ sP1(X0)), inference(resolution, [], [f14984, f287])).
fof(f14984, plain, ! [X0, X1] : (~ r1(X1, sK61(X0)) | ~ sP3(sK60(X0)) | ~ sP1(X0) | ~ sP3(X1)), inference(subsumption_resolution, [], [f14983, f288])).
fof(f14983, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP3(sK60(X0)) | p2(sK61(X0)) | ~ r1(X1, sK61(X0)) | ~ sP3(X1)), inference(resolution, [], [f14840, f280])).
fof(f280, plain, ! [X0, X1] : (~ p2(sK53(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0] : ((! [X1] : ((p2(sK52(X1)) & (~ p2(sK53(X1)) & r1(sK52(X1), sK53(X1))) & r1(X1, sK52(X1))) | p2(X1) | ~ r1(X0, X1)) & ((! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK55(X0), X6)) & ~ p2(sK55(X0)) & r1(sK54(X0), sK55(X0))) & r1(X0, sK54(X0)))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52, sK53, sK54, sK55])], [f114, f118, f117, f116, f115])).
fof(f115, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK52(X1)) & ? [X3] : (~ p2(X3) & r1(sK52(X1), X3)) & r1(X1, sK52(X1)))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK52(X1), X3)) => (~ p2(sK53(X1)) & r1(sK52(X1), sK53(X1)))), introduced(choice_axiom, [])).
fof(f117, plain, ! [X0] : (? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4)) => (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK54(X0), X5)) & r1(X0, sK54(X0)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X0] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(sK54(X0), X5)) => (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(sK55(X0), X6)) & ~ p2(sK55(X0)) & r1(sK54(X0), sK55(X0)))), introduced(choice_axiom, [])).
fof(f114, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (? [X5] : (! [X6] : (~ p2(X6) | ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) & ~ p2(X5) & r1(X4, X5)) & r1(X0, X4))) | ~ sP3(X0)), inference(rectify, [], [f113])).
fof(f113, plain, ! [X93] : ((! [X96] : (? [X97] : (p2(X97) & ? [X98] : (~ p2(X98) & r1(X97, X98)) & r1(X96, X97)) | p2(X96) | ~ r1(X93, X96)) & ? [X99] : (? [X100] : (! [X101] : (~ p2(X101) | ! [X102] : (p2(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) & ~ p2(X100) & r1(X99, X100)) & r1(X93, X99))) | ~ sP3(X93)), inference(nnf_transformation, [], [f11])).
fof(f14840, plain, ! [X0] : (p2(sK53(sK61(X0))) | ~ sP1(X0) | ~ sP3(sK60(X0))), inference(duplicate_literal_removal, [], [f14822])).
fof(f14822, plain, ! [X0] : (~ sP1(X0) | p2(sK53(sK61(X0))) | ~ sP3(sK60(X0)) | ~ sP3(sK60(X0)) | ~ sP1(X0)), inference(resolution, [], [f6089, f1536])).
fof(f1536, plain, ! [X27] : (r1(sK52(sK61(X27)), sK53(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1515, f288])).
fof(f1515, plain, ! [X27] : (p2(sK61(X27)) | r1(sK52(sK61(X27)), sK53(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f279, f287])).
fof(f279, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK52(X1), sK53(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f119])).
fof(f6089, plain, ! [X0, X1] : (~ r1(sK52(sK61(X0)), X1) | ~ sP1(X0) | p2(X1) | ~ sP3(sK60(X0))), inference(subsumption_resolution, [], [f6088, f1204])).
fof(f1204, plain, ! [X27] : (p2(sK52(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1183, f288])).
fof(f1183, plain, ! [X27] : (p2(sK61(X27)) | p2(sK52(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f281, f287])).
fof(f281, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK52(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f119])).
fof(f6088, plain, ! [X0, X1] : (~ sP3(sK60(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK52(sK61(X0)), X1) | ~ p2(sK52(sK61(X0)))), inference(duplicate_literal_removal, [], [f6043])).
fof(f6043, plain, ! [X0, X1] : (~ sP3(sK60(X0)) | ~ sP1(X0) | p2(X1) | ~ r1(sK52(sK61(X0)), X1) | ~ p2(sK52(sK61(X0))) | ~ sP1(X0)), inference(resolution, [], [f1387, f289])).
fof(f1387, plain, ! [X27] : (r1(sK61(X27), sK52(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(subsumption_resolution, [], [f1366, f288])).
fof(f1366, plain, ! [X27] : (p2(sK61(X27)) | r1(sK61(X27), sK52(sK61(X27))) | ~ sP3(sK60(X27)) | ~ sP1(X27)), inference(resolution, [], [f278, f287])).
fof(f278, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK52(X1)) | ~ sP3(X0)), inference(cnf_transformation, [], [f119])).
fof(f42972, plain, (~ spl91_104 | ~ spl91_266), inference(avatar_split_clause, [], [f42969, f24248, f914])).
fof(f42969, plain, (~ sP0(sK64) | ~ spl91_266), inference(resolution, [], [f24249, f304])).
fof(f24249, plain, (! [X1] : (~ r1(X1, sK87) | ~ sP0(X1)) | ~ spl91_266), inference(avatar_component_clause, [], [f24248])).
fof(f42968, plain, (spl91_286 | spl91_266), inference(avatar_split_clause, [], [f42963, f24248, f42965])).
fof(f42963, plain, ! [X1] : (~ r1(X1, sK87) | ~ sP0(X1) | r1(sK65(sK87), sK62(sK65(sK87)))), inference(subsumption_resolution, [], [f42944, f305])).
fof(f42944, plain, ! [X1] : (~ r1(X1, sK87) | ~ sP0(X1) | p2(sK87) | r1(sK65(sK87), sK62(sK65(sK87)))), inference(resolution, [], [f1906, f304])).
fof(f1906, plain, ! [X26, X27] : (~ r1(sK64, X26) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | r1(sK65(X26), sK62(sK65(X26)))), inference(subsumption_resolution, [], [f1875, f373])).
fof(f1875, plain, ! [X26, X27] : (p2(sK65(X26)) | r1(sK65(X26), sK62(sK65(X26))) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | ~ r1(sK64, X26)), inference(resolution, [], [f294, f372])).
fof(f294, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | r1(X2, sK62(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f136])).
fof(f24250, plain, (spl91_265 | spl91_266), inference(avatar_split_clause, [], [f24242, f24248, f24244])).
fof(f24242, plain, ! [X1] : (~ r1(X1, sK87) | ~ sP0(X1) | p2(sK62(sK65(sK87)))), inference(subsumption_resolution, [], [f24223, f305])).
fof(f24223, plain, ! [X1] : (~ r1(X1, sK87) | ~ sP0(X1) | p2(sK87) | p2(sK62(sK65(sK87)))), inference(resolution, [], [f1742, f304])).
fof(f1742, plain, ! [X26, X27] : (~ r1(sK64, X26) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | p2(sK62(sK65(X26)))), inference(subsumption_resolution, [], [f1711, f373])).
fof(f1711, plain, ! [X26, X27] : (p2(sK65(X26)) | p2(sK62(sK65(X26))) | ~ r1(X27, X26) | ~ sP0(X27) | p2(X26) | ~ r1(sK64, X26)), inference(resolution, [], [f297, f372])).
fof(f297, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p2(X2) | p2(sK62(X2)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f136])).
fof(f2221, plain, (~ spl91_62 | spl91_130 | spl91_131 | spl91_188 | ~ spl91_60), inference(avatar_split_clause, [], [f2063, f655, f2218, f1079, f1075, f663])).
fof(f655, plain, (spl91_60 <=> ! [X42, X44, X43] : (~ p2(X43) | ~ r1(sK81, X42) | sP4(X42) | sP3(X42) | ~ r1(X42, X43) | ~ r1(X43, X44) | p2(X44))), introduced(avatar_definition, [new_symbols(naming, [spl91_60])])).
fof(f2063, plain, (! [X24, X25] : (~ p2(X24) | sP4(sK60(sK81)) | sP3(sK60(sK81)) | ~ r1(sK60(sK81), X24) | ~ r1(X24, X25) | p2(X25) | ~ sP1(sK81)) | ~ spl91_60), inference(resolution, [], [f656, f286])).
fof(f656, plain, (! [X43, X44, X42] : (~ r1(sK81, X42) | ~ p2(X43) | sP4(X42) | sP3(X42) | ~ r1(X42, X43) | ~ r1(X43, X44) | p2(X44)) | ~ spl91_60), inference(avatar_component_clause, [], [f655])).
fof(f958, plain, (spl91_111 | spl91_64 | ~ spl91_65), inference(avatar_split_clause, [], [f943, f676, f671, f955])).
fof(f943, plain, (p2(sK81) | p2(sK85(sK81)) | ~ spl91_65), inference(resolution, [], [f678, f309])).
fof(f679, plain, (spl91_59 | spl91_65), inference(avatar_split_clause, [], [f316, f676, f651])).
fof(f316, plain, (r1(sK64, sK81) | sP5(sK64)), inference(cnf_transformation, [], [f165])).
fof(f674, plain, (spl91_59 | spl91_62 | ~ spl91_64), inference(avatar_split_clause, [], [f317, f671, f663, f651])).
fof(f317, plain, (~ p2(sK81) | sP1(sK81) | sP5(sK64)), inference(cnf_transformation, [], [f165])).
fof(f669, plain, (spl91_59 | spl91_62 | spl91_63), inference(avatar_split_clause, [], [f318, f667, f663, f651])).
fof(f318, plain, ! [X45, X46] : (~ p2(X45) | p2(X46) | ~ r1(X45, X46) | ~ r1(sK81, X45) | sP1(sK81) | sP5(sK64)), inference(cnf_transformation, [], [f165])).
fof(f661, plain, (spl91_59 | spl91_61), inference(avatar_split_clause, [], [f319, f659, f651])).
fof(f319, plain, ! [X42] : (~ p2(X42) | sP3(X42) | sP4(X42) | ~ r1(sK81, X42) | sP5(sK64)), inference(cnf_transformation, [], [f165])).
fof(f657, plain, (spl91_59 | spl91_60), inference(avatar_split_clause, [], [f320, f655, f651])).
fof(f320, plain, ! [X43, X44, X42] : (~ p2(X43) | p2(X44) | ~ r1(X43, X44) | ~ r1(X42, X43) | sP3(X42) | sP4(X42) | ~ r1(sK81, X42) | sP5(sK64)), inference(cnf_transformation, [], [f165])).