fof(f23473, plain, $false, inference(avatar_sat_refutation, [], [f4859, f6058, f23472])).
fof(f23472, plain, (spl83_59 | ~ spl83_60), inference(avatar_contradiction_clause, [], [f23471])).
fof(f23471, plain, ($false | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23470, f202])).
fof(f202, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.020.p', reflexivity)).
fof(f23470, plain, (~ r1(sK82, sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23469, f467])).
fof(f467, plain, ~ p101(sK82), inference(cnf_transformation, [], [f201])).
fof(f201, plain, (p100(sK82) & ~ p101(sK82) & ! [X1] : (sP41(X1) | ~ r1(sK82, X1)) & ! [X2] : (p8(X2) | ~ r1(sK82, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK82])], [f199, f200])).
fof(f200, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP41(X1) | ~ r1(X0, X1)) & ! [X2] : (p8(X2) | ~ r1(X0, X2))) => (p100(sK82) & ~ p101(sK82) & ! [X1] : (sP41(X1) | ~ r1(sK82, X1)) & ! [X2] : (p8(X2) | ~ r1(sK82, X2)))), introduced(choice_axiom, [])).
fof(f199, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP41(X1) | ~ r1(X0, X1)) & ! [X2] : (p8(X2) | ~ r1(X0, X2))), inference(rectify, [], [f54])).
fof(f54, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP41(X1) | ~ r1(X0, X1)) & ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(definition_folding, [], [f11, e53, e52, e51, e50, e49, e48, e47, e46, e45, e44, e43, e42, e41, e40, e39, e38, e37, e36, e35, e34, e33, e32, e31, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12])).
fof(f12, plain, ! [X1] : (~ p118(X1) | p119(X1) | (? [X80] : (p119(X80) & ~ p120(X80) & p20(X80) & r1(X1, X80)) & ? [X81] : (p119(X81) & ~ p120(X81) & ~ p20(X81) & r1(X1, X81))) | ~ sP0(X1)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X1] : (sP0(X1) <=> (~ p118(X1) | p119(X1) | (? [X80] : (p119(X80) & ~ p120(X80) & p20(X80) & r1(X1, X80)) & ? [X81] : (p119(X81) & ~ p120(X81) & ~ p20(X81) & r1(X1, X81))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f13, plain, ! [X1] : (~ p117(X1) | p118(X1) | (? [X78] : (p118(X78) & ~ p119(X78) & p19(X78) & r1(X1, X78)) & ? [X79] : (p118(X79) & ~ p119(X79) & ~ p19(X79) & r1(X1, X79))) | ~ sP1(X1)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X1] : (sP1(X1) <=> (~ p117(X1) | p118(X1) | (? [X78] : (p118(X78) & ~ p119(X78) & p19(X78) & r1(X1, X78)) & ? [X79] : (p118(X79) & ~ p119(X79) & ~ p19(X79) & r1(X1, X79))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f14, plain, ! [X1] : (~ p116(X1) | p117(X1) | (? [X76] : (p117(X76) & ~ p118(X76) & p18(X76) & r1(X1, X76)) & ? [X77] : (p117(X77) & ~ p118(X77) & ~ p18(X77) & r1(X1, X77))) | ~ sP2(X1)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X1] : (sP2(X1) <=> (~ p116(X1) | p117(X1) | (? [X76] : (p117(X76) & ~ p118(X76) & p18(X76) & r1(X1, X76)) & ? [X77] : (p117(X77) & ~ p118(X77) & ~ p18(X77) & r1(X1, X77))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f15, plain, ! [X1] : (~ p115(X1) | p116(X1) | (? [X74] : (p116(X74) & ~ p117(X74) & p17(X74) & r1(X1, X74)) & ? [X75] : (p116(X75) & ~ p117(X75) & ~ p17(X75) & r1(X1, X75))) | ~ sP3(X1)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X1] : (sP3(X1) <=> (~ p115(X1) | p116(X1) | (? [X74] : (p116(X74) & ~ p117(X74) & p17(X74) & r1(X1, X74)) & ? [X75] : (p116(X75) & ~ p117(X75) & ~ p17(X75) & r1(X1, X75))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X1] : (~ p114(X1) | p115(X1) | (? [X72] : (p115(X72) & ~ p116(X72) & p16(X72) & r1(X1, X72)) & ? [X73] : (p115(X73) & ~ p116(X73) & ~ p16(X73) & r1(X1, X73))) | ~ sP4(X1)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X1] : (sP4(X1) <=> (~ p114(X1) | p115(X1) | (? [X72] : (p115(X72) & ~ p116(X72) & p16(X72) & r1(X1, X72)) & ? [X73] : (p115(X73) & ~ p116(X73) & ~ p16(X73) & r1(X1, X73))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X1] : (~ p113(X1) | p114(X1) | (? [X70] : (p114(X70) & ~ p115(X70) & p15(X70) & r1(X1, X70)) & ? [X71] : (p114(X71) & ~ p115(X71) & ~ p15(X71) & r1(X1, X71))) | ~ sP5(X1)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X1] : (sP5(X1) <=> (~ p113(X1) | p114(X1) | (? [X70] : (p114(X70) & ~ p115(X70) & p15(X70) & r1(X1, X70)) & ? [X71] : (p114(X71) & ~ p115(X71) & ~ p15(X71) & r1(X1, X71))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f18, plain, ! [X1] : (~ p112(X1) | p113(X1) | (? [X68] : (p113(X68) & ~ p114(X68) & p14(X68) & r1(X1, X68)) & ? [X69] : (p113(X69) & ~ p114(X69) & ~ p14(X69) & r1(X1, X69))) | ~ sP6(X1)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X1] : (sP6(X1) <=> (~ p112(X1) | p113(X1) | (? [X68] : (p113(X68) & ~ p114(X68) & p14(X68) & r1(X1, X68)) & ? [X69] : (p113(X69) & ~ p114(X69) & ~ p14(X69) & r1(X1, X69))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f19, plain, ! [X1] : (~ p111(X1) | p112(X1) | (? [X66] : (p112(X66) & ~ p113(X66) & p13(X66) & r1(X1, X66)) & ? [X67] : (p112(X67) & ~ p113(X67) & ~ p13(X67) & r1(X1, X67))) | ~ sP7(X1)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X1] : (sP7(X1) <=> (~ p111(X1) | p112(X1) | (? [X66] : (p112(X66) & ~ p113(X66) & p13(X66) & r1(X1, X66)) & ? [X67] : (p112(X67) & ~ p113(X67) & ~ p13(X67) & r1(X1, X67))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f20, plain, ! [X1] : (~ p110(X1) | p111(X1) | (? [X64] : (p111(X64) & ~ p112(X64) & p12(X64) & r1(X1, X64)) & ? [X65] : (p111(X65) & ~ p112(X65) & ~ p12(X65) & r1(X1, X65))) | ~ sP8(X1)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X1] : (sP8(X1) <=> (~ p110(X1) | p111(X1) | (? [X64] : (p111(X64) & ~ p112(X64) & p12(X64) & r1(X1, X64)) & ? [X65] : (p111(X65) & ~ p112(X65) & ~ p12(X65) & r1(X1, X65))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f21, plain, ! [X1] : (~ p109(X1) | p110(X1) | (? [X62] : (p110(X62) & ~ p111(X62) & p11(X62) & r1(X1, X62)) & ? [X63] : (p110(X63) & ~ p111(X63) & ~ p11(X63) & r1(X1, X63))) | ~ sP9(X1)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X1] : (sP9(X1) <=> (~ p109(X1) | p110(X1) | (? [X62] : (p110(X62) & ~ p111(X62) & p11(X62) & r1(X1, X62)) & ? [X63] : (p110(X63) & ~ p111(X63) & ~ p11(X63) & r1(X1, X63))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f22, plain, ! [X1] : (~ p108(X1) | p109(X1) | (? [X60] : (p109(X60) & ~ p110(X60) & p10(X60) & r1(X1, X60)) & ? [X61] : (p109(X61) & ~ p110(X61) & ~ p10(X61) & r1(X1, X61))) | ~ sP10(X1)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X1] : (sP10(X1) <=> (~ p108(X1) | p109(X1) | (? [X60] : (p109(X60) & ~ p110(X60) & p10(X60) & r1(X1, X60)) & ? [X61] : (p109(X61) & ~ p110(X61) & ~ p10(X61) & r1(X1, X61))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f23, plain, ! [X1] : (~ p107(X1) | p108(X1) | (? [X58] : (p108(X58) & ~ p109(X58) & p9(X58) & r1(X1, X58)) & ? [X59] : (p108(X59) & ~ p109(X59) & ~ p9(X59) & r1(X1, X59))) | ~ sP11(X1)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X1] : (sP11(X1) <=> (~ p107(X1) | p108(X1) | (? [X58] : (p108(X58) & ~ p109(X58) & p9(X58) & r1(X1, X58)) & ? [X59] : (p108(X59) & ~ p109(X59) & ~ p9(X59) & r1(X1, X59))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f24, plain, ! [X1] : (~ p106(X1) | p107(X1) | (? [X56] : (p107(X56) & ~ p108(X56) & p8(X56) & r1(X1, X56)) & ? [X57] : (p107(X57) & ~ p108(X57) & ~ p8(X57) & r1(X1, X57))) | ~ sP12(X1)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X1] : (sP12(X1) <=> (~ p106(X1) | p107(X1) | (? [X56] : (p107(X56) & ~ p108(X56) & p8(X56) & r1(X1, X56)) & ? [X57] : (p107(X57) & ~ p108(X57) & ~ p8(X57) & r1(X1, X57))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f25, plain, ! [X1] : (~ p105(X1) | p106(X1) | (? [X54] : (p106(X54) & ~ p107(X54) & p7(X54) & r1(X1, X54)) & ? [X55] : (p106(X55) & ~ p107(X55) & ~ p7(X55) & r1(X1, X55))) | ~ sP13(X1)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X1] : (sP13(X1) <=> (~ p105(X1) | p106(X1) | (? [X54] : (p106(X54) & ~ p107(X54) & p7(X54) & r1(X1, X54)) & ? [X55] : (p106(X55) & ~ p107(X55) & ~ p7(X55) & r1(X1, X55))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f26, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X52] : (p105(X52) & ~ p106(X52) & p6(X52) & r1(X1, X52)) & ? [X53] : (p105(X53) & ~ p106(X53) & ~ p6(X53) & r1(X1, X53))) | ~ sP14(X1)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X1] : (sP14(X1) <=> (~ p104(X1) | p105(X1) | (? [X52] : (p105(X52) & ~ p106(X52) & p6(X52) & r1(X1, X52)) & ? [X53] : (p105(X53) & ~ p106(X53) & ~ p6(X53) & r1(X1, X53))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f27, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X50] : (p104(X50) & ~ p105(X50) & p5(X50) & r1(X1, X50)) & ? [X51] : (p104(X51) & ~ p105(X51) & ~ p5(X51) & r1(X1, X51))) | ~ sP15(X1)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X1] : (sP15(X1) <=> (~ p103(X1) | p104(X1) | (? [X50] : (p104(X50) & ~ p105(X50) & p5(X50) & r1(X1, X50)) & ? [X51] : (p104(X51) & ~ p105(X51) & ~ p5(X51) & r1(X1, X51))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f28, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X48] : (p103(X48) & ~ p104(X48) & p4(X48) & r1(X1, X48)) & ? [X49] : (p103(X49) & ~ p104(X49) & ~ p4(X49) & r1(X1, X49))) | ~ sP16(X1)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X1] : (sP16(X1) <=> (~ p102(X1) | p103(X1) | (? [X48] : (p103(X48) & ~ p104(X48) & p4(X48) & r1(X1, X48)) & ? [X49] : (p103(X49) & ~ p104(X49) & ~ p4(X49) & r1(X1, X49))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f29, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X46] : (p102(X46) & ~ p103(X46) & p3(X46) & r1(X1, X46)) & ? [X47] : (p102(X47) & ~ p103(X47) & ~ p3(X47) & r1(X1, X47))) | ~ sP17(X1)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X1] : (sP17(X1) <=> (~ p101(X1) | p102(X1) | (? [X46] : (p102(X46) & ~ p103(X46) & p3(X46) & r1(X1, X46)) & ? [X47] : (p102(X47) & ~ p103(X47) & ~ p3(X47) & r1(X1, X47))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f30, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X44] : (p101(X44) & ~ p102(X44) & p2(X44) & r1(X1, X44)) & ? [X45] : (p101(X45) & ~ p102(X45) & ~ p2(X45) & r1(X1, X45))) | ~ sP18(X1)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X1] : (sP18(X1) <=> (~ p100(X1) | p101(X1) | (? [X44] : (p101(X44) & ~ p102(X44) & p2(X44) & r1(X1, X44)) & ? [X45] : (p101(X45) & ~ p102(X45) & ~ p2(X45) & r1(X1, X45))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f31, plain, ! [X1] : (~ p119(X1) | p120(X1) | (? [X82] : (p120(X82) & p21(X82) & r1(X1, X82)) & ? [X83] : (p120(X83) & ~ p21(X83) & r1(X1, X83))) | ~ sP19(X1)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X1] : (sP19(X1) <=> (~ p119(X1) | p120(X1) | (? [X82] : (p120(X82) & p21(X82) & r1(X1, X82)) & ? [X83] : (p120(X83) & ~ p21(X83) & r1(X1, X83))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f32, plain, ! [X1] : (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43)))) | ~ sP20(X1)), inference(usedef, [], [e32])).
fof(e32, plain, ! [X1] : (sP20(X1) <=> (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f33, plain, ! [X1] : (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41)))) | ~ sP21(X1)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X1] : (sP21(X1) <=> (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f34, plain, ! [X1] : (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39)))) | ~ sP22(X1)), inference(usedef, [], [e34])).
fof(e34, plain, ! [X1] : (sP22(X1) <=> (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f35, plain, ! [X1] : (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37)))) | ~ sP23(X1)), inference(usedef, [], [e35])).
fof(e35, plain, ! [X1] : (sP23(X1) <=> (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f36, plain, ! [X1] : (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35)))) | ~ sP24(X1)), inference(usedef, [], [e36])).
fof(e36, plain, ! [X1] : (sP24(X1) <=> (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f37, plain, ! [X1] : (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33)))) | ~ sP25(X1)), inference(usedef, [], [e37])).
fof(e37, plain, ! [X1] : (sP25(X1) <=> (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f38, plain, ! [X1] : (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31)))) | ~ sP26(X1)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X1] : (sP26(X1) <=> (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f39, plain, ! [X1] : (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29)))) | ~ sP27(X1)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X1] : (sP27(X1) <=> (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f40, plain, ! [X1] : (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27)))) | ~ sP28(X1)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X1] : (sP28(X1) <=> (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f41, plain, ! [X1] : (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25)))) | ~ sP29(X1)), inference(usedef, [], [e41])).
fof(e41, plain, ! [X1] : (sP29(X1) <=> (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f42, plain, ! [X1] : (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))) | ~ sP30(X1)), inference(usedef, [], [e42])).
fof(e42, plain, ! [X1] : (sP30(X1) <=> (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f43, plain, ! [X1] : (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))) | ~ sP31(X1)), inference(usedef, [], [e43])).
fof(e43, plain, ! [X1] : (sP31(X1) <=> (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f44, plain, ! [X1] : (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))) | ~ sP32(X1)), inference(usedef, [], [e44])).
fof(e44, plain, ! [X1] : (sP32(X1) <=> (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f45, plain, ! [X1] : (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))) | ~ sP33(X1)), inference(usedef, [], [e45])).
fof(e45, plain, ! [X1] : (sP33(X1) <=> (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f46, plain, ! [X1] : (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))) | ~ sP34(X1)), inference(usedef, [], [e46])).
fof(e46, plain, ! [X1] : (sP34(X1) <=> (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f47, plain, ! [X1] : (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))) | ~ sP35(X1)), inference(usedef, [], [e47])).
fof(e47, plain, ! [X1] : (sP35(X1) <=> (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f48, plain, ! [X1] : (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))) | ~ sP36(X1)), inference(usedef, [], [e48])).
fof(e48, plain, ! [X1] : (sP36(X1) <=> (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f49, plain, ! [X1] : (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))) | ~ sP37(X1)), inference(usedef, [], [e49])).
fof(e49, plain, ! [X1] : (sP37(X1) <=> (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f50, plain, ! [X1] : (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))) | ~ sP38(X1)), inference(usedef, [], [e50])).
fof(e50, plain, ! [X1] : (sP38(X1) <=> (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f51, plain, ! [X1] : (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))) | ~ sP39(X1)), inference(usedef, [], [e51])).
fof(e51, plain, ! [X1] : (sP39(X1) <=> (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f52, plain, ! [X1] : (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))) | ~ sP40(X1)), inference(usedef, [], [e52])).
fof(e52, plain, ! [X1] : (sP40(X1) <=> (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f53, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & sP40(X1) & sP39(X1) & sP38(X1) & sP37(X1) & sP36(X1) & sP35(X1) & sP34(X1) & sP33(X1) & sP32(X1) & sP31(X1) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP19(X1)) | ~ sP41(X1)), inference(usedef, [], [e53])).
fof(e53, plain, ! [X1] : (sP41(X1) <=> ((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & sP40(X1) & sP39(X1) & sP38(X1) & sP37(X1) & sP36(X1) & sP35(X1) & sP34(X1) & sP33(X1) & sP32(X1) & sP31(X1) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP19(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f11, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35))))) & (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37))))) & (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39))))) & (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41))))) & (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43))))) & (~ p100(X1) | p101(X1) | (? [X44] : (p101(X44) & ~ p102(X44) & p2(X44) & r1(X1, X44)) & ? [X45] : (p101(X45) & ~ p102(X45) & ~ p2(X45) & r1(X1, X45)))) & (~ p101(X1) | p102(X1) | (? [X46] : (p102(X46) & ~ p103(X46) & p3(X46) & r1(X1, X46)) & ? [X47] : (p102(X47) & ~ p103(X47) & ~ p3(X47) & r1(X1, X47)))) & (~ p102(X1) | p103(X1) | (? [X48] : (p103(X48) & ~ p104(X48) & p4(X48) & r1(X1, X48)) & ? [X49] : (p103(X49) & ~ p104(X49) & ~ p4(X49) & r1(X1, X49)))) & (~ p103(X1) | p104(X1) | (? [X50] : (p104(X50) & ~ p105(X50) & p5(X50) & r1(X1, X50)) & ? [X51] : (p104(X51) & ~ p105(X51) & ~ p5(X51) & r1(X1, X51)))) & (~ p104(X1) | p105(X1) | (? [X52] : (p105(X52) & ~ p106(X52) & p6(X52) & r1(X1, X52)) & ? [X53] : (p105(X53) & ~ p106(X53) & ~ p6(X53) & r1(X1, X53)))) & (~ p105(X1) | p106(X1) | (? [X54] : (p106(X54) & ~ p107(X54) & p7(X54) & r1(X1, X54)) & ? [X55] : (p106(X55) & ~ p107(X55) & ~ p7(X55) & r1(X1, X55)))) & (~ p106(X1) | p107(X1) | (? [X56] : (p107(X56) & ~ p108(X56) & p8(X56) & r1(X1, X56)) & ? [X57] : (p107(X57) & ~ p108(X57) & ~ p8(X57) & r1(X1, X57)))) & (~ p107(X1) | p108(X1) | (? [X58] : (p108(X58) & ~ p109(X58) & p9(X58) & r1(X1, X58)) & ? [X59] : (p108(X59) & ~ p109(X59) & ~ p9(X59) & r1(X1, X59)))) & (~ p108(X1) | p109(X1) | (? [X60] : (p109(X60) & ~ p110(X60) & p10(X60) & r1(X1, X60)) & ? [X61] : (p109(X61) & ~ p110(X61) & ~ p10(X61) & r1(X1, X61)))) & (~ p109(X1) | p110(X1) | (? [X62] : (p110(X62) & ~ p111(X62) & p11(X62) & r1(X1, X62)) & ? [X63] : (p110(X63) & ~ p111(X63) & ~ p11(X63) & r1(X1, X63)))) & (~ p110(X1) | p111(X1) | (? [X64] : (p111(X64) & ~ p112(X64) & p12(X64) & r1(X1, X64)) & ? [X65] : (p111(X65) & ~ p112(X65) & ~ p12(X65) & r1(X1, X65)))) & (~ p111(X1) | p112(X1) | (? [X66] : (p112(X66) & ~ p113(X66) & p13(X66) & r1(X1, X66)) & ? [X67] : (p112(X67) & ~ p113(X67) & ~ p13(X67) & r1(X1, X67)))) & (~ p112(X1) | p113(X1) | (? [X68] : (p113(X68) & ~ p114(X68) & p14(X68) & r1(X1, X68)) & ? [X69] : (p113(X69) & ~ p114(X69) & ~ p14(X69) & r1(X1, X69)))) & (~ p113(X1) | p114(X1) | (? [X70] : (p114(X70) & ~ p115(X70) & p15(X70) & r1(X1, X70)) & ? [X71] : (p114(X71) & ~ p115(X71) & ~ p15(X71) & r1(X1, X71)))) & (~ p114(X1) | p115(X1) | (? [X72] : (p115(X72) & ~ p116(X72) & p16(X72) & r1(X1, X72)) & ? [X73] : (p115(X73) & ~ p116(X73) & ~ p16(X73) & r1(X1, X73)))) & (~ p115(X1) | p116(X1) | (? [X74] : (p116(X74) & ~ p117(X74) & p17(X74) & r1(X1, X74)) & ? [X75] : (p116(X75) & ~ p117(X75) & ~ p17(X75) & r1(X1, X75)))) & (~ p116(X1) | p117(X1) | (? [X76] : (p117(X76) & ~ p118(X76) & p18(X76) & r1(X1, X76)) & ? [X77] : (p117(X77) & ~ p118(X77) & ~ p18(X77) & r1(X1, X77)))) & (~ p117(X1) | p118(X1) | (? [X78] : (p118(X78) & ~ p119(X78) & p19(X78) & r1(X1, X78)) & ? [X79] : (p118(X79) & ~ p119(X79) & ~ p19(X79) & r1(X1, X79)))) & (~ p118(X1) | p119(X1) | (? [X80] : (p119(X80) & ~ p120(X80) & p20(X80) & r1(X1, X80)) & ? [X81] : (p119(X81) & ~ p120(X81) & ~ p20(X81) & r1(X1, X81)))) & (~ p119(X1) | p120(X1) | (? [X82] : (p120(X82) & p21(X82) & r1(X1, X82)) & ? [X83] : (p120(X83) & ~ p21(X83) & r1(X1, X83))))) | ~ r1(X0, X1)) & ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35))))) & (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37))))) & (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39))))) & (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41))))) & (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43))))) & ((~ p100(X1) | p101(X1)) | (? [X44] : ((p101(X44) & ~ p102(X44) & p2(X44)) & r1(X1, X44)) & ? [X45] : ((p101(X45) & ~ p102(X45) & ~ p2(X45)) & r1(X1, X45)))) & ((~ p101(X1) | p102(X1)) | (? [X46] : ((p102(X46) & ~ p103(X46) & p3(X46)) & r1(X1, X46)) & ? [X47] : ((p102(X47) & ~ p103(X47) & ~ p3(X47)) & r1(X1, X47)))) & ((~ p102(X1) | p103(X1)) | (? [X48] : ((p103(X48) & ~ p104(X48) & p4(X48)) & r1(X1, X48)) & ? [X49] : ((p103(X49) & ~ p104(X49) & ~ p4(X49)) & r1(X1, X49)))) & ((~ p103(X1) | p104(X1)) | (? [X50] : ((p104(X50) & ~ p105(X50) & p5(X50)) & r1(X1, X50)) & ? [X51] : ((p104(X51) & ~ p105(X51) & ~ p5(X51)) & r1(X1, X51)))) & ((~ p104(X1) | p105(X1)) | (? [X52] : ((p105(X52) & ~ p106(X52) & p6(X52)) & r1(X1, X52)) & ? [X53] : ((p105(X53) & ~ p106(X53) & ~ p6(X53)) & r1(X1, X53)))) & ((~ p105(X1) | p106(X1)) | (? [X54] : ((p106(X54) & ~ p107(X54) & p7(X54)) & r1(X1, X54)) & ? [X55] : ((p106(X55) & ~ p107(X55) & ~ p7(X55)) & r1(X1, X55)))) & ((~ p106(X1) | p107(X1)) | (? [X56] : ((p107(X56) & ~ p108(X56) & p8(X56)) & r1(X1, X56)) & ? [X57] : ((p107(X57) & ~ p108(X57) & ~ p8(X57)) & r1(X1, X57)))) & ((~ p107(X1) | p108(X1)) | (? [X58] : ((p108(X58) & ~ p109(X58) & p9(X58)) & r1(X1, X58)) & ? [X59] : ((p108(X59) & ~ p109(X59) & ~ p9(X59)) & r1(X1, X59)))) & ((~ p108(X1) | p109(X1)) | (? [X60] : ((p109(X60) & ~ p110(X60) & p10(X60)) & r1(X1, X60)) & ? [X61] : ((p109(X61) & ~ p110(X61) & ~ p10(X61)) & r1(X1, X61)))) & ((~ p109(X1) | p110(X1)) | (? [X62] : ((p110(X62) & ~ p111(X62) & p11(X62)) & r1(X1, X62)) & ? [X63] : ((p110(X63) & ~ p111(X63) & ~ p11(X63)) & r1(X1, X63)))) & ((~ p110(X1) | p111(X1)) | (? [X64] : ((p111(X64) & ~ p112(X64) & p12(X64)) & r1(X1, X64)) & ? [X65] : ((p111(X65) & ~ p112(X65) & ~ p12(X65)) & r1(X1, X65)))) & ((~ p111(X1) | p112(X1)) | (? [X66] : ((p112(X66) & ~ p113(X66) & p13(X66)) & r1(X1, X66)) & ? [X67] : ((p112(X67) & ~ p113(X67) & ~ p13(X67)) & r1(X1, X67)))) & ((~ p112(X1) | p113(X1)) | (? [X68] : ((p113(X68) & ~ p114(X68) & p14(X68)) & r1(X1, X68)) & ? [X69] : ((p113(X69) & ~ p114(X69) & ~ p14(X69)) & r1(X1, X69)))) & ((~ p113(X1) | p114(X1)) | (? [X70] : ((p114(X70) & ~ p115(X70) & p15(X70)) & r1(X1, X70)) & ? [X71] : ((p114(X71) & ~ p115(X71) & ~ p15(X71)) & r1(X1, X71)))) & ((~ p114(X1) | p115(X1)) | (? [X72] : ((p115(X72) & ~ p116(X72) & p16(X72)) & r1(X1, X72)) & ? [X73] : ((p115(X73) & ~ p116(X73) & ~ p16(X73)) & r1(X1, X73)))) & ((~ p115(X1) | p116(X1)) | (? [X74] : ((p116(X74) & ~ p117(X74) & p17(X74)) & r1(X1, X74)) & ? [X75] : ((p116(X75) & ~ p117(X75) & ~ p17(X75)) & r1(X1, X75)))) & ((~ p116(X1) | p117(X1)) | (? [X76] : ((p117(X76) & ~ p118(X76) & p18(X76)) & r1(X1, X76)) & ? [X77] : ((p117(X77) & ~ p118(X77) & ~ p18(X77)) & r1(X1, X77)))) & ((~ p117(X1) | p118(X1)) | (? [X78] : ((p118(X78) & ~ p119(X78) & p19(X78)) & r1(X1, X78)) & ? [X79] : ((p118(X79) & ~ p119(X79) & ~ p19(X79)) & r1(X1, X79)))) & ((~ p118(X1) | p119(X1)) | (? [X80] : ((p119(X80) & ~ p120(X80) & p20(X80)) & r1(X1, X80)) & ? [X81] : ((p119(X81) & ~ p120(X81) & ~ p20(X81)) & r1(X1, X81)))) & ((~ p119(X1) | p120(X1)) | (? [X82] : ((p120(X82) & p21(X82)) & r1(X1, X82)) & ? [X83] : ((p120(X83) & ~ p21(X83)) & r1(X1, X83))))) | ~ r1(X0, X1))) & ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35))))) & (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37))))) & (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39))))) & (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41))))) & (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X44] : (~ (p101(X44) & ~ p102(X44) & p2(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p101(X45) & ~ p102(X45) & ~ p2(X45)) | ~ r1(X1, X45)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X46] : (~ (p102(X46) & ~ p103(X46) & p3(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p102(X47) & ~ p103(X47) & ~ p3(X47)) | ~ r1(X1, X47)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X48] : (~ (p103(X48) & ~ p104(X48) & p4(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p103(X49) & ~ p104(X49) & ~ p4(X49)) | ~ r1(X1, X49)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X50] : (~ (p104(X50) & ~ p105(X50) & p5(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p104(X51) & ~ p105(X51) & ~ p5(X51)) | ~ r1(X1, X51)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X52] : (~ (p105(X52) & ~ p106(X52) & p6(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p105(X53) & ~ p106(X53) & ~ p6(X53)) | ~ r1(X1, X53)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X54] : (~ (p106(X54) & ~ p107(X54) & p7(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p106(X55) & ~ p107(X55) & ~ p7(X55)) | ~ r1(X1, X55)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X56] : (~ (p107(X56) & ~ p108(X56) & p8(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p107(X57) & ~ p108(X57) & ~ p8(X57)) | ~ r1(X1, X57)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X58] : (~ (p108(X58) & ~ p109(X58) & p9(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p108(X59) & ~ p109(X59) & ~ p9(X59)) | ~ r1(X1, X59)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X60] : (~ (p109(X60) & ~ p110(X60) & p10(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p109(X61) & ~ p110(X61) & ~ p10(X61)) | ~ r1(X1, X61)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X62] : (~ (p110(X62) & ~ p111(X62) & p11(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p110(X63) & ~ p111(X63) & ~ p11(X63)) | ~ r1(X1, X63)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X64] : (~ (p111(X64) & ~ p112(X64) & p12(X64)) | ~ r1(X1, X64)) & ~ ! [X65] : (~ (p111(X65) & ~ p112(X65) & ~ p12(X65)) | ~ r1(X1, X65)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X66] : (~ (p112(X66) & ~ p113(X66) & p13(X66)) | ~ r1(X1, X66)) & ~ ! [X67] : (~ (p112(X67) & ~ p113(X67) & ~ p13(X67)) | ~ r1(X1, X67)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X68] : (~ (p113(X68) & ~ p114(X68) & p14(X68)) | ~ r1(X1, X68)) & ~ ! [X69] : (~ (p113(X69) & ~ p114(X69) & ~ p14(X69)) | ~ r1(X1, X69)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X70] : (~ (p114(X70) & ~ p115(X70) & p15(X70)) | ~ r1(X1, X70)) & ~ ! [X71] : (~ (p114(X71) & ~ p115(X71) & ~ p15(X71)) | ~ r1(X1, X71)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X72] : (~ (p115(X72) & ~ p116(X72) & p16(X72)) | ~ r1(X1, X72)) & ~ ! [X73] : (~ (p115(X73) & ~ p116(X73) & ~ p16(X73)) | ~ r1(X1, X73)))) & (~ (p115(X1) & ~ p116(X1)) | (~ ! [X74] : (~ (p116(X74) & ~ p117(X74) & p17(X74)) | ~ r1(X1, X74)) & ~ ! [X75] : (~ (p116(X75) & ~ p117(X75) & ~ p17(X75)) | ~ r1(X1, X75)))) & (~ (p116(X1) & ~ p117(X1)) | (~ ! [X76] : (~ (p117(X76) & ~ p118(X76) & p18(X76)) | ~ r1(X1, X76)) & ~ ! [X77] : (~ (p117(X77) & ~ p118(X77) & ~ p18(X77)) | ~ r1(X1, X77)))) & (~ (p117(X1) & ~ p118(X1)) | (~ ! [X78] : (~ (p118(X78) & ~ p119(X78) & p19(X78)) | ~ r1(X1, X78)) & ~ ! [X79] : (~ (p118(X79) & ~ p119(X79) & ~ p19(X79)) | ~ r1(X1, X79)))) & (~ (p118(X1) & ~ p119(X1)) | (~ ! [X80] : (~ (p119(X80) & ~ p120(X80) & p20(X80)) | ~ r1(X1, X80)) & ~ ! [X81] : (~ (p119(X81) & ~ p120(X81) & ~ p20(X81)) | ~ r1(X1, X81)))) & (~ (p119(X1) & ~ p120(X1)) | (~ ! [X82] : (~ (p120(X82) & p21(X82)) | ~ r1(X1, X82)) & ~ ! [X83] : (~ (p120(X83) & ~ p21(X83)) | ~ r1(X1, X83))))) | ~ r1(X0, X1))) | ~ ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p121(X1) | p120(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35))))) & (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37))))) & (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39))))) & (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41))))) & (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X44] : (~ (p101(X44) & ~ p102(X44) & p2(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p101(X45) & ~ p102(X45) & ~ p2(X45)) | ~ r1(X1, X45)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X46] : (~ (p102(X46) & ~ p103(X46) & p3(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p102(X47) & ~ p103(X47) & ~ p3(X47)) | ~ r1(X1, X47)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X48] : (~ (p103(X48) & ~ p104(X48) & p4(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p103(X49) & ~ p104(X49) & ~ p4(X49)) | ~ r1(X1, X49)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X50] : (~ (p104(X50) & ~ p105(X50) & p5(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p104(X51) & ~ p105(X51) & ~ p5(X51)) | ~ r1(X1, X51)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X52] : (~ (p105(X52) & ~ p106(X52) & p6(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p105(X53) & ~ p106(X53) & ~ p6(X53)) | ~ r1(X1, X53)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X54] : (~ (p106(X54) & ~ p107(X54) & p7(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p106(X55) & ~ p107(X55) & ~ p7(X55)) | ~ r1(X1, X55)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X56] : (~ (p107(X56) & ~ p108(X56) & p8(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p107(X57) & ~ p108(X57) & ~ p8(X57)) | ~ r1(X1, X57)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X58] : (~ (p108(X58) & ~ p109(X58) & p9(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p108(X59) & ~ p109(X59) & ~ p9(X59)) | ~ r1(X1, X59)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X60] : (~ (p109(X60) & ~ p110(X60) & p10(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p109(X61) & ~ p110(X61) & ~ p10(X61)) | ~ r1(X1, X61)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X62] : (~ (p110(X62) & ~ p111(X62) & p11(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p110(X63) & ~ p111(X63) & ~ p11(X63)) | ~ r1(X1, X63)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X64] : (~ (p111(X64) & ~ p112(X64) & p12(X64)) | ~ r1(X1, X64)) & ~ ! [X65] : (~ (p111(X65) & ~ p112(X65) & ~ p12(X65)) | ~ r1(X1, X65)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X66] : (~ (p112(X66) & ~ p113(X66) & p13(X66)) | ~ r1(X1, X66)) & ~ ! [X67] : (~ (p112(X67) & ~ p113(X67) & ~ p13(X67)) | ~ r1(X1, X67)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X68] : (~ (p113(X68) & ~ p114(X68) & p14(X68)) | ~ r1(X1, X68)) & ~ ! [X69] : (~ (p113(X69) & ~ p114(X69) & ~ p14(X69)) | ~ r1(X1, X69)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X70] : (~ (p114(X70) & ~ p115(X70) & p15(X70)) | ~ r1(X1, X70)) & ~ ! [X71] : (~ (p114(X71) & ~ p115(X71) & ~ p15(X71)) | ~ r1(X1, X71)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X72] : (~ (p115(X72) & ~ p116(X72) & p16(X72)) | ~ r1(X1, X72)) & ~ ! [X73] : (~ (p115(X73) & ~ p116(X73) & ~ p16(X73)) | ~ r1(X1, X73)))) & (~ (p115(X1) & ~ p116(X1)) | (~ ! [X74] : (~ (p116(X74) & ~ p117(X74) & p17(X74)) | ~ r1(X1, X74)) & ~ ! [X75] : (~ (p116(X75) & ~ p117(X75) & ~ p17(X75)) | ~ r1(X1, X75)))) & (~ (p116(X1) & ~ p117(X1)) | (~ ! [X76] : (~ (p117(X76) & ~ p118(X76) & p18(X76)) | ~ r1(X1, X76)) & ~ ! [X77] : (~ (p117(X77) & ~ p118(X77) & ~ p18(X77)) | ~ r1(X1, X77)))) & (~ (p117(X1) & ~ p118(X1)) | (~ ! [X78] : (~ (p118(X78) & ~ p119(X78) & p19(X78)) | ~ r1(X1, X78)) & ~ ! [X79] : (~ (p118(X79) & ~ p119(X79) & ~ p19(X79)) | ~ r1(X1, X79)))) & (~ (p118(X1) & ~ p119(X1)) | (~ ! [X80] : (~ (p119(X80) & ~ p120(X80) & p20(X80)) | ~ r1(X1, X80)) & ~ ! [X81] : (~ (p119(X81) & ~ p120(X81) & ~ p20(X81)) | ~ r1(X1, X81)))) & (~ (p119(X1) & ~ p120(X1)) | (~ ! [X82] : (~ (p120(X82) & ~ p121(X82) & p21(X82)) | ~ r1(X1, X82)) & ~ ! [X83] : (~ (p120(X83) & ~ p121(X83) & ~ p21(X83)) | ~ r1(X1, X83))))) | ~ r1(X0, X1))) | ~ ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p121(X1) | p120(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p106(X1) | ((~ p7(X1) | ! [X14] : (~ p106(X14) | p7(X14) | ~ r1(X1, X14))) & (p7(X1) | ! [X15] : (~ p106(X15) | ~ p7(X15) | ~ r1(X1, X15))))) & (~ p107(X1) | ((~ p8(X1) | ! [X16] : (~ p107(X16) | p8(X16) | ~ r1(X1, X16))) & (p8(X1) | ! [X17] : (~ p107(X17) | ~ p8(X17) | ~ r1(X1, X17))))) & (~ p108(X1) | ((~ p9(X1) | ! [X18] : (~ p108(X18) | p9(X18) | ~ r1(X1, X18))) & (p9(X1) | ! [X19] : (~ p108(X19) | ~ p9(X19) | ~ r1(X1, X19))))) & (~ p109(X1) | ((~ p10(X1) | ! [X20] : (~ p109(X20) | p10(X20) | ~ r1(X1, X20))) & (p10(X1) | ! [X21] : (~ p109(X21) | ~ p10(X21) | ~ r1(X1, X21))))) & (~ p110(X1) | ((~ p11(X1) | ! [X22] : (~ p110(X22) | p11(X22) | ~ r1(X1, X22))) & (p11(X1) | ! [X23] : (~ p110(X23) | ~ p11(X23) | ~ r1(X1, X23))))) & (~ p111(X1) | ((~ p12(X1) | ! [X24] : (~ p111(X24) | p12(X24) | ~ r1(X1, X24))) & (p12(X1) | ! [X25] : (~ p111(X25) | ~ p12(X25) | ~ r1(X1, X25))))) & (~ p112(X1) | ((~ p13(X1) | ! [X26] : (~ p112(X26) | p13(X26) | ~ r1(X1, X26))) & (p13(X1) | ! [X27] : (~ p112(X27) | ~ p13(X27) | ~ r1(X1, X27))))) & (~ p113(X1) | ((~ p14(X1) | ! [X28] : (~ p113(X28) | p14(X28) | ~ r1(X1, X28))) & (p14(X1) | ! [X29] : (~ p113(X29) | ~ p14(X29) | ~ r1(X1, X29))))) & (~ p114(X1) | ((~ p15(X1) | ! [X30] : (~ p114(X30) | p15(X30) | ~ r1(X1, X30))) & (p15(X1) | ! [X31] : (~ p114(X31) | ~ p15(X31) | ~ r1(X1, X31))))) & (~ p115(X1) | ((~ p16(X1) | ! [X32] : (~ p115(X32) | p16(X32) | ~ r1(X1, X32))) & (p16(X1) | ! [X33] : (~ p115(X33) | ~ p16(X33) | ~ r1(X1, X33))))) & (~ p116(X1) | ((~ p17(X1) | ! [X34] : (~ p116(X34) | p17(X34) | ~ r1(X1, X34))) & (p17(X1) | ! [X35] : (~ p116(X35) | ~ p17(X35) | ~ r1(X1, X35))))) & (~ p117(X1) | ((~ p18(X1) | ! [X36] : (~ p117(X36) | p18(X36) | ~ r1(X1, X36))) & (p18(X1) | ! [X37] : (~ p117(X37) | ~ p18(X37) | ~ r1(X1, X37))))) & (~ p118(X1) | ((~ p19(X1) | ! [X38] : (~ p118(X38) | p19(X38) | ~ r1(X1, X38))) & (p19(X1) | ! [X39] : (~ p118(X39) | ~ p19(X39) | ~ r1(X1, X39))))) & (~ p119(X1) | ((~ p20(X1) | ! [X40] : (~ p119(X40) | p20(X40) | ~ r1(X1, X40))) & (p20(X1) | ! [X41] : (~ p119(X41) | ~ p20(X41) | ~ r1(X1, X41))))) & (~ p120(X1) | ((~ p21(X1) | ! [X42] : (~ p120(X42) | p21(X42) | ~ r1(X1, X42))) & (p21(X1) | ! [X43] : (~ p120(X43) | ~ p21(X43) | ~ r1(X1, X43))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X44] : (~ (p101(X44) & ~ p102(X44) & p2(X44)) | ~ r1(X1, X44)) & ~ ! [X45] : (~ (p101(X45) & ~ p102(X45) & ~ p2(X45)) | ~ r1(X1, X45)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X46] : (~ (p102(X46) & ~ p103(X46) & p3(X46)) | ~ r1(X1, X46)) & ~ ! [X47] : (~ (p102(X47) & ~ p103(X47) & ~ p3(X47)) | ~ r1(X1, X47)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X48] : (~ (p103(X48) & ~ p104(X48) & p4(X48)) | ~ r1(X1, X48)) & ~ ! [X49] : (~ (p103(X49) & ~ p104(X49) & ~ p4(X49)) | ~ r1(X1, X49)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X50] : (~ (p104(X50) & ~ p105(X50) & p5(X50)) | ~ r1(X1, X50)) & ~ ! [X51] : (~ (p104(X51) & ~ p105(X51) & ~ p5(X51)) | ~ r1(X1, X51)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X52] : (~ (p105(X52) & ~ p106(X52) & p6(X52)) | ~ r1(X1, X52)) & ~ ! [X53] : (~ (p105(X53) & ~ p106(X53) & ~ p6(X53)) | ~ r1(X1, X53)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X54] : (~ (p106(X54) & ~ p107(X54) & p7(X54)) | ~ r1(X1, X54)) & ~ ! [X55] : (~ (p106(X55) & ~ p107(X55) & ~ p7(X55)) | ~ r1(X1, X55)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X56] : (~ (p107(X56) & ~ p108(X56) & p8(X56)) | ~ r1(X1, X56)) & ~ ! [X57] : (~ (p107(X57) & ~ p108(X57) & ~ p8(X57)) | ~ r1(X1, X57)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X58] : (~ (p108(X58) & ~ p109(X58) & p9(X58)) | ~ r1(X1, X58)) & ~ ! [X59] : (~ (p108(X59) & ~ p109(X59) & ~ p9(X59)) | ~ r1(X1, X59)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X60] : (~ (p109(X60) & ~ p110(X60) & p10(X60)) | ~ r1(X1, X60)) & ~ ! [X61] : (~ (p109(X61) & ~ p110(X61) & ~ p10(X61)) | ~ r1(X1, X61)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X62] : (~ (p110(X62) & ~ p111(X62) & p11(X62)) | ~ r1(X1, X62)) & ~ ! [X63] : (~ (p110(X63) & ~ p111(X63) & ~ p11(X63)) | ~ r1(X1, X63)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X64] : (~ (p111(X64) & ~ p112(X64) & p12(X64)) | ~ r1(X1, X64)) & ~ ! [X65] : (~ (p111(X65) & ~ p112(X65) & ~ p12(X65)) | ~ r1(X1, X65)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X66] : (~ (p112(X66) & ~ p113(X66) & p13(X66)) | ~ r1(X1, X66)) & ~ ! [X67] : (~ (p112(X67) & ~ p113(X67) & ~ p13(X67)) | ~ r1(X1, X67)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X68] : (~ (p113(X68) & ~ p114(X68) & p14(X68)) | ~ r1(X1, X68)) & ~ ! [X69] : (~ (p113(X69) & ~ p114(X69) & ~ p14(X69)) | ~ r1(X1, X69)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X70] : (~ (p114(X70) & ~ p115(X70) & p15(X70)) | ~ r1(X1, X70)) & ~ ! [X71] : (~ (p114(X71) & ~ p115(X71) & ~ p15(X71)) | ~ r1(X1, X71)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X72] : (~ (p115(X72) & ~ p116(X72) & p16(X72)) | ~ r1(X1, X72)) & ~ ! [X73] : (~ (p115(X73) & ~ p116(X73) & ~ p16(X73)) | ~ r1(X1, X73)))) & (~ (p115(X1) & ~ p116(X1)) | (~ ! [X74] : (~ (p116(X74) & ~ p117(X74) & p17(X74)) | ~ r1(X1, X74)) & ~ ! [X75] : (~ (p116(X75) & ~ p117(X75) & ~ p17(X75)) | ~ r1(X1, X75)))) & (~ (p116(X1) & ~ p117(X1)) | (~ ! [X76] : (~ (p117(X76) & ~ p118(X76) & p18(X76)) | ~ r1(X1, X76)) & ~ ! [X77] : (~ (p117(X77) & ~ p118(X77) & ~ p18(X77)) | ~ r1(X1, X77)))) & (~ (p117(X1) & ~ p118(X1)) | (~ ! [X78] : (~ (p118(X78) & ~ p119(X78) & p19(X78)) | ~ r1(X1, X78)) & ~ ! [X79] : (~ (p118(X79) & ~ p119(X79) & ~ p19(X79)) | ~ r1(X1, X79)))) & (~ (p118(X1) & ~ p119(X1)) | (~ ! [X80] : (~ (p119(X80) & ~ p120(X80) & p20(X80)) | ~ r1(X1, X80)) & ~ ! [X81] : (~ (p119(X81) & ~ p120(X81) & ~ p20(X81)) | ~ r1(X1, X81)))) & (~ (p119(X1) & ~ p120(X1)) | (~ ! [X82] : (~ (p120(X82) & ~ p121(X82) & p21(X82)) | ~ r1(X1, X82)) & ~ ! [X83] : (~ (p120(X83) & ~ p121(X83) & ~ p21(X83)) | ~ r1(X1, X83))))) | ~ r1(X0, X1))) | ~ ! [X84] : (p8(X84) | ~ r1(X0, X84))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p121(X1) | p120(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ p116(X1) | ((~ p17(X1) | ! [X0] : (~ p116(X0) | p17(X0) | ~ r1(X1, X0))) & (p17(X1) | ! [X0] : (~ p116(X0) | ~ p17(X0) | ~ r1(X1, X0))))) & (~ p117(X1) | ((~ p18(X1) | ! [X0] : (~ p117(X0) | p18(X0) | ~ r1(X1, X0))) & (p18(X1) | ! [X0] : (~ p117(X0) | ~ p18(X0) | ~ r1(X1, X0))))) & (~ p118(X1) | ((~ p19(X1) | ! [X0] : (~ p118(X0) | p19(X0) | ~ r1(X1, X0))) & (p19(X1) | ! [X0] : (~ p118(X0) | ~ p19(X0) | ~ r1(X1, X0))))) & (~ p119(X1) | ((~ p20(X1) | ! [X0] : (~ p119(X0) | p20(X0) | ~ r1(X1, X0))) & (p20(X1) | ! [X0] : (~ p119(X0) | ~ p20(X0) | ~ r1(X1, X0))))) & (~ p120(X1) | ((~ p21(X1) | ! [X0] : (~ p120(X0) | p21(X0) | ~ r1(X1, X0))) & (p21(X1) | ! [X0] : (~ p120(X0) | ~ p21(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0)))) & (~ (p115(X1) & ~ p116(X1)) | (~ ! [X0] : (~ (p116(X0) & ~ p117(X0) & p17(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p116(X0) & ~ p117(X0) & ~ p17(X0)) | ~ r1(X1, X0)))) & (~ (p116(X1) & ~ p117(X1)) | (~ ! [X0] : (~ (p117(X0) & ~ p118(X0) & p18(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p117(X0) & ~ p118(X0) & ~ p18(X0)) | ~ r1(X1, X0)))) & (~ (p117(X1) & ~ p118(X1)) | (~ ! [X0] : (~ (p118(X0) & ~ p119(X0) & p19(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p118(X0) & ~ p119(X0) & ~ p19(X0)) | ~ r1(X1, X0)))) & (~ (p118(X1) & ~ p119(X1)) | (~ ! [X0] : (~ (p119(X0) & ~ p120(X0) & p20(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p119(X0) & ~ p120(X0) & ~ p20(X0)) | ~ r1(X1, X0)))) & (~ (p119(X1) & ~ p120(X1)) | (~ ! [X0] : (~ (p120(X0) & ~ p121(X0) & p21(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p120(X0) & ~ p121(X0) & ~ p21(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p8(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & (~ p121(X1) | p120(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ p106(X1) | ((~ p7(X1) | ! [X0] : (~ p106(X0) | p7(X0) | ~ r1(X1, X0))) & (p7(X1) | ! [X0] : (~ p106(X0) | ~ p7(X0) | ~ r1(X1, X0))))) & (~ p107(X1) | ((~ p8(X1) | ! [X0] : (~ p107(X0) | p8(X0) | ~ r1(X1, X0))) & (p8(X1) | ! [X0] : (~ p107(X0) | ~ p8(X0) | ~ r1(X1, X0))))) & (~ p108(X1) | ((~ p9(X1) | ! [X0] : (~ p108(X0) | p9(X0) | ~ r1(X1, X0))) & (p9(X1) | ! [X0] : (~ p108(X0) | ~ p9(X0) | ~ r1(X1, X0))))) & (~ p109(X1) | ((~ p10(X1) | ! [X0] : (~ p109(X0) | p10(X0) | ~ r1(X1, X0))) & (p10(X1) | ! [X0] : (~ p109(X0) | ~ p10(X0) | ~ r1(X1, X0))))) & (~ p110(X1) | ((~ p11(X1) | ! [X0] : (~ p110(X0) | p11(X0) | ~ r1(X1, X0))) & (p11(X1) | ! [X0] : (~ p110(X0) | ~ p11(X0) | ~ r1(X1, X0))))) & (~ p111(X1) | ((~ p12(X1) | ! [X0] : (~ p111(X0) | p12(X0) | ~ r1(X1, X0))) & (p12(X1) | ! [X0] : (~ p111(X0) | ~ p12(X0) | ~ r1(X1, X0))))) & (~ p112(X1) | ((~ p13(X1) | ! [X0] : (~ p112(X0) | p13(X0) | ~ r1(X1, X0))) & (p13(X1) | ! [X0] : (~ p112(X0) | ~ p13(X0) | ~ r1(X1, X0))))) & (~ p113(X1) | ((~ p14(X1) | ! [X0] : (~ p113(X0) | p14(X0) | ~ r1(X1, X0))) & (p14(X1) | ! [X0] : (~ p113(X0) | ~ p14(X0) | ~ r1(X1, X0))))) & (~ p114(X1) | ((~ p15(X1) | ! [X0] : (~ p114(X0) | p15(X0) | ~ r1(X1, X0))) & (p15(X1) | ! [X0] : (~ p114(X0) | ~ p15(X0) | ~ r1(X1, X0))))) & (~ p115(X1) | ((~ p16(X1) | ! [X0] : (~ p115(X0) | p16(X0) | ~ r1(X1, X0))) & (p16(X1) | ! [X0] : (~ p115(X0) | ~ p16(X0) | ~ r1(X1, X0))))) & (~ p116(X1) | ((~ p17(X1) | ! [X0] : (~ p116(X0) | p17(X0) | ~ r1(X1, X0))) & (p17(X1) | ! [X0] : (~ p116(X0) | ~ p17(X0) | ~ r1(X1, X0))))) & (~ p117(X1) | ((~ p18(X1) | ! [X0] : (~ p117(X0) | p18(X0) | ~ r1(X1, X0))) & (p18(X1) | ! [X0] : (~ p117(X0) | ~ p18(X0) | ~ r1(X1, X0))))) & (~ p118(X1) | ((~ p19(X1) | ! [X0] : (~ p118(X0) | p19(X0) | ~ r1(X1, X0))) & (p19(X1) | ! [X0] : (~ p118(X0) | ~ p19(X0) | ~ r1(X1, X0))))) & (~ p119(X1) | ((~ p20(X1) | ! [X0] : (~ p119(X0) | p20(X0) | ~ r1(X1, X0))) & (p20(X1) | ! [X0] : (~ p119(X0) | ~ p20(X0) | ~ r1(X1, X0))))) & (~ p120(X1) | ((~ p21(X1) | ! [X0] : (~ p120(X0) | p21(X0) | ~ r1(X1, X0))) & (p21(X1) | ! [X0] : (~ p120(X0) | ~ p21(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0)))) & (~ (p105(X1) & ~ p106(X1)) | (~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & p7(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p106(X0) & ~ p107(X0) & ~ p7(X0)) | ~ r1(X1, X0)))) & (~ (p106(X1) & ~ p107(X1)) | (~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & p8(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p107(X0) & ~ p108(X0) & ~ p8(X0)) | ~ r1(X1, X0)))) & (~ (p107(X1) & ~ p108(X1)) | (~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & p9(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p108(X0) & ~ p109(X0) & ~ p9(X0)) | ~ r1(X1, X0)))) & (~ (p108(X1) & ~ p109(X1)) | (~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & p10(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p109(X0) & ~ p110(X0) & ~ p10(X0)) | ~ r1(X1, X0)))) & (~ (p109(X1) & ~ p110(X1)) | (~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & p11(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p110(X0) & ~ p111(X0) & ~ p11(X0)) | ~ r1(X1, X0)))) & (~ (p110(X1) & ~ p111(X1)) | (~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & p12(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p111(X0) & ~ p112(X0) & ~ p12(X0)) | ~ r1(X1, X0)))) & (~ (p111(X1) & ~ p112(X1)) | (~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & p13(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p112(X0) & ~ p113(X0) & ~ p13(X0)) | ~ r1(X1, X0)))) & (~ (p112(X1) & ~ p113(X1)) | (~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & p14(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p113(X0) & ~ p114(X0) & ~ p14(X0)) | ~ r1(X1, X0)))) & (~ (p113(X1) & ~ p114(X1)) | (~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & p15(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p114(X0) & ~ p115(X0) & ~ p15(X0)) | ~ r1(X1, X0)))) & (~ (p114(X1) & ~ p115(X1)) | (~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & p16(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p115(X0) & ~ p116(X0) & ~ p16(X0)) | ~ r1(X1, X0)))) & (~ (p115(X1) & ~ p116(X1)) | (~ ! [X0] : (~ (p116(X0) & ~ p117(X0) & p17(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p116(X0) & ~ p117(X0) & ~ p17(X0)) | ~ r1(X1, X0)))) & (~ (p116(X1) & ~ p117(X1)) | (~ ! [X0] : (~ (p117(X0) & ~ p118(X0) & p18(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p117(X0) & ~ p118(X0) & ~ p18(X0)) | ~ r1(X1, X0)))) & (~ (p117(X1) & ~ p118(X1)) | (~ ! [X0] : (~ (p118(X0) & ~ p119(X0) & p19(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p118(X0) & ~ p119(X0) & ~ p19(X0)) | ~ r1(X1, X0)))) & (~ (p118(X1) & ~ p119(X1)) | (~ ! [X0] : (~ (p119(X0) & ~ p120(X0) & p20(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p119(X0) & ~ p120(X0) & ~ p20(X0)) | ~ r1(X1, X0)))) & (~ (p119(X1) & ~ p120(X1)) | (~ ! [X0] : (~ (p120(X0) & ~ p121(X0) & p21(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p120(X0) & ~ p121(X0) & ~ p21(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p8(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.020.p', main)).
fof(f23469, plain, (p101(sK82) | ~ r1(sK82, sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23468, f490])).
fof(f490, plain, sP18(sK82), inference(resolution, [], [f470, f223])).
fof(f223, plain, ! [X0] : (~ sP41(X0) | sP18(X0)), inference(cnf_transformation, [], [f56])).
fof(f56, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & (~ p111(X0) | p110(X0)) & (~ p112(X0) | p111(X0)) & (~ p113(X0) | p112(X0)) & (~ p114(X0) | p113(X0)) & (~ p115(X0) | p114(X0)) & (~ p116(X0) | p115(X0)) & (~ p117(X0) | p116(X0)) & (~ p118(X0) | p117(X0)) & (~ p119(X0) | p118(X0)) & (~ p120(X0) | p119(X0)) & sP40(X0) & sP39(X0) & sP38(X0) & sP37(X0) & sP36(X0) & sP35(X0) & sP34(X0) & sP33(X0) & sP32(X0) & sP31(X0) & sP30(X0) & sP29(X0) & sP28(X0) & sP27(X0) & sP26(X0) & sP25(X0) & sP24(X0) & sP23(X0) & sP22(X0) & sP21(X0) & sP20(X0) & sP18(X0) & sP17(X0) & sP16(X0) & sP15(X0) & sP14(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP4(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP19(X0)) | ~ sP41(X0)), inference(rectify, [], [f55])).
fof(f55, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p107(X1) | p106(X1)) & (~ p108(X1) | p107(X1)) & (~ p109(X1) | p108(X1)) & (~ p110(X1) | p109(X1)) & (~ p111(X1) | p110(X1)) & (~ p112(X1) | p111(X1)) & (~ p113(X1) | p112(X1)) & (~ p114(X1) | p113(X1)) & (~ p115(X1) | p114(X1)) & (~ p116(X1) | p115(X1)) & (~ p117(X1) | p116(X1)) & (~ p118(X1) | p117(X1)) & (~ p119(X1) | p118(X1)) & (~ p120(X1) | p119(X1)) & sP40(X1) & sP39(X1) & sP38(X1) & sP37(X1) & sP36(X1) & sP35(X1) & sP34(X1) & sP33(X1) & sP32(X1) & sP31(X1) & sP30(X1) & sP29(X1) & sP28(X1) & sP27(X1) & sP26(X1) & sP25(X1) & sP24(X1) & sP23(X1) & sP22(X1) & sP21(X1) & sP20(X1) & sP18(X1) & sP17(X1) & sP16(X1) & sP15(X1) & sP14(X1) & sP13(X1) & sP12(X1) & sP11(X1) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP4(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP19(X1)) | ~ sP41(X1)), inference(nnf_transformation, [], [f53])).
fof(f470, plain, sP41(sK82), inference(resolution, [], [f466, f202])).
fof(f466, plain, ! [X1] : (~ r1(sK82, X1) | sP41(X1)), inference(cnf_transformation, [], [f201])).
fof(f23468, plain, (~ sP18(sK82) | p101(sK82) | ~ r1(sK82, sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23463, f468])).
fof(f468, plain, p100(sK82), inference(cnf_transformation, [], [f201])).
fof(f23463, plain, (~ p100(sK82) | ~ sP18(sK82) | p101(sK82) | ~ r1(sK82, sK82) | (spl83_59 | ~ spl83_60)), inference(resolution, [], [f23452, f811])).
fof(f811, plain, ! [X0, X1] : (r1(X1, sK44(X0)) | ~ p100(X0) | ~ sP18(X0) | p101(X0) | ~ r1(X1, X0)), inference(resolution, [], [f317, f203])).
fof(f203, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.020.p', transitivity)).
fof(f317, plain, ! [X0] : (r1(X0, sK44(X0)) | p101(X0) | ~ p100(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK44(X0)) & ~ p102(sK44(X0)) & p2(sK44(X0)) & r1(X0, sK44(X0))) & (p101(sK45(X0)) & ~ p102(sK45(X0)) & ~ p2(sK45(X0)) & r1(X0, sK45(X0)))) | ~ sP18(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44, sK45])], [f105, f107, f106])).
fof(f106, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK44(X0)) & ~ p102(sK44(X0)) & p2(sK44(X0)) & r1(X0, sK44(X0)))), introduced(choice_axiom, [])).
fof(f107, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK45(X0)) & ~ p102(sK45(X0)) & ~ p2(sK45(X0)) & r1(X0, sK45(X0)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP18(X0)), inference(rectify, [], [f104])).
fof(f104, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X44] : (p101(X44) & ~ p102(X44) & p2(X44) & r1(X1, X44)) & ? [X45] : (p101(X45) & ~ p102(X45) & ~ p2(X45) & r1(X1, X45))) | ~ sP18(X1)), inference(nnf_transformation, [], [f30])).
fof(f23452, plain, (~ r1(sK82, sK44(sK82)) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23451, f973])).
fof(f973, plain, (~ p102(sK44(sK82)) | spl83_59), inference(avatar_component_clause, [], [f972])).
fof(f972, plain, (spl83_59 <=> p102(sK44(sK82))), introduced(avatar_definition, [new_symbols(naming, [spl83_59])])).
fof(f23451, plain, (p102(sK44(sK82)) | ~ r1(sK82, sK44(sK82)) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23450, f839])).
fof(f839, plain, sP17(sK44(sK82)), inference(resolution, [], [f814, f222])).
fof(f222, plain, ! [X0] : (~ sP41(X0) | sP17(X0)), inference(cnf_transformation, [], [f56])).
fof(f814, plain, sP41(sK44(sK82)), inference(subsumption_resolution, [], [f813, f490])).
fof(f813, plain, (~ sP18(sK82) | sP41(sK44(sK82))), inference(subsumption_resolution, [], [f812, f468])).
fof(f812, plain, (~ p100(sK82) | ~ sP18(sK82) | sP41(sK44(sK82))), inference(subsumption_resolution, [], [f809, f467])).
fof(f809, plain, (p101(sK82) | ~ p100(sK82) | ~ sP18(sK82) | sP41(sK44(sK82))), inference(resolution, [], [f317, f466])).
fof(f23450, plain, (~ sP17(sK44(sK82)) | p102(sK44(sK82)) | ~ r1(sK82, sK44(sK82)) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f23444, f979])).
fof(f979, plain, (p101(sK44(sK82)) | ~ spl83_60), inference(avatar_component_clause, [], [f977])).
fof(f977, plain, (spl83_60 <=> p101(sK44(sK82))), introduced(avatar_definition, [new_symbols(naming, [spl83_60])])).
fof(f23444, plain, (~ p101(sK44(sK82)) | ~ sP17(sK44(sK82)) | p102(sK44(sK82)) | ~ r1(sK82, sK44(sK82)) | (spl83_59 | ~ spl83_60)), inference(resolution, [], [f23436, f6474])).
fof(f6474, plain, (sP16(sK46(sK44(sK82))) | (spl83_59 | ~ spl83_60)), inference(resolution, [], [f6446, f221])).
fof(f221, plain, ! [X0] : (~ sP41(X0) | sP16(X0)), inference(cnf_transformation, [], [f56])).
fof(f6446, plain, (sP41(sK46(sK44(sK82))) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f6445, f490])).
fof(f6445, plain, (sP41(sK46(sK44(sK82))) | ~ sP18(sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f6444, f468])).
fof(f6444, plain, (sP41(sK46(sK44(sK82))) | ~ p100(sK82) | ~ sP18(sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f6443, f467])).
fof(f6443, plain, (sP41(sK46(sK44(sK82))) | p101(sK82) | ~ p100(sK82) | ~ sP18(sK82) | (spl83_59 | ~ spl83_60)), inference(subsumption_resolution, [], [f6442, f979])).
fof(f6442, plain, (~ p101(sK44(sK82)) | sP41(sK46(sK44(sK82))) | p101(sK82) | ~ p100(sK82) | ~ sP18(sK82) | spl83_59), inference(subsumption_resolution, [], [f6441, f973])).
fof(f6441, plain, (p102(sK44(sK82)) | ~ p101(sK44(sK82)) | sP41(sK46(sK44(sK82))) | p101(sK82) | ~ p100(sK82) | ~ sP18(sK82)), inference(subsumption_resolution, [], [f6365, f839])).
fof(f6365, plain, (~ sP17(sK44(sK82)) | p102(sK44(sK82)) | ~ p101(sK44(sK82)) | sP41(sK46(sK44(sK82))) | p101(sK82) | ~ p100(sK82) | ~ sP18(sK82)), inference(resolution, [], [f3177, f317])).
fof(f3177, plain, ! [X0] : (~ r1(sK82, X0) | ~ sP17(X0) | p102(X0) | ~ p101(X0) | sP41(sK46(X0))), inference(resolution, [], [f988, f466])).
fof(f988, plain, ! [X0, X1] : (r1(X1, sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(X1, X0)), inference(resolution, [], [f325, f203])).
fof(f325, plain, ! [X0] : (r1(X0, sK46(X0)) | p102(X0) | ~ p101(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK46(X0)) & ~ p103(sK46(X0)) & p3(sK46(X0)) & r1(X0, sK46(X0))) & (p102(sK47(X0)) & ~ p103(sK47(X0)) & ~ p3(sK47(X0)) & r1(X0, sK47(X0)))) | ~ sP17(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46, sK47])], [f110, f112, f111])).
fof(f111, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK46(X0)) & ~ p103(sK46(X0)) & p3(sK46(X0)) & r1(X0, sK46(X0)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK47(X0)) & ~ p103(sK47(X0)) & ~ p3(sK47(X0)) & r1(X0, sK47(X0)))), introduced(choice_axiom, [])).
fof(f110, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP17(X0)), inference(rectify, [], [f109])).
fof(f109, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X46] : (p102(X46) & ~ p103(X46) & p3(X46) & r1(X1, X46)) & ? [X47] : (p102(X47) & ~ p103(X47) & ~ p3(X47) & r1(X1, X47))) | ~ sP17(X1)), inference(nnf_transformation, [], [f29])).
fof(f23436, plain, ! [X0] : (~ sP16(sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23435, f988])).
fof(f23435, plain, ! [X0] : (~ sP16(sK46(X0)) | ~ r1(sK82, sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23434, f327])).
fof(f327, plain, ! [X0] : (~ p103(sK46(X0)) | p102(X0) | ~ p101(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f113])).
fof(f23434, plain, ! [X0] : (~ sP16(sK46(X0)) | p103(sK46(X0)) | ~ r1(sK82, sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23433, f328])).
fof(f328, plain, ! [X0] : (p102(sK46(X0)) | p102(X0) | ~ p101(X0) | ~ sP17(X0)), inference(cnf_transformation, [], [f113])).
fof(f23433, plain, ! [X0] : (~ p102(sK46(X0)) | ~ sP16(sK46(X0)) | p103(sK46(X0)) | ~ r1(sK82, sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(sK82, X0)), inference(duplicate_literal_removal, [], [f23430])).
fof(f23430, plain, ! [X0] : (~ p102(sK46(X0)) | ~ sP16(sK46(X0)) | p103(sK46(X0)) | ~ r1(sK82, sK46(X0)) | ~ p101(X0) | ~ sP17(X0) | p102(X0) | ~ r1(sK82, X0) | ~ sP16(sK46(X0))), inference(resolution, [], [f23422, f19749])).
fof(f19749, plain, ! [X16] : (sP15(sK48(sK46(X16))) | ~ p101(X16) | ~ sP17(X16) | p102(X16) | ~ r1(sK82, X16) | ~ sP16(sK46(X16))), inference(resolution, [], [f7156, f220])).
fof(f220, plain, ! [X0] : (~ sP41(X0) | sP15(X0)), inference(cnf_transformation, [], [f56])).
fof(f7156, plain, ! [X4] : (sP41(sK48(sK46(X4))) | ~ sP16(sK46(X4)) | ~ p101(X4) | ~ sP17(X4) | p102(X4) | ~ r1(sK82, X4)), inference(subsumption_resolution, [], [f7155, f328])).
fof(f7155, plain, ! [X4] : (~ sP16(sK46(X4)) | ~ p102(sK46(X4)) | sP41(sK48(sK46(X4))) | ~ p101(X4) | ~ sP17(X4) | p102(X4) | ~ r1(sK82, X4)), inference(subsumption_resolution, [], [f7084, f327])).
fof(f7084, plain, ! [X4] : (~ sP16(sK46(X4)) | p103(sK46(X4)) | ~ p102(sK46(X4)) | sP41(sK48(sK46(X4))) | ~ p101(X4) | ~ sP17(X4) | p102(X4) | ~ r1(sK82, X4)), inference(resolution, [], [f3272, f988])).
fof(f3272, plain, ! [X0] : (~ r1(sK82, X0) | ~ sP16(X0) | p103(X0) | ~ p102(X0) | sP41(sK48(X0))), inference(resolution, [], [f994, f466])).
fof(f994, plain, ! [X0, X1] : (r1(X1, sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(X1, X0)), inference(resolution, [], [f333, f203])).
fof(f333, plain, ! [X0] : (r1(X0, sK48(X0)) | p103(X0) | ~ p102(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f118])).
fof(f118, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK48(X0)) & ~ p104(sK48(X0)) & p4(sK48(X0)) & r1(X0, sK48(X0))) & (p103(sK49(X0)) & ~ p104(sK49(X0)) & ~ p4(sK49(X0)) & r1(X0, sK49(X0)))) | ~ sP16(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48, sK49])], [f115, f117, f116])).
fof(f116, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK48(X0)) & ~ p104(sK48(X0)) & p4(sK48(X0)) & r1(X0, sK48(X0)))), introduced(choice_axiom, [])).
fof(f117, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK49(X0)) & ~ p104(sK49(X0)) & ~ p4(sK49(X0)) & r1(X0, sK49(X0)))), introduced(choice_axiom, [])).
fof(f115, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP16(X0)), inference(rectify, [], [f114])).
fof(f114, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X48] : (p103(X48) & ~ p104(X48) & p4(X48) & r1(X1, X48)) & ? [X49] : (p103(X49) & ~ p104(X49) & ~ p4(X49) & r1(X1, X49))) | ~ sP16(X1)), inference(nnf_transformation, [], [f28])).
fof(f23422, plain, ! [X0] : (~ sP15(sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23421, f994])).
fof(f23421, plain, ! [X0] : (~ sP15(sK48(X0)) | ~ r1(sK82, sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23420, f335])).
fof(f335, plain, ! [X0] : (~ p104(sK48(X0)) | p103(X0) | ~ p102(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f118])).
fof(f23420, plain, ! [X0] : (~ sP15(sK48(X0)) | p104(sK48(X0)) | ~ r1(sK82, sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23419, f336])).
fof(f336, plain, ! [X0] : (p103(sK48(X0)) | p103(X0) | ~ p102(X0) | ~ sP16(X0)), inference(cnf_transformation, [], [f118])).
fof(f23419, plain, ! [X0] : (~ p103(sK48(X0)) | ~ sP15(sK48(X0)) | p104(sK48(X0)) | ~ r1(sK82, sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(sK82, X0)), inference(duplicate_literal_removal, [], [f23416])).
fof(f23416, plain, ! [X0] : (~ p103(sK48(X0)) | ~ sP15(sK48(X0)) | p104(sK48(X0)) | ~ r1(sK82, sK48(X0)) | ~ p102(X0) | ~ sP16(X0) | p103(X0) | ~ r1(sK82, X0) | ~ sP15(sK48(X0))), inference(resolution, [], [f23412, f19992])).
fof(f19992, plain, ! [X15] : (sP14(sK50(sK48(X15))) | ~ p102(X15) | ~ sP16(X15) | p103(X15) | ~ r1(sK82, X15) | ~ sP15(sK48(X15))), inference(resolution, [], [f7496, f219])).
fof(f219, plain, ! [X0] : (~ sP41(X0) | sP14(X0)), inference(cnf_transformation, [], [f56])).
fof(f7496, plain, ! [X6] : (sP41(sK50(sK48(X6))) | ~ sP15(sK48(X6)) | ~ p102(X6) | ~ sP16(X6) | p103(X6) | ~ r1(sK82, X6)), inference(subsumption_resolution, [], [f7495, f336])).
fof(f7495, plain, ! [X6] : (~ sP15(sK48(X6)) | ~ p103(sK48(X6)) | sP41(sK50(sK48(X6))) | ~ p102(X6) | ~ sP16(X6) | p103(X6) | ~ r1(sK82, X6)), inference(subsumption_resolution, [], [f7428, f335])).
fof(f7428, plain, ! [X6] : (~ sP15(sK48(X6)) | p104(sK48(X6)) | ~ p103(sK48(X6)) | sP41(sK50(sK48(X6))) | ~ p102(X6) | ~ sP16(X6) | p103(X6) | ~ r1(sK82, X6)), inference(resolution, [], [f3367, f994])).
fof(f3367, plain, ! [X0] : (~ r1(sK82, X0) | ~ sP15(X0) | p104(X0) | ~ p103(X0) | sP41(sK50(X0))), inference(resolution, [], [f1000, f466])).
fof(f1000, plain, ! [X0, X1] : (r1(X1, sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(X1, X0)), inference(resolution, [], [f341, f203])).
fof(f341, plain, ! [X0] : (r1(X0, sK50(X0)) | p104(X0) | ~ p103(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f123])).
fof(f123, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK50(X0)) & ~ p105(sK50(X0)) & p5(sK50(X0)) & r1(X0, sK50(X0))) & (p104(sK51(X0)) & ~ p105(sK51(X0)) & ~ p5(sK51(X0)) & r1(X0, sK51(X0)))) | ~ sP15(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50, sK51])], [f120, f122, f121])).
fof(f121, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK50(X0)) & ~ p105(sK50(X0)) & p5(sK50(X0)) & r1(X0, sK50(X0)))), introduced(choice_axiom, [])).
fof(f122, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK51(X0)) & ~ p105(sK51(X0)) & ~ p5(sK51(X0)) & r1(X0, sK51(X0)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP15(X0)), inference(rectify, [], [f119])).
fof(f119, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X50] : (p104(X50) & ~ p105(X50) & p5(X50) & r1(X1, X50)) & ? [X51] : (p104(X51) & ~ p105(X51) & ~ p5(X51) & r1(X1, X51))) | ~ sP15(X1)), inference(nnf_transformation, [], [f27])).
fof(f23412, plain, ! [X0] : (~ sP14(sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23411, f344])).
fof(f344, plain, ! [X0] : (p104(sK50(X0)) | p104(X0) | ~ p103(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f123])).
fof(f23411, plain, ! [X0] : (~ sP14(sK50(X0)) | ~ p104(sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23410, f1000])).
fof(f23410, plain, ! [X0] : (~ sP14(sK50(X0)) | ~ r1(sK82, sK50(X0)) | ~ p104(sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(sK82, X0)), inference(subsumption_resolution, [], [f23409, f343])).
fof(f343, plain, ! [X0] : (~ p105(sK50(X0)) | p104(X0) | ~ p103(X0) | ~ sP15(X0)), inference(cnf_transformation, [], [f123])).
fof(f23409, plain, ! [X0] : (~ sP14(sK50(X0)) | p105(sK50(X0)) | ~ r1(sK82, sK50(X0)) | ~ p104(sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(sK82, X0)), inference(duplicate_literal_removal, [], [f23406])).
fof(f23406, plain, ! [X0] : (~ sP14(sK50(X0)) | p105(sK50(X0)) | ~ r1(sK82, sK50(X0)) | ~ p104(sK50(X0)) | ~ p103(X0) | ~ sP15(X0) | p104(X0) | ~ r1(sK82, X0) | ~ sP14(sK50(X0))), inference(resolution, [], [f23405, f20235])).
fof(f20235, plain, ! [X14] : (sP13(sK52(sK50(X14))) | ~ p103(X14) | ~ sP15(X14) | p104(X14) | ~ r1(sK82, X14) | ~ sP14(sK50(X14))), inference(resolution, [], [f7836, f218])).
fof(f218, plain, ! [X0] : (~ sP41(X0) | sP13(X0)), inference(cnf_transformation, [], [f56])).
fof(f7836, plain, ! [X8] : (sP41(sK52(sK50(X8))) | ~ sP14(sK50(X8)) | ~ p103(X8) | ~ sP15(X8) | p104(X8) | ~ r1(sK82, X8)), inference(subsumption_resolution, [], [f7835, f344])).
fof(f7835, plain, ! [X8] : (~ sP14(sK50(X8)) | ~ p104(sK50(X8)) | sP41(sK52(sK50(X8))) | ~ p103(X8) | ~ sP15(X8) | p104(X8) | ~ r1(sK82, X8)), inference(subsumption_resolution, [], [f7772, f343])).
fof(f7772, plain, ! [X8] : (~ sP14(sK50(X8)) | p105(sK50(X8)) | ~ p104(sK50(X8)) | sP41(sK52(sK50(X8))) | ~ p103(X8) | ~ sP15(X8) | p104(X8) | ~ r1(sK82, X8)), inference(resolution, [], [f3462, f1000])).
fof(f3462, plain, ! [X0] : (~ r1(sK82, X0) | ~ sP14(X0) | p105(X0) | ~ p104(X0) | sP41(sK52(X0))), inference(resolution, [], [f1006, f466])).
fof(f1006, plain, ! [X0, X1] : (r1(X1, sK52(X0)) | ~ p104(X0) | ~ sP14(X0) | p105(X0) | ~ r1(X1, X0)), inference(resolution, [], [f349, f203])).
fof(f349, plain, ! [X0] : (r1(X0, sK52(X0)) | p105(X0) | ~ p104(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ! [X0] : (~ p104(X0) | p105(X0) | ((p105(sK52(X0)) & ~ p106(sK52(X0)) & p6(sK52(X0)) & r1(X0, sK52(X0))) & (p105(sK53(X0)) & ~ p106(sK53(X0)) & ~ p6(sK53(X0)) & r1(X0, sK53(X0)))) | ~ sP14(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52, sK53])], [f125, f127, f126])).
fof(f126, plain, ! [X0] : (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) => (p105(sK52(X0)) & ~ p106(sK52(X0)) & p6(sK52(X0)) & r1(X0, sK52(X0)))), introduced(choice_axiom, [])).
fof(f127, plain, ! [X0] : (? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2)) => (p105(sK53(X0)) & ~ p106(sK53(X0)) & ~ p6(sK53(X0)) & r1(X0, sK53(X0)))), introduced(choice_axiom, [])).
fof(f125, plain, ! [X0] : (~ p104(X0) | p105(X0) | (? [X1] : (p105(X1) & ~ p106(X1) & p6(X1) & r1(X0, X1)) & ? [X2] : (p105(X2) & ~ p106(X2) & ~ p6(X2) & r1(X0, X2))) | ~ sP14(X0)), inference(rectify, [], [f124])).
fof(f124, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X52] : (p105(X52) & ~ p106(X52) & p6(X52) & r1(X1, X52)) & ? [X53] : (p105(X53) & ~ p106(X53) & ~ p6(X53) & r1(X1, X53))) | ~ sP14(X1)), inference(nnf_transformation, [], [f26])).
fof(f23405, plain, ! [X3] : (~ sP13(sK52(X3)) | ~ sP14(X3) | p105(X3) | ~ r1(sK82, X3) | ~ p104(X3)), inference(subsumption_resolution, [], [f23404, f1006])).
fof(f23404, plain, ! [X3] : (~ p104(X3) | ~ sP14(X3) | p105(X3) | ~ r1(sK82, X3) | ~ sP13(sK52(X3)) | ~ r1(sK82, sK52(X3))), inference(subsumption_resolution, [], [f23403, f351])).
fof(f351, plain, ! [X0] : (~ p106(sK52(X0)) | p105(X0) | ~ p104(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f128])).
fof(f23403, plain, ! [X3] : (~ p104(X3) | ~ sP14(X3) | p105(X3) | ~ r1(sK82, X3) | ~ sP13(sK52(X3)) | p106(sK52(X3)) | ~ r1(sK82, sK52(X3))), inference(subsumption_resolution, [], [f23396, f352])).
fof(f352, plain, ! [X0] : (p105(sK52(X0)) | p105(X0) | ~ p104(X0) | ~ sP14(X0)), inference(cnf_transformation, [], [f128])).
fof(f23396, plain, ! [X3] : (~ p104(X3) | ~ sP14(X3) | p105(X3) | ~ r1(sK82, X3) | ~ sP13(sK52(X3)) | ~ p105(sK52(X3)) | p106(sK52(X3)) | ~ r1(sK82, sK52(X3))), inference(duplicate_literal_removal, [], [f23395])).
fof(f23395, plain, ! [X3] : (~ p104(X3) | ~ sP14(X3) | p105(X3) | ~ r1(sK82, X3) | ~ sP13(sK52(X3)) | ~ p105(sK52(X3)) | ~ sP13(sK52(X3)) | p106(sK52(X3)) | ~ r1(sK82, sK52(X3))), inference(resolution, [], [f20356, f4945])).
fof(f4945, plain, ! [X13] : (~ sP12(sK55(X13)) | ~ p105(X13) | ~ sP13(X13) | p106(X13) | ~ r1(sK82, X13)), inference(subsumption_resolution, [], [f4944, f356])).
fof(f356, plain, ! [X0] : (p106(sK55(X0)) | p106(X0) | ~ p105(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f133])).
fof(f133, plain, ! [X0] : (~ p105(X0) | p106(X0) | ((p106(sK54(X0)) & ~ p107(sK54(X0)) & p7(sK54(X0)) & r1(X0, sK54(X0))) & (p106(sK55(X0)) & ~ p107(sK55(X0)) & ~ p7(sK55(X0)) & r1(X0, sK55(X0)))) | ~ sP13(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55])], [f130, f132, f131])).
fof(f131, plain, ! [X0] : (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) => (p106(sK54(X0)) & ~ p107(sK54(X0)) & p7(sK54(X0)) & r1(X0, sK54(X0)))), introduced(choice_axiom, [])).
fof(f132, plain, ! [X0] : (? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2)) => (p106(sK55(X0)) & ~ p107(sK55(X0)) & ~ p7(sK55(X0)) & r1(X0, sK55(X0)))), introduced(choice_axiom, [])).
fof(f130, plain, ! [X0] : (~ p105(X0) | p106(X0) | (? [X1] : (p106(X1) & ~ p107(X1) & p7(X1) & r1(X0, X1)) & ? [X2] : (p106(X2) & ~ p107(X2) & ~ p7(X2) & r1(X0, X2))) | ~ sP13(X0)), inference(rectify, [], [f129])).
fof(f129, plain, ! [X1] : (~ p105(X1) | p106(X1) | (? [X54] : (p106(X54) & ~ p107(X54) & p7(X54) & r1(X1, X54)) & ? [X55] : (p106(X55) & ~ p107(X55) & ~ p7(X55) & r1(X1, X55))) | ~ sP13(X1)), inference(nnf_transformation, [], [f25])).
fof(f4944, plain, ! [X13] : (~ sP12(sK55(X13)) | ~ p106(sK55(X13)) | ~ p105(X13) | ~ sP13(X13) | p106(X13) | ~ r1(sK82, X13)), inference(subsumption_resolution, [], [f4889, f355])).
fof(f355, plain, ! [X0] : (~ p107(sK55(X0)) | p106(X0) | ~ p105(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f133])).
fof(f4889, plain, ! [X13] : (~ sP12(sK55(X13)) | p107(sK55(X13)) | ~ p106(sK55(X13)) | ~ p105(X13) | ~ sP13(X13) | p106(X13) | ~ r1(sK82, X13)), inference(resolution, [], [f3650, f1009])).
fof(f1009, plain, ! [X0, X1] : (r1(X1, sK55(X0)) | ~ p105(X0) | ~ sP13(X0) | p106(X0) | ~ r1(X1, X0)), inference(resolution, [], [f353, f203])).
fof(f353, plain, ! [X0] : (r1(X0, sK55(X0)) | p106(X0) | ~ p105(X0) | ~ sP13(X0)), inference(cnf_transformation, [], [f133])).
fof(f3650, plain, ! [X1] : (~ r1(sK82, X1) | ~ sP12(X1) | p107(X1) | ~ p106(X1)), inference(subsumption_resolution, [], [f3606, f362])).
fof(f362, plain, ! [X0] : (~ p8(sK57(X0)) | p107(X0) | ~ p106(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0] : (~ p106(X0) | p107(X0) | ((p107(sK56(X0)) & ~ p108(sK56(X0)) & p8(sK56(X0)) & r1(X0, sK56(X0))) & (p107(sK57(X0)) & ~ p108(sK57(X0)) & ~ p8(sK57(X0)) & r1(X0, sK57(X0)))) | ~ sP12(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK56, sK57])], [f135, f137, f136])).
fof(f136, plain, ! [X0] : (? [X1] : (p107(X1) & ~ p108(X1) & p8(X1) & r1(X0, X1)) => (p107(sK56(X0)) & ~ p108(sK56(X0)) & p8(sK56(X0)) & r1(X0, sK56(X0)))), introduced(choice_axiom, [])).
fof(f137, plain, ! [X0] : (? [X2] : (p107(X2) & ~ p108(X2) & ~ p8(X2) & r1(X0, X2)) => (p107(sK57(X0)) & ~ p108(sK57(X0)) & ~ p8(sK57(X0)) & r1(X0, sK57(X0)))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X0] : (~ p106(X0) | p107(X0) | (? [X1] : (p107(X1) & ~ p108(X1) & p8(X1) & r1(X0, X1)) & ? [X2] : (p107(X2) & ~ p108(X2) & ~ p8(X2) & r1(X0, X2))) | ~ sP12(X0)), inference(rectify, [], [f134])).
fof(f134, plain, ! [X1] : (~ p106(X1) | p107(X1) | (? [X56] : (p107(X56) & ~ p108(X56) & p8(X56) & r1(X1, X56)) & ? [X57] : (p107(X57) & ~ p108(X57) & ~ p8(X57) & r1(X1, X57))) | ~ sP12(X1)), inference(nnf_transformation, [], [f24])).
fof(f3606, plain, ! [X1] : (~ p106(X1) | ~ sP12(X1) | p107(X1) | ~ r1(sK82, X1) | p8(sK57(X1))), inference(resolution, [], [f1015, f465])).
fof(f465, plain, ! [X2] : (~ r1(sK82, X2) | p8(X2)), inference(cnf_transformation, [], [f201])).
fof(f1015, plain, ! [X0, X1] : (r1(X1, sK57(X0)) | ~ p106(X0) | ~ sP12(X0) | p107(X0) | ~ r1(X1, X0)), inference(resolution, [], [f361, f203])).
fof(f361, plain, ! [X0] : (r1(X0, sK57(X0)) | p107(X0) | ~ p106(X0) | ~ sP12(X0)), inference(cnf_transformation, [], [f138])).
fof(f20356, plain, ! [X13] : (sP12(sK55(sK52(X13))) | ~ p104(X13) | ~ sP14(X13) | p105(X13) | ~ r1(sK82, X13) | ~ sP13(sK52(X13))), inference(resolution, [], [f8006, f217])).
fof(f217, plain, ! [X0] : (~ sP41(X0) | sP12(X0)), inference(cnf_transformation, [], [f56])).
fof(f8006, plain, ! [X10] : (sP41(sK55(sK52(X10))) | ~ sP13(sK52(X10)) | ~ p104(X10) | ~ sP14(X10) | p105(X10) | ~ r1(sK82, X10)), inference(subsumption_resolution, [], [f8005, f352])).
fof(f8005, plain, ! [X10] : (~ sP13(sK52(X10)) | ~ p105(sK52(X10)) | sP41(sK55(sK52(X10))) | ~ p104(X10) | ~ sP14(X10) | p105(X10) | ~ r1(sK82, X10)), inference(subsumption_resolution, [], [f7946, f351])).
fof(f7946, plain, ! [X10] : (~ sP13(sK52(X10)) | p106(sK52(X10)) | ~ p105(sK52(X10)) | sP41(sK55(sK52(X10))) | ~ p104(X10) | ~ sP14(X10) | p105(X10) | ~ r1(sK82, X10)), inference(resolution, [], [f3510, f1006])).
fof(f3510, plain, ! [X0] : (~ r1(sK82, X0) | ~ sP13(X0) | p106(X0) | ~ p105(X0) | sP41(sK55(X0))), inference(resolution, [], [f1009, f466])).
fof(f6058, plain, ~ spl83_59, inference(avatar_contradiction_clause, [], [f6057])).
fof(f6057, plain, ($false | ~ spl83_59), inference(subsumption_resolution, [], [f6056, f490])).
fof(f6056, plain, (~ sP18(sK82) | ~ spl83_59), inference(subsumption_resolution, [], [f6055, f468])).
fof(f6055, plain, (~ p100(sK82) | ~ sP18(sK82) | ~ spl83_59), inference(subsumption_resolution, [], [f6054, f467])).
fof(f6054, plain, (p101(sK82) | ~ p100(sK82) | ~ sP18(sK82) | ~ spl83_59), inference(resolution, [], [f974, f319])).
fof(f319, plain, ! [X0] : (~ p102(sK44(X0)) | p101(X0) | ~ p100(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f108])).
fof(f974, plain, (p102(sK44(sK82)) | ~ spl83_59), inference(avatar_component_clause, [], [f972])).
fof(f4859, plain, spl83_60, inference(avatar_contradiction_clause, [], [f4858])).
fof(f4858, plain, ($false | spl83_60), inference(subsumption_resolution, [], [f4857, f490])).
fof(f4857, plain, (~ sP18(sK82) | spl83_60), inference(subsumption_resolution, [], [f4856, f468])).
fof(f4856, plain, (~ p100(sK82) | ~ sP18(sK82) | spl83_60), inference(subsumption_resolution, [], [f4855, f467])).
fof(f4855, plain, (p101(sK82) | ~ p100(sK82) | ~ sP18(sK82) | spl83_60), inference(resolution, [], [f978, f320])).
fof(f320, plain, ! [X0] : (p101(sK44(X0)) | p101(X0) | ~ p100(X0) | ~ sP18(X0)), inference(cnf_transformation, [], [f108])).
fof(f978, plain, (~ p101(sK44(sK82)) | spl83_60), inference(avatar_component_clause, [], [f977])).