fof(f63338, plain, $false, inference(avatar_sat_refutation, [], [f1927, f5157, f5212, f36487, f61775, f62069, f62177, f62282, f62292, f63337])).
fof(f63337, plain, (spl16_1422 | ~ spl16_3835 | ~ spl16_3837), inference(avatar_split_clause, [], [f63336, f62289, f62279, f20349])).
fof(f20349, plain, (spl16_1422 <=> ! [X3, X5, X4] : (~ r1(X3, sK14(sK6)) | ~ r1(sK4, X5) | ~ r1(X5, X3) | ~ r1(X3, X4) | p2(X4) | ~ p2(X3))), introduced(avatar_definition, [new_symbols(naming, [spl16_1422])])).
fof(f62279, plain, (spl16_3835 <=> p2(sK2(sK12(sK14(sK6))))), introduced(avatar_definition, [new_symbols(naming, [spl16_3835])])).
fof(f62289, plain, (spl16_3837 <=> r1(sK12(sK14(sK6)), sK2(sK12(sK14(sK6))))), introduced(avatar_definition, [new_symbols(naming, [spl16_3837])])).
fof(f63336, plain, (! [X4, X5, X3] : (~ r1(X3, sK14(sK6)) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ r1(sK4, X5)) | (~ spl16_3835 | ~ spl16_3837)), inference(subsumption_resolution, [], [f63320, f62281])).
fof(f62281, plain, (p2(sK2(sK12(sK14(sK6)))) | ~ spl16_3835), inference(avatar_component_clause, [], [f62279])).
fof(f63320, plain, (! [X4, X5, X3] : (~ p2(sK2(sK12(sK14(sK6)))) | ~ r1(X3, sK14(sK6)) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ r1(sK4, X5)) | ~ spl16_3837), inference(resolution, [], [f62291, f45])).
fof(f45, plain, ! [X21, X17, X20, X18, X16] : (~ r1(sK12(X18), X20) | ~ p2(X20) | ~ r1(X17, X18) | ~ p2(X17) | p2(X21) | ~ r1(X17, X21) | ~ r1(X16, X17) | ~ r1(sK4, X16)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, (! [X1] : ((p1(sK5(X1)) & r1(X1, sK5(X1))) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(sK4, X1)) & (! [X6] : ((~ p2(sK7(X6)) & r1(X6, sK7(X6))) | ~ r1(sK6, X6)) & r1(sK4, sK6)) & ((~ p1(sK9) & r1(sK8, sK9)) & (! [X11] : (p1(X11) | ~ r1(sK10, X11)) & r1(sK8, sK10)) & r1(sK4, sK8)) & ! [X12] : ((~ p1(sK11(X12)) & r1(X12, sK11(X12))) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(sK4, X12)) & ! [X16] : ((sP0(X16) & ! [X17] : (! [X18] : ((! [X20] : (~ p2(X20) | ~ r1(sK12(X18), X20)) & r1(X18, sK12(X18))) | ~ r1(X17, X18)) | ~ p2(X17) | ! [X21] : ((p2(X21) & r1(X21, sK13(X21))) | ~ r1(X17, X21)) | ~ r1(X16, X17)) & ! [X23] : ((p2(sK14(X23)) & r1(sK14(X23), sK15(X23)) & r1(X23, sK14(X23))) | ~ r1(X16, X23))) | ~ r1(sK4, X16))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6, sK7, sK8, sK9, sK10, sK11, sK12, sK13, sK14, sK15])], [f18, f30, f29, f28, f27, f26, f25, f24, f23, f22, f21, f20, f19])).
fof(f19, plain, (? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP0(X16) & ! [X17] : (! [X18] : (? [X19] : (! [X20] : (~ p2(X20) | ~ r1(X19, X20)) & r1(X18, X19)) | ~ r1(X17, X18)) | ~ p2(X17) | ! [X21] : ((p2(X21) & ? [X22] : r1(X21, X22)) | ~ r1(X17, X21)) | ~ r1(X16, X17)) & ! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X16, X23))) | ~ r1(X0, X16))) => (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(sK4, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(sK4, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(sK4, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(sK4, X12)) & ! [X16] : ((sP0(X16) & ! [X17] : (! [X18] : (? [X19] : (! [X20] : (~ p2(X20) | ~ r1(X19, X20)) & r1(X18, X19)) | ~ r1(X17, X18)) | ~ p2(X17) | ! [X21] : ((p2(X21) & ? [X22] : r1(X21, X22)) | ~ r1(X17, X21)) | ~ r1(X16, X17)) & ! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X16, X23))) | ~ r1(sK4, X16)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) => (p1(sK5(X1)) & r1(X1, sK5(X1)))), introduced(choice_axiom, [])).
fof(f21, plain, (? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(sK4, X5)) => (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(sK6, X6)) & r1(sK4, sK6))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) => (~ p2(sK7(X6)) & r1(X6, sK7(X6)))), introduced(choice_axiom, [])).
fof(f23, plain, (? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(sK4, X8)) => (? [X9] : (~ p1(X9) & r1(sK8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(sK8, X10)) & r1(sK4, sK8))), introduced(choice_axiom, [])).
fof(f24, plain, (? [X9] : (~ p1(X9) & r1(sK8, X9)) => (~ p1(sK9) & r1(sK8, sK9))), introduced(choice_axiom, [])).
fof(f25, plain, (? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(sK8, X10)) => (! [X11] : (p1(X11) | ~ r1(sK10, X11)) & r1(sK8, sK10))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) => (~ p1(sK11(X12)) & r1(X12, sK11(X12)))), introduced(choice_axiom, [])).
fof(f27, plain, ! [X18] : (? [X19] : (! [X20] : (~ p2(X20) | ~ r1(X19, X20)) & r1(X18, X19)) => (! [X20] : (~ p2(X20) | ~ r1(sK12(X18), X20)) & r1(X18, sK12(X18)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X21] : (? [X22] : r1(X21, X22) => r1(X21, sK13(X21))), introduced(choice_axiom, [])).
fof(f29, plain, ! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) => (p2(sK14(X23)) & ? [X25] : r1(sK14(X23), X25) & r1(X23, sK14(X23)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X23] : (? [X25] : r1(sK14(X23), X25) => r1(sK14(X23), sK15(X23))), introduced(choice_axiom, [])).
fof(f18, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP0(X16) & ! [X17] : (! [X18] : (? [X19] : (! [X20] : (~ p2(X20) | ~ r1(X19, X20)) & r1(X18, X19)) | ~ r1(X17, X18)) | ~ p2(X17) | ! [X21] : ((p2(X21) & ? [X22] : r1(X21, X22)) | ~ r1(X17, X21)) | ~ r1(X16, X17)) & ! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X16, X23))) | ~ r1(X0, X16))), inference(rectify, [], [f11])).
fof(f11, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP0(X16) & ! [X23] : (! [X24] : (? [X25] : (! [X26] : (~ p2(X26) | ~ r1(X25, X26)) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ p2(X23) | ! [X27] : ((p2(X27) & ? [X28] : r1(X27, X28)) | ~ r1(X23, X27)) | ~ r1(X16, X23)) & ! [X29] : (? [X30] : (p2(X30) & ? [X31] : r1(X30, X31) & r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16))), inference(definition_folding, [], [f9, e10])).
fof(f10, plain, ! [X16] : (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ! [X19] : (! [X20] : (? [X21] : (p2(X21) & ? [X22] : r1(X21, X22) & r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19)) | ~ sP0(X16)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X16] : (sP0(X16) <=> (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ! [X19] : (! [X20] : (? [X21] : (p2(X21) & ? [X22] : r1(X21, X22) & r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f9, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : (((? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ! [X19] : (! [X20] : (? [X21] : (p2(X21) & ? [X22] : r1(X21, X22) & r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) & ! [X23] : (! [X24] : (? [X25] : (! [X26] : (~ p2(X26) | ~ r1(X25, X26)) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ p2(X23) | ! [X27] : ((p2(X27) & ? [X28] : r1(X27, X28)) | ~ r1(X23, X27)) | ~ r1(X16, X23)) & ! [X29] : (? [X30] : (p2(X30) & ? [X31] : r1(X30, X31) & r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16))), inference(flattening, [], [f8])).
fof(f8, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : ((? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) & ! [X16] : (((? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ! [X19] : (! [X20] : (? [X21] : (p2(X21) & ? [X22] : r1(X21, X22) & r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) & ! [X23] : ((! [X24] : (? [X25] : (! [X26] : (~ p2(X26) | ~ r1(X25, X26)) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ p2(X23) | ! [X27] : ((p2(X27) & ? [X28] : r1(X27, X28)) | ~ r1(X23, X27))) | ~ r1(X16, X23)) & ! [X29] : (? [X30] : (p2(X30) & ? [X31] : r1(X30, X31) & r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ~ ! [X19] : (! [X20] : (~ ! [X21] : (~ p2(X21) | ! [X22] : ~ r1(X21, X22) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X23] : (~ (~ ! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p2(X26) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) & p2(X23) & ~ ! [X27] : (~ (~ p2(X27) | ! [X28] : ~ r1(X27, X28)) | ~ r1(X23, X27))) | ~ r1(X16, X23)) | ~ ! [X29] : (~ ! [X30] : (~ p2(X30) | ! [X31] : ~ r1(X30, X31) | ~ r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ~ ! [X19] : (! [X20] : (~ ! [X21] : (~ p2(X21) | ! [X22] : ~ r1(X21, X22) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X23] : (~ (~ ! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p2(X26) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) & p2(X23) & ~ ! [X27] : (~ (~ p2(X27) | ! [X28] : ~ r1(X27, X28)) | ~ r1(X23, X27))) | ~ r1(X16, X23)) | ~ ! [X29] : (~ ! [X30] : (~ p2(X30) | ! [X31] : ~ r1(X30, X31) | ~ r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16)) | ~ ! [X32] : (~ p4(X32) | ~ r1(X0, X32))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ~ ! [X19] : (! [X20] : (~ ! [X21] : (~ p2(X21) | ! [X22] : (p3(X22) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X23] : (~ (~ ! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p2(X26) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) & p2(X23) & ~ ! [X27] : (~ (~ p2(X27) | ! [X28] : (p3(X28) | ~ r1(X27, X28))) | ~ r1(X23, X27))) | ~ r1(X16, X23)) | ~ ! [X29] : (~ ! [X30] : (~ p2(X30) | ! [X31] : (p3(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16)) | ~ ! [X32] : (~ p4(X32) | ~ r1(X0, X32))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ~ ! [X19] : (! [X20] : (~ ! [X21] : (~ p2(X21) | ! [X22] : (p3(X22) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X23] : (~ (~ ! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p2(X26) | ~ r1(X25, X26)) | ~ r1(X24, X25)) | ~ r1(X23, X24)) & p2(X23) & ~ ! [X27] : (~ (~ p2(X27) | ! [X28] : (p3(X28) | ~ r1(X27, X28))) | ~ r1(X23, X27))) | ~ r1(X16, X23)) | ~ ! [X29] : (~ ! [X30] : (~ p2(X30) | ! [X31] : (p3(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X16, X29))) | ~ r1(X0, X16)) | ~ ! [X32] : (~ p4(X32) | ~ r1(X0, X32))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (p1(X0) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (~ ! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & p2(X0) & ~ ! [X1] : (~ (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0))) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ p4(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (p1(X0) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (~ ! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & p2(X0) & ~ ! [X1] : (~ (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0))) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ p4(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL670+1.001.p', main)).
fof(f62291, plain, (r1(sK12(sK14(sK6)), sK2(sK12(sK14(sK6)))) | ~ spl16_3837), inference(avatar_component_clause, [], [f62289])).
fof(f62292, plain, (spl16_88 | spl16_3837 | ~ spl16_115), inference(avatar_split_clause, [], [f5879, f1085, f62289, f956])).
fof(f956, plain, (spl16_88 <=> ! [X5] : (r1(X5, sK1(X5)) | ~ sP0(X5) | ~ r1(X5, sK14(sK6)))), introduced(avatar_definition, [new_symbols(naming, [spl16_88])])).
fof(f1085, plain, (spl16_115 <=> ! [X11] : (~ r1(sK14(sK6), X11) | r1(X11, sK12(X11)))), introduced(avatar_definition, [new_symbols(naming, [spl16_115])])).
fof(f5879, plain, (! [X0] : (r1(sK12(sK14(sK6)), sK2(sK12(sK14(sK6)))) | r1(X0, sK1(X0)) | ~ r1(X0, sK14(sK6)) | ~ sP0(X0)) | ~ spl16_115), inference(resolution, [], [f5226, f33])).
fof(f33, plain, ! [X4, X0, X3] : (~ r1(X3, X4) | r1(X4, sK2(X4)) | r1(X0, sK1(X0)) | ~ r1(X0, X3) | ~ sP0(X0)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : ((! [X2] : (~ p2(X2) | ~ r1(sK1(X0), X2)) & r1(X0, sK1(X0))) | ! [X3] : (! [X4] : ((p2(sK2(X4)) & r1(sK2(X4), sK3(X4)) & r1(X4, sK2(X4))) | ~ r1(X3, X4)) | ~ r1(X0, X3)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3])], [f13, f16, f15, f14])).
fof(f14, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) => (! [X2] : (~ p2(X2) | ~ r1(sK1(X0), X2)) & r1(X0, sK1(X0)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X4] : (? [X5] : (p2(X5) & ? [X6] : r1(X5, X6) & r1(X4, X5)) => (p2(sK2(X4)) & ? [X6] : r1(sK2(X4), X6) & r1(X4, sK2(X4)))), introduced(choice_axiom, [])).
fof(f16, plain, ! [X4] : (? [X6] : r1(sK2(X4), X6) => r1(sK2(X4), sK3(X4))), introduced(choice_axiom, [])).
fof(f13, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) | ! [X3] : (! [X4] : (? [X5] : (p2(X5) & ? [X6] : r1(X5, X6) & r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X0, X3)) | ~ sP0(X0)), inference(rectify, [], [f12])).
fof(f12, plain, ! [X16] : (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ! [X19] : (! [X20] : (? [X21] : (p2(X21) & ? [X22] : r1(X21, X22) & r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X16, X19)) | ~ sP0(X16)), inference(nnf_transformation, [], [f10])).
fof(f5226, plain, (r1(sK14(sK6), sK12(sK14(sK6))) | ~ spl16_115), inference(resolution, [], [f1086, f32])).
fof(f32, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL670+1.001.p', reflexivity)).
fof(f1086, plain, (! [X11] : (~ r1(sK14(sK6), X11) | r1(X11, sK12(X11))) | ~ spl16_115), inference(avatar_component_clause, [], [f1085])).
fof(f62282, plain, (spl16_88 | spl16_3835 | ~ spl16_115), inference(avatar_split_clause, [], [f5881, f1085, f62279, f956])).
fof(f5881, plain, (! [X2] : (p2(sK2(sK12(sK14(sK6)))) | r1(X2, sK1(X2)) | ~ r1(X2, sK14(sK6)) | ~ sP0(X2)) | ~ spl16_115), inference(resolution, [], [f5226, f35])).
fof(f35, plain, ! [X4, X0, X3] : (~ r1(X3, X4) | p2(sK2(X4)) | r1(X0, sK1(X0)) | ~ r1(X0, X3) | ~ sP0(X0)), inference(cnf_transformation, [], [f17])).
fof(f62177, plain, (spl16_79 | ~ spl16_88), inference(avatar_split_clause, [], [f62176, f956, f844])).
fof(f844, plain, (spl16_79 <=> r1(sK6, sK1(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl16_79])])).
fof(f62176, plain, (r1(sK6, sK1(sK6)) | ~ spl16_88), inference(subsumption_resolution, [], [f62175, f54])).
fof(f54, plain, r1(sK4, sK6), inference(cnf_transformation, [], [f31])).
fof(f62175, plain, (r1(sK6, sK1(sK6)) | ~ r1(sK4, sK6) | ~ spl16_88), inference(subsumption_resolution, [], [f2505, f60])).
fof(f60, plain, sP0(sK6), inference(resolution, [], [f46, f54])).
fof(f46, plain, ! [X16] : (~ r1(sK4, X16) | sP0(X16)), inference(cnf_transformation, [], [f31])).
fof(f2505, plain, (~ sP0(sK6) | r1(sK6, sK1(sK6)) | ~ r1(sK4, sK6) | ~ spl16_88), inference(resolution, [], [f957, f89])).
fof(f89, plain, ! [X2] : (r1(X2, sK14(X2)) | ~ r1(sK4, X2)), inference(resolution, [], [f39, f32])).
fof(f39, plain, ! [X23, X16] : (~ r1(sK4, X16) | ~ r1(X16, X23) | r1(X23, sK14(X23))), inference(cnf_transformation, [], [f31])).
fof(f957, plain, (! [X5] : (~ r1(X5, sK14(sK6)) | ~ sP0(X5) | r1(X5, sK1(X5))) | ~ spl16_88), inference(avatar_component_clause, [], [f956])).
fof(f62069, plain, ~ spl16_3812, inference(avatar_contradiction_clause, [], [f62068])).
fof(f62068, plain, ($false | ~ spl16_3812), inference(subsumption_resolution, [], [f62067, f32])).
fof(f62067, plain, (~ r1(sK6, sK6) | ~ spl16_3812), inference(subsumption_resolution, [], [f62060, f54])).
fof(f62060, plain, (~ r1(sK4, sK6) | ~ r1(sK6, sK6) | ~ spl16_3812), inference(resolution, [], [f61774, f88])).
fof(f88, plain, ! [X1] : (r1(X1, sK14(X1)) | ~ r1(sK6, X1)), inference(resolution, [], [f39, f54])).
fof(f61774, plain, (! [X6] : (~ r1(X6, sK14(sK6)) | ~ r1(sK4, X6)) | ~ spl16_3812), inference(avatar_component_clause, [], [f61773])).
fof(f61773, plain, (spl16_3812 <=> ! [X6] : (~ r1(sK4, X6) | ~ r1(X6, sK14(sK6)))), introduced(avatar_definition, [new_symbols(naming, [spl16_3812])])).
fof(f61775, plain, (spl16_116 | spl16_3812 | ~ spl16_1422), inference(avatar_split_clause, [], [f61771, f20349, f61773, f1088])).
fof(f1088, plain, (spl16_116 <=> ! [X10] : (p2(X10) | ~ r1(sK14(sK6), X10))), introduced(avatar_definition, [new_symbols(naming, [spl16_116])])).
fof(f61771, plain, (! [X6, X7] : (~ r1(sK4, X6) | ~ r1(X6, sK14(sK6)) | ~ r1(sK14(sK6), X7) | p2(X7)) | ~ spl16_1422), inference(subsumption_resolution, [], [f61769, f305])).
fof(f305, plain, p2(sK14(sK6)), inference(resolution, [], [f84, f32])).
fof(f84, plain, ! [X1] : (~ r1(sK6, X1) | p2(sK14(X1))), inference(resolution, [], [f41, f54])).
fof(f41, plain, ! [X23, X16] : (~ r1(sK4, X16) | ~ r1(X16, X23) | p2(sK14(X23))), inference(cnf_transformation, [], [f31])).
fof(f61769, plain, (! [X6, X7] : (~ r1(sK4, X6) | ~ r1(X6, sK14(sK6)) | ~ r1(sK14(sK6), X7) | p2(X7) | ~ p2(sK14(sK6))) | ~ spl16_1422), inference(resolution, [], [f20350, f32])).
fof(f20350, plain, (! [X4, X5, X3] : (~ r1(X3, sK14(sK6)) | ~ r1(sK4, X5) | ~ r1(X5, X3) | ~ r1(X3, X4) | p2(X4) | ~ p2(X3)) | ~ spl16_1422), inference(avatar_component_clause, [], [f20349])).
fof(f36487, plain, (spl16_1422 | ~ spl16_79 | ~ spl16_115 | ~ spl16_179), inference(avatar_split_clause, [], [f36486, f1566, f1085, f844, f20349])).
fof(f1566, plain, (spl16_179 <=> r1(sK6, sK14(sK6))), introduced(avatar_definition, [new_symbols(naming, [spl16_179])])).
fof(f36486, plain, (! [X4, X5, X3] : (~ r1(X3, sK14(sK6)) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ r1(sK4, X5)) | (~ spl16_79 | ~ spl16_115 | ~ spl16_179)), inference(subsumption_resolution, [], [f36466, f13000])).
fof(f13000, plain, (p2(sK2(sK12(sK14(sK6)))) | (~ spl16_79 | ~ spl16_115 | ~ spl16_179)), inference(resolution, [], [f8695, f5226])).
fof(f8695, plain, (! [X3] : (~ r1(sK14(sK6), X3) | p2(sK2(X3))) | (~ spl16_79 | ~ spl16_179)), inference(resolution, [], [f3191, f1567])).
fof(f1567, plain, (r1(sK6, sK14(sK6)) | ~ spl16_179), inference(avatar_component_clause, [], [f1566])).
fof(f3191, plain, (! [X0, X1] : (~ r1(sK6, X1) | ~ r1(X1, X0) | p2(sK2(X0))) | ~ spl16_79), inference(subsumption_resolution, [], [f3175, f60])).
fof(f3175, plain, (! [X0, X1] : (p2(sK2(X0)) | ~ r1(X1, X0) | ~ r1(sK6, X1) | ~ sP0(sK6)) | ~ spl16_79), inference(resolution, [], [f504, f846])).
fof(f846, plain, (r1(sK6, sK1(sK6)) | ~ spl16_79), inference(avatar_component_clause, [], [f844])).
fof(f504, plain, ! [X28, X29, X27] : (~ r1(sK6, sK1(X27)) | p2(sK2(X28)) | ~ r1(X29, X28) | ~ r1(X27, X29) | ~ sP0(X27)), inference(subsumption_resolution, [], [f497, f84])).
fof(f497, plain, ! [X28, X29, X27] : (~ r1(sK6, sK1(X27)) | ~ p2(sK14(sK1(X27))) | p2(sK2(X28)) | ~ r1(X29, X28) | ~ r1(X27, X29) | ~ sP0(X27)), inference(resolution, [], [f88, f38])).
fof(f38, plain, ! [X4, X2, X0, X3] : (~ r1(sK1(X0), X2) | ~ p2(X2) | p2(sK2(X4)) | ~ r1(X3, X4) | ~ r1(X0, X3) | ~ sP0(X0)), inference(cnf_transformation, [], [f17])).
fof(f36466, plain, (! [X4, X5, X3] : (~ p2(sK2(sK12(sK14(sK6)))) | ~ r1(X3, sK14(sK6)) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ r1(sK4, X5)) | (~ spl16_79 | ~ spl16_115 | ~ spl16_179)), inference(resolution, [], [f21511, f45])).
fof(f21511, plain, (r1(sK12(sK14(sK6)), sK2(sK12(sK14(sK6)))) | (~ spl16_79 | ~ spl16_115 | ~ spl16_179)), inference(resolution, [], [f9470, f5226])).
fof(f9470, plain, (! [X3] : (~ r1(sK14(sK6), X3) | r1(X3, sK2(X3))) | (~ spl16_79 | ~ spl16_179)), inference(resolution, [], [f3720, f1567])).
fof(f3720, plain, (! [X0, X1] : (~ r1(sK6, X1) | ~ r1(X1, X0) | r1(X0, sK2(X0))) | ~ spl16_79), inference(subsumption_resolution, [], [f3704, f60])).
fof(f3704, plain, (! [X0, X1] : (r1(X0, sK2(X0)) | ~ r1(X1, X0) | ~ r1(sK6, X1) | ~ sP0(sK6)) | ~ spl16_79), inference(resolution, [], [f503, f846])).
fof(f503, plain, ! [X26, X24, X25] : (~ r1(sK6, sK1(X24)) | r1(X25, sK2(X25)) | ~ r1(X26, X25) | ~ r1(X24, X26) | ~ sP0(X24)), inference(subsumption_resolution, [], [f496, f84])).
fof(f496, plain, ! [X26, X24, X25] : (~ r1(sK6, sK1(X24)) | ~ p2(sK14(sK1(X24))) | r1(X25, sK2(X25)) | ~ r1(X26, X25) | ~ r1(X24, X26) | ~ sP0(X24)), inference(resolution, [], [f88, f36])).
fof(f36, plain, ! [X4, X2, X0, X3] : (~ r1(sK1(X0), X2) | ~ p2(X2) | r1(X4, sK2(X4)) | ~ r1(X3, X4) | ~ r1(X0, X3) | ~ sP0(X0)), inference(cnf_transformation, [], [f17])).
fof(f5212, plain, (~ spl16_116 | ~ spl16_179), inference(avatar_contradiction_clause, [], [f5211])).
fof(f5211, plain, ($false | (~ spl16_116 | ~ spl16_179)), inference(subsumption_resolution, [], [f5210, f1567])).
fof(f5210, plain, (~ r1(sK6, sK14(sK6)) | (~ spl16_116 | ~ spl16_179)), inference(resolution, [], [f5192, f56])).
fof(f56, plain, ! [X6] : (~ p2(sK7(X6)) | ~ r1(sK6, X6)), inference(cnf_transformation, [], [f31])).
fof(f5192, plain, (p2(sK7(sK14(sK6))) | (~ spl16_116 | ~ spl16_179)), inference(subsumption_resolution, [], [f5187, f1567])).
fof(f5187, plain, (p2(sK7(sK14(sK6))) | ~ r1(sK6, sK14(sK6)) | ~ spl16_116), inference(resolution, [], [f1089, f55])).
fof(f55, plain, ! [X6] : (r1(X6, sK7(X6)) | ~ r1(sK6, X6)), inference(cnf_transformation, [], [f31])).
fof(f1089, plain, (! [X10] : (~ r1(sK14(sK6), X10) | p2(X10)) | ~ spl16_116), inference(avatar_component_clause, [], [f1088])).
fof(f5157, plain, (spl16_115 | spl16_116 | ~ spl16_179), inference(avatar_split_clause, [], [f5156, f1566, f1088, f1085])).
fof(f5156, plain, (! [X4, X5] : (p2(X4) | ~ r1(sK14(sK6), X4) | ~ r1(sK14(sK6), X5) | r1(X5, sK12(X5))) | ~ spl16_179), inference(subsumption_resolution, [], [f1953, f305])).
fof(f1953, plain, (! [X4, X5] : (~ p2(sK14(sK6)) | p2(X4) | ~ r1(sK14(sK6), X4) | ~ r1(sK14(sK6), X5) | r1(X5, sK12(X5))) | ~ spl16_179), inference(resolution, [], [f1567, f260])).
fof(f260, plain, ! [X4, X5, X3] : (~ r1(sK6, X3) | ~ p2(X3) | p2(X5) | ~ r1(X3, X5) | ~ r1(X3, X4) | r1(X4, sK12(X4))), inference(resolution, [], [f43, f54])).
fof(f43, plain, ! [X21, X17, X18, X16] : (~ r1(sK4, X16) | ~ r1(X17, X18) | ~ p2(X17) | p2(X21) | ~ r1(X17, X21) | ~ r1(X16, X17) | r1(X18, sK12(X18))), inference(cnf_transformation, [], [f31])).
fof(f1927, plain, spl16_179, inference(avatar_contradiction_clause, [], [f1926])).
fof(f1926, plain, ($false | spl16_179), inference(subsumption_resolution, [], [f1922, f32])).
fof(f1922, plain, (~ r1(sK6, sK6) | spl16_179), inference(resolution, [], [f1568, f88])).
fof(f1568, plain, (~ r1(sK6, sK14(sK6)) | spl16_179), inference(avatar_component_clause, [], [f1566])).