fof(f6344, plain, $false, inference(avatar_sat_refutation, [], [f325, f406, f496, f502, f1293, f1591, f1633, f4172, f4225, f4457, f4965, f5011, f5014, f5049, f5050, f5060, f5063, f5138, f5264, f6343])).
fof(f6343, plain, ~ spl32_122, inference(avatar_contradiction_clause, [], [f6342])).
fof(f6342, plain, ($false | ~ spl32_122), inference(subsumption_resolution, [], [f6331, f95])).
fof(f95, plain, r1(sK22, sK25), inference(cnf_transformation, [], [f53])).
fof(f53, plain, (! [X1] : ((! [X2] : ((~ p1(sK14(X2)) & r1(X2, sK14(X2))) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & ((! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1))) | ! [X9] : ((~ p1(sK16(X9)) & r1(X9, sK16(X9))) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : ((~ p1(sK17(X13)) & r1(X13, sK17(X13))) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ((! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12))) | ! [X20] : ((~ p1(sK19(X20)) & r1(X20, sK19(X20))) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & (((! [X25] : ((p1(sK23(X25)) & (~ p1(sK24(X25)) & r1(sK23(X25), sK24(X25))) & r1(X25, sK23(X25))) | p1(X25) | ~ r1(sK22, X25)) & (! [X29] : (p1(X29) | ~ r1(sK25, X29)) & r1(sK22, sK25)) & (~ p1(sK26) & r1(sK22, sK26)) & r1(sK21, sK22)) & r1(sK20, sK21)) & r1(sK13, sK20)) & ! [X31] : (! [X32] : (! [X33] : ((((! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(sK27(X33), X35)) & ~ p1(sK27(X33)) & r1(X33, sK27(X33))) | sP1(X33)) & ! [X37] : ((~ p1(sK28(X37)) & r1(X37, sK28(X37))) | ! [X39] : (! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X37, X39)) | ~ r1(X33, X37)) & sP2(X33) & sP3(X33)) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(sK13, X31)) & ! [X41] : (! [X42] : (! [X43] : (! [X44] : ((! [X45] : ((~ p1(sK29(X45)) & r1(X45, sK29(X45))) | ! [X47] : (! [X48] : (p1(X48) | ~ r1(X47, X48)) | ~ r1(X45, X47)) | ~ r1(X44, X45)) & ((! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(sK30(X44), X50)) & ~ p1(sK30(X44)) & r1(X44, sK30(X44))) | ! [X52] : ((~ p1(sK31(X52)) & r1(X52, sK31(X52))) | ~ r1(X44, X52)) | p1(X44))) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | ~ r1(sK13, X41))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15, sK16, sK17, sK18, sK19, sK20, sK21, sK22, sK23, sK24, sK25, sK26, sK27, sK28, sK29, sK30, sK31])], [f33, f52, f51, f50, f49, f48, f47, f46, f45, f44, f43, f42, f41, f40, f39, f38, f37, f36, f35, f34])).
fof(f34, plain, (? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(X0, X22)) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | sP1(X33)) & ! [X37] : (? [X38] : (~ p1(X38) & r1(X37, X38)) | ! [X39] : (! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X37, X39)) | ~ r1(X33, X37)) & sP2(X33) & sP3(X33)) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) & ! [X41] : (! [X42] : (! [X43] : (! [X44] : ((! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ! [X47] : (! [X48] : (p1(X48) | ~ r1(X47, X48)) | ~ r1(X45, X47)) | ~ r1(X44, X45)) & (? [X49] : (! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(X49, X50)) & ~ p1(X49) & r1(X44, X49)) | ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ~ r1(X44, X52)) | p1(X44))) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | ~ r1(X0, X41))) => (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(sK13, X22)) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | sP1(X33)) & ! [X37] : (? [X38] : (~ p1(X38) & r1(X37, X38)) | ! [X39] : (! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X37, X39)) | ~ r1(X33, X37)) & sP2(X33) & sP3(X33)) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(sK13, X31)) & ! [X41] : (! [X42] : (! [X43] : (! [X44] : ((! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ! [X47] : (! [X48] : (p1(X48) | ~ r1(X47, X48)) | ~ r1(X45, X47)) | ~ r1(X44, X45)) & (? [X49] : (! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(X49, X50)) & ~ p1(X49) & r1(X44, X49)) | ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ~ r1(X44, X52)) | p1(X44))) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | ~ r1(sK13, X41)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK14(X2)) & r1(X2, sK14(X2)))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) => (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) => (~ p1(sK16(X9)) & r1(X9, sK16(X9)))), introduced(choice_axiom, [])).
fof(f38, plain, ! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) => (~ p1(sK17(X13)) & r1(X13, sK17(X13)))), introduced(choice_axiom, [])).
fof(f39, plain, ! [X12] : (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) => (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) => (~ p1(sK19(X20)) & r1(X20, sK19(X20)))), introduced(choice_axiom, [])).
fof(f41, plain, (? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(sK13, X22)) => (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(sK20, X23)) & r1(sK13, sK20))), introduced(choice_axiom, [])).
fof(f42, plain, (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(sK20, X23)) => (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(sK21, X24)) & r1(sK20, sK21))), introduced(choice_axiom, [])).
fof(f43, plain, (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(sK21, X24)) => (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(sK22, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(sK22, X28)) & ? [X30] : (~ p1(X30) & r1(sK22, X30)) & r1(sK21, sK22))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) => (p1(sK23(X25)) & ? [X27] : (~ p1(X27) & r1(sK23(X25), X27)) & r1(X25, sK23(X25)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X25] : (? [X27] : (~ p1(X27) & r1(sK23(X25), X27)) => (~ p1(sK24(X25)) & r1(sK23(X25), sK24(X25)))), introduced(choice_axiom, [])).
fof(f46, plain, (? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(sK22, X28)) => (! [X29] : (p1(X29) | ~ r1(sK25, X29)) & r1(sK22, sK25))), introduced(choice_axiom, [])).
fof(f47, plain, (? [X30] : (~ p1(X30) & r1(sK22, X30)) => (~ p1(sK26) & r1(sK22, sK26))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X33] : (? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) => (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(sK27(X33), X35)) & ~ p1(sK27(X33)) & r1(X33, sK27(X33)))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X37] : (? [X38] : (~ p1(X38) & r1(X37, X38)) => (~ p1(sK28(X37)) & r1(X37, sK28(X37)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) => (~ p1(sK29(X45)) & r1(X45, sK29(X45)))), introduced(choice_axiom, [])).
fof(f51, plain, ! [X44] : (? [X49] : (! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(X49, X50)) & ~ p1(X49) & r1(X44, X49)) => (! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(sK30(X44), X50)) & ~ p1(sK30(X44)) & r1(X44, sK30(X44)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) => (~ p1(sK31(X52)) & r1(X52, sK31(X52)))), introduced(choice_axiom, [])).
fof(f33, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(X0, X22)) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | sP1(X33)) & ! [X37] : (? [X38] : (~ p1(X38) & r1(X37, X38)) | ! [X39] : (! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X37, X39)) | ~ r1(X33, X37)) & sP2(X33) & sP3(X33)) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) & ! [X41] : (! [X42] : (! [X43] : (! [X44] : ((! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ! [X47] : (! [X48] : (p1(X48) | ~ r1(X47, X48)) | ~ r1(X45, X47)) | ~ r1(X44, X45)) & (? [X49] : (! [X50] : (~ p1(X50) | ! [X51] : (p1(X51) | ~ r1(X50, X51)) | ~ r1(X49, X50)) & ~ p1(X49) & r1(X44, X49)) | ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ~ r1(X44, X52)) | p1(X44))) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | ~ r1(X0, X41))), inference(rectify, [], [f11])).
fof(f11, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(X0, X22)) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | sP1(X33)) & ! [X41] : (? [X42] : (~ p1(X42) & r1(X41, X42)) | ! [X43] : (! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X41, X43)) | ~ r1(X33, X41)) & sP2(X33) & sP3(X33)) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) & ! [X60] : (! [X61] : (! [X62] : (! [X63] : ((! [X64] : (? [X65] : (~ p1(X65) & r1(X64, X65)) | ! [X66] : (! [X67] : (p1(X67) | ~ r1(X66, X67)) | ~ r1(X64, X66)) | ~ r1(X63, X64)) & (? [X68] : (! [X69] : (~ p1(X69) | ! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X68, X69)) & ~ p1(X68) & r1(X63, X68)) | ! [X71] : (? [X72] : (~ p1(X72) & r1(X71, X72)) | ~ r1(X63, X71)) | p1(X63))) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X0, X60))), inference(definition_folding, [], [f6, e10, e9, e8, e7])).
fof(f7, plain, ! [X33] : (? [X50] : (! [X51] : ((p1(X51) & ? [X52] : (~ p1(X52) & r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) & p1(X50) & ? [X55] : (~ p1(X55) & r1(X50, X55)) & r1(X33, X50)) | ~ sP0(X33)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X33] : (sP0(X33) <=> ? [X50] : (! [X51] : ((p1(X51) & ? [X52] : (~ p1(X52) & r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) & p1(X50) & ? [X55] : (~ p1(X55) & r1(X50, X55)) & r1(X33, X50))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X33] : (! [X37] : (! [X38] : (? [X39] : (p1(X39) & ? [X40] : (~ p1(X40) & r1(X39, X40)) & r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37)) | ~ sP1(X33)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X33] : (sP1(X33) <=> ! [X37] : (! [X38] : (? [X39] : (p1(X39) & ? [X40] : (~ p1(X40) & r1(X39, X40)) & r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ! [X33] : (? [X45] : (! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) & ~ p1(X45) & r1(X33, X45)) | ! [X48] : (? [X49] : (~ p1(X49) & r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33) | ~ sP2(X33)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X33] : (sP2(X33) <=> (? [X45] : (! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) & ~ p1(X45) & r1(X33, X45)) | ! [X48] : (? [X49] : (~ p1(X49) & r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f10, plain, ! [X33] : (sP0(X33) | ! [X56] : (? [X57] : (p1(X57) & ? [X58] : (~ p1(X58) & r1(X57, X58)) & r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)) | ~ sP3(X33)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X33] : (sP3(X33) <=> (sP0(X33) | ! [X56] : (? [X57] : (p1(X57) & ? [X58] : (~ p1(X58) & r1(X57, X58)) & r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f6, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(X0, X22)) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | ! [X37] : (! [X38] : (? [X39] : (p1(X39) & ? [X40] : (~ p1(X40) & r1(X39, X40)) & r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37))) & ! [X41] : (? [X42] : (~ p1(X42) & r1(X41, X42)) | ! [X43] : (! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X41, X43)) | ~ r1(X33, X41)) & (? [X45] : (! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) & ~ p1(X45) & r1(X33, X45)) | ! [X48] : (? [X49] : (~ p1(X49) & r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33)) & (? [X50] : (! [X51] : ((p1(X51) & ? [X52] : (~ p1(X52) & r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) & p1(X50) & ? [X55] : (~ p1(X55) & r1(X50, X55)) & r1(X33, X50)) | ! [X56] : (? [X57] : (p1(X57) & ? [X58] : (~ p1(X58) & r1(X57, X58)) & r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)))) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) & ! [X60] : (! [X61] : (! [X62] : (! [X63] : ((! [X64] : (? [X65] : (~ p1(X65) & r1(X64, X65)) | ! [X66] : (! [X67] : (p1(X67) | ~ r1(X66, X67)) | ~ r1(X64, X66)) | ~ r1(X63, X64)) & (? [X68] : (! [X69] : (~ p1(X69) | ! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X68, X69)) & ~ p1(X68) & r1(X63, X68)) | ! [X71] : (? [X72] : (~ p1(X72) & r1(X71, X72)) | ~ r1(X63, X71)) | p1(X63))) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X0, X60))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : ((! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ? [X22] : (? [X23] : (? [X24] : (! [X25] : (? [X26] : (p1(X26) & ? [X27] : (~ p1(X27) & r1(X26, X27)) & r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) & ? [X28] : (! [X29] : (p1(X29) | ~ r1(X28, X29)) & r1(X24, X28)) & ? [X30] : (~ p1(X30) & r1(X24, X30)) & r1(X23, X24)) & r1(X22, X23)) & r1(X0, X22))) & ! [X31] : (! [X32] : (! [X33] : (((? [X34] : (! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & ~ p1(X34) & r1(X33, X34)) | ! [X37] : (! [X38] : (? [X39] : (p1(X39) & ? [X40] : (~ p1(X40) & r1(X39, X40)) & r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37))) & ! [X41] : (? [X42] : (~ p1(X42) & r1(X41, X42)) | ! [X43] : (! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X41, X43)) | ~ r1(X33, X41)) & (? [X45] : (! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) & ~ p1(X45) & r1(X33, X45)) | ! [X48] : (? [X49] : (~ p1(X49) & r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33)) & (? [X50] : (! [X51] : ((p1(X51) & ? [X52] : (~ p1(X52) & r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) & p1(X50) & ? [X55] : (~ p1(X55) & r1(X50, X55)) & r1(X33, X50)) | ! [X56] : (? [X57] : (p1(X57) & ? [X58] : (~ p1(X58) & r1(X57, X58)) & r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)))) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) & ! [X60] : (! [X61] : (! [X62] : (! [X63] : ((! [X64] : (? [X65] : (~ p1(X65) & r1(X64, X65)) | ! [X66] : (! [X67] : (p1(X67) | ~ r1(X66, X67)) | ~ r1(X64, X66)) | ~ r1(X63, X64)) & (? [X68] : (! [X69] : (~ p1(X69) | ! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X68, X69)) & ~ p1(X68) & r1(X63, X68)) | ! [X71] : (? [X72] : (~ p1(X72) & r1(X71, X72)) | ~ r1(X63, X71)) | p1(X63))) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X0, X60))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ~ ! [X22] : (! [X23] : (! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p1(X26) | ! [X27] : (p1(X27) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) | ! [X28] : (~ ! [X29] : (p1(X29) | ~ r1(X28, X29)) | ~ r1(X24, X28)) | ! [X30] : (p1(X30) | ~ r1(X24, X30)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22))) | ~ ! [X31] : (! [X32] : (! [X33] : (((~ ! [X34] : (~ ! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p1(X34) | ~ r1(X33, X34)) | ! [X37] : (! [X38] : (~ ! [X39] : (~ p1(X39) | ! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37))) & ! [X41] : (~ ! [X42] : (p1(X42) | ~ r1(X41, X42)) | ! [X43] : (! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X41, X43)) | ~ r1(X33, X41)) & (~ ! [X45] : (~ ! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) | p1(X45) | ~ r1(X33, X45)) | ! [X48] : (~ ! [X49] : (p1(X49) | ~ r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33)) & (~ ! [X50] : (~ ! [X51] : (~ (~ p1(X51) | ! [X52] : (p1(X52) | ~ r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) | ~ p1(X50) | ! [X55] : (p1(X55) | ~ r1(X50, X55)) | ~ r1(X33, X50)) | ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)))) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) | ~ ! [X60] : (! [X61] : (! [X62] : (! [X63] : ((! [X64] : (~ ! [X65] : (p1(X65) | ~ r1(X64, X65)) | ! [X66] : (! [X67] : (p1(X67) | ~ r1(X66, X67)) | ~ r1(X64, X66)) | ~ r1(X63, X64)) & (~ ! [X68] : (~ ! [X69] : (~ p1(X69) | ! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X68, X69)) | p1(X68) | ~ r1(X63, X68)) | ! [X71] : (~ ! [X72] : (p1(X72) | ~ r1(X71, X72)) | ~ r1(X63, X71)) | p1(X63))) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X0, X60))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ~ ! [X22] : (! [X23] : (! [X24] : (~ ! [X25] : (~ ! [X26] : (~ p1(X26) | ! [X27] : (p1(X27) | ~ r1(X26, X27)) | ~ r1(X25, X26)) | p1(X25) | ~ r1(X24, X25)) | ! [X28] : (~ ! [X29] : (p1(X29) | ~ r1(X28, X29)) | ~ r1(X24, X28)) | ! [X30] : (p1(X30) | ~ r1(X24, X30)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22))) | ~ ! [X31] : (! [X32] : (! [X33] : (((~ ! [X34] : (~ ! [X35] : (~ p1(X35) | ! [X36] : (p1(X36) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | p1(X34) | ~ r1(X33, X34)) | ! [X37] : (! [X38] : (~ ! [X39] : (~ p1(X39) | ! [X40] : (p1(X40) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37))) & ! [X41] : (~ ! [X42] : (p1(X42) | ~ r1(X41, X42)) | ! [X43] : (! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X41, X43)) | ~ r1(X33, X41)) & (~ ! [X45] : (~ ! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) | p1(X45) | ~ r1(X33, X45)) | ! [X48] : (~ ! [X49] : (p1(X49) | ~ r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33)) & (~ ! [X50] : (~ ! [X51] : (~ (~ p1(X51) | ! [X52] : (p1(X52) | ~ r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) | ~ p1(X50) | ! [X55] : (p1(X55) | ~ r1(X50, X55)) | ~ r1(X33, X50)) | ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)))) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X0, X31)) | ~ ! [X60] : (! [X61] : (! [X62] : (! [X63] : ((! [X64] : (~ ! [X65] : (p1(X65) | ~ r1(X64, X65)) | ! [X66] : (! [X67] : (p1(X67) | ~ r1(X66, X67)) | ~ r1(X64, X66)) | ~ r1(X63, X64)) & (~ ! [X68] : (~ ! [X69] : (~ p1(X69) | ! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X68, X69)) | p1(X68) | ~ r1(X63, X68)) | ! [X71] : (~ ! [X72] : (p1(X72) | ~ r1(X71, X72)) | ~ r1(X63, X71)) | p1(X63))) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X0, X60))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (((~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) & ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1)) & (~ ! [X0] : (~ ! [X1] : (~ (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0))) | ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (((~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) & ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1)) & (~ ! [X0] : (~ ! [X1] : (~ (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0))) | ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL640+1.005.p', main)).
fof(f6331, plain, (~ r1(sK22, sK25) | ~ spl32_122), inference(resolution, [], [f6306, f92])).
fof(f92, plain, r1(sK21, sK22), inference(cnf_transformation, [], [f53])).
fof(f6306, plain, (! [X0] : (~ r1(sK21, X0) | ~ r1(X0, sK25)) | ~ spl32_122), inference(resolution, [], [f6284, f91])).
fof(f91, plain, r1(sK20, sK21), inference(cnf_transformation, [], [f53])).
fof(f6284, plain, (! [X0, X1] : (~ r1(sK20, X1) | ~ r1(X0, sK25) | ~ r1(X1, X0)) | ~ spl32_122), inference(resolution, [], [f1632, f90])).
fof(f90, plain, r1(sK13, sK20), inference(cnf_transformation, [], [f53])).
fof(f1632, plain, (! [X4, X2, X3] : (~ r1(sK13, X4) | ~ r1(X2, sK25) | ~ r1(X4, X3) | ~ r1(X3, X2)) | ~ spl32_122), inference(avatar_component_clause, [], [f1631])).
fof(f1631, plain, (spl32_122 <=> ! [X3, X2, X4] : (~ r1(X2, sK25) | ~ r1(sK13, X4) | ~ r1(X4, X3) | ~ r1(X3, X2))), introduced(avatar_definition, [new_symbols(naming, [spl32_122])])).
fof(f5264, plain, (spl32_36 | ~ spl32_418), inference(avatar_contradiction_clause, [], [f5263])).
fof(f5263, plain, ($false | (spl32_36 | ~ spl32_418)), inference(subsumption_resolution, [], [f5147, f95])).
fof(f5147, plain, (~ r1(sK22, sK25) | (spl32_36 | ~ spl32_418)), inference(resolution, [], [f4498, f323])).
fof(f323, plain, (~ r1(sK25, sK7(sK25)) | spl32_36), inference(avatar_component_clause, [], [f322])).
fof(f322, plain, (spl32_36 <=> r1(sK25, sK7(sK25))), introduced(avatar_definition, [new_symbols(naming, [spl32_36])])).
fof(f4498, plain, (! [X1] : (r1(X1, sK7(X1)) | ~ r1(sK22, X1)) | ~ spl32_418), inference(avatar_component_clause, [], [f4497])).
fof(f4497, plain, (spl32_418 <=> ! [X1] : (r1(X1, sK7(X1)) | ~ r1(sK22, X1))), introduced(avatar_definition, [new_symbols(naming, [spl32_418])])).
fof(f5138, plain, (~ spl32_36 | ~ spl32_420), inference(avatar_contradiction_clause, [], [f5137])).
fof(f5137, plain, ($false | (~ spl32_36 | ~ spl32_420)), inference(subsumption_resolution, [], [f5134, f95])).
fof(f5134, plain, (~ r1(sK22, sK25) | (~ spl32_36 | ~ spl32_420)), inference(resolution, [], [f4507, f4529])).
fof(f4529, plain, (p1(sK7(sK25)) | ~ spl32_36), inference(resolution, [], [f324, f96])).
fof(f96, plain, ! [X29] : (~ r1(sK25, X29) | p1(X29)), inference(cnf_transformation, [], [f53])).
fof(f324, plain, (r1(sK25, sK7(sK25)) | ~ spl32_36), inference(avatar_component_clause, [], [f322])).
fof(f4507, plain, (! [X1] : (~ p1(sK7(X1)) | ~ r1(sK22, X1)) | ~ spl32_420), inference(avatar_component_clause, [], [f4506])).
fof(f4506, plain, (spl32_420 <=> ! [X1] : (~ p1(sK7(X1)) | ~ r1(sK22, X1))), introduced(avatar_definition, [new_symbols(naming, [spl32_420])])).
fof(f5063, plain, (spl32_418 | spl32_20 | ~ spl32_33 | ~ spl32_44), inference(avatar_split_clause, [], [f5062, f427, f309, f220, f4497])).
fof(f220, plain, (spl32_20 <=> p1(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl32_20])])).
fof(f309, plain, (spl32_33 <=> sP2(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl32_33])])).
fof(f427, plain, (spl32_44 <=> p1(sK6(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl32_44])])).
fof(f5062, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK22, X0)) | (spl32_20 | ~ spl32_33 | ~ spl32_44)), inference(subsumption_resolution, [], [f5061, f310])).
fof(f310, plain, (sP2(sK22) | ~ spl32_33), inference(avatar_component_clause, [], [f309])).
fof(f5061, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK22, X0) | ~ sP2(sK22)) | (spl32_20 | ~ spl32_44)), inference(subsumption_resolution, [], [f5056, f221])).
fof(f221, plain, (~ p1(sK22) | spl32_20), inference(avatar_component_clause, [], [f220])).
fof(f5056, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK22, X0) | p1(sK22) | ~ sP2(sK22)) | ~ spl32_44), inference(resolution, [], [f429, f60])).
fof(f60, plain, ! [X4, X0] : (~ p1(sK6(X0)) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : ((! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0))) | ! [X4] : ((~ p1(sK7(X4)) & r1(X4, sK7(X4))) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7])], [f18, f20, f19])).
fof(f19, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) => (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) => (~ p1(sK7(X4)) & r1(X4, sK7(X4)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) | ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X33] : (? [X45] : (! [X46] : (~ p1(X46) | ! [X47] : (p1(X47) | ~ r1(X46, X47)) | ~ r1(X45, X46)) & ~ p1(X45) & r1(X33, X45)) | ! [X48] : (? [X49] : (~ p1(X49) & r1(X48, X49)) | ~ r1(X33, X48)) | p1(X33) | ~ sP2(X33)), inference(nnf_transformation, [], [f9])).
fof(f429, plain, (p1(sK6(sK22)) | ~ spl32_44), inference(avatar_component_clause, [], [f427])).
fof(f5060, plain, (spl32_20 | ~ spl32_33 | ~ spl32_36 | ~ spl32_44), inference(avatar_contradiction_clause, [], [f5059])).
fof(f5059, plain, ($false | (spl32_20 | ~ spl32_33 | ~ spl32_36 | ~ spl32_44)), inference(subsumption_resolution, [], [f5058, f310])).
fof(f5058, plain, (~ sP2(sK22) | (spl32_20 | ~ spl32_36 | ~ spl32_44)), inference(subsumption_resolution, [], [f5057, f221])).
fof(f5057, plain, (p1(sK22) | ~ sP2(sK22) | (~ spl32_36 | ~ spl32_44)), inference(subsumption_resolution, [], [f5055, f95])).
fof(f5055, plain, (~ r1(sK22, sK25) | p1(sK22) | ~ sP2(sK22) | (~ spl32_36 | ~ spl32_44)), inference(resolution, [], [f429, f4540])).
fof(f4540, plain, (! [X1] : (~ p1(sK6(X1)) | ~ r1(X1, sK25) | p1(X1) | ~ sP2(X1)) | ~ spl32_36), inference(resolution, [], [f4529, f61])).
fof(f61, plain, ! [X4, X0] : (~ p1(sK7(X4)) | ~ p1(sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f5050, plain, (spl32_43 | spl32_44 | ~ spl32_34), inference(avatar_split_clause, [], [f5039, f313, f427, f423])).
fof(f423, plain, (spl32_43 <=> p1(sK23(sK6(sK22)))), introduced(avatar_definition, [new_symbols(naming, [spl32_43])])).
fof(f313, plain, (spl32_34 <=> r1(sK22, sK6(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl32_34])])).
fof(f5039, plain, (p1(sK6(sK22)) | p1(sK23(sK6(sK22))) | ~ spl32_34), inference(resolution, [], [f315, f100])).
fof(f100, plain, ! [X25] : (~ r1(sK22, X25) | p1(X25) | p1(sK23(X25))), inference(cnf_transformation, [], [f53])).
fof(f315, plain, (r1(sK22, sK6(sK22)) | ~ spl32_34), inference(avatar_component_clause, [], [f313])).
fof(f5049, plain, (spl32_44 | ~ spl32_34 | ~ spl32_449), inference(avatar_split_clause, [], [f5048, f5008, f313, f427])).
fof(f5008, plain, (spl32_449 <=> p1(sK24(sK6(sK22)))), introduced(avatar_definition, [new_symbols(naming, [spl32_449])])).
fof(f5048, plain, (p1(sK6(sK22)) | (~ spl32_34 | ~ spl32_449)), inference(subsumption_resolution, [], [f5047, f315])).
fof(f5047, plain, (p1(sK6(sK22)) | ~ r1(sK22, sK6(sK22)) | ~ spl32_449), inference(resolution, [], [f5010, f99])).
fof(f99, plain, ! [X25] : (~ p1(sK24(X25)) | p1(X25) | ~ r1(sK22, X25)), inference(cnf_transformation, [], [f53])).
fof(f5010, plain, (p1(sK24(sK6(sK22))) | ~ spl32_449), inference(avatar_component_clause, [], [f5008])).
fof(f5014, plain, (~ spl32_34 | spl32_418 | spl32_449 | spl32_20 | ~ spl32_33 | ~ spl32_43), inference(avatar_split_clause, [], [f5013, f423, f309, f220, f5008, f4497, f313])).
fof(f5013, plain, (! [X0] : (p1(sK24(sK6(sK22))) | r1(X0, sK7(X0)) | ~ r1(sK22, X0) | ~ r1(sK22, sK6(sK22))) | (spl32_20 | ~ spl32_33 | ~ spl32_43)), inference(subsumption_resolution, [], [f5012, f310])).
fof(f5012, plain, (! [X0] : (p1(sK24(sK6(sK22))) | r1(X0, sK7(X0)) | ~ r1(sK22, X0) | ~ sP2(sK22) | ~ r1(sK22, sK6(sK22))) | (spl32_20 | ~ spl32_43)), inference(subsumption_resolution, [], [f2058, f221])).
fof(f2058, plain, (! [X0] : (p1(sK24(sK6(sK22))) | r1(X0, sK7(X0)) | ~ r1(sK22, X0) | p1(sK22) | ~ sP2(sK22) | ~ r1(sK22, sK6(sK22))) | ~ spl32_43), inference(resolution, [], [f1239, f425])).
fof(f425, plain, (p1(sK23(sK6(sK22))) | ~ spl32_43), inference(avatar_component_clause, [], [f423])).
fof(f1239, plain, ! [X0, X1] : (~ p1(sK23(sK6(X0))) | p1(sK24(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0))), inference(subsumption_resolution, [], [f1238, f60])).
fof(f1238, plain, ! [X0, X1] : (p1(sK24(sK6(X0))) | ~ p1(sK23(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f1231])).
fof(f1231, plain, ! [X0, X1] : (p1(sK24(sK6(X0))) | ~ p1(sK23(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK22, sK6(X0))), inference(resolution, [], [f536, f98])).
fof(f98, plain, ! [X25] : (r1(sK23(X25), sK24(X25)) | p1(X25) | ~ r1(sK22, X25)), inference(cnf_transformation, [], [f53])).
fof(f536, plain, ! [X4, X5, X3] : (~ r1(sK23(sK6(X4)), X3) | p1(X3) | ~ p1(sK23(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK22, sK6(X4))), inference(subsumption_resolution, [], [f534, f60])).
fof(f534, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK23(sK6(X4)), X3) | ~ p1(sK23(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK22, sK6(X4))), inference(resolution, [], [f62, f97])).
fof(f97, plain, ! [X25] : (r1(X25, sK23(X25)) | p1(X25) | ~ r1(sK22, X25)), inference(cnf_transformation, [], [f53])).
fof(f62, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f5011, plain, (~ spl32_34 | spl32_420 | spl32_449 | spl32_20 | ~ spl32_33 | ~ spl32_43), inference(avatar_split_clause, [], [f5006, f423, f309, f220, f5008, f4506, f313])).
fof(f5006, plain, (! [X0] : (p1(sK24(sK6(sK22))) | ~ p1(sK7(X0)) | ~ r1(sK22, X0) | ~ r1(sK22, sK6(sK22))) | (spl32_20 | ~ spl32_33 | ~ spl32_43)), inference(subsumption_resolution, [], [f5005, f310])).
fof(f5005, plain, (! [X0] : (p1(sK24(sK6(sK22))) | ~ p1(sK7(X0)) | ~ r1(sK22, X0) | ~ sP2(sK22) | ~ r1(sK22, sK6(sK22))) | (spl32_20 | ~ spl32_43)), inference(subsumption_resolution, [], [f1994, f221])).
fof(f1994, plain, (! [X0] : (p1(sK24(sK6(sK22))) | ~ p1(sK7(X0)) | ~ r1(sK22, X0) | p1(sK22) | ~ sP2(sK22) | ~ r1(sK22, sK6(sK22))) | ~ spl32_43), inference(resolution, [], [f1132, f425])).
fof(f1132, plain, ! [X0, X1] : (~ p1(sK23(sK6(X0))) | p1(sK24(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0))), inference(subsumption_resolution, [], [f1131, f61])).
fof(f1131, plain, ! [X0, X1] : (p1(sK24(sK6(X0))) | ~ p1(sK23(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f1124])).
fof(f1124, plain, ! [X0, X1] : (p1(sK24(sK6(X0))) | ~ p1(sK23(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK22, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK22, sK6(X0))), inference(resolution, [], [f513, f98])).
fof(f513, plain, ! [X4, X5, X3] : (~ r1(sK23(sK6(X4)), X3) | p1(X3) | ~ p1(sK23(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK22, sK6(X4))), inference(subsumption_resolution, [], [f511, f61])).
fof(f511, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK23(sK6(X4)), X3) | ~ p1(sK23(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK22, sK6(X4))), inference(resolution, [], [f63, f97])).
fof(f63, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | ~ p1(sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f4965, plain, (spl32_20 | ~ spl32_33 | spl32_34 | ~ spl32_36), inference(avatar_contradiction_clause, [], [f4964])).
fof(f4964, plain, ($false | (spl32_20 | ~ spl32_33 | spl32_34 | ~ spl32_36)), inference(subsumption_resolution, [], [f4963, f310])).
fof(f4963, plain, (~ sP2(sK22) | (spl32_20 | spl32_34 | ~ spl32_36)), inference(subsumption_resolution, [], [f4962, f221])).
fof(f4962, plain, (p1(sK22) | ~ sP2(sK22) | (spl32_34 | ~ spl32_36)), inference(subsumption_resolution, [], [f4840, f95])).
fof(f4840, plain, (~ r1(sK22, sK25) | p1(sK22) | ~ sP2(sK22) | (spl32_34 | ~ spl32_36)), inference(resolution, [], [f4539, f314])).
fof(f314, plain, (~ r1(sK22, sK6(sK22)) | spl32_34), inference(avatar_component_clause, [], [f313])).
fof(f4539, plain, (! [X0] : (r1(X0, sK6(X0)) | ~ r1(X0, sK25) | p1(X0) | ~ sP2(X0)) | ~ spl32_36), inference(resolution, [], [f4529, f59])).
fof(f59, plain, ! [X4, X0] : (~ p1(sK7(X4)) | r1(X0, sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f4457, plain, ~ spl32_376, inference(avatar_contradiction_clause, [], [f4456])).
fof(f4456, plain, ($false | ~ spl32_376), inference(subsumption_resolution, [], [f4446, f92])).
fof(f4446, plain, (~ r1(sK21, sK22) | ~ spl32_376), inference(resolution, [], [f4232, f91])).
fof(f4232, plain, (! [X0] : (~ r1(sK20, X0) | ~ r1(X0, sK22)) | ~ spl32_376), inference(resolution, [], [f4224, f90])).
fof(f4224, plain, (! [X0, X1] : (~ r1(sK13, X1) | ~ r1(X0, sK22) | ~ r1(X1, X0)) | ~ spl32_376), inference(avatar_component_clause, [], [f4223])).
fof(f4223, plain, (spl32_376 <=> ! [X1, X0] : (~ r1(X0, sK22) | ~ r1(sK13, X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl32_376])])).
fof(f4225, plain, (spl32_74 | spl32_376), inference(avatar_split_clause, [], [f4221, f4223, f693])).
fof(f693, plain, (spl32_74 <=> sP1(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl32_74])])).
fof(f4221, plain, ! [X0, X1] : (~ r1(X0, sK22) | ~ r1(X1, X0) | ~ r1(sK13, X1) | sP1(sK22)), inference(subsumption_resolution, [], [f1497, f92])).
fof(f1497, plain, ! [X0, X1] : (~ r1(X0, sK22) | ~ r1(X1, X0) | ~ r1(sK13, X1) | sP1(sK22) | ~ r1(sK21, sK22)), inference(duplicate_literal_removal, [], [f1496])).
fof(f1496, plain, ! [X0, X1] : (~ r1(X0, sK22) | ~ r1(X1, X0) | ~ r1(sK13, X1) | sP1(sK22) | sP1(sK22) | ~ r1(sK21, sK22)), inference(resolution, [], [f1041, f569])).
fof(f569, plain, ! [X0] : (r1(X0, sK27(X0)) | sP1(X0) | ~ r1(sK21, X0)), inference(resolution, [], [f247, f91])).
fof(f247, plain, ! [X0, X1] : (~ r1(sK20, X1) | ~ r1(X1, X0) | sP1(X0) | r1(X0, sK27(X0))), inference(resolution, [], [f87, f90])).
fof(f87, plain, ! [X33, X31, X32] : (~ r1(sK13, X31) | sP1(X33) | ~ r1(X32, X33) | ~ r1(X31, X32) | r1(X33, sK27(X33))), inference(cnf_transformation, [], [f53])).
fof(f1041, plain, ! [X2, X0, X1] : (~ r1(sK22, sK27(X0)) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(sK13, X2) | sP1(X0)), inference(subsumption_resolution, [], [f1040, f88])).
fof(f88, plain, ! [X33, X31, X32] : (~ p1(sK27(X33)) | sP1(X33) | ~ r1(X32, X33) | ~ r1(X31, X32) | ~ r1(sK13, X31)), inference(cnf_transformation, [], [f53])).
fof(f1040, plain, ! [X2, X0, X1] : (sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(sK13, X2) | ~ r1(sK22, sK27(X0)) | p1(sK27(X0))), inference(subsumption_resolution, [], [f1039, f99])).
fof(f1039, plain, ! [X2, X0, X1] : (p1(sK24(sK27(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(sK13, X2) | ~ r1(sK22, sK27(X0)) | p1(sK27(X0))), inference(duplicate_literal_removal, [], [f1032])).
fof(f1032, plain, ! [X2, X0, X1] : (p1(sK24(sK27(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(sK13, X2) | ~ r1(sK22, sK27(X0)) | p1(sK27(X0)) | ~ r1(sK22, sK27(X0))), inference(resolution, [], [f526, f98])).
fof(f526, plain, ! [X6, X4, X7, X5] : (~ r1(sK23(sK27(X5)), X4) | p1(X4) | sP1(X5) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | ~ r1(sK22, sK27(X5))), inference(subsumption_resolution, [], [f525, f88])).
fof(f525, plain, ! [X6, X4, X7, X5] : (p1(X4) | ~ r1(sK23(sK27(X5)), X4) | sP1(X5) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | p1(sK27(X5)) | ~ r1(sK22, sK27(X5))), inference(subsumption_resolution, [], [f523, f100])).
fof(f523, plain, ! [X6, X4, X7, X5] : (p1(X4) | ~ r1(sK23(sK27(X5)), X4) | ~ p1(sK23(sK27(X5))) | sP1(X5) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | p1(sK27(X5)) | ~ r1(sK22, sK27(X5))), inference(resolution, [], [f89, f97])).
fof(f89, plain, ! [X35, X33, X31, X36, X32] : (~ r1(sK27(X33), X35) | p1(X36) | ~ r1(X35, X36) | ~ p1(X35) | sP1(X33) | ~ r1(X32, X33) | ~ r1(X31, X32) | ~ r1(sK13, X31)), inference(cnf_transformation, [], [f53])).
fof(f4172, plain, (~ spl32_20 | spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119), inference(avatar_contradiction_clause, [], [f4171])).
fof(f4171, plain, ($false | (~ spl32_20 | spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f4159, f94])).
fof(f94, plain, ~ p1(sK26), inference(cnf_transformation, [], [f53])).
fof(f4159, plain, (p1(sK26) | (~ spl32_20 | spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(resolution, [], [f4145, f93])).
fof(f93, plain, r1(sK22, sK26), inference(cnf_transformation, [], [f53])).
fof(f4145, plain, (! [X0] : (~ r1(sK22, X0) | p1(X0)) | (~ spl32_20 | spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f4144, f251])).
fof(f251, plain, sP3(sK22), inference(resolution, [], [f159, f92])).
fof(f159, plain, ! [X0] : (~ r1(sK21, X0) | sP3(X0)), inference(resolution, [], [f133, f91])).
fof(f133, plain, ! [X0, X1] : (~ r1(sK20, X0) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f83, f90])).
fof(f83, plain, ! [X33, X31, X32] : (~ r1(sK13, X31) | ~ r1(X32, X33) | ~ r1(X31, X32) | sP3(X33)), inference(cnf_transformation, [], [f53])).
fof(f4144, plain, (! [X0] : (p1(X0) | ~ r1(sK22, X0) | ~ sP3(sK22)) | (~ spl32_20 | spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f4143, f222])).
fof(f222, plain, (p1(sK22) | ~ spl32_20), inference(avatar_component_clause, [], [f220])).
fof(f4143, plain, (! [X0] : (~ p1(sK22) | p1(X0) | ~ r1(sK22, X0) | ~ sP3(sK22)) | (spl32_41 | ~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f4142, f374])).
fof(f374, plain, (~ sP0(sK22) | spl32_41), inference(avatar_component_clause, [], [f373])).
fof(f373, plain, (spl32_41 <=> sP0(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl32_41])])).
fof(f4142, plain, (! [X0] : (sP0(sK22) | ~ p1(sK22) | p1(X0) | ~ r1(sK22, X0) | ~ sP3(sK22)) | (~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(resolution, [], [f2059, f95])).
fof(f2059, plain, (! [X0, X1] : (~ r1(X0, sK25) | sP0(X0) | ~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ sP3(X0)) | (~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(resolution, [], [f2044, f56])).
fof(f56, plain, ! [X4, X0, X1] : (~ p1(sK5(X1)) | sP0(X0) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | ~ r1(X0, X4) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (sP0(X0) | ! [X1] : ((p1(sK4(X1)) & (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1))) & r1(X1, sK4(X1))) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f13, f15, f14])).
fof(f14, plain, ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) => (p1(sK4(X1)) & ? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) & r1(X1, sK4(X1)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X1] : (? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) => (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1)))), introduced(choice_axiom, [])).
fof(f13, plain, ! [X0] : (sP0(X0) | ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(rectify, [], [f12])).
fof(f12, plain, ! [X33] : (sP0(X33) | ! [X56] : (? [X57] : (p1(X57) & ? [X58] : (~ p1(X58) & r1(X57, X58)) & r1(X56, X57)) | ~ r1(X33, X56)) | ~ p1(X33) | ! [X59] : (p1(X59) | ~ r1(X33, X59)) | ~ sP3(X33)), inference(nnf_transformation, [], [f10])).
fof(f2044, plain, (p1(sK5(sK25)) | (~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f2036, f95])).
fof(f2036, plain, (p1(sK5(sK25)) | ~ r1(sK22, sK25) | (~ spl32_53 | ~ spl32_54 | ~ spl32_119)), inference(resolution, [], [f1985, f495])).
fof(f495, plain, (! [X9] : (r1(sK4(X9), sK5(X9)) | ~ r1(sK22, X9)) | ~ spl32_53), inference(avatar_component_clause, [], [f494])).
fof(f494, plain, (spl32_53 <=> ! [X9] : (r1(sK4(X9), sK5(X9)) | ~ r1(sK22, X9))), introduced(avatar_definition, [new_symbols(naming, [spl32_53])])).
fof(f1985, plain, (! [X1] : (~ r1(sK4(sK25), X1) | p1(X1)) | (~ spl32_54 | ~ spl32_119)), inference(subsumption_resolution, [], [f1979, f95])).
fof(f1979, plain, (! [X1] : (p1(X1) | ~ r1(sK4(sK25), X1) | ~ r1(sK22, sK25)) | (~ spl32_54 | ~ spl32_119)), inference(resolution, [], [f1590, f501])).
fof(f501, plain, (! [X9] : (r1(X9, sK4(X9)) | ~ r1(sK22, X9)) | ~ spl32_54), inference(avatar_component_clause, [], [f500])).
fof(f500, plain, (spl32_54 <=> ! [X9] : (r1(X9, sK4(X9)) | ~ r1(sK22, X9))), introduced(avatar_definition, [new_symbols(naming, [spl32_54])])).
fof(f1590, plain, (! [X2, X3] : (~ r1(sK25, X2) | p1(X3) | ~ r1(X2, X3)) | ~ spl32_119), inference(avatar_component_clause, [], [f1589])).
fof(f1589, plain, (spl32_119 <=> ! [X3, X2] : (~ r1(sK25, X2) | p1(X3) | ~ r1(X2, X3))), introduced(avatar_definition, [new_symbols(naming, [spl32_119])])).
fof(f1633, plain, (spl32_122 | spl32_119 | ~ spl32_118), inference(avatar_split_clause, [], [f1629, f1585, f1589, f1631])).
fof(f1585, plain, (spl32_118 <=> r1(sK25, sK28(sK25))), introduced(avatar_definition, [new_symbols(naming, [spl32_118])])).
fof(f1629, plain, (! [X4, X2, X0, X3, X1] : (p1(X0) | ~ r1(X1, X0) | ~ r1(sK25, X1) | ~ r1(X2, sK25) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(sK13, X4)) | ~ spl32_118), inference(resolution, [], [f1615, f86])).
fof(f86, plain, ! [X39, X37, X33, X31, X32, X40] : (~ p1(sK28(X37)) | p1(X40) | ~ r1(X39, X40) | ~ r1(X37, X39) | ~ r1(X33, X37) | ~ r1(X32, X33) | ~ r1(X31, X32) | ~ r1(sK13, X31)), inference(cnf_transformation, [], [f53])).
fof(f1615, plain, (p1(sK28(sK25)) | ~ spl32_118), inference(resolution, [], [f1587, f96])).
fof(f1587, plain, (r1(sK25, sK28(sK25)) | ~ spl32_118), inference(avatar_component_clause, [], [f1585])).
fof(f1591, plain, (spl32_118 | spl32_119), inference(avatar_split_clause, [], [f1571, f1589, f1585])).
fof(f1571, plain, ! [X2, X3] : (~ r1(sK25, X2) | ~ r1(X2, X3) | p1(X3) | r1(sK25, sK28(sK25))), inference(resolution, [], [f1564, f95])).
fof(f1564, plain, ! [X2, X0, X1] : (~ r1(sK22, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | p1(X2) | r1(X0, sK28(X0))), inference(resolution, [], [f1018, f92])).
fof(f1018, plain, ! [X2, X0, X3, X1] : (~ r1(sK21, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | p1(X1) | r1(X2, sK28(X2))), inference(resolution, [], [f720, f91])).
fof(f720, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK20, X4) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | p1(X0) | r1(X2, sK28(X2))), inference(resolution, [], [f85, f90])).
fof(f85, plain, ! [X39, X37, X33, X31, X32, X40] : (~ r1(sK13, X31) | p1(X40) | ~ r1(X39, X40) | ~ r1(X37, X39) | ~ r1(X33, X37) | ~ r1(X32, X33) | ~ r1(X31, X32) | r1(X37, sK28(X37))), inference(cnf_transformation, [], [f53])).
fof(f1293, plain, (~ spl32_41 | ~ spl32_74), inference(avatar_split_clause, [], [f1273, f693, f373])).
fof(f1273, plain, (~ sP0(sK22) | ~ spl32_74), inference(resolution, [], [f695, f946])).
fof(f946, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f945])).
fof(f945, plain, ! [X0] : (~ sP0(X0) | ~ sP1(X0) | ~ sP1(X0) | ~ sP0(X0)), inference(resolution, [], [f935, f68])).
fof(f68, plain, ! [X0] : (r1(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((! [X2] : ((p1(X2) & (~ p1(sK11(X2)) & r1(X2, sK11(X2)))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0))) & r1(X0, sK10(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12])], [f28, f31, f30, f29])).
fof(f29, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) => (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & ? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) & r1(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK11(X2)) & r1(X2, sK11(X2)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) => (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f27])).
fof(f27, plain, ! [X33] : (? [X50] : (! [X51] : ((p1(X51) & ? [X52] : (~ p1(X52) & r1(X51, X52))) | ! [X53] : (~ p1(X53) | ! [X54] : (p1(X54) | ~ r1(X53, X54)) | ~ r1(X51, X53)) | ~ r1(X50, X51)) & p1(X50) & ? [X55] : (~ p1(X55) & r1(X50, X55)) & r1(X33, X50)) | ~ sP0(X33)), inference(nnf_transformation, [], [f7])).
fof(f935, plain, ! [X0, X1] : (~ r1(X1, sK10(X0)) | ~ sP0(X0) | ~ sP1(X0) | ~ sP1(X1)), inference(duplicate_literal_removal, [], [f934])).
fof(f934, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | ~ r1(X1, sK10(X0)) | ~ sP1(X1) | ~ sP0(X0)), inference(resolution, [], [f905, f69])).
fof(f69, plain, ! [X0] : (r1(sK10(X0), sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f905, plain, ! [X2, X0, X1] : (~ r1(X1, sK12(X0)) | ~ sP1(X0) | ~ sP0(X0) | ~ r1(X2, X1) | ~ sP1(X2)), inference(subsumption_resolution, [], [f904, f70])).
fof(f70, plain, ! [X0] : (~ p1(sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f904, plain, ! [X2, X0, X1] : (~ sP0(X0) | ~ sP1(X0) | p1(sK12(X0)) | ~ r1(X1, sK12(X0)) | ~ r1(X2, X1) | ~ sP1(X2)), inference(resolution, [], [f901, f66])).
fof(f66, plain, ! [X2, X0, X1] : (~ p1(sK9(X2)) | p1(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (! [X1] : (! [X2] : ((p1(sK8(X2)) & (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2))) & r1(X2, sK8(X2))) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9])], [f23, f25, f24])).
fof(f24, plain, ! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) => (p1(sK8(X2)) & ? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) & r1(X2, sK8(X2)))), introduced(choice_axiom, [])).
fof(f25, plain, ! [X2] : (? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) => (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f22])).
fof(f22, plain, ! [X33] : (! [X37] : (! [X38] : (? [X39] : (p1(X39) & ? [X40] : (~ p1(X40) & r1(X39, X40)) & r1(X38, X39)) | p1(X38) | ~ r1(X37, X38)) | ~ r1(X33, X37)) | ~ sP1(X33)), inference(nnf_transformation, [], [f8])).
fof(f901, plain, ! [X0] : (p1(sK9(sK12(X0))) | ~ sP0(X0) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f896])).
fof(f896, plain, ! [X0] : (~ sP0(X0) | p1(sK9(sK12(X0))) | ~ sP1(X0) | ~ sP1(X0) | ~ sP0(X0)), inference(resolution, [], [f796, f821])).
fof(f821, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f820])).
fof(f820, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f335, f68])).
fof(f335, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f328, f70])).
fof(f328, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f65, f69])).
fof(f65, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(sK8(X2), sK9(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f796, plain, ! [X0, X1] : (~ r1(sK8(sK12(X0)), X1) | ~ sP0(X0) | p1(X1) | ~ sP1(X0)), inference(subsumption_resolution, [], [f795, f737])).
fof(f737, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f736])).
fof(f736, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f191, f68])).
fof(f191, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | p1(sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f184, f70])).
fof(f184, plain, ! [X4, X5] : (p1(sK12(X4)) | p1(sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f67, f69])).
fof(f67, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | p1(sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f795, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0)))), inference(duplicate_literal_removal, [], [f787])).
fof(f787, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0))) | ~ sP0(X0)), inference(resolution, [], [f784, f411])).
fof(f411, plain, ! [X2, X0, X1] : (~ r1(sK12(X2), X0) | p1(X1) | ~ r1(X0, X1) | ~ p1(X0) | ~ sP0(X2)), inference(subsumption_resolution, [], [f410, f70])).
fof(f410, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2)), inference(duplicate_literal_removal, [], [f407])).
fof(f407, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2) | ~ sP0(X2)), inference(resolution, [], [f74, f69])).
fof(f74, plain, ! [X4, X2, X0, X5] : (~ r1(sK10(X0), X2) | ~ p1(X4) | p1(X5) | ~ r1(X4, X5) | ~ r1(X2, X4) | p1(X2) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f784, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f783])).
fof(f783, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f245, f68])).
fof(f245, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f238, f70])).
fof(f238, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f64, f69])).
fof(f64, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(X2, sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f695, plain, (sP1(sK22) | ~ spl32_74), inference(avatar_component_clause, [], [f693])).
fof(f502, plain, (spl32_41 | ~ spl32_20 | spl32_54), inference(avatar_split_clause, [], [f498, f500, f220, f373])).
fof(f498, plain, ! [X9] : (r1(X9, sK4(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | sP0(sK22)), inference(subsumption_resolution, [], [f497, f251])).
fof(f497, plain, ! [X9] : (r1(X9, sK4(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | sP0(sK22) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f394, f94])).
fof(f394, plain, ! [X9] : (r1(X9, sK4(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | p1(sK26) | sP0(sK22) | ~ sP3(sK22)), inference(resolution, [], [f54, f93])).
fof(f54, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(X1, sK4(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f496, plain, (spl32_41 | ~ spl32_20 | spl32_53), inference(avatar_split_clause, [], [f492, f494, f220, f373])).
fof(f492, plain, ! [X9] : (r1(sK4(X9), sK5(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | sP0(sK22)), inference(subsumption_resolution, [], [f491, f251])).
fof(f491, plain, ! [X9] : (r1(sK4(X9), sK5(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | sP0(sK22) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f437, f94])).
fof(f437, plain, ! [X9] : (r1(sK4(X9), sK5(X9)) | ~ r1(sK22, X9) | ~ p1(sK22) | p1(sK26) | sP0(sK22) | ~ sP3(sK22)), inference(resolution, [], [f55, f93])).
fof(f55, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(sK4(X1), sK5(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f406, plain, spl32_33, inference(avatar_split_clause, [], [f401, f309])).
fof(f401, plain, sP2(sK22), inference(resolution, [], [f385, f92])).
fof(f385, plain, ! [X0] : (~ r1(sK21, X0) | sP2(X0)), inference(resolution, [], [f156, f91])).
fof(f156, plain, ! [X0, X1] : (~ r1(sK20, X0) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f84, f90])).
fof(f84, plain, ! [X33, X31, X32] : (~ r1(sK13, X31) | ~ r1(X32, X33) | ~ r1(X31, X32) | sP2(X33)), inference(cnf_transformation, [], [f53])).
fof(f325, plain, (~ spl32_33 | spl32_20 | spl32_34 | spl32_36), inference(avatar_split_clause, [], [f279, f322, f313, f220, f309])).
fof(f279, plain, (r1(sK25, sK7(sK25)) | r1(sK22, sK6(sK22)) | p1(sK22) | ~ sP2(sK22)), inference(resolution, [], [f58, f95])).
fof(f58, plain, ! [X4, X0] : (~ r1(X0, X4) | r1(X4, sK7(X4)) | r1(X0, sK6(X0)) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).