fof(f528, plain, $false, inference(avatar_sat_refutation, [], [f239, f403, f408, f409, f410, f503, f527])).
fof(f527, plain, (~ spl13_38 | ~ spl13_41), inference(avatar_contradiction_clause, [], [f526])).
fof(f526, plain, ($false | (~ spl13_38 | ~ spl13_41)), inference(subsumption_resolution, [], [f525, f234])).
fof(f234, plain, (p3(sK12) | ~ spl13_38), inference(avatar_component_clause, [], [f232])).
fof(f232, plain, (spl13_38 <=> p3(sK12)), introduced(avatar_definition, [new_symbols(naming, [spl13_38])])).
fof(f525, plain, (~ p3(sK12) | ~ spl13_41), inference(subsumption_resolution, [], [f523, f406])).
fof(f406, plain, (p1(sK12) | ~ spl13_41), inference(avatar_component_clause, [], [f405])).
fof(f405, plain, (spl13_41 <=> p1(sK12)), introduced(avatar_definition, [new_symbols(naming, [spl13_41])])).
fof(f523, plain, (~ p1(sK12) | ~ p3(sK12)), inference(resolution, [], [f479, f37])).
fof(f37, plain, r1(sK11, sK12), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((((r1(sK5, sK6) & r1(sK4, sK5)) & r1(sK3, sK4)) & r1(sK2, sK3)) & ! [X5] : ((sP0(X5) & r1(X5, sK7(X5)) & ! [X7] : (! [X8] : (! [X9] : (((~ p3(X9) | ~ p1(X9)) & (p1(X9) | p3(X9))) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X5, X7))) | ~ r1(sK2, X5)) & r1(sK2, sK8) & (((r1(sK11, sK12) & r1(sK10, sK11)) & r1(sK9, sK10)) & r1(sK2, sK9))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3, sK4, sK5, sK6, sK7, sK8, sK9, sK10, sK11, sK12])], [f15, f26, f25, f24, f23, f22, f21, f20, f19, f18, f17, f16])).
fof(f16, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X5] : ((sP0(X5) & ? [X6] : r1(X5, X6) & ! [X7] : (! [X8] : (! [X9] : (((~ p3(X9) | ~ p1(X9)) & (p1(X9) | p3(X9))) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X5, X7))) | ~ r1(X0, X5)) & ? [X10] : r1(X0, X10) & ? [X11] : (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(X11, X12)) & r1(X0, X11))) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(sK2, X1)) & ! [X5] : ((sP0(X5) & ? [X6] : r1(X5, X6) & ! [X7] : (! [X8] : (! [X9] : (((~ p3(X9) | ~ p1(X9)) & (p1(X9) | p3(X9))) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X5, X7))) | ~ r1(sK2, X5)) & ? [X10] : r1(sK2, X10) & ? [X11] : (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(X11, X12)) & r1(sK2, X11)))), introduced(choice_axiom, [])).
fof(f17, plain, (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(sK2, X1)) => (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(sK3, X2)) & r1(sK2, sK3))), introduced(choice_axiom, [])).
fof(f18, plain, (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(sK3, X2)) => (? [X3] : (? [X4] : r1(X3, X4) & r1(sK4, X3)) & r1(sK3, sK4))), introduced(choice_axiom, [])).
fof(f19, plain, (? [X3] : (? [X4] : r1(X3, X4) & r1(sK4, X3)) => (? [X4] : r1(sK5, X4) & r1(sK4, sK5))), introduced(choice_axiom, [])).
fof(f20, plain, (? [X4] : r1(sK5, X4) => r1(sK5, sK6)), introduced(choice_axiom, [])).
fof(f21, plain, ! [X5] : (? [X6] : r1(X5, X6) => r1(X5, sK7(X5))), introduced(choice_axiom, [])).
fof(f22, plain, (? [X10] : r1(sK2, X10) => r1(sK2, sK8)), introduced(choice_axiom, [])).
fof(f23, plain, (? [X11] : (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(X11, X12)) & r1(sK2, X11)) => (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(sK9, X12)) & r1(sK2, sK9))), introduced(choice_axiom, [])).
fof(f24, plain, (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(sK9, X12)) => (? [X13] : (? [X14] : r1(X13, X14) & r1(sK10, X13)) & r1(sK9, sK10))), introduced(choice_axiom, [])).
fof(f25, plain, (? [X13] : (? [X14] : r1(X13, X14) & r1(sK10, X13)) => (? [X14] : r1(sK11, X14) & r1(sK10, sK11))), introduced(choice_axiom, [])).
fof(f26, plain, (? [X14] : r1(sK11, X14) => r1(sK11, sK12)), introduced(choice_axiom, [])).
fof(f15, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X5] : ((sP0(X5) & ? [X6] : r1(X5, X6) & ! [X7] : (! [X8] : (! [X9] : (((~ p3(X9) | ~ p1(X9)) & (p1(X9) | p3(X9))) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X5, X7))) | ~ r1(X0, X5)) & ? [X10] : r1(X0, X10) & ? [X11] : (? [X12] : (? [X13] : (? [X14] : r1(X13, X14) & r1(X12, X13)) & r1(X11, X12)) & r1(X0, X11))), inference(rectify, [], [f10])).
fof(f10, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X5] : ((sP0(X5) & ? [X12] : r1(X5, X12) & ! [X13] : (! [X14] : (! [X15] : (((~ p3(X15) | ~ p1(X15)) & (p1(X15) | p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) & ? [X16] : r1(X0, X16) & ? [X17] : (? [X18] : (? [X19] : (? [X20] : r1(X19, X20) & r1(X18, X19)) & r1(X17, X18)) & r1(X0, X17))), inference(definition_folding, [], [f8, e9])).
fof(f9, plain, ! [X5] : (! [X6] : ((! [X7] : (! [X8] : (((~ p1(X8) | ~ p2(X8)) & (p2(X8) | p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ? [X9] : (~ p3(X9) & r1(X6, X9)) & ! [X10] : (! [X11] : (((~ p2(X11) | ~ p3(X11)) & (p3(X11) | p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ~ sP0(X5)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X5] : (sP0(X5) <=> ! [X6] : ((! [X7] : (! [X8] : (((~ p1(X8) | ~ p2(X8)) & (p2(X8) | p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ? [X9] : (~ p3(X9) & r1(X6, X9)) & ! [X10] : (! [X11] : (((~ p2(X11) | ~ p3(X11)) & (p3(X11) | p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : r1(X3, X4) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X5] : ((! [X6] : ((! [X7] : (! [X8] : (((~ p1(X8) | ~ p2(X8)) & (p2(X8) | p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ? [X9] : (~ p3(X9) & r1(X6, X9)) & ! [X10] : (! [X11] : (((~ p2(X11) | ~ p3(X11)) & (p3(X11) | p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) & ? [X12] : r1(X5, X12) & ! [X13] : (! [X14] : (! [X15] : (((~ p3(X15) | ~ p1(X15)) & (p1(X15) | p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) & ? [X16] : r1(X0, X16) & ? [X17] : (? [X18] : (? [X19] : (? [X20] : r1(X19, X20) & r1(X18, X19)) & r1(X17, X18)) & r1(X0, X17))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : ~ r1(X3, X4) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X5] : (~ (~ ! [X6] : (~ (~ ! [X7] : (! [X8] : (~ ((p1(X8) & p2(X8)) | (~ p2(X8) & ~ p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ! [X9] : (p3(X9) | ~ r1(X6, X9)) | ~ ! [X10] : (! [X11] : (~ ((p2(X11) & p3(X11)) | (~ p3(X11) & ~ p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ! [X12] : ~ r1(X5, X12) | ~ ! [X13] : (! [X14] : (! [X15] : (~ ((p3(X15) & p1(X15)) | (~ p1(X15) & ~ p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) | ! [X16] : ~ r1(X0, X16) | ! [X17] : (! [X18] : (! [X19] : (! [X20] : ~ r1(X19, X20) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X0, X17))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((p1(X4) & p2(X4) & p3(X4) & p4(X4)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X5] : (~ (~ ! [X6] : (~ (~ ! [X7] : (! [X8] : (~ ((p1(X8) & p2(X8)) | (~ p2(X8) & ~ p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ! [X9] : (p3(X9) | ~ r1(X6, X9)) | ~ ! [X10] : (! [X11] : (~ ((p2(X11) & p3(X11)) | (~ p3(X11) & ~ p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ! [X12] : (p4(X12) | ~ r1(X5, X12)) | ~ ! [X13] : (! [X14] : (! [X15] : (~ ((p3(X15) & p1(X15)) | (~ p1(X15) & ~ p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) | ! [X16] : ~ r1(X0, X16) | ! [X17] : (! [X18] : (! [X19] : (! [X20] : ~ r1(X19, X20) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X0, X17))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((p1(X4) & p2(X4) & p3(X4) & p4(X4)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X5] : (~ (~ ! [X6] : (~ (~ ! [X7] : (! [X8] : (~ ((p1(X8) & p2(X8)) | (~ p2(X8) & ~ p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ! [X9] : (p3(X9) | ~ r1(X6, X9)) | ~ ! [X10] : (! [X11] : (~ ((p2(X11) & p3(X11)) | (~ p3(X11) & ~ p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ! [X12] : (p4(X12) | ~ r1(X5, X12)) | ~ ! [X13] : (! [X14] : (! [X15] : (~ ((p3(X15) & p1(X15)) | (~ p1(X15) & ~ p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) | ! [X16] : ~ r1(X0, X16) | ! [X17] : (! [X18] : (! [X19] : (! [X20] : ((~ p2(X20) & ~ p4(X20) & ~ p6(X20) & ~ p8(X20)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X0, X17))), inference(pure_predicate_removal, [], [f4])).
fof(f4, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((p1(X4) & p2(X4) & p3(X4) & p4(X4)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X5] : (~ (~ ! [X6] : (~ (~ ! [X7] : (! [X8] : (~ ((p1(X8) & p2(X8)) | (~ p2(X8) & ~ p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ! [X9] : (p3(X9) | ~ r1(X6, X9)) | ~ ! [X10] : (! [X11] : (~ ((p2(X11) & p3(X11)) | (~ p3(X11) & ~ p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ! [X12] : (p4(X12) | ~ r1(X5, X12)) | ~ ! [X13] : (! [X14] : (! [X15] : (~ ((p3(X15) & p1(X15)) | (~ p1(X15) & ~ p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) | ! [X16] : (p5(X16) | ~ r1(X0, X16)) | ! [X17] : (! [X18] : (! [X19] : (! [X20] : ((~ p2(X20) & ~ p4(X20) & ~ p6(X20) & ~ p8(X20)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X0, X17))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((p1(X4) & p2(X4) & p3(X4) & p4(X4)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X5] : (~ (~ ! [X6] : (~ (~ ! [X7] : (~ ~ ! [X8] : (~ ((p1(X8) & p2(X8)) | (~ p2(X8) & ~ p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ! [X9] : (p3(X9) | ~ r1(X6, X9)) | ~ ! [X10] : (! [X11] : (~ ((p2(X11) & p3(X11)) | (~ p3(X11) & ~ p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ! [X12] : (p4(X12) | ~ r1(X5, X12)) | ~ ! [X13] : (! [X14] : (! [X15] : (~ ((p3(X15) & p1(X15)) | (~ p1(X15) & ~ p3(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X5, X13))) | ~ r1(X0, X5)) | ! [X16] : (p5(X16) | ~ r1(X0, X16)) | ! [X17] : (! [X18] : (! [X19] : (! [X20] : ((~ p2(X20) & ~ p4(X20) & ~ p6(X20) & ~ p8(X20)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X0, X17))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((p1(X0) & p2(X0) & p3(X0) & p4(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ ~ ! [X0] : (~ ((p1(X0) & p2(X0)) | (~ p2(X0) & ~ p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (~ ((p2(X0) & p3(X0)) | (~ p3(X0) & ~ p2(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p4(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (~ ((p3(X0) & p1(X0)) | (~ p1(X0) & ~ p3(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p5(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((~ p2(X0) & ~ p4(X0) & ~ p6(X0) & ~ p8(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((p1(X0) & p2(X0) & p3(X0) & p4(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ ~ ! [X0] : (~ ((p1(X0) & p2(X0)) | (~ p2(X0) & ~ p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (~ ((p2(X0) & p3(X0)) | (~ p3(X0) & ~ p2(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p4(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (~ ((p3(X0) & p1(X0)) | (~ p1(X0) & ~ p3(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p5(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((~ p2(X0) & ~ p4(X0) & ~ p6(X0) & ~ p8(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL650+1.001.p', main)).
fof(f479, plain, ! [X0] : (~ r1(sK11, X0) | ~ p1(X0) | ~ p3(X0)), inference(resolution, [], [f429, f36])).
fof(f36, plain, r1(sK10, sK11), inference(cnf_transformation, [], [f27])).
fof(f429, plain, ! [X0, X1] : (~ r1(sK10, X0) | ~ r1(X0, X1) | ~ p1(X1) | ~ p3(X1)), inference(resolution, [], [f274, f35])).
fof(f35, plain, r1(sK9, sK10), inference(cnf_transformation, [], [f27])).
fof(f274, plain, ! [X2, X0, X1] : (~ r1(sK9, X2) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ p1(X0) | ~ p3(X0)), inference(resolution, [], [f40, f34])).
fof(f34, plain, r1(sK2, sK9), inference(cnf_transformation, [], [f27])).
fof(f40, plain, ! [X8, X7, X5, X9] : (~ r1(sK2, X5) | ~ p1(X9) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X5, X7) | ~ p3(X9)), inference(cnf_transformation, [], [f27])).
fof(f503, plain, (spl13_41 | spl13_38), inference(avatar_split_clause, [], [f485, f232, f405])).
fof(f485, plain, (p3(sK12) | p1(sK12)), inference(resolution, [], [f457, f37])).
fof(f457, plain, ! [X0] : (~ r1(sK11, X0) | p3(X0) | p1(X0)), inference(resolution, [], [f421, f36])).
fof(f421, plain, ! [X0, X1] : (~ r1(sK10, X0) | ~ r1(X0, X1) | p3(X1) | p1(X1)), inference(resolution, [], [f270, f35])).
fof(f270, plain, ! [X2, X0, X1] : (~ r1(sK9, X2) | ~ r1(X1, X0) | ~ r1(X2, X1) | p3(X0) | p1(X0)), inference(resolution, [], [f39, f34])).
fof(f39, plain, ! [X8, X7, X5, X9] : (~ r1(sK2, X5) | p3(X9) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X5, X7) | p1(X9)), inference(cnf_transformation, [], [f27])).
fof(f410, plain, (spl13_37 | ~ spl13_39 | ~ spl13_38), inference(avatar_split_clause, [], [f249, f232, f236, f229])).
fof(f229, plain, (spl13_37 <=> ! [X20, X19] : (~ r1(X19, sK11) | ~ sP0(X20) | ~ r1(X20, X19))), introduced(avatar_definition, [new_symbols(naming, [spl13_37])])).
fof(f236, plain, (spl13_39 <=> p2(sK12)), introduced(avatar_definition, [new_symbols(naming, [spl13_39])])).
fof(f249, plain, ! [X19, X20] : (~ p3(sK12) | ~ p2(sK12) | ~ r1(X19, sK11) | ~ r1(X20, X19) | ~ sP0(X20)), inference(resolution, [], [f29, f37])).
fof(f29, plain, ! [X6, X0, X5, X1] : (~ r1(X5, X6) | ~ p3(X6) | ~ p2(X6) | ~ r1(X1, X5) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ! [X0] : (! [X1] : ((! [X2] : (! [X3] : (((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3))) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & (~ p3(sK1(X1)) & r1(X1, sK1(X1))) & ! [X5] : (! [X6] : (((~ p2(X6) | ~ p3(X6)) & (p3(X6) | p2(X6))) | ~ r1(X5, X6)) | ~ r1(X1, X5))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1])], [f12, f13])).
fof(f13, plain, ! [X1] : (? [X4] : (~ p3(X4) & r1(X1, X4)) => (~ p3(sK1(X1)) & r1(X1, sK1(X1)))), introduced(choice_axiom, [])).
fof(f12, plain, ! [X0] : (! [X1] : ((! [X2] : (! [X3] : (((~ p1(X3) | ~ p2(X3)) & (p2(X3) | p1(X3))) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ? [X4] : (~ p3(X4) & r1(X1, X4)) & ! [X5] : (! [X6] : (((~ p2(X6) | ~ p3(X6)) & (p3(X6) | p2(X6))) | ~ r1(X5, X6)) | ~ r1(X1, X5))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f11])).
fof(f11, plain, ! [X5] : (! [X6] : ((! [X7] : (! [X8] : (((~ p1(X8) | ~ p2(X8)) & (p2(X8) | p1(X8))) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ? [X9] : (~ p3(X9) & r1(X6, X9)) & ! [X10] : (! [X11] : (((~ p2(X11) | ~ p3(X11)) & (p3(X11) | p2(X11))) | ~ r1(X10, X11)) | ~ r1(X6, X10))) | ~ r1(X5, X6)) | ~ sP0(X5)), inference(nnf_transformation, [], [f9])).
fof(f409, plain, (spl13_37 | spl13_39 | spl13_41), inference(avatar_split_clause, [], [f259, f405, f236, f229])).
fof(f259, plain, ! [X19, X20] : (p1(sK12) | p2(sK12) | ~ r1(X19, sK11) | ~ r1(X20, X19) | ~ sP0(X20)), inference(resolution, [], [f32, f37])).
fof(f32, plain, ! [X2, X0, X3, X1] : (~ r1(X2, X3) | p1(X3) | p2(X3) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f14])).
fof(f408, plain, (spl13_37 | ~ spl13_41 | ~ spl13_39), inference(avatar_split_clause, [], [f269, f236, f405, f229])).
fof(f269, plain, ! [X19, X20] : (~ p2(sK12) | ~ p1(sK12) | ~ r1(X19, sK11) | ~ r1(X20, X19) | ~ sP0(X20)), inference(resolution, [], [f33, f37])).
fof(f33, plain, ! [X2, X0, X3, X1] : (~ r1(X2, X3) | ~ p2(X3) | ~ p1(X3) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f14])).
fof(f403, plain, ~ spl13_37, inference(avatar_contradiction_clause, [], [f402])).
fof(f402, plain, ($false | ~ spl13_37), inference(subsumption_resolution, [], [f401, f36])).
fof(f401, plain, (~ r1(sK10, sK11) | ~ spl13_37), inference(subsumption_resolution, [], [f395, f47])).
fof(f47, plain, sP0(sK9), inference(resolution, [], [f42, f34])).
fof(f42, plain, ! [X5] : (~ r1(sK2, X5) | sP0(X5)), inference(cnf_transformation, [], [f27])).
fof(f395, plain, (~ sP0(sK9) | ~ r1(sK10, sK11) | ~ spl13_37), inference(resolution, [], [f230, f35])).
fof(f230, plain, (! [X19, X20] : (~ r1(X20, X19) | ~ sP0(X20) | ~ r1(X19, sK11)) | ~ spl13_37), inference(avatar_component_clause, [], [f229])).
fof(f239, plain, (spl13_37 | spl13_38 | spl13_39), inference(avatar_split_clause, [], [f137, f236, f232, f229])).
fof(f137, plain, ! [X19, X20] : (p2(sK12) | p3(sK12) | ~ r1(X19, sK11) | ~ r1(X20, X19) | ~ sP0(X20)), inference(resolution, [], [f28, f37])).
fof(f28, plain, ! [X6, X0, X5, X1] : (~ r1(X5, X6) | p2(X6) | p3(X6) | ~ r1(X1, X5) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f14])).