fof(f125069, plain, $false, inference(avatar_sat_refutation, [], [f5580, f5604, f5614, f9602, f9604, f21389, f21408, f21451, f21879, f21881, f23757, f23777, f31544, f123668, f123762, f123793, f123819, f123821, f124634, f124635, f124636, f124637, f124648, f124733, f125067])).
fof(f125067, plain, (spl28_173 | ~ spl28_540 | ~ spl28_1916), inference(avatar_contradiction_clause, [], [f125066])).
fof(f125066, plain, ($false | (spl28_173 | ~ spl28_540 | ~ spl28_1916)), inference(subsumption_resolution, [], [f125065, f1490])).
fof(f1490, plain, (~ r1(sK18, sK9(sK18)) | spl28_173), inference(avatar_component_clause, [], [f1489])).
fof(f1489, plain, (spl28_173 <=> r1(sK18, sK9(sK18))), introduced(avatar_definition, [new_symbols(naming, [spl28_173])])).
fof(f125065, plain, (r1(sK18, sK9(sK18)) | (~ spl28_540 | ~ spl28_1916)), inference(subsumption_resolution, [], [f124984, f102])).
fof(f102, plain, sP1(sK18), inference(resolution, [], [f84, f94])).
fof(f94, plain, r1(sK16, sK18), inference(cnf_transformation, [], [f51])).
fof(f51, plain, (! [X1] : ((p1(sK17(X1)) & r1(X1, sK17(X1))) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(sK16, X1)) & (! [X6] : ((~ p2(sK19(X6)) & r1(X6, sK19(X6))) | ~ r1(sK18, X6)) & r1(sK16, sK18)) & ((~ p1(sK21) & r1(sK20, sK21)) & (! [X11] : (p1(X11) | ~ r1(sK22, X11)) & r1(sK20, sK22)) & r1(sK16, sK20)) & ! [X12] : ((~ p1(sK23(X12)) & r1(X12, sK23(X12))) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(sK16, X12)) & ! [X16] : ((sP2(X16) & sP0(X16) & sP1(X16) & sP3(X16) & ! [X17] : ((p2(sK24(X17)) & r1(sK24(X17), sK25(X17)) & r1(X17, sK24(X17))) | ~ r1(X16, X17))) | ~ r1(sK16, X16)) & ! [X20] : ((! [X22] : (~ p1(X22) | ~ r1(sK26(X20), X22)) & r1(X20, sK26(X20))) | ! [X23] : (! [X24] : ((p1(sK27(X24)) & r1(X24, sK27(X24))) | ~ r1(X23, X24)) | ~ r1(X20, X23)) | ~ r1(sK16, X20))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17, sK18, sK19, sK20, sK21, sK22, sK23, sK24, sK25, sK26, sK27])], [f38, f50, f49, f48, f47, f46, f45, f44, f43, f42, f41, f40, f39])).
fof(f39, plain, (? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP2(X16) & sP0(X16) & sP1(X16) & sP3(X16) & ! [X17] : (? [X18] : (p2(X18) & ? [X19] : r1(X18, X19) & r1(X17, X18)) | ~ r1(X16, X17))) | ~ r1(X0, X16)) & ! [X20] : (? [X21] : (! [X22] : (~ p1(X22) | ~ r1(X21, X22)) & r1(X20, X21)) | ! [X23] : (! [X24] : (? [X25] : (p1(X25) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X20, X23)) | ~ r1(X0, X20))) => (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(sK16, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(sK16, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(sK16, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(sK16, X12)) & ! [X16] : ((sP2(X16) & sP0(X16) & sP1(X16) & sP3(X16) & ! [X17] : (? [X18] : (p2(X18) & ? [X19] : r1(X18, X19) & r1(X17, X18)) | ~ r1(X16, X17))) | ~ r1(sK16, X16)) & ! [X20] : (? [X21] : (! [X22] : (~ p1(X22) | ~ r1(X21, X22)) & r1(X20, X21)) | ! [X23] : (! [X24] : (? [X25] : (p1(X25) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X20, X23)) | ~ r1(sK16, X20)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) => (p1(sK17(X1)) & r1(X1, sK17(X1)))), introduced(choice_axiom, [])).
fof(f41, plain, (? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(sK16, X5)) => (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(sK18, X6)) & r1(sK16, sK18))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) => (~ p2(sK19(X6)) & r1(X6, sK19(X6)))), introduced(choice_axiom, [])).
fof(f43, plain, (? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(sK16, X8)) => (? [X9] : (~ p1(X9) & r1(sK20, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(sK20, X10)) & r1(sK16, sK20))), introduced(choice_axiom, [])).
fof(f44, plain, (? [X9] : (~ p1(X9) & r1(sK20, X9)) => (~ p1(sK21) & r1(sK20, sK21))), introduced(choice_axiom, [])).
fof(f45, plain, (? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(sK20, X10)) => (! [X11] : (p1(X11) | ~ r1(sK22, X11)) & r1(sK20, sK22))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) => (~ p1(sK23(X12)) & r1(X12, sK23(X12)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X17] : (? [X18] : (p2(X18) & ? [X19] : r1(X18, X19) & r1(X17, X18)) => (p2(sK24(X17)) & ? [X19] : r1(sK24(X17), X19) & r1(X17, sK24(X17)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X17] : (? [X19] : r1(sK24(X17), X19) => r1(sK24(X17), sK25(X17))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X20] : (? [X21] : (! [X22] : (~ p1(X22) | ~ r1(X21, X22)) & r1(X20, X21)) => (! [X22] : (~ p1(X22) | ~ r1(sK26(X20), X22)) & r1(X20, sK26(X20)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X24] : (? [X25] : (p1(X25) & r1(X24, X25)) => (p1(sK27(X24)) & r1(X24, sK27(X24)))), introduced(choice_axiom, [])).
fof(f38, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP2(X16) & sP0(X16) & sP1(X16) & sP3(X16) & ! [X17] : (? [X18] : (p2(X18) & ? [X19] : r1(X18, X19) & r1(X17, X18)) | ~ r1(X16, X17))) | ~ r1(X0, X16)) & ! [X20] : (? [X21] : (! [X22] : (~ p1(X22) | ~ r1(X21, X22)) & r1(X20, X21)) | ! [X23] : (! [X24] : (? [X25] : (p1(X25) & r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X20, X23)) | ~ r1(X0, X20))), inference(rectify, [], [f13])).
fof(f13, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : ((sP2(X16) & sP0(X16) & sP1(X16) & sP3(X16) & ! [X40] : (? [X41] : (p2(X41) & ? [X42] : r1(X41, X42) & r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) & ! [X43] : (? [X44] : (! [X45] : (~ p1(X45) | ~ r1(X44, X45)) & r1(X43, X44)) | ! [X46] : (! [X47] : (? [X48] : (p1(X48) & r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46)) | ~ r1(X0, X43))), inference(definition_folding, [], [f8, e12, e11, e10, e9])).
fof(f9, plain, ! [X16] : (! [X21] : (? [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X22, X23)) & r1(X21, X22)) | ? [X26] : (! [X27] : (~ p2(X27) | ~ r1(X26, X27)) & r1(X21, X26)) | ~ r1(X16, X21)) | ~ sP0(X16)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X16] : (sP0(X16) <=> ! [X21] : (? [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X22, X23)) & r1(X21, X22)) | ? [X26] : (! [X27] : (~ p2(X27) | ~ r1(X26, X27)) & r1(X21, X26)) | ~ r1(X16, X21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X16] : (? [X28] : (! [X29] : (~ p2(X29) | ~ r1(X28, X29)) & r1(X16, X28)) | ! [X30] : (! [X31] : (? [X32] : (p2(X32) & ? [X33] : r1(X32, X33) & r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30)) | ~ sP1(X16)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X16] : (sP1(X16) <=> (? [X28] : (! [X29] : (~ p2(X29) | ~ r1(X28, X29)) & r1(X16, X28)) | ! [X30] : (! [X31] : (? [X32] : (p2(X32) & ? [X33] : r1(X32, X33) & r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f11, plain, ! [X16] : (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ? [X19] : (p2(X19) & ? [X20] : r1(X19, X20) & r1(X16, X19)) | ~ sP2(X16)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X16] : (sP2(X16) <=> (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ? [X19] : (p2(X19) & ? [X20] : r1(X19, X20) & r1(X16, X19)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f12, plain, ! [X16] : (! [X34] : (! [X35] : (? [X36] : (! [X37] : (~ p2(X37) | ~ r1(X36, X37)) & r1(X35, X36)) | ~ r1(X34, X35)) | ~ p2(X34) | ! [X38] : ((p2(X38) & ? [X39] : r1(X38, X39)) | ~ r1(X34, X38)) | ~ r1(X16, X34)) | ~ sP3(X16)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X16] : (sP3(X16) <=> ! [X34] : (! [X35] : (? [X36] : (! [X37] : (~ p2(X37) | ~ r1(X36, X37)) & r1(X35, X36)) | ~ r1(X34, X35)) | ~ p2(X34) | ! [X38] : ((p2(X38) & ? [X39] : r1(X38, X39)) | ~ r1(X34, X38)) | ~ r1(X16, X34))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f8, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : (? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14)) | ~ r1(X0, X12)) & ! [X16] : (((? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ? [X19] : (p2(X19) & ? [X20] : r1(X19, X20) & r1(X16, X19))) & ! [X21] : (? [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X22, X23)) & r1(X21, X22)) | ? [X26] : (! [X27] : (~ p2(X27) | ~ r1(X26, X27)) & r1(X21, X26)) | ~ r1(X16, X21)) & (? [X28] : (! [X29] : (~ p2(X29) | ~ r1(X28, X29)) & r1(X16, X28)) | ! [X30] : (! [X31] : (? [X32] : (p2(X32) & ? [X33] : r1(X32, X33) & r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) & ! [X34] : (! [X35] : (? [X36] : (! [X37] : (~ p2(X37) | ~ r1(X36, X37)) & r1(X35, X36)) | ~ r1(X34, X35)) | ~ p2(X34) | ! [X38] : ((p2(X38) & ? [X39] : r1(X38, X39)) | ~ r1(X34, X38)) | ~ r1(X16, X34)) & ! [X40] : (? [X41] : (p2(X41) & ? [X42] : r1(X41, X42) & r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) & ! [X43] : (? [X44] : (! [X45] : (~ p1(X45) | ~ r1(X44, X45)) & r1(X43, X44)) | ! [X46] : (! [X47] : (? [X48] : (p1(X48) & r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46)) | ~ r1(X0, X43))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : (! [X1] : (? [X2] : (p1(X2) & r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) & ? [X5] : (! [X6] : (? [X7] : (~ p2(X7) & r1(X6, X7)) | ~ r1(X5, X6)) & r1(X0, X5)) & ? [X8] : (? [X9] : (~ p1(X9) & r1(X8, X9)) & ? [X10] : (! [X11] : (p1(X11) | ~ r1(X10, X11)) & r1(X8, X10)) & r1(X0, X8)) & ! [X12] : ((? [X13] : (~ p1(X13) & r1(X12, X13)) | ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) & ! [X16] : (((? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ? [X19] : (p2(X19) & ? [X20] : r1(X19, X20) & r1(X16, X19))) & ! [X21] : ((? [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X22, X23)) & r1(X21, X22)) | ? [X26] : (! [X27] : (~ p2(X27) | ~ r1(X26, X27)) & r1(X21, X26))) | ~ r1(X16, X21)) & (? [X28] : (! [X29] : (~ p2(X29) | ~ r1(X28, X29)) & r1(X16, X28)) | ! [X30] : (! [X31] : (? [X32] : (p2(X32) & ? [X33] : r1(X32, X33) & r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) & ! [X34] : ((! [X35] : (? [X36] : (! [X37] : (~ p2(X37) | ~ r1(X36, X37)) & r1(X35, X36)) | ~ r1(X34, X35)) | ~ p2(X34) | ! [X38] : ((p2(X38) & ? [X39] : r1(X38, X39)) | ~ r1(X34, X38))) | ~ r1(X16, X34)) & ! [X40] : (? [X41] : (p2(X41) & ? [X42] : r1(X41, X42) & r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) & ! [X43] : ((? [X44] : (! [X45] : (~ p1(X45) | ~ r1(X44, X45)) & r1(X43, X44)) | ! [X46] : (! [X47] : (? [X48] : (p1(X48) & r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46))) | ~ r1(X0, X43))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ! [X19] : (~ p2(X19) | ! [X20] : ~ r1(X19, X20) | ~ r1(X16, X19))) | ~ ! [X21] : (~ (! [X22] : (~ ! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : ~ r1(X24, X25) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) & ! [X26] : (~ ! [X27] : (~ p2(X27) | ~ r1(X26, X27)) | ~ r1(X21, X26))) | ~ r1(X16, X21)) | (! [X28] : (~ ! [X29] : (~ p2(X29) | ~ r1(X28, X29)) | ~ r1(X16, X28)) & ~ ! [X30] : (! [X31] : (~ ! [X32] : (~ p2(X32) | ! [X33] : ~ r1(X32, X33) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) | ~ ! [X34] : (~ (~ ! [X35] : (~ ! [X36] : (~ ! [X37] : (~ p2(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & p2(X34) & ~ ! [X38] : (~ (~ p2(X38) | ! [X39] : ~ r1(X38, X39)) | ~ r1(X34, X38))) | ~ r1(X16, X34)) | ~ ! [X40] : (~ ! [X41] : (~ p2(X41) | ! [X42] : ~ r1(X41, X42) | ~ r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) | ~ ! [X43] : (~ (! [X44] : (~ ! [X45] : (~ p1(X45) | ~ r1(X44, X45)) | ~ r1(X43, X44)) & ~ ! [X46] : (! [X47] : (~ ! [X48] : (~ p1(X48) | ~ r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46))) | ~ r1(X0, X43))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ! [X19] : (~ p2(X19) | ! [X20] : ~ r1(X19, X20) | ~ r1(X16, X19))) | ~ ! [X21] : (~ (! [X22] : (~ ! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : ~ r1(X24, X25) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) & ! [X26] : (~ ! [X27] : (~ p2(X27) | ~ r1(X26, X27)) | ~ r1(X21, X26))) | ~ r1(X16, X21)) | (! [X28] : (~ ! [X29] : (~ p2(X29) | ~ r1(X28, X29)) | ~ r1(X16, X28)) & ~ ! [X30] : (! [X31] : (~ ! [X32] : (~ p2(X32) | ! [X33] : ~ r1(X32, X33) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) | ~ ! [X34] : (~ (~ ! [X35] : (~ ! [X36] : (~ ! [X37] : (~ p2(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & p2(X34) & ~ ! [X38] : (~ (~ p2(X38) | ! [X39] : ~ r1(X38, X39)) | ~ r1(X34, X38))) | ~ r1(X16, X34)) | ~ ! [X40] : (~ ! [X41] : (~ p2(X41) | ! [X42] : ~ r1(X41, X42) | ~ r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) | ~ ! [X43] : (~ (! [X44] : (~ ! [X45] : (~ p1(X45) | ~ r1(X44, X45)) | ~ r1(X43, X44)) & ~ ! [X46] : (! [X47] : (~ ! [X48] : (~ p1(X48) | ~ r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46))) | ~ r1(X0, X43)) | ~ ! [X49] : (~ p4(X49) | ~ r1(X0, X49))), inference(pure_predicate_removal, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ! [X19] : (~ p2(X19) | ! [X20] : (p3(X20) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X21] : (~ (! [X22] : (~ ! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p3(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) & ! [X26] : (~ ! [X27] : (~ p2(X27) | ~ r1(X26, X27)) | ~ r1(X21, X26))) | ~ r1(X16, X21)) | (! [X28] : (~ ! [X29] : (~ p2(X29) | ~ r1(X28, X29)) | ~ r1(X16, X28)) & ~ ! [X30] : (! [X31] : (~ ! [X32] : (~ p2(X32) | ! [X33] : (p3(X33) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) | ~ ! [X34] : (~ (~ ! [X35] : (~ ! [X36] : (~ ! [X37] : (~ p2(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & p2(X34) & ~ ! [X38] : (~ (~ p2(X38) | ! [X39] : (p3(X39) | ~ r1(X38, X39))) | ~ r1(X34, X38))) | ~ r1(X16, X34)) | ~ ! [X40] : (~ ! [X41] : (~ p2(X41) | ! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) | ~ ! [X43] : (~ (! [X44] : (~ ! [X45] : (~ p1(X45) | ~ r1(X44, X45)) | ~ r1(X43, X44)) & ~ ! [X46] : (! [X47] : (~ ! [X48] : (~ p1(X48) | ~ r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46))) | ~ r1(X0, X43)) | ~ ! [X49] : (~ p4(X49) | ~ r1(X0, X49))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X2] : (~ p1(X2) | ~ r1(X1, X2)) | ! [X3] : (! [X4] : (~ p1(X4) | ~ r1(X3, X4)) | ~ r1(X1, X3)) | ~ r1(X0, X1)) | ! [X5] : (~ ! [X6] : (~ ! [X7] : (p2(X7) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X0, X5)) | ! [X8] : (! [X9] : (p1(X9) | ~ r1(X8, X9)) | ! [X10] : (~ ! [X11] : (p1(X11) | ~ r1(X10, X11)) | ~ r1(X8, X10)) | ~ r1(X0, X8)) | ~ ! [X12] : (~ (! [X13] : (p1(X13) | ~ r1(X12, X13)) & ~ ! [X14] : (! [X15] : (p1(X15) | ~ r1(X14, X15)) | ~ r1(X12, X14))) | ~ r1(X0, X12)) | ~ ! [X16] : (~ ((! [X17] : (~ ! [X18] : (~ p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) & ! [X19] : (~ p2(X19) | ! [X20] : (p3(X20) | ~ r1(X19, X20)) | ~ r1(X16, X19))) | ~ ! [X21] : (~ (! [X22] : (~ ! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p3(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) & ! [X26] : (~ ! [X27] : (~ p2(X27) | ~ r1(X26, X27)) | ~ r1(X21, X26))) | ~ r1(X16, X21)) | (! [X28] : (~ ! [X29] : (~ p2(X29) | ~ r1(X28, X29)) | ~ r1(X16, X28)) & ~ ! [X30] : (! [X31] : (~ ! [X32] : (~ p2(X32) | ! [X33] : (p3(X33) | ~ r1(X32, X33)) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30))) | ~ ! [X34] : (~ (~ ! [X35] : (~ ! [X36] : (~ ! [X37] : (~ p2(X37) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) & p2(X34) & ~ ! [X38] : (~ (~ p2(X38) | ! [X39] : (p3(X39) | ~ r1(X38, X39))) | ~ r1(X34, X38))) | ~ r1(X16, X34)) | ~ ! [X40] : (~ ! [X41] : (~ p2(X41) | ! [X42] : (p3(X42) | ~ r1(X41, X42)) | ~ r1(X40, X41)) | ~ r1(X16, X40))) | ~ r1(X0, X16)) | ~ ! [X43] : (~ (! [X44] : (~ ! [X45] : (~ p1(X45) | ~ r1(X44, X45)) | ~ r1(X43, X44)) & ~ ! [X46] : (! [X47] : (~ ! [X48] : (~ p1(X48) | ~ r1(X47, X48)) | ~ r1(X46, X47)) | ~ r1(X43, X46))) | ~ r1(X0, X43)) | ~ ! [X49] : (~ p4(X49) | ~ r1(X0, X49))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (p1(X0) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (~ ! [X0] : (~ p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | (! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (~ ! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & p2(X0) & ~ ! [X1] : (~ (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0))) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (~ ! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ p4(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (p1(X0) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ ((! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (~ ! [X0] : (~ p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | (! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ ! [X0] : (~ (~ ! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & p2(X0) & ~ ! [X1] : (~ (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0))) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (! [X0] : (~ ! [X1] : (~ p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & ~ ! [X0] : (! [X1] : (~ ! [X0] : (~ p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ~ ! [X1] : (~ p4(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL652+1.001.p', main)).
fof(f84, plain, ! [X16] : (~ r1(sK16, X16) | sP1(X16)), inference(cnf_transformation, [], [f51])).
fof(f124984, plain, (~ sP1(sK18) | r1(sK18, sK9(sK18)) | (~ spl28_540 | ~ spl28_1916)), inference(resolution, [], [f21411, f4216])).
fof(f4216, plain, (! [X1] : (~ r1(X1, sK7(sK18)) | ~ sP1(X1) | r1(X1, sK9(X1))) | ~ spl28_540), inference(avatar_component_clause, [], [f4215])).
fof(f4215, plain, (spl28_540 <=> ! [X1] : (r1(X1, sK9(X1)) | ~ sP1(X1) | ~ r1(X1, sK7(sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl28_540])])).
fof(f21411, plain, (r1(sK18, sK7(sK18)) | ~ spl28_1916), inference(avatar_component_clause, [], [f21410])).
fof(f21410, plain, (spl28_1916 <=> r1(sK18, sK7(sK18))), introduced(avatar_definition, [new_symbols(naming, [spl28_1916])])).
fof(f124733, plain, (~ spl28_109 | spl28_1916), inference(avatar_contradiction_clause, [], [f124732])).
fof(f124732, plain, ($false | (~ spl28_109 | spl28_1916)), inference(subsumption_resolution, [], [f124731, f106])).
fof(f106, plain, sP2(sK18), inference(resolution, [], [f86, f94])).
fof(f86, plain, ! [X16] : (~ r1(sK16, X16) | sP2(X16)), inference(cnf_transformation, [], [f51])).
fof(f124731, plain, (~ sP2(sK18) | (~ spl28_109 | spl28_1916)), inference(subsumption_resolution, [], [f124650, f21412])).
fof(f21412, plain, (~ r1(sK18, sK7(sK18)) | spl28_1916), inference(avatar_component_clause, [], [f21410])).
fof(f124650, plain, (r1(sK18, sK7(sK18)) | ~ sP2(sK18) | ~ spl28_109), inference(resolution, [], [f830, f1015])).
fof(f1015, plain, ! [X35] : (~ r1(sK18, sK6(X35)) | r1(X35, sK7(X35)) | ~ sP2(X35)), inference(subsumption_resolution, [], [f998, f146])).
fof(f146, plain, ! [X1] : (~ r1(sK18, X1) | p2(sK24(X1))), inference(resolution, [], [f82, f94])).
fof(f82, plain, ! [X17, X16] : (~ r1(sK16, X16) | ~ r1(X16, X17) | p2(sK24(X17))), inference(cnf_transformation, [], [f51])).
fof(f998, plain, ! [X35] : (~ r1(sK18, sK6(X35)) | ~ p2(sK24(sK6(X35))) | r1(X35, sK7(X35)) | ~ sP2(X35)), inference(resolution, [], [f200, f59])).
fof(f59, plain, ! [X2, X0] : (~ r1(sK6(X0), X2) | ~ p2(X2) | r1(X0, sK7(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0] : ((! [X2] : (~ p2(X2) | ~ r1(sK6(X0), X2)) & r1(X0, sK6(X0))) | (p2(sK7(X0)) & r1(sK7(X0), sK8(X0)) & r1(X0, sK7(X0))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7, sK8])], [f20, f23, f22, f21])).
fof(f21, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) => (! [X2] : (~ p2(X2) | ~ r1(sK6(X0), X2)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X0] : (? [X3] : (p2(X3) & ? [X4] : r1(X3, X4) & r1(X0, X3)) => (p2(sK7(X0)) & ? [X4] : r1(sK7(X0), X4) & r1(X0, sK7(X0)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X0] : (? [X4] : r1(sK7(X0), X4) => r1(sK7(X0), sK8(X0))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) | ? [X3] : (p2(X3) & ? [X4] : r1(X3, X4) & r1(X0, X3)) | ~ sP2(X0)), inference(rectify, [], [f19])).
fof(f19, plain, ! [X16] : (? [X17] : (! [X18] : (~ p2(X18) | ~ r1(X17, X18)) & r1(X16, X17)) | ? [X19] : (p2(X19) & ? [X20] : r1(X19, X20) & r1(X16, X19)) | ~ sP2(X16)), inference(nnf_transformation, [], [f11])).
fof(f200, plain, ! [X1] : (r1(X1, sK24(X1)) | ~ r1(sK18, X1)), inference(resolution, [], [f80, f94])).
fof(f80, plain, ! [X17, X16] : (~ r1(sK16, X16) | ~ r1(X16, X17) | r1(X17, sK24(X17))), inference(cnf_transformation, [], [f51])).
fof(f830, plain, (r1(sK18, sK6(sK18)) | ~ spl28_109), inference(avatar_component_clause, [], [f828])).
fof(f828, plain, (spl28_109 <=> r1(sK18, sK6(sK18))), introduced(avatar_definition, [new_symbols(naming, [spl28_109])])).
fof(f124648, plain, (spl28_1921 | ~ spl28_173 | ~ spl28_1916), inference(avatar_split_clause, [], [f21661, f21410, f1489, f21436])).
fof(f21436, plain, (spl28_1921 <=> ! [X7] : (~ r1(sK7(sK18), X7) | p2(sK10(X7)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1921])])).
fof(f21661, plain, (! [X12] : (~ r1(sK7(sK18), X12) | p2(sK10(X12))) | (~ spl28_173 | ~ spl28_1916)), inference(resolution, [], [f21411, f4807])).
fof(f4807, plain, (! [X0, X1] : (~ r1(sK18, X1) | ~ r1(X1, X0) | p2(sK10(X0))) | ~ spl28_173), inference(subsumption_resolution, [], [f4796, f102])).
fof(f4796, plain, (! [X0, X1] : (p2(sK10(X0)) | ~ r1(X1, X0) | ~ r1(sK18, X1) | ~ sP1(sK18)) | ~ spl28_173), inference(resolution, [], [f1019, f1491])).
fof(f1491, plain, (r1(sK18, sK9(sK18)) | ~ spl28_173), inference(avatar_component_clause, [], [f1489])).
fof(f1019, plain, ! [X45, X43, X44] : (~ r1(sK18, sK9(X43)) | p2(sK10(X44)) | ~ r1(X45, X44) | ~ r1(X43, X45) | ~ sP1(X43)), inference(subsumption_resolution, [], [f1002, f146])).
fof(f1002, plain, ! [X45, X43, X44] : (~ r1(sK18, sK9(X43)) | ~ p2(sK24(sK9(X43))) | p2(sK10(X44)) | ~ r1(X45, X44) | ~ r1(X43, X45) | ~ sP1(X43)), inference(resolution, [], [f200, f67])).
fof(f67, plain, ! [X4, X2, X0, X3] : (~ r1(sK9(X0), X2) | ~ p2(X2) | p2(sK10(X4)) | ~ r1(X3, X4) | ~ r1(X0, X3) | ~ sP1(X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : ((! [X2] : (~ p2(X2) | ~ r1(sK9(X0), X2)) & r1(X0, sK9(X0))) | ! [X3] : (! [X4] : ((p2(sK10(X4)) & r1(sK10(X4), sK11(X4)) & r1(X4, sK10(X4))) | ~ r1(X3, X4)) | ~ r1(X0, X3)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f26, f29, f28, f27])).
fof(f27, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) => (! [X2] : (~ p2(X2) | ~ r1(sK9(X0), X2)) & r1(X0, sK9(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X4] : (? [X5] : (p2(X5) & ? [X6] : r1(X5, X6) & r1(X4, X5)) => (p2(sK10(X4)) & ? [X6] : r1(sK10(X4), X6) & r1(X4, sK10(X4)))), introduced(choice_axiom, [])).
fof(f29, plain, ! [X4] : (? [X6] : r1(sK10(X4), X6) => r1(sK10(X4), sK11(X4))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X0] : (? [X1] : (! [X2] : (~ p2(X2) | ~ r1(X1, X2)) & r1(X0, X1)) | ! [X3] : (! [X4] : (? [X5] : (p2(X5) & ? [X6] : r1(X5, X6) & r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X0, X3)) | ~ sP1(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X16] : (? [X28] : (! [X29] : (~ p2(X29) | ~ r1(X28, X29)) & r1(X16, X28)) | ! [X30] : (! [X31] : (? [X32] : (p2(X32) & ? [X33] : r1(X32, X33) & r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X16, X30)) | ~ sP1(X16)), inference(nnf_transformation, [], [f10])).
fof(f124637, plain, (spl28_2535 | ~ spl28_544 | ~ spl28_545 | ~ spl28_1047), inference(avatar_split_clause, [], [f124288, f9600, f4238, f4233, f31541])).
fof(f31541, plain, (spl28_2535 <=> p2(sK13(sK4(sK12(sK7(sK18)))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2535])])).
fof(f4233, plain, (spl28_544 <=> ! [X4] : (r1(X4, sK4(X4)) | ~ r1(sK7(sK18), X4))), introduced(avatar_definition, [new_symbols(naming, [spl28_544])])).
fof(f4238, plain, (spl28_545 <=> r1(sK7(sK18), sK12(sK7(sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl28_545])])).
fof(f9600, plain, (spl28_1047 <=> ! [X4] : (~ r1(sK12(sK7(sK18)), X4) | p2(sK13(X4)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1047])])).
fof(f124288, plain, (p2(sK13(sK4(sK12(sK7(sK18))))) | (~ spl28_544 | ~ spl28_545 | ~ spl28_1047)), inference(resolution, [], [f9601, f23781])).
fof(f23781, plain, (r1(sK12(sK7(sK18)), sK4(sK12(sK7(sK18)))) | (~ spl28_544 | ~ spl28_545)), inference(resolution, [], [f4240, f4234])).
fof(f4234, plain, (! [X4] : (~ r1(sK7(sK18), X4) | r1(X4, sK4(X4))) | ~ spl28_544), inference(avatar_component_clause, [], [f4233])).
fof(f4240, plain, (r1(sK7(sK18), sK12(sK7(sK18))) | ~ spl28_545), inference(avatar_component_clause, [], [f4238])).
fof(f9601, plain, (! [X4] : (~ r1(sK12(sK7(sK18)), X4) | p2(sK13(X4))) | ~ spl28_1047), inference(avatar_component_clause, [], [f9600])).
fof(f124636, plain, (spl28_1044 | ~ spl28_675 | spl28_1046 | ~ spl28_673), inference(avatar_split_clause, [], [f124155, f5601, f9595, f5611, f9587])).
fof(f9587, plain, (spl28_1044 <=> ! [X1] : (~ r1(X1, sK7(sK18)) | ~ sP0(X1))), introduced(avatar_definition, [new_symbols(naming, [spl28_1044])])).
fof(f5611, plain, (spl28_675 <=> p2(sK10(sK15(sK7(sK18))))), introduced(avatar_definition, [new_symbols(naming, [spl28_675])])).
fof(f9595, plain, (spl28_1046 <=> ! [X2] : (~ r1(sK12(sK7(sK18)), X2) | r1(X2, sK13(X2)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1046])])).
fof(f5601, plain, (spl28_673 <=> r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18))))), introduced(avatar_definition, [new_symbols(naming, [spl28_673])])).
fof(f124155, plain, (! [X2, X3] : (~ r1(sK12(sK7(sK18)), X2) | ~ p2(sK10(sK15(sK7(sK18)))) | r1(X2, sK13(X2)) | ~ r1(X3, sK7(sK18)) | ~ sP0(X3)) | ~ spl28_673), inference(resolution, [], [f5603, f71])).
fof(f71, plain, ! [X0, X7, X3, X1] : (~ r1(sK15(X1), X7) | ~ r1(sK12(X1), X3) | ~ p2(X7) | r1(X3, sK13(X3)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ! [X0] : (! [X1] : ((! [X3] : ((p2(sK13(X3)) & r1(sK13(X3), sK14(X3)) & r1(X3, sK13(X3))) | ~ r1(sK12(X1), X3)) & r1(X1, sK12(X1))) | (! [X7] : (~ p2(X7) | ~ r1(sK15(X1), X7)) & r1(X1, sK15(X1))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12, sK13, sK14, sK15])], [f32, f36, f35, f34, f33])).
fof(f33, plain, ! [X1] : (? [X2] : (! [X3] : (? [X4] : (p2(X4) & ? [X5] : r1(X4, X5) & r1(X3, X4)) | ~ r1(X2, X3)) & r1(X1, X2)) => (! [X3] : (? [X4] : (p2(X4) & ? [X5] : r1(X4, X5) & r1(X3, X4)) | ~ r1(sK12(X1), X3)) & r1(X1, sK12(X1)))), introduced(choice_axiom, [])).
fof(f34, plain, ! [X3] : (? [X4] : (p2(X4) & ? [X5] : r1(X4, X5) & r1(X3, X4)) => (p2(sK13(X3)) & ? [X5] : r1(sK13(X3), X5) & r1(X3, sK13(X3)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X3] : (? [X5] : r1(sK13(X3), X5) => r1(sK13(X3), sK14(X3))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X6] : (! [X7] : (~ p2(X7) | ~ r1(X6, X7)) & r1(X1, X6)) => (! [X7] : (~ p2(X7) | ~ r1(sK15(X1), X7)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f32, plain, ! [X0] : (! [X1] : (? [X2] : (! [X3] : (? [X4] : (p2(X4) & ? [X5] : r1(X4, X5) & r1(X3, X4)) | ~ r1(X2, X3)) & r1(X1, X2)) | ? [X6] : (! [X7] : (~ p2(X7) | ~ r1(X6, X7)) & r1(X1, X6)) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f31])).
fof(f31, plain, ! [X16] : (! [X21] : (? [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : r1(X24, X25) & r1(X23, X24)) | ~ r1(X22, X23)) & r1(X21, X22)) | ? [X26] : (! [X27] : (~ p2(X27) | ~ r1(X26, X27)) & r1(X21, X26)) | ~ r1(X16, X21)) | ~ sP0(X16)), inference(nnf_transformation, [], [f9])).
fof(f5603, plain, (r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18)))) | ~ spl28_673), inference(avatar_component_clause, [], [f5601])).
fof(f124635, plain, (spl28_675 | ~ spl28_546 | ~ spl28_1921), inference(avatar_split_clause, [], [f123827, f21436, f4242, f5611])).
fof(f4242, plain, (spl28_546 <=> r1(sK7(sK18), sK15(sK7(sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl28_546])])).
fof(f123827, plain, (p2(sK10(sK15(sK7(sK18)))) | (~ spl28_546 | ~ spl28_1921)), inference(resolution, [], [f4244, f21437])).
fof(f21437, plain, (! [X7] : (~ r1(sK7(sK18), X7) | p2(sK10(X7))) | ~ spl28_1921), inference(avatar_component_clause, [], [f21436])).
fof(f4244, plain, (r1(sK7(sK18), sK15(sK7(sK18))) | ~ spl28_546), inference(avatar_component_clause, [], [f4242])).
fof(f124634, plain, (~ spl28_544 | ~ spl28_545 | ~ spl28_1046 | spl28_2534), inference(avatar_contradiction_clause, [], [f124633])).
fof(f124633, plain, ($false | (~ spl28_544 | ~ spl28_545 | ~ spl28_1046 | spl28_2534)), inference(subsumption_resolution, [], [f124604, f31536])).
fof(f31536, plain, (~ r1(sK4(sK12(sK7(sK18))), sK13(sK4(sK12(sK7(sK18))))) | spl28_2534), inference(avatar_component_clause, [], [f31535])).
fof(f31535, plain, (spl28_2534 <=> r1(sK4(sK12(sK7(sK18))), sK13(sK4(sK12(sK7(sK18)))))), introduced(avatar_definition, [new_symbols(naming, [spl28_2534])])).
fof(f124604, plain, (r1(sK4(sK12(sK7(sK18))), sK13(sK4(sK12(sK7(sK18))))) | (~ spl28_544 | ~ spl28_545 | ~ spl28_1046)), inference(resolution, [], [f9596, f23781])).
fof(f9596, plain, (! [X2] : (~ r1(sK12(sK7(sK18)), X2) | r1(X2, sK13(X2))) | ~ spl28_1046), inference(avatar_component_clause, [], [f9595])).
fof(f123821, plain, (spl28_1044 | spl28_546 | ~ spl28_544 | ~ spl28_545 | spl28_2534), inference(avatar_split_clause, [], [f123820, f31535, f4238, f4233, f4242, f9587])).
fof(f123820, plain, (! [X4] : (r1(sK7(sK18), sK15(sK7(sK18))) | ~ r1(X4, sK7(sK18)) | ~ sP0(X4)) | (~ spl28_544 | ~ spl28_545 | spl28_2534)), inference(subsumption_resolution, [], [f31492, f31536])).
fof(f31492, plain, (! [X4] : (r1(sK4(sK12(sK7(sK18))), sK13(sK4(sK12(sK7(sK18))))) | r1(sK7(sK18), sK15(sK7(sK18))) | ~ r1(X4, sK7(sK18)) | ~ sP0(X4)) | (~ spl28_544 | ~ spl28_545)), inference(resolution, [], [f23781, f70])).
fof(f70, plain, ! [X0, X3, X1] : (~ r1(sK12(X1), X3) | r1(X3, sK13(X3)) | r1(X1, sK15(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f123819, plain, (~ spl28_1044 | ~ spl28_1916), inference(avatar_contradiction_clause, [], [f123818])).
fof(f123818, plain, ($false | (~ spl28_1044 | ~ spl28_1916)), inference(subsumption_resolution, [], [f123813, f104])).
fof(f104, plain, sP0(sK18), inference(resolution, [], [f85, f94])).
fof(f85, plain, ! [X16] : (~ r1(sK16, X16) | sP0(X16)), inference(cnf_transformation, [], [f51])).
fof(f123813, plain, (~ sP0(sK18) | (~ spl28_1044 | ~ spl28_1916)), inference(resolution, [], [f9588, f21411])).
fof(f9588, plain, (! [X1] : (~ r1(X1, sK7(sK18)) | ~ sP0(X1)) | ~ spl28_1044), inference(avatar_component_clause, [], [f9587])).
fof(f123793, plain, (~ spl28_902 | ~ spl28_1916), inference(avatar_contradiction_clause, [], [f123792])).
fof(f123792, plain, ($false | (~ spl28_902 | ~ spl28_1916)), inference(subsumption_resolution, [], [f123791, f21411])).
fof(f123791, plain, (~ r1(sK18, sK7(sK18)) | ~ spl28_902), inference(resolution, [], [f8264, f96])).
fof(f96, plain, ! [X6] : (~ p2(sK19(X6)) | ~ r1(sK18, X6)), inference(cnf_transformation, [], [f51])).
fof(f8264, plain, (p2(sK19(sK7(sK18))) | ~ spl28_902), inference(avatar_component_clause, [], [f8263])).
fof(f8263, plain, (spl28_902 <=> p2(sK19(sK7(sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl28_902])])).
fof(f123762, plain, (spl28_902 | ~ spl28_1916 | ~ spl28_7912), inference(avatar_split_clause, [], [f123761, f123666, f21410, f8263])).
fof(f123666, plain, (spl28_7912 <=> ! [X10] : (p2(X10) | ~ r1(sK7(sK18), X10))), introduced(avatar_definition, [new_symbols(naming, [spl28_7912])])).
fof(f123761, plain, (p2(sK19(sK7(sK18))) | (~ spl28_1916 | ~ spl28_7912)), inference(subsumption_resolution, [], [f123743, f21411])).
fof(f123743, plain, (p2(sK19(sK7(sK18))) | ~ r1(sK18, sK7(sK18)) | ~ spl28_7912), inference(resolution, [], [f123667, f95])).
fof(f95, plain, ! [X6] : (r1(X6, sK19(X6)) | ~ r1(sK18, X6)), inference(cnf_transformation, [], [f51])).
fof(f123667, plain, (! [X10] : (~ r1(sK7(sK18), X10) | p2(X10)) | ~ spl28_7912), inference(avatar_component_clause, [], [f123666])).
fof(f123668, plain, (spl28_543 | spl28_7912 | ~ spl28_107 | ~ spl28_545 | ~ spl28_2534 | ~ spl28_2535), inference(avatar_split_clause, [], [f123664, f31541, f31535, f4238, f818, f123666, f4230])).
fof(f4230, plain, (spl28_543 <=> ! [X5] : (~ r1(X5, sK7(sK18)) | ~ sP3(X5))), introduced(avatar_definition, [new_symbols(naming, [spl28_543])])).
fof(f818, plain, (spl28_107 <=> p2(sK7(sK18))), introduced(avatar_definition, [new_symbols(naming, [spl28_107])])).
fof(f123664, plain, (! [X10, X11] : (p2(X10) | ~ r1(sK7(sK18), X10) | ~ r1(X11, sK7(sK18)) | ~ sP3(X11)) | (~ spl28_107 | ~ spl28_545 | ~ spl28_2534 | ~ spl28_2535)), inference(subsumption_resolution, [], [f123663, f820])).
fof(f820, plain, (p2(sK7(sK18)) | ~ spl28_107), inference(avatar_component_clause, [], [f818])).
fof(f123663, plain, (! [X10, X11] : (~ p2(sK7(sK18)) | p2(X10) | ~ r1(sK7(sK18), X10) | ~ r1(X11, sK7(sK18)) | ~ sP3(X11)) | (~ spl28_545 | ~ spl28_2534 | ~ spl28_2535)), inference(resolution, [], [f58005, f4240])).
fof(f58005, plain, (! [X4, X5, X3] : (~ r1(X3, sK12(sK7(sK18))) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ sP3(X5)) | (~ spl28_2534 | ~ spl28_2535)), inference(subsumption_resolution, [], [f57969, f31543])).
fof(f31543, plain, (p2(sK13(sK4(sK12(sK7(sK18))))) | ~ spl28_2535), inference(avatar_component_clause, [], [f31541])).
fof(f57969, plain, (! [X4, X5, X3] : (~ p2(sK13(sK4(sK12(sK7(sK18))))) | ~ r1(X3, sK12(sK7(sK18))) | ~ p2(X3) | p2(X4) | ~ r1(X3, X4) | ~ r1(X5, X3) | ~ sP3(X5)) | ~ spl28_2534), inference(resolution, [], [f31537, f55])).
fof(f55, plain, ! [X4, X2, X0, X5, X1] : (~ r1(sK4(X2), X4) | ~ p2(X4) | ~ r1(X1, X2) | ~ p2(X1) | p2(X5) | ~ r1(X1, X5) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (! [X1] : (! [X2] : ((! [X4] : (~ p2(X4) | ~ r1(sK4(X2), X4)) & r1(X2, sK4(X2))) | ~ r1(X1, X2)) | ~ p2(X1) | ! [X5] : ((p2(X5) & r1(X5, sK5(X5))) | ~ r1(X1, X5)) | ~ r1(X0, X1)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f15, f17, f16])).
fof(f16, plain, ! [X2] : (? [X3] : (! [X4] : (~ p2(X4) | ~ r1(X3, X4)) & r1(X2, X3)) => (! [X4] : (~ p2(X4) | ~ r1(sK4(X2), X4)) & r1(X2, sK4(X2)))), introduced(choice_axiom, [])).
fof(f17, plain, ! [X5] : (? [X6] : r1(X5, X6) => r1(X5, sK5(X5))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (! [X4] : (~ p2(X4) | ~ r1(X3, X4)) & r1(X2, X3)) | ~ r1(X1, X2)) | ~ p2(X1) | ! [X5] : ((p2(X5) & ? [X6] : r1(X5, X6)) | ~ r1(X1, X5)) | ~ r1(X0, X1)) | ~ sP3(X0)), inference(rectify, [], [f14])).
fof(f14, plain, ! [X16] : (! [X34] : (! [X35] : (? [X36] : (! [X37] : (~ p2(X37) | ~ r1(X36, X37)) & r1(X35, X36)) | ~ r1(X34, X35)) | ~ p2(X34) | ! [X38] : ((p2(X38) & ? [X39] : r1(X38, X39)) | ~ r1(X34, X38)) | ~ r1(X16, X34)) | ~ sP3(X16)), inference(nnf_transformation, [], [f12])).
fof(f31537, plain, (r1(sK4(sK12(sK7(sK18))), sK13(sK4(sK12(sK7(sK18))))) | ~ spl28_2534), inference(avatar_component_clause, [], [f31535])).
fof(f31544, plain, (spl28_1044 | spl28_2535 | ~ spl28_544 | ~ spl28_545 | spl28_546), inference(avatar_split_clause, [], [f31539, f4242, f4238, f4233, f31541, f9587])).
fof(f31539, plain, (! [X5] : (p2(sK13(sK4(sK12(sK7(sK18))))) | ~ r1(X5, sK7(sK18)) | ~ sP0(X5)) | (~ spl28_544 | ~ spl28_545 | spl28_546)), inference(subsumption_resolution, [], [f31493, f4243])).
fof(f4243, plain, (~ r1(sK7(sK18), sK15(sK7(sK18))) | spl28_546), inference(avatar_component_clause, [], [f4242])).
fof(f31493, plain, (! [X5] : (p2(sK13(sK4(sK12(sK7(sK18))))) | r1(sK7(sK18), sK15(sK7(sK18))) | ~ r1(X5, sK7(sK18)) | ~ sP0(X5)) | (~ spl28_544 | ~ spl28_545)), inference(resolution, [], [f23781, f74])).
fof(f74, plain, ! [X0, X3, X1] : (~ r1(sK12(X1), X3) | p2(sK13(X3)) | r1(X1, sK15(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f23777, plain, (spl28_545 | spl28_546 | ~ spl28_1916), inference(avatar_split_clause, [], [f23776, f21410, f4242, f4238])).
fof(f23776, plain, (r1(sK7(sK18), sK15(sK7(sK18))) | r1(sK7(sK18), sK12(sK7(sK18))) | ~ spl28_1916), inference(subsumption_resolution, [], [f21669, f104])).
fof(f21669, plain, (r1(sK7(sK18), sK15(sK7(sK18))) | r1(sK7(sK18), sK12(sK7(sK18))) | ~ sP0(sK18) | ~ spl28_1916), inference(resolution, [], [f21411, f68])).
fof(f68, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK15(X1)) | r1(X1, sK12(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f23757, plain, (~ spl28_546 | spl28_673 | ~ spl28_1915), inference(avatar_contradiction_clause, [], [f23756])).
fof(f23756, plain, ($false | (~ spl28_546 | spl28_673 | ~ spl28_1915)), inference(subsumption_resolution, [], [f23718, f5602])).
fof(f5602, plain, (~ r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18)))) | spl28_673), inference(avatar_component_clause, [], [f5601])).
fof(f23718, plain, (r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18)))) | (~ spl28_546 | ~ spl28_1915)), inference(resolution, [], [f21407, f4244])).
fof(f21407, plain, (! [X7] : (~ r1(sK7(sK18), X7) | r1(X7, sK10(X7))) | ~ spl28_1915), inference(avatar_component_clause, [], [f21406])).
fof(f21406, plain, (spl28_1915 <=> ! [X7] : (~ r1(sK7(sK18), X7) | r1(X7, sK10(X7)))), introduced(avatar_definition, [new_symbols(naming, [spl28_1915])])).
fof(f21881, plain, (spl28_543 | spl28_544 | ~ spl28_107 | ~ spl28_1916), inference(avatar_split_clause, [], [f21880, f21410, f818, f4233, f4230])).
fof(f21880, plain, (! [X8, X7] : (r1(X7, sK4(X7)) | ~ r1(X8, sK7(sK18)) | ~ sP3(X8) | ~ r1(sK7(sK18), X7)) | (~ spl28_107 | ~ spl28_1916)), inference(subsumption_resolution, [], [f21657, f820])).
fof(f21657, plain, (! [X8, X7] : (~ p2(sK7(sK18)) | r1(X7, sK4(X7)) | ~ r1(X8, sK7(sK18)) | ~ sP3(X8) | ~ r1(sK7(sK18), X7)) | ~ spl28_1916), inference(resolution, [], [f21411, f663])).
fof(f663, plain, ! [X2, X0, X1] : (~ r1(sK18, X0) | ~ p2(X0) | r1(X1, sK4(X1)) | ~ r1(X2, X0) | ~ sP3(X2) | ~ r1(X0, X1)), inference(subsumption_resolution, [], [f649, f96])).
fof(f649, plain, ! [X2, X0, X1] : (~ r1(X0, X1) | ~ p2(X0) | p2(sK19(X0)) | r1(X1, sK4(X1)) | ~ r1(X2, X0) | ~ sP3(X2) | ~ r1(sK18, X0)), inference(resolution, [], [f53, f95])).
fof(f53, plain, ! [X2, X0, X5, X1] : (~ r1(X1, X5) | ~ r1(X1, X2) | ~ p2(X1) | p2(X5) | r1(X2, sK4(X2)) | ~ r1(X0, X1) | ~ sP3(X0)), inference(cnf_transformation, [], [f18])).
fof(f21879, plain, (~ spl28_543 | ~ spl28_1916), inference(avatar_contradiction_clause, [], [f21878])).
fof(f21878, plain, ($false | (~ spl28_543 | ~ spl28_1916)), inference(subsumption_resolution, [], [f21873, f100])).
fof(f100, plain, sP3(sK18), inference(resolution, [], [f83, f94])).
fof(f83, plain, ! [X16] : (~ r1(sK16, X16) | sP3(X16)), inference(cnf_transformation, [], [f51])).
fof(f21873, plain, (~ sP3(sK18) | (~ spl28_543 | ~ spl28_1916)), inference(resolution, [], [f4231, f21411])).
fof(f4231, plain, (! [X5] : (~ r1(X5, sK7(sK18)) | ~ sP3(X5)) | ~ spl28_543), inference(avatar_component_clause, [], [f4230])).
fof(f21451, plain, (spl28_109 | spl28_1916), inference(avatar_contradiction_clause, [], [f21450])).
fof(f21450, plain, ($false | (spl28_109 | spl28_1916)), inference(subsumption_resolution, [], [f21449, f106])).
fof(f21449, plain, (~ sP2(sK18) | (spl28_109 | spl28_1916)), inference(subsumption_resolution, [], [f21448, f829])).
fof(f829, plain, (~ r1(sK18, sK6(sK18)) | spl28_109), inference(avatar_component_clause, [], [f828])).
fof(f21448, plain, (r1(sK18, sK6(sK18)) | ~ sP2(sK18) | spl28_1916), inference(resolution, [], [f21412, f56])).
fof(f56, plain, ! [X0] : (r1(X0, sK7(X0)) | r1(X0, sK6(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f24])).
fof(f21408, plain, (spl28_109 | spl28_1915 | ~ spl28_173), inference(avatar_split_clause, [], [f21404, f1489, f21406, f828])).
fof(f21404, plain, (! [X7] : (~ r1(sK7(sK18), X7) | r1(X7, sK10(X7)) | r1(sK18, sK6(sK18))) | ~ spl28_173), inference(subsumption_resolution, [], [f21284, f106])).
fof(f21284, plain, (! [X7] : (~ r1(sK7(sK18), X7) | r1(X7, sK10(X7)) | r1(sK18, sK6(sK18)) | ~ sP2(sK18)) | ~ spl28_173), inference(resolution, [], [f4894, f56])).
fof(f4894, plain, (! [X0, X1] : (~ r1(sK18, X1) | ~ r1(X1, X0) | r1(X0, sK10(X0))) | ~ spl28_173), inference(subsumption_resolution, [], [f4883, f102])).
fof(f4883, plain, (! [X0, X1] : (r1(X0, sK10(X0)) | ~ r1(X1, X0) | ~ r1(sK18, X1) | ~ sP1(sK18)) | ~ spl28_173), inference(resolution, [], [f1018, f1491])).
fof(f1018, plain, ! [X41, X42, X40] : (~ r1(sK18, sK9(X40)) | r1(X41, sK10(X41)) | ~ r1(X42, X41) | ~ r1(X40, X42) | ~ sP1(X40)), inference(subsumption_resolution, [], [f1001, f146])).
fof(f1001, plain, ! [X41, X42, X40] : (~ r1(sK18, sK9(X40)) | ~ p2(sK24(sK9(X40))) | r1(X41, sK10(X41)) | ~ r1(X42, X41) | ~ r1(X40, X42) | ~ sP1(X40)), inference(resolution, [], [f200, f65])).
fof(f65, plain, ! [X4, X2, X0, X3] : (~ r1(sK9(X0), X2) | ~ p2(X2) | r1(X4, sK10(X4)) | ~ r1(X3, X4) | ~ r1(X0, X3) | ~ sP1(X0)), inference(cnf_transformation, [], [f30])).
fof(f21389, plain, (~ spl28_109 | ~ spl28_173 | ~ spl28_546 | spl28_673), inference(avatar_contradiction_clause, [], [f21388])).
fof(f21388, plain, ($false | (~ spl28_109 | ~ spl28_173 | ~ spl28_546 | spl28_673)), inference(subsumption_resolution, [], [f21356, f5602])).
fof(f21356, plain, (r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18)))) | (~ spl28_109 | ~ spl28_173 | ~ spl28_546)), inference(resolution, [], [f21278, f4244])).
fof(f21278, plain, (! [X1] : (~ r1(sK7(sK18), X1) | r1(X1, sK10(X1))) | (~ spl28_109 | ~ spl28_173)), inference(resolution, [], [f4894, f3895])).
fof(f3895, plain, (r1(sK18, sK7(sK18)) | ~ spl28_109), inference(subsumption_resolution, [], [f3892, f106])).
fof(f3892, plain, (r1(sK18, sK7(sK18)) | ~ sP2(sK18) | ~ spl28_109), inference(resolution, [], [f1015, f830])).
fof(f9604, plain, (spl28_1044 | spl28_545 | ~ spl28_673 | ~ spl28_675), inference(avatar_split_clause, [], [f9603, f5611, f5601, f4238, f9587])).
fof(f9603, plain, (! [X6] : (r1(sK7(sK18), sK12(sK7(sK18))) | ~ r1(X6, sK7(sK18)) | ~ sP0(X6)) | (~ spl28_673 | ~ spl28_675)), inference(subsumption_resolution, [], [f9561, f5613])).
fof(f5613, plain, (p2(sK10(sK15(sK7(sK18)))) | ~ spl28_675), inference(avatar_component_clause, [], [f5611])).
fof(f9561, plain, (! [X6] : (~ p2(sK10(sK15(sK7(sK18)))) | r1(sK7(sK18), sK12(sK7(sK18))) | ~ r1(X6, sK7(sK18)) | ~ sP0(X6)) | ~ spl28_673), inference(resolution, [], [f5603, f69])).
fof(f69, plain, ! [X0, X7, X1] : (~ r1(sK15(X1), X7) | ~ p2(X7) | r1(X1, sK12(X1)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f9602, plain, (spl28_1044 | spl28_1047 | ~ spl28_673 | ~ spl28_675), inference(avatar_split_clause, [], [f9598, f5611, f5601, f9600, f9587])).
fof(f9598, plain, (! [X4, X5] : (~ r1(sK12(sK7(sK18)), X4) | p2(sK13(X4)) | ~ r1(X5, sK7(sK18)) | ~ sP0(X5)) | (~ spl28_673 | ~ spl28_675)), inference(subsumption_resolution, [], [f9560, f5613])).
fof(f9560, plain, (! [X4, X5] : (~ r1(sK12(sK7(sK18)), X4) | ~ p2(sK10(sK15(sK7(sK18)))) | p2(sK13(X4)) | ~ r1(X5, sK7(sK18)) | ~ sP0(X5)) | ~ spl28_673), inference(resolution, [], [f5603, f75])).
fof(f75, plain, ! [X0, X7, X3, X1] : (~ r1(sK15(X1), X7) | ~ r1(sK12(X1), X3) | ~ p2(X7) | p2(sK13(X3)) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f37])).
fof(f5614, plain, (spl28_540 | spl28_675 | ~ spl28_546), inference(avatar_split_clause, [], [f5589, f4242, f5611, f4215])).
fof(f5589, plain, (! [X8] : (p2(sK10(sK15(sK7(sK18)))) | r1(X8, sK9(X8)) | ~ r1(X8, sK7(sK18)) | ~ sP1(X8)) | ~ spl28_546), inference(resolution, [], [f4244, f64])).
fof(f64, plain, ! [X4, X0, X3] : (~ r1(X3, X4) | p2(sK10(X4)) | r1(X0, sK9(X0)) | ~ r1(X0, X3) | ~ sP1(X0)), inference(cnf_transformation, [], [f30])).
fof(f5604, plain, (spl28_540 | spl28_673 | ~ spl28_546), inference(avatar_split_clause, [], [f5587, f4242, f5601, f4215])).
fof(f5587, plain, (! [X6] : (r1(sK15(sK7(sK18)), sK10(sK15(sK7(sK18)))) | r1(X6, sK9(X6)) | ~ r1(X6, sK7(sK18)) | ~ sP1(X6)) | ~ spl28_546), inference(resolution, [], [f4244, f62])).
fof(f62, plain, ! [X4, X0, X3] : (~ r1(X3, X4) | r1(X4, sK10(X4)) | r1(X0, sK9(X0)) | ~ r1(X0, X3) | ~ sP1(X0)), inference(cnf_transformation, [], [f30])).
fof(f5580, plain, spl28_107, inference(avatar_split_clause, [], [f5579, f818])).
fof(f5579, plain, p2(sK7(sK18)), inference(subsumption_resolution, [], [f3627, f106])).
fof(f3627, plain, (p2(sK7(sK18)) | ~ sP2(sK18)), inference(duplicate_literal_removal, [], [f3626])).
fof(f3626, plain, (p2(sK7(sK18)) | ~ sP2(sK18) | p2(sK7(sK18)) | ~ sP2(sK18)), inference(resolution, [], [f1016, f58])).
fof(f58, plain, ! [X0] : (r1(X0, sK6(X0)) | p2(sK7(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f24])).
fof(f1016, plain, ! [X36] : (~ r1(sK18, sK6(X36)) | p2(sK7(X36)) | ~ sP2(X36)), inference(subsumption_resolution, [], [f999, f146])).
fof(f999, plain, ! [X36] : (~ r1(sK18, sK6(X36)) | ~ p2(sK24(sK6(X36))) | p2(sK7(X36)) | ~ sP2(X36)), inference(resolution, [], [f200, f61])).
fof(f61, plain, ! [X2, X0] : (~ r1(sK6(X0), X2) | ~ p2(X2) | p2(sK7(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f24])).