fof(f266292, plain, $false, inference(avatar_sat_refutation, [], [f197, f201, f209, f214, f224, f10928, f18274, f18445, f19399, f40782, f61210, f65893, f89221, f96286, f97791, f97808, f97812, f232796, f232802, f232806, f232993, f260760, f265027, f266291])).
fof(f266291, plain, (~ spl44_8 | ~ spl44_10 | ~ spl44_11 | ~ spl44_76), inference(avatar_contradiction_clause, [], [f266290])).
fof(f266290, plain, ($false | (~ spl44_8 | ~ spl44_10 | ~ spl44_11 | ~ spl44_76)), inference(subsumption_resolution, [], [f266289, f265247])).
fof(f265247, plain, (p3(sK38(sK43)) | (~ spl44_8 | ~ spl44_76)), inference(resolution, [], [f192, f592])).
fof(f592, plain, (r1(sK32, sK43) | ~ spl44_76), inference(avatar_component_clause, [], [f591])).
fof(f591, plain, (spl44_76 <=> r1(sK32, sK43)), introduced(avatar_definition, [new_symbols(naming, [spl44_76])])).
fof(f192, plain, (! [X15] : (~ r1(sK32, X15) | p3(sK38(X15))) | ~ spl44_8), inference(avatar_component_clause, [], [f191])).
fof(f191, plain, (spl44_8 <=> ! [X15] : (p3(sK38(X15)) | ~ r1(sK32, X15))), introduced(avatar_definition, [new_symbols(naming, [spl44_8])])).
fof(f266289, plain, (~ p3(sK38(sK43)) | (~ spl44_10 | ~ spl44_11 | ~ spl44_76)), inference(subsumption_resolution, [], [f265802, f592])).
fof(f265802, plain, (~ r1(sK32, sK43) | ~ p3(sK38(sK43)) | (~ spl44_10 | ~ spl44_11)), inference(resolution, [], [f200, f204])).
fof(f204, plain, (! [X23] : (~ r1(sK43, X23) | ~ p3(X23)) | ~ spl44_11), inference(avatar_component_clause, [], [f203])).
fof(f203, plain, (spl44_11 <=> ! [X23] : (~ p3(X23) | ~ r1(sK43, X23))), introduced(avatar_definition, [new_symbols(naming, [spl44_11])])).
fof(f200, plain, (! [X15] : (r1(X15, sK38(X15)) | ~ r1(sK32, X15)) | ~ spl44_10), inference(avatar_component_clause, [], [f199])).
fof(f199, plain, (spl44_10 <=> ! [X15] : (r1(X15, sK38(X15)) | ~ r1(sK32, X15))), introduced(avatar_definition, [new_symbols(naming, [spl44_10])])).
fof(f265027, plain, (spl44_635 | spl44_214 | ~ spl44_227 | spl44_560 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531), inference(avatar_split_clause, [], [f265026, f96261, f65866, f40762, f28413, f19033, f18085, f35437])).
fof(f35437, plain, (spl44_635 <=> ! [X88] : (~ r1(X88, sK28(sK32)) | ~ sP1(X88))), introduced(avatar_definition, [new_symbols(naming, [spl44_635])])).
fof(f18085, plain, (spl44_214 <=> p2(sK28(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl44_214])])).
fof(f19033, plain, (spl44_227 <=> r1(sK32, sK28(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl44_227])])).
fof(f28413, plain, (spl44_560 <=> p2(sK33(sK28(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl44_560])])).
fof(f40762, plain, (spl44_729 <=> p2(sK26(sK33(sK28(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_729])])).
fof(f65866, plain, (spl44_1654 <=> r1(sK33(sK28(sK32)), sK26(sK33(sK28(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_1654])])).
fof(f96261, plain, (spl44_2531 <=> r1(sK26(sK33(sK28(sK32))), sK27(sK33(sK28(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_2531])])).
fof(f265026, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK28(sK32))) | (spl44_214 | ~ spl44_227 | spl44_560 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(subsumption_resolution, [], [f265025, f19034])).
fof(f19034, plain, (r1(sK32, sK28(sK32)) | ~ spl44_227), inference(avatar_component_clause, [], [f19033])).
fof(f265025, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK28(sK32)) | ~ r1(sK32, sK28(sK32))) | (spl44_214 | ~ spl44_227 | spl44_560 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(subsumption_resolution, [], [f265023, f18086])).
fof(f18086, plain, (~ p2(sK28(sK32)) | spl44_214), inference(avatar_component_clause, [], [f18085])).
fof(f265023, plain, (! [X0] : (~ sP1(X0) | ~ r1(X0, sK28(sK32)) | p2(sK28(sK32)) | ~ r1(sK32, sK28(sK32))) | (spl44_214 | ~ spl44_227 | spl44_560 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(resolution, [], [f264076, f744])).
fof(f744, plain, ! [X14, X15] : (r1(X14, sK33(X15)) | ~ r1(X14, X15) | p2(X15) | ~ r1(sK32, X15)), inference(resolution, [], [f83, f157])).
fof(f157, plain, ! [X1] : (r1(X1, sK33(X1)) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f81])).
fof(f81, plain, (! [X1] : ((! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK33(X1), X3)) & ~ p2(sK33(X1)) & r1(X1, sK33(X1))) | p2(X1) | ~ r1(sK32, X1)) & ((! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(sK34, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(sK34, X9)) & ~ p2(sK34)) | sP3(sK34)) & r1(sK32, sK34)) | sP7(sK32)) & ! [X11] : ((p1(sK35(X11)) & (~ p1(sK36(X11)) & r1(sK35(X11), sK36(X11))) & r1(X11, sK35(X11))) | p1(X11) | ~ r1(sK32, X11)) & (~ p1(sK37) & r1(sK32, sK37)) & (sP1(sK32) | ! [X15] : ((p3(sK38(X15)) & r1(X15, sK38(X15))) | ~ r1(sK32, X15))) & ! [X17] : ((p4(sK39(X17)) & (~ p4(sK40(X17)) & r1(sK39(X17), sK40(X17))) & r1(X17, sK39(X17))) | p4(X17) | ~ r1(sK32, X17)) & (~ p4(sK41) & r1(sK32, sK41)) & (sP0(sK32) | ((! [X23] : (~ p3(X23) | ~ r1(sK43, X23)) & r1(sK42, sK43)) & ~ p1(sK42) & r1(sK32, sK42)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43])], [f68, f80, f79, f78, f77, f76, f75, f74, f73, f72, f71, f70, f69])).
fof(f69, plain, (? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP3(X5)) & r1(X0, X5)) | sP7(X0)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(X0, X11)) & ? [X14] : (~ p1(X14) & r1(X0, X14)) & (sP1(X0) | ! [X15] : (? [X16] : (p3(X16) & r1(X15, X16)) | ~ r1(X0, X15))) & ! [X17] : (? [X18] : (p4(X18) & ? [X19] : (~ p4(X19) & r1(X18, X19)) & r1(X17, X18)) | p4(X17) | ~ r1(X0, X17)) & ? [X20] : (~ p4(X20) & r1(X0, X20)) & (sP0(X0) | ? [X21] : (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(X21, X22)) & ~ p1(X21) & r1(X0, X21)))) => (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(sK32, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP3(X5)) & r1(sK32, X5)) | sP7(sK32)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(sK32, X11)) & ? [X14] : (~ p1(X14) & r1(sK32, X14)) & (sP1(sK32) | ! [X15] : (? [X16] : (p3(X16) & r1(X15, X16)) | ~ r1(sK32, X15))) & ! [X17] : (? [X18] : (p4(X18) & ? [X19] : (~ p4(X19) & r1(X18, X19)) & r1(X17, X18)) | p4(X17) | ~ r1(sK32, X17)) & ? [X20] : (~ p4(X20) & r1(sK32, X20)) & (sP0(sK32) | ? [X21] : (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(X21, X22)) & ~ p1(X21) & r1(sK32, X21))))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) => (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(sK33(X1), X3)) & ~ p2(sK33(X1)) & r1(X1, sK33(X1)))), introduced(choice_axiom, [])).
fof(f71, plain, (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP3(X5)) & r1(sK32, X5)) => (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(sK34, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(sK34, X9)) & ~ p2(sK34)) | sP3(sK34)) & r1(sK32, sK34))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) => (p1(sK35(X11)) & ? [X13] : (~ p1(X13) & r1(sK35(X11), X13)) & r1(X11, sK35(X11)))), introduced(choice_axiom, [])).
fof(f73, plain, ! [X11] : (? [X13] : (~ p1(X13) & r1(sK35(X11), X13)) => (~ p1(sK36(X11)) & r1(sK35(X11), sK36(X11)))), introduced(choice_axiom, [])).
fof(f74, plain, (? [X14] : (~ p1(X14) & r1(sK32, X14)) => (~ p1(sK37) & r1(sK32, sK37))), introduced(choice_axiom, [])).
fof(f75, plain, ! [X15] : (? [X16] : (p3(X16) & r1(X15, X16)) => (p3(sK38(X15)) & r1(X15, sK38(X15)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X17] : (? [X18] : (p4(X18) & ? [X19] : (~ p4(X19) & r1(X18, X19)) & r1(X17, X18)) => (p4(sK39(X17)) & ? [X19] : (~ p4(X19) & r1(sK39(X17), X19)) & r1(X17, sK39(X17)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X17] : (? [X19] : (~ p4(X19) & r1(sK39(X17), X19)) => (~ p4(sK40(X17)) & r1(sK39(X17), sK40(X17)))), introduced(choice_axiom, [])).
fof(f78, plain, (? [X20] : (~ p4(X20) & r1(sK32, X20)) => (~ p4(sK41) & r1(sK32, sK41))), introduced(choice_axiom, [])).
fof(f79, plain, (? [X21] : (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(X21, X22)) & ~ p1(X21) & r1(sK32, X21)) => (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(sK42, X22)) & ~ p1(sK42) & r1(sK32, sK42))), introduced(choice_axiom, [])).
fof(f80, plain, (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(sK42, X22)) => (! [X23] : (~ p3(X23) | ~ r1(sK43, X23)) & r1(sK42, sK43))), introduced(choice_axiom, [])).
fof(f68, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(X5, X6)) & ((! [X9] : (~ p2(X9) | ! [X10] : (p2(X10) | ~ r1(X9, X10)) | ~ r1(X5, X9)) & ~ p2(X5)) | sP3(X5)) & r1(X0, X5)) | sP7(X0)) & ! [X11] : (? [X12] : (p1(X12) & ? [X13] : (~ p1(X13) & r1(X12, X13)) & r1(X11, X12)) | p1(X11) | ~ r1(X0, X11)) & ? [X14] : (~ p1(X14) & r1(X0, X14)) & (sP1(X0) | ! [X15] : (? [X16] : (p3(X16) & r1(X15, X16)) | ~ r1(X0, X15))) & ! [X17] : (? [X18] : (p4(X18) & ? [X19] : (~ p4(X19) & r1(X18, X19)) & r1(X17, X18)) | p4(X17) | ~ r1(X0, X17)) & ? [X20] : (~ p4(X20) & r1(X0, X20)) & (sP0(X0) | ? [X21] : (? [X22] : (! [X23] : (~ p3(X23) | ~ r1(X22, X23)) & r1(X21, X22)) & ~ p1(X21) & r1(X0, X21)))), inference(rectify, [], [f19])).
fof(f19, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | sP5(X6) | sP6(X6) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | sP3(X5)) & r1(X0, X5)) | sP7(X0)) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & (sP1(X0) | ! [X52] : (? [X53] : (p3(X53) & r1(X52, X53)) | ~ r1(X0, X52))) & ! [X54] : (? [X55] : (p4(X55) & ? [X56] : (~ p4(X56) & r1(X55, X56)) & r1(X54, X55)) | p4(X54) | ~ r1(X0, X54)) & ? [X57] : (~ p4(X57) & r1(X0, X57)) & (sP0(X0) | ? [X62] : (? [X63] : (! [X64] : (~ p3(X64) | ~ r1(X63, X64)) & r1(X62, X63)) & ~ p1(X62) & r1(X0, X62)))), inference(definition_folding, [], [f10, e18, e17, e16, e15, e14, e13, e12, e11])).
fof(f11, plain, ! [X0] : ((! [X58] : (? [X59] : (p2(X59) & ? [X60] : (~ p2(X60) & r1(X59, X60)) & r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) & ? [X61] : (~ p2(X61) & r1(X0, X61))) | ~ sP0(X0)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X0] : (sP0(X0) <=> (! [X58] : (? [X59] : (p2(X59) & ? [X60] : (~ p2(X60) & r1(X59, X60)) & r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) & ? [X61] : (~ p2(X61) & r1(X0, X61)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f12, plain, ! [X0] : ((! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51))) | ~ sP1(X0)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X0] : (sP1(X0) <=> (! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f13, plain, ! [X0] : (! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)) | ~ sP2(X0)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X0] : (sP2(X0) <=> ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f14, plain, ! [X5] : ((! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31))) | ~ sP3(X5)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X5] : (sP3(X5) <=> (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f15, plain, ! [X16] : (! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)) | ~ sP4(X16)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X16] : (sP4(X16) <=> ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f16, plain, ! [X6] : ((! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12))) | ~ sP5(X6)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X6] : (sP5(X6) <=> (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f17, plain, ! [X6] : (! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | sP4(X16))) | ~ r1(X6, X16)) | ~ sP6(X6)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X6] : (sP6(X6) <=> ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | sP4(X16))) | ~ r1(X6, X16))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f18, plain, ! [X0] : (((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | sP2(X0))) | ~ sP7(X0)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X0] : (sP7(X0) <=> ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | sP2(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f10, plain, ? [X0] : (! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : ((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12))) | ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))) & r1(X0, X5)) | ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40))))) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & ((! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51))) | ! [X52] : (? [X53] : (p3(X53) & r1(X52, X53)) | ~ r1(X0, X52))) & ! [X54] : (? [X55] : (p4(X55) & ? [X56] : (~ p4(X56) & r1(X55, X56)) & r1(X54, X55)) | p4(X54) | ~ r1(X0, X54)) & ? [X57] : (~ p4(X57) & r1(X0, X57)) & ((! [X58] : (? [X59] : (p2(X59) & ? [X60] : (~ p2(X60) & r1(X59, X60)) & r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) & ? [X61] : (~ p2(X61) & r1(X0, X61))) | ? [X62] : (? [X63] : (! [X64] : (~ p3(X64) | ~ r1(X63, X64)) & r1(X62, X63)) & ~ p1(X62) & r1(X0, X62)))), inference(flattening, [], [f9])).
fof(f9, plain, ? [X0] : ((! [X1] : (? [X2] : (! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) & ~ p2(X2) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (? [X5] : (! [X6] : (((! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p2(X6)) | (! [X9] : (? [X10] : (p2(X10) & ? [X11] : (~ p2(X11) & r1(X10, X11)) & r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) & ? [X12] : (? [X13] : (! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) & ~ p2(X13) & r1(X12, X13)) & r1(X6, X12)))) | ! [X16] : (((? [X17] : (p2(X17) & ? [X18] : (~ p2(X18) & r1(X17, X18)) & r1(X16, X17)) | p2(X16)) & (? [X19] : (! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) & ~ p2(X19) & r1(X16, X19)) | ! [X22] : (! [X23] : (? [X24] : (p2(X24) & ? [X25] : (~ p2(X25) & r1(X24, X25)) & r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) & ((! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) & ~ p2(X5)) | (! [X28] : (? [X29] : (p2(X29) & ? [X30] : (~ p2(X30) & r1(X29, X30)) & r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) & ? [X31] : (? [X32] : (! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) & ~ p2(X32) & r1(X31, X32)) & r1(X5, X31)))) & r1(X0, X5)) | ((? [X35] : (p2(X35) & ? [X36] : (~ p2(X36) & r1(X35, X36)) & r1(X0, X35)) | p2(X0)) & (? [X37] : (! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) & ~ p2(X37) & r1(X0, X37)) | ! [X40] : (! [X41] : (? [X42] : (p2(X42) & ? [X43] : (~ p2(X43) & r1(X42, X43)) & r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) & ! [X44] : (? [X45] : (p1(X45) & ? [X46] : (~ p1(X46) & r1(X45, X46)) & r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) & ? [X47] : (~ p1(X47) & r1(X0, X47)) & ((! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51))) | ! [X52] : (? [X53] : (p3(X53) & r1(X52, X53)) | ~ r1(X0, X52))) & ! [X54] : (? [X55] : (p4(X55) & ? [X56] : (~ p4(X56) & r1(X55, X56)) & r1(X54, X55)) | p4(X54) | ~ r1(X0, X54)) & ? [X57] : (~ p4(X57) & r1(X0, X57)) & ((! [X58] : (? [X59] : (p2(X59) & ? [X60] : (~ p2(X60) & r1(X59, X60)) & r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) & ? [X61] : (~ p2(X61) & r1(X0, X61))) | ? [X62] : (? [X63] : (! [X64] : (~ p3(X64) | ~ r1(X63, X64)) & r1(X62, X63)) & ~ p1(X62) & r1(X0, X62)))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ ((~ ! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p2(X6)) & (~ ! [X9] : (~ ! [X10] : (~ p2(X10) | ! [X11] : (p2(X11) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) | ! [X12] : (! [X13] : (~ ! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | p2(X13) | ~ r1(X12, X13)) | ~ r1(X6, X12)))) | ! [X16] : (((~ ! [X17] : (~ p2(X17) | ! [X18] : (p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | p2(X16)) & (~ ! [X19] : (~ ! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | p2(X19) | ~ r1(X16, X19)) | ! [X22] : (! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p2(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) | ((~ ! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) | p2(X5)) & (~ ! [X28] : (~ ! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) | ! [X31] : (! [X32] : (~ ! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) | p2(X32) | ~ r1(X31, X32)) | ~ r1(X5, X31)))) | ~ r1(X0, X5)) | ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X0, X35)) | p2(X0)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X0, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) | ~ ! [X44] : (~ ! [X45] : (~ p1(X45) | ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) | ! [X47] : (p1(X47) | ~ r1(X0, X47)) | ((~ ! [X48] : (~ ! [X49] : (~ p2(X49) | ! [X50] : (p2(X50) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) | ! [X51] : (p2(X51) | ~ r1(X0, X51))) & ~ ! [X52] : (~ ! [X53] : (~ p3(X53) | ~ r1(X52, X53)) | ~ r1(X0, X52))) | ~ ! [X54] : (~ ! [X55] : (~ p4(X55) | ! [X56] : (p4(X56) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | p4(X54) | ~ r1(X0, X54)) | ! [X57] : (p4(X57) | ~ r1(X0, X57)) | ((~ ! [X58] : (~ ! [X59] : (~ p2(X59) | ! [X60] : (p2(X60) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) | ! [X61] : (p2(X61) | ~ r1(X0, X61))) & ! [X62] : (! [X63] : (~ ! [X64] : (~ p3(X64) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | p1(X62) | ~ r1(X0, X62)))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X2] : (~ ! [X3] : (~ p2(X3) | ! [X4] : (p2(X4) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | p2(X2) | ~ r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X5] : (~ ! [X6] : (~ ((~ ! [X7] : (~ p2(X7) | ! [X8] : (p2(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p2(X6)) & (~ ! [X9] : (~ ! [X10] : (~ p2(X10) | ! [X11] : (p2(X11) | ~ r1(X10, X11)) | ~ r1(X9, X10)) | p2(X9) | ~ r1(X6, X9)) | ! [X12] : (! [X13] : (~ ! [X14] : (~ p2(X14) | ! [X15] : (p2(X15) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | p2(X13) | ~ r1(X12, X13)) | ~ r1(X6, X12)))) | ! [X16] : (((~ ! [X17] : (~ p2(X17) | ! [X18] : (p2(X18) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | p2(X16)) & (~ ! [X19] : (~ ! [X20] : (~ p2(X20) | ! [X21] : (p2(X21) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | p2(X19) | ~ r1(X16, X19)) | ! [X22] : (! [X23] : (~ ! [X24] : (~ p2(X24) | ! [X25] : (p2(X25) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | p2(X23) | ~ r1(X22, X23)) | ~ r1(X16, X22)))) | ~ r1(X6, X16)) | ~ r1(X5, X6)) | ((~ ! [X26] : (~ p2(X26) | ! [X27] : (p2(X27) | ~ r1(X26, X27)) | ~ r1(X5, X26)) | p2(X5)) & (~ ! [X28] : (~ ! [X29] : (~ p2(X29) | ! [X30] : (p2(X30) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | p2(X28) | ~ r1(X5, X28)) | ! [X31] : (! [X32] : (~ ! [X33] : (~ p2(X33) | ! [X34] : (p2(X34) | ~ r1(X33, X34)) | ~ r1(X32, X33)) | p2(X32) | ~ r1(X31, X32)) | ~ r1(X5, X31)))) | ~ r1(X0, X5)) | ((~ ! [X35] : (~ p2(X35) | ! [X36] : (p2(X36) | ~ r1(X35, X36)) | ~ r1(X0, X35)) | p2(X0)) & (~ ! [X37] : (~ ! [X38] : (~ p2(X38) | ! [X39] : (p2(X39) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | p2(X37) | ~ r1(X0, X37)) | ! [X40] : (! [X41] : (~ ! [X42] : (~ p2(X42) | ! [X43] : (p2(X43) | ~ r1(X42, X43)) | ~ r1(X41, X42)) | p2(X41) | ~ r1(X40, X41)) | ~ r1(X0, X40)))))) | ~ ! [X44] : (~ ! [X45] : (~ p1(X45) | ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | p1(X44) | ~ r1(X0, X44)) | ! [X47] : (p1(X47) | ~ r1(X0, X47)) | ((~ ! [X48] : (~ ! [X49] : (~ p2(X49) | ! [X50] : (p2(X50) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) | ! [X51] : (p2(X51) | ~ r1(X0, X51))) & ~ ! [X52] : (~ ! [X53] : (~ p3(X53) | ~ r1(X52, X53)) | ~ r1(X0, X52))) | ~ ! [X54] : (~ ! [X55] : (~ p4(X55) | ! [X56] : (p4(X56) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | p4(X54) | ~ r1(X0, X54)) | ! [X57] : (p4(X57) | ~ r1(X0, X57)) | ((~ ! [X58] : (~ ! [X59] : (~ p2(X59) | ! [X60] : (p2(X60) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) | ! [X61] : (p2(X61) | ~ r1(X0, X61))) & ! [X62] : (! [X63] : (~ ! [X64] : (~ p3(X64) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | p1(X62) | ~ r1(X0, X62)))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p4(X0) | ! [X1] : (p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (p4(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ! [X1] : (! [X0] : (~ ! [X1] : (~ p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))) | ! [X1] : (((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ((~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1)) & (~ ! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)))) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0)) & (~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p2(X1) | ! [X0] : (p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p2(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)))))) | ~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ~ ! [X1] : (~ ! [X0] : (~ p3(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (~ ! [X0] : (~ p4(X0) | ! [X1] : (p4(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p4(X1) | ~ r1(X0, X1)) | ! [X1] : (p4(X1) | ~ r1(X0, X1)) | ((~ ! [X1] : (~ ! [X0] : (~ p2(X0) | ! [X1] : (p2(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p2(X1) | ~ r1(X0, X1)) | ! [X1] : (p2(X1) | ~ r1(X0, X1))) & ! [X1] : (! [X0] : (~ ! [X1] : (~ p3(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.001.p', main)).
fof(f83, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f7])).
fof(f7, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.001.p', transitivity)).
fof(f264076, plain, (! [X0] : (~ r1(X0, sK33(sK28(sK32))) | ~ sP1(X0)) | (spl44_214 | ~ spl44_227 | spl44_560 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(subsumption_resolution, [], [f264075, f28414])).
fof(f28414, plain, (~ p2(sK33(sK28(sK32))) | spl44_560), inference(avatar_component_clause, [], [f28413])).
fof(f264075, plain, (! [X0] : (p2(sK33(sK28(sK32))) | ~ r1(X0, sK33(sK28(sK32))) | ~ sP1(X0)) | (spl44_214 | ~ spl44_227 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(resolution, [], [f263181, f126])).
fof(f126, plain, ! [X0, X1] : (~ p2(sK27(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : ((! [X1] : ((p2(sK26(X1)) & (~ p2(sK27(X1)) & r1(sK26(X1), sK27(X1))) & r1(X1, sK26(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK28(X0)) & r1(X0, sK28(X0)))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27, sK28])], [f57, f60, f59, f58])).
fof(f58, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK26(X1)) & ? [X3] : (~ p2(X3) & r1(sK26(X1), X3)) & r1(X1, sK26(X1)))), introduced(choice_axiom, [])).
fof(f59, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK26(X1), X3)) => (~ p2(sK27(X1)) & r1(sK26(X1), sK27(X1)))), introduced(choice_axiom, [])).
fof(f60, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK28(X0)) & r1(X0, sK28(X0)))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP1(X0)), inference(rectify, [], [f56])).
fof(f56, plain, ! [X0] : ((! [X48] : (? [X49] : (p2(X49) & ? [X50] : (~ p2(X50) & r1(X49, X50)) & r1(X48, X49)) | p2(X48) | ~ r1(X0, X48)) & ? [X51] : (~ p2(X51) & r1(X0, X51))) | ~ sP1(X0)), inference(nnf_transformation, [], [f12])).
fof(f263181, plain, (p2(sK27(sK33(sK28(sK32)))) | (spl44_214 | ~ spl44_227 | ~ spl44_729 | ~ spl44_1654 | ~ spl44_2531)), inference(resolution, [], [f96263, f72855])).
fof(f72855, plain, (! [X0] : (~ r1(sK26(sK33(sK28(sK32))), X0) | p2(X0)) | (spl44_214 | ~ spl44_227 | ~ spl44_729 | ~ spl44_1654)), inference(subsumption_resolution, [], [f72854, f19034])).
fof(f72854, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK33(sK28(sK32))), X0) | ~ r1(sK32, sK28(sK32))) | (spl44_214 | ~ spl44_729 | ~ spl44_1654)), inference(subsumption_resolution, [], [f72853, f18086])).
fof(f72853, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK33(sK28(sK32))), X0) | p2(sK28(sK32)) | ~ r1(sK32, sK28(sK32))) | (~ spl44_729 | ~ spl44_1654)), inference(subsumption_resolution, [], [f72718, f40764])).
fof(f40764, plain, (p2(sK26(sK33(sK28(sK32)))) | ~ spl44_729), inference(avatar_component_clause, [], [f40762])).
fof(f72718, plain, (! [X0] : (p2(X0) | ~ r1(sK26(sK33(sK28(sK32))), X0) | ~ p2(sK26(sK33(sK28(sK32)))) | p2(sK28(sK32)) | ~ r1(sK32, sK28(sK32))) | ~ spl44_1654), inference(resolution, [], [f65868, f159])).
fof(f159, plain, ! [X4, X3, X1] : (~ r1(sK33(X1), X3) | p2(X4) | ~ r1(X3, X4) | ~ p2(X3) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f81])).
fof(f65868, plain, (r1(sK33(sK28(sK32)), sK26(sK33(sK28(sK32)))) | ~ spl44_1654), inference(avatar_component_clause, [], [f65866])).
fof(f96263, plain, (r1(sK26(sK33(sK28(sK32))), sK27(sK33(sK28(sK32)))) | ~ spl44_2531), inference(avatar_component_clause, [], [f96261])).
fof(f260760, plain, (spl44_214 | ~ spl44_227 | ~ spl44_560), inference(avatar_contradiction_clause, [], [f260759])).
fof(f260759, plain, ($false | (spl44_214 | ~ spl44_227 | ~ spl44_560)), inference(subsumption_resolution, [], [f260758, f19034])).
fof(f260758, plain, (~ r1(sK32, sK28(sK32)) | (spl44_214 | ~ spl44_560)), inference(subsumption_resolution, [], [f260757, f18086])).
fof(f260757, plain, (p2(sK28(sK32)) | ~ r1(sK32, sK28(sK32)) | ~ spl44_560), inference(resolution, [], [f28415, f158])).
fof(f158, plain, ! [X1] : (~ p2(sK33(X1)) | p2(X1) | ~ r1(sK32, X1)), inference(cnf_transformation, [], [f81])).
fof(f28415, plain, (p2(sK33(sK28(sK32))) | ~ spl44_560), inference(avatar_component_clause, [], [f28413])).
fof(f232993, plain, (spl44_2000 | spl44_1977 | spl44_200 | ~ spl44_2036), inference(avatar_split_clause, [], [f232992, f88866, f10832, f88415, f88526])).
fof(f88526, plain, (spl44_2000 <=> p2(sK29(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_2000])])).
fof(f88415, plain, (spl44_1977 <=> ! [X126] : (~ r1(X126, sK31(sK32)) | ~ sP0(X126))), introduced(avatar_definition, [new_symbols(naming, [spl44_1977])])).
fof(f10832, plain, (spl44_200 <=> p2(sK31(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl44_200])])).
fof(f88866, plain, (spl44_2036 <=> r1(sK32, sK31(sK32))), introduced(avatar_definition, [new_symbols(naming, [spl44_2036])])).
fof(f232992, plain, (! [X310] : (~ r1(X310, sK31(sK32)) | p2(sK29(sK33(sK31(sK32)))) | ~ sP0(X310)) | (spl44_200 | ~ spl44_2036)), inference(subsumption_resolution, [], [f232991, f10833])).
fof(f10833, plain, (~ p2(sK31(sK32)) | spl44_200), inference(avatar_component_clause, [], [f10832])).
fof(f232991, plain, (! [X310] : (p2(sK31(sK32)) | ~ r1(X310, sK31(sK32)) | p2(sK29(sK33(sK31(sK32)))) | ~ sP0(X310)) | ~ spl44_2036), inference(subsumption_resolution, [], [f111215, f82])).
fof(f82, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL676+1.001.p', reflexivity)).
fof(f111215, plain, (! [X310] : (~ r1(sK32, sK32) | p2(sK31(sK32)) | ~ r1(X310, sK31(sK32)) | p2(sK29(sK33(sK31(sK32)))) | ~ sP0(X310)) | ~ spl44_2036), inference(resolution, [], [f89911, f4437])).
fof(f4437, plain, ! [X52, X51] : (~ r1(sK32, X52) | p2(X52) | ~ r1(X51, X52) | p2(sK29(sK33(X52))) | ~ sP0(X51)), inference(subsumption_resolution, [], [f4395, f158])).
fof(f4395, plain, ! [X52, X51] : (~ r1(X51, X52) | p2(X52) | ~ r1(sK32, X52) | p2(sK33(X52)) | p2(sK29(sK33(X52))) | ~ sP0(X51)), inference(resolution, [], [f744, f133])).
fof(f133, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK29(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ! [X0] : ((! [X1] : ((p2(sK29(X1)) & (~ p2(sK30(X1)) & r1(sK29(X1), sK30(X1))) & r1(X1, sK29(X1))) | p2(X1) | ~ r1(X0, X1)) & (~ p2(sK31(X0)) & r1(X0, sK31(X0)))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30, sK31])], [f63, f66, f65, f64])).
fof(f64, plain, ! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) => (p2(sK29(X1)) & ? [X3] : (~ p2(X3) & r1(sK29(X1), X3)) & r1(X1, sK29(X1)))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X1] : (? [X3] : (~ p2(X3) & r1(sK29(X1), X3)) => (~ p2(sK30(X1)) & r1(sK29(X1), sK30(X1)))), introduced(choice_axiom, [])).
fof(f66, plain, ! [X0] : (? [X4] : (~ p2(X4) & r1(X0, X4)) => (~ p2(sK31(X0)) & r1(X0, sK31(X0)))), introduced(choice_axiom, [])).
fof(f63, plain, ! [X0] : ((! [X1] : (? [X2] : (p2(X2) & ? [X3] : (~ p2(X3) & r1(X2, X3)) & r1(X1, X2)) | p2(X1) | ~ r1(X0, X1)) & ? [X4] : (~ p2(X4) & r1(X0, X4))) | ~ sP0(X0)), inference(rectify, [], [f62])).
fof(f62, plain, ! [X0] : ((! [X58] : (? [X59] : (p2(X59) & ? [X60] : (~ p2(X60) & r1(X59, X60)) & r1(X58, X59)) | p2(X58) | ~ r1(X0, X58)) & ? [X61] : (~ p2(X61) & r1(X0, X61))) | ~ sP0(X0)), inference(nnf_transformation, [], [f11])).
fof(f89911, plain, (! [X32] : (r1(X32, sK31(sK32)) | ~ r1(X32, sK32)) | ~ spl44_2036), inference(resolution, [], [f88867, f83])).
fof(f88867, plain, (r1(sK32, sK31(sK32)) | ~ spl44_2036), inference(avatar_component_clause, [], [f88866])).
fof(f232806, plain, (spl44_200 | ~ spl44_2036 | ~ spl44_5402), inference(avatar_contradiction_clause, [], [f232805])).
fof(f232805, plain, ($false | (spl44_200 | ~ spl44_2036 | ~ spl44_5402)), inference(subsumption_resolution, [], [f232804, f88867])).
fof(f232804, plain, (~ r1(sK32, sK31(sK32)) | (spl44_200 | ~ spl44_5402)), inference(subsumption_resolution, [], [f232803, f10833])).
fof(f232803, plain, (p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32)) | ~ spl44_5402), inference(resolution, [], [f179114, f158])).
fof(f179114, plain, (p2(sK33(sK31(sK32))) | ~ spl44_5402), inference(avatar_component_clause, [], [f179112])).
fof(f179112, plain, (spl44_5402 <=> p2(sK33(sK31(sK32)))), introduced(avatar_definition, [new_symbols(naming, [spl44_5402])])).
fof(f232802, plain, (spl44_1977 | spl44_200 | ~ spl44_2036 | ~ spl44_7330), inference(avatar_split_clause, [], [f232801, f232794, f88866, f10832, f88415])).
fof(f232794, plain, (spl44_7330 <=> ! [X0] : (~ r1(X0, sK33(sK31(sK32))) | ~ sP0(X0))), introduced(avatar_definition, [new_symbols(naming, [spl44_7330])])).
fof(f232801, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK31(sK32))) | (spl44_200 | ~ spl44_2036 | ~ spl44_7330)), inference(subsumption_resolution, [], [f232800, f88867])).
fof(f232800, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | (spl44_200 | ~ spl44_7330)), inference(subsumption_resolution, [], [f232798, f10833])).
fof(f232798, plain, (! [X0] : (~ sP0(X0) | ~ r1(X0, sK31(sK32)) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | ~ spl44_7330), inference(resolution, [], [f232795, f744])).
fof(f232795, plain, (! [X0] : (~ r1(X0, sK33(sK31(sK32))) | ~ sP0(X0)) | ~ spl44_7330), inference(avatar_component_clause, [], [f232794])).
fof(f232796, plain, (spl44_7330 | spl44_5402 | spl44_200 | ~ spl44_1976 | ~ spl44_2000 | ~ spl44_2036 | ~ spl44_2540), inference(avatar_split_clause, [], [f232792, f97803, f88866, f88526, f88411, f10832, f179112, f232794])).
fof(f88411, plain, (spl44_1976 <=> r1(sK33(sK31(sK32)), sK29(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_1976])])).
fof(f97803, plain, (spl44_2540 <=> r1(sK29(sK33(sK31(sK32))), sK30(sK33(sK31(sK32))))), introduced(avatar_definition, [new_symbols(naming, [spl44_2540])])).
fof(f232792, plain, (! [X0] : (p2(sK33(sK31(sK32))) | ~ r1(X0, sK33(sK31(sK32))) | ~ sP0(X0)) | (spl44_200 | ~ spl44_1976 | ~ spl44_2000 | ~ spl44_2036 | ~ spl44_2540)), inference(resolution, [], [f232456, f132])).
fof(f132, plain, ! [X0, X1] : (~ p2(sK30(X1)) | p2(X1) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f232456, plain, (p2(sK30(sK33(sK31(sK32)))) | (spl44_200 | ~ spl44_1976 | ~ spl44_2000 | ~ spl44_2036 | ~ spl44_2540)), inference(resolution, [], [f97805, f177616])).
fof(f177616, plain, (! [X1] : (~ r1(sK29(sK33(sK31(sK32))), X1) | p2(X1)) | (spl44_200 | ~ spl44_1976 | ~ spl44_2000 | ~ spl44_2036)), inference(subsumption_resolution, [], [f177615, f88867])).
fof(f177615, plain, (! [X1] : (p2(X1) | ~ r1(sK29(sK33(sK31(sK32))), X1) | ~ r1(sK32, sK31(sK32))) | (spl44_200 | ~ spl44_1976 | ~ spl44_2000)), inference(subsumption_resolution, [], [f177614, f10833])).
fof(f177614, plain, (! [X1] : (p2(X1) | ~ r1(sK29(sK33(sK31(sK32))), X1) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | (~ spl44_1976 | ~ spl44_2000)), inference(subsumption_resolution, [], [f177429, f88528])).
fof(f88528, plain, (p2(sK29(sK33(sK31(sK32)))) | ~ spl44_2000), inference(avatar_component_clause, [], [f88526])).
fof(f177429, plain, (! [X1] : (p2(X1) | ~ r1(sK29(sK33(sK31(sK32))), X1) | ~ p2(sK29(sK33(sK31(sK32)))) | p2(sK31(sK32)) | ~ r1(sK32, sK31(sK32))) | ~ spl44_1976), inference(resolution, [], [f88413, f159])).
fof(f88413, plain, (r1(sK33(sK31(sK32)), sK29(sK33(sK31(sK32)))) | ~ spl44_1976), inference(avatar_component_clause, [], [f88411])).
fof(f97805, plain, (r1(sK29(sK33(sK31(sK32))), sK30(sK33(sK31(sK32)))) | ~ spl44_2540), inference(avatar_component_clause, [], [f97803])).
fof(f97812, plain, (spl44_1976 | spl44_1977 | spl44_200 | ~ spl44_2036), inference(avatar_split_clause, [], [f97811, f88866, f10832, f88415, f88411])).
fof(f97811, plain, (! [X26] : (~ r1(X26, sK31(sK32)) | r1(sK33(sK31(sK32)), sK29(sK33(sK31(sK32)))) | ~ sP0(X26)) | (spl44_200 | ~ spl44_2036)), inference(subsumption_resolution, [], [f89903, f10833])).
fof(f89903, plain, (! [X26] : (p2(sK31(sK32)) | ~ r1(X26, sK31(sK32)) | r1(sK33(sK31(sK32)), sK29(sK33(sK31(sK32)))) | ~ sP0(X26)) | ~ spl44_2036), inference(resolution, [], [f88867, f4435])).
fof(f4435, plain, ! [X47, X48] : (~ r1(sK32, X48) | p2(X48) | ~ r1(X47, X48) | r1(sK33(X48), sK29(sK33(X48))) | ~ sP0(X47)), inference(subsumption_resolution, [], [f4393, f158])).
fof(f4393, plain, ! [X47, X48] : (~ r1(X47, X48) | p2(X48) | ~ r1(sK32, X48) | p2(sK33(X48)) | r1(sK33(X48), sK29(sK33(X48))) | ~ sP0(X47)), inference(resolution, [], [f744, f130])).
fof(f130, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK29(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f97808, plain, (spl44_2540 | spl44_1977 | spl44_200 | ~ spl44_2036), inference(avatar_split_clause, [], [f97807, f88866, f10832, f88415, f97803])).
fof(f97807, plain, (! [X14] : (~ r1(X14, sK31(sK32)) | r1(sK29(sK33(sK31(sK32))), sK30(sK33(sK31(sK32)))) | ~ sP0(X14)) | (spl44_200 | ~ spl44_2036)), inference(subsumption_resolution, [], [f96300, f10833])).
fof(f96300, plain, (! [X14] : (p2(sK31(sK32)) | ~ r1(X14, sK31(sK32)) | r1(sK29(sK33(sK31(sK32))), sK30(sK33(sK31(sK32)))) | ~ sP0(X14)) | ~ spl44_2036), inference(resolution, [], [f4436, f88867])).
fof(f4436, plain, ! [X50, X49] : (~ r1(sK32, X50) | p2(X50) | ~ r1(X49, X50) | r1(sK29(sK33(X50)), sK30(sK33(X50))) | ~ sP0(X49)), inference(subsumption_resolution, [], [f4394, f158])).
fof(f4394, plain, ! [X50, X49] : (~ r1(X49, X50) | p2(X50) | ~ r1(sK32, X50) | p2(sK33(X50)) | r1(sK29(sK33(X50)), sK30(sK33(X50))) | ~ sP0(X49)), inference(resolution, [], [f744, f131])).
fof(f131, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK29(X1), sK30(X1)) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f97791, plain, (~ spl44_12 | ~ spl44_1977), inference(avatar_contradiction_clause, [], [f97790])).
fof(f97790, plain, ($false | (~ spl44_12 | ~ spl44_1977)), inference(subsumption_resolution, [], [f97789, f208])).
fof(f208, plain, (sP0(sK32) | ~ spl44_12), inference(avatar_component_clause, [], [f206])).
fof(f206, plain, (spl44_12 <=> sP0(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl44_12])])).
fof(f97789, plain, (~ sP0(sK32) | ~ spl44_1977), inference(duplicate_literal_removal, [], [f97779])).
fof(f97779, plain, (~ sP0(sK32) | ~ sP0(sK32) | ~ spl44_1977), inference(resolution, [], [f88416, f128])).
fof(f128, plain, ! [X0] : (r1(X0, sK31(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f88416, plain, (! [X126] : (~ r1(X126, sK31(sK32)) | ~ sP0(X126)) | ~ spl44_1977), inference(avatar_component_clause, [], [f88415])).
fof(f96286, plain, (spl44_2531 | spl44_635 | spl44_214 | ~ spl44_227), inference(avatar_split_clause, [], [f96285, f19033, f18085, f35437, f96261])).
fof(f96285, plain, (! [X120] : (~ r1(X120, sK28(sK32)) | r1(sK26(sK33(sK28(sK32))), sK27(sK33(sK28(sK32)))) | ~ sP1(X120)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f96284, f82])).
fof(f96284, plain, (! [X120] : (~ r1(X120, sK28(sK32)) | r1(sK26(sK33(sK28(sK32))), sK27(sK33(sK28(sK32)))) | ~ sP1(X120) | ~ r1(sK32, sK32)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f96228, f18086])).
fof(f96228, plain, (! [X120] : (p2(sK28(sK32)) | ~ r1(X120, sK28(sK32)) | r1(sK26(sK33(sK28(sK32))), sK27(sK33(sK28(sK32)))) | ~ sP1(X120) | ~ r1(sK32, sK32)) | ~ spl44_227), inference(resolution, [], [f4433, f19412])).
fof(f19412, plain, (! [X4] : (r1(X4, sK28(sK32)) | ~ r1(X4, sK32)) | ~ spl44_227), inference(resolution, [], [f19034, f83])).
fof(f4433, plain, ! [X43, X44] : (~ r1(sK32, X44) | p2(X44) | ~ r1(X43, X44) | r1(sK26(sK33(X44)), sK27(sK33(X44))) | ~ sP1(X43)), inference(subsumption_resolution, [], [f4391, f158])).
fof(f4391, plain, ! [X43, X44] : (~ r1(X43, X44) | p2(X44) | ~ r1(sK32, X44) | p2(sK33(X44)) | r1(sK26(sK33(X44)), sK27(sK33(X44))) | ~ sP1(X43)), inference(resolution, [], [f744, f125])).
fof(f125, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(sK26(X1), sK27(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f89221, plain, (~ spl44_12 | spl44_2036), inference(avatar_contradiction_clause, [], [f89220])).
fof(f89220, plain, ($false | (~ spl44_12 | spl44_2036)), inference(subsumption_resolution, [], [f89219, f208])).
fof(f89219, plain, (~ sP0(sK32) | spl44_2036), inference(subsumption_resolution, [], [f89212, f82])).
fof(f89212, plain, (~ r1(sK32, sK32) | ~ sP0(sK32) | spl44_2036), inference(resolution, [], [f88868, f743])).
fof(f743, plain, ! [X12, X13] : (r1(X12, sK31(X13)) | ~ r1(X12, X13) | ~ sP0(X13)), inference(resolution, [], [f83, f128])).
fof(f88868, plain, (~ r1(sK32, sK31(sK32)) | spl44_2036), inference(avatar_component_clause, [], [f88866])).
fof(f65893, plain, (spl44_1654 | spl44_635 | spl44_214 | ~ spl44_227), inference(avatar_split_clause, [], [f65892, f19033, f18085, f35437, f65866])).
fof(f65892, plain, (! [X111] : (~ r1(X111, sK28(sK32)) | r1(sK33(sK28(sK32)), sK26(sK33(sK28(sK32)))) | ~ sP1(X111)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f65891, f82])).
fof(f65891, plain, (! [X111] : (~ r1(X111, sK28(sK32)) | r1(sK33(sK28(sK32)), sK26(sK33(sK28(sK32)))) | ~ sP1(X111) | ~ r1(sK32, sK32)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f65838, f18086])).
fof(f65838, plain, (! [X111] : (p2(sK28(sK32)) | ~ r1(X111, sK28(sK32)) | r1(sK33(sK28(sK32)), sK26(sK33(sK28(sK32)))) | ~ sP1(X111) | ~ r1(sK32, sK32)) | ~ spl44_227), inference(resolution, [], [f4432, f19412])).
fof(f4432, plain, ! [X41, X42] : (~ r1(sK32, X42) | p2(X42) | ~ r1(X41, X42) | r1(sK33(X42), sK26(sK33(X42))) | ~ sP1(X41)), inference(subsumption_resolution, [], [f4390, f158])).
fof(f4390, plain, ! [X41, X42] : (~ r1(X41, X42) | p2(X42) | ~ r1(sK32, X42) | p2(sK33(X42)) | r1(sK33(X42), sK26(sK33(X42))) | ~ sP1(X41)), inference(resolution, [], [f744, f124])).
fof(f124, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | r1(X1, sK26(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f61210, plain, (~ spl44_9 | ~ spl44_635), inference(avatar_contradiction_clause, [], [f61209])).
fof(f61209, plain, ($false | (~ spl44_9 | ~ spl44_635)), inference(subsumption_resolution, [], [f61207, f196])).
fof(f196, plain, (sP1(sK32) | ~ spl44_9), inference(avatar_component_clause, [], [f194])).
fof(f194, plain, (spl44_9 <=> sP1(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl44_9])])).
fof(f61207, plain, (~ sP1(sK32) | ~ spl44_635), inference(duplicate_literal_removal, [], [f61194])).
fof(f61194, plain, (~ sP1(sK32) | ~ sP1(sK32) | ~ spl44_635), inference(resolution, [], [f35438, f122])).
fof(f122, plain, ! [X0] : (r1(X0, sK28(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f35438, plain, (! [X88] : (~ r1(X88, sK28(sK32)) | ~ sP1(X88)) | ~ spl44_635), inference(avatar_component_clause, [], [f35437])).
fof(f40782, plain, (spl44_729 | spl44_635 | spl44_214 | ~ spl44_227), inference(avatar_split_clause, [], [f40781, f19033, f18085, f35437, f40762])).
fof(f40781, plain, (! [X88] : (~ r1(X88, sK28(sK32)) | p2(sK26(sK33(sK28(sK32)))) | ~ sP1(X88)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f40780, f82])).
fof(f40780, plain, (! [X88] : (~ r1(X88, sK28(sK32)) | p2(sK26(sK33(sK28(sK32)))) | ~ sP1(X88) | ~ r1(sK32, sK32)) | (spl44_214 | ~ spl44_227)), inference(subsumption_resolution, [], [f40737, f18086])).
fof(f40737, plain, (! [X88] : (p2(sK28(sK32)) | ~ r1(X88, sK28(sK32)) | p2(sK26(sK33(sK28(sK32)))) | ~ sP1(X88) | ~ r1(sK32, sK32)) | ~ spl44_227), inference(resolution, [], [f4434, f19412])).
fof(f4434, plain, ! [X45, X46] : (~ r1(sK32, X46) | p2(X46) | ~ r1(X45, X46) | p2(sK26(sK33(X46))) | ~ sP1(X45)), inference(subsumption_resolution, [], [f4392, f158])).
fof(f4392, plain, ! [X45, X46] : (~ r1(X45, X46) | p2(X46) | ~ r1(sK32, X46) | p2(sK33(X46)) | p2(sK26(sK33(X46))) | ~ sP1(X45)), inference(resolution, [], [f744, f127])).
fof(f127, plain, ! [X0, X1] : (~ r1(X0, X1) | p2(X1) | p2(sK26(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f19399, plain, (~ spl44_9 | spl44_227), inference(avatar_contradiction_clause, [], [f19398])).
fof(f19398, plain, ($false | (~ spl44_9 | spl44_227)), inference(subsumption_resolution, [], [f19397, f196])).
fof(f19397, plain, (~ sP1(sK32) | spl44_227), inference(subsumption_resolution, [], [f19394, f82])).
fof(f19394, plain, (~ r1(sK32, sK32) | ~ sP1(sK32) | spl44_227), inference(resolution, [], [f19035, f742])).
fof(f742, plain, ! [X10, X11] : (r1(X10, sK28(X11)) | ~ r1(X10, X11) | ~ sP1(X11)), inference(resolution, [], [f83, f122])).
fof(f19035, plain, (~ r1(sK32, sK28(sK32)) | spl44_227), inference(avatar_component_clause, [], [f19033])).
fof(f18445, plain, (~ spl44_9 | ~ spl44_214), inference(avatar_contradiction_clause, [], [f18444])).
fof(f18444, plain, ($false | (~ spl44_9 | ~ spl44_214)), inference(subsumption_resolution, [], [f18443, f196])).
fof(f18443, plain, (~ sP1(sK32) | ~ spl44_214), inference(resolution, [], [f18087, f123])).
fof(f123, plain, ! [X0] : (~ p2(sK28(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f61])).
fof(f18087, plain, (p2(sK28(sK32)) | ~ spl44_214), inference(avatar_component_clause, [], [f18085])).
fof(f18274, plain, (spl44_76 | ~ spl44_13 | ~ spl44_15), inference(avatar_split_clause, [], [f18271, f221, f211, f591])).
fof(f211, plain, (spl44_13 <=> r1(sK42, sK43)), introduced(avatar_definition, [new_symbols(naming, [spl44_13])])).
fof(f221, plain, (spl44_15 <=> r1(sK32, sK42)), introduced(avatar_definition, [new_symbols(naming, [spl44_15])])).
fof(f18271, plain, (r1(sK32, sK43) | (~ spl44_13 | ~ spl44_15)), inference(resolution, [], [f18105, f223])).
fof(f223, plain, (r1(sK32, sK42) | ~ spl44_15), inference(avatar_component_clause, [], [f221])).
fof(f18105, plain, (! [X0] : (~ r1(X0, sK42) | r1(X0, sK43)) | ~ spl44_13), inference(resolution, [], [f213, f83])).
fof(f213, plain, (r1(sK42, sK43) | ~ spl44_13), inference(avatar_component_clause, [], [f211])).
fof(f10928, plain, (~ spl44_12 | ~ spl44_200), inference(avatar_contradiction_clause, [], [f10927])).
fof(f10927, plain, ($false | (~ spl44_12 | ~ spl44_200)), inference(subsumption_resolution, [], [f10926, f208])).
fof(f10926, plain, (~ sP0(sK32) | ~ spl44_200), inference(resolution, [], [f10834, f129])).
fof(f129, plain, ! [X0] : (~ p2(sK31(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f67])).
fof(f10834, plain, (p2(sK31(sK32)) | ~ spl44_200), inference(avatar_component_clause, [], [f10832])).
fof(f224, plain, (spl44_15 | spl44_12), inference(avatar_split_clause, [], [f134, f206, f221])).
fof(f134, plain, (sP0(sK32) | r1(sK32, sK42)), inference(cnf_transformation, [], [f81])).
fof(f214, plain, (spl44_13 | spl44_12), inference(avatar_split_clause, [], [f136, f206, f211])).
fof(f136, plain, (sP0(sK32) | r1(sK42, sK43)), inference(cnf_transformation, [], [f81])).
fof(f209, plain, (spl44_11 | spl44_12), inference(avatar_split_clause, [], [f137, f206, f203])).
fof(f137, plain, ! [X23] : (sP0(sK32) | ~ p3(X23) | ~ r1(sK43, X23)), inference(cnf_transformation, [], [f81])).
fof(f201, plain, (spl44_10 | spl44_9), inference(avatar_split_clause, [], [f144, f194, f199])).
fof(f144, plain, ! [X15] : (sP1(sK32) | r1(X15, sK38(X15)) | ~ r1(sK32, X15)), inference(cnf_transformation, [], [f81])).
fof(f197, plain, (spl44_8 | spl44_9), inference(avatar_split_clause, [], [f145, f194, f191])).
fof(f145, plain, ! [X15] : (sP1(sK32) | p3(sK38(X15)) | ~ r1(sK32, X15)), inference(cnf_transformation, [], [f81])).