fof(f730, plain, $false, inference(avatar_sat_refutation, [], [f479, f719, f729])).
fof(f729, plain, ~ spl23_8, inference(avatar_contradiction_clause, [], [f728])).
fof(f728, plain, ($false | ~ spl23_8), inference(subsumption_resolution, [], [f727, f145])).
fof(f145, plain, sP3(sK22), inference(resolution, [], [f140, f73])).
fof(f73, plain, ! [X0] : (~ sP11(X0) | sP3(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & sP10(X0) & sP9(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP4(X0)) | ~ sP11(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP4(X1)) | ~ sP11(X1)), inference(nnf_transformation, [], [f23])).
fof(f23, plain, ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP4(X1)) | ~ sP11(X1)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X1] : (sP11(X1) <=> ((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & sP10(X1) & sP9(X1) & sP8(X1) & sP7(X1) & sP6(X1) & sP5(X1) & sP3(X1) & sP2(X1) & sP1(X1) & sP0(X1) & sP4(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f140, plain, sP11(sK22), inference(resolution, [], [f136, f67])).
fof(f67, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.005.p', reflexivity)).
fof(f136, plain, ! [X1] : (~ r1(sK22, X1) | sP11(X1)), inference(cnf_transformation, [], [f66])).
fof(f66, plain, (p100(sK22) & ~ p101(sK22) & ! [X1] : (sP11(X1) | ~ r1(sK22, X1)) & ! [X2] : (p3(X2) | ~ r1(sK22, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22])], [f64, f65])).
fof(f65, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP11(X1) | ~ r1(X0, X1)) & ! [X2] : (p3(X2) | ~ r1(X0, X2))) => (p100(sK22) & ~ p101(sK22) & ! [X1] : (sP11(X1) | ~ r1(sK22, X1)) & ! [X2] : (p3(X2) | ~ r1(sK22, X2)))), introduced(choice_axiom, [])).
fof(f64, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP11(X1) | ~ r1(X0, X1)) & ! [X2] : (p3(X2) | ~ r1(X0, X2))), inference(rectify, [], [f24])).
fof(f24, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (sP11(X1) | ~ r1(X0, X1)) & ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(definition_folding, [], [f11, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12])).
fof(f12, plain, ! [X1] : (~ p103(X1) | p104(X1) | (? [X20] : (p104(X20) & ~ p105(X20) & p5(X20) & r1(X1, X20)) & ? [X21] : (p104(X21) & ~ p105(X21) & ~ p5(X21) & r1(X1, X21))) | ~ sP0(X1)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X1] : (sP0(X1) <=> (~ p103(X1) | p104(X1) | (? [X20] : (p104(X20) & ~ p105(X20) & p5(X20) & r1(X1, X20)) & ? [X21] : (p104(X21) & ~ p105(X21) & ~ p5(X21) & r1(X1, X21))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f13, plain, ! [X1] : (~ p102(X1) | p103(X1) | (? [X18] : (p103(X18) & ~ p104(X18) & p4(X18) & r1(X1, X18)) & ? [X19] : (p103(X19) & ~ p104(X19) & ~ p4(X19) & r1(X1, X19))) | ~ sP1(X1)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X1] : (sP1(X1) <=> (~ p102(X1) | p103(X1) | (? [X18] : (p103(X18) & ~ p104(X18) & p4(X18) & r1(X1, X18)) & ? [X19] : (p103(X19) & ~ p104(X19) & ~ p4(X19) & r1(X1, X19))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f14, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X16] : (p102(X16) & ~ p103(X16) & p3(X16) & r1(X1, X16)) & ? [X17] : (p102(X17) & ~ p103(X17) & ~ p3(X17) & r1(X1, X17))) | ~ sP2(X1)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X1] : (sP2(X1) <=> (~ p101(X1) | p102(X1) | (? [X16] : (p102(X16) & ~ p103(X16) & p3(X16) & r1(X1, X16)) & ? [X17] : (p102(X17) & ~ p103(X17) & ~ p3(X17) & r1(X1, X17))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f15, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X14] : (p101(X14) & ~ p102(X14) & p2(X14) & r1(X1, X14)) & ? [X15] : (p101(X15) & ~ p102(X15) & ~ p2(X15) & r1(X1, X15))) | ~ sP3(X1)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X1] : (sP3(X1) <=> (~ p100(X1) | p101(X1) | (? [X14] : (p101(X14) & ~ p102(X14) & p2(X14) & r1(X1, X14)) & ? [X15] : (p101(X15) & ~ p102(X15) & ~ p2(X15) & r1(X1, X15))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f16, plain, ! [X1] : (~ p104(X1) | p105(X1) | (? [X22] : (p105(X22) & p6(X22) & r1(X1, X22)) & ? [X23] : (p105(X23) & ~ p6(X23) & r1(X1, X23))) | ~ sP4(X1)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X1] : (sP4(X1) <=> (~ p104(X1) | p105(X1) | (? [X22] : (p105(X22) & p6(X22) & r1(X1, X22)) & ? [X23] : (p105(X23) & ~ p6(X23) & r1(X1, X23))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X1] : (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))) | ~ sP5(X1)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X1] : (sP5(X1) <=> (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f18, plain, ! [X1] : (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))) | ~ sP6(X1)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X1] : (sP6(X1) <=> (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f19, plain, ! [X1] : (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))) | ~ sP7(X1)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X1] : (sP7(X1) <=> (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f20, plain, ! [X1] : (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))) | ~ sP8(X1)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X1] : (sP8(X1) <=> (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f21, plain, ! [X1] : (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))) | ~ sP9(X1)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X1] : (sP9(X1) <=> (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f22, plain, ! [X1] : (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))) | ~ sP10(X1)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X1] : (sP10(X1) <=> (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f11, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ p100(X1) | p101(X1) | (? [X14] : (p101(X14) & ~ p102(X14) & p2(X14) & r1(X1, X14)) & ? [X15] : (p101(X15) & ~ p102(X15) & ~ p2(X15) & r1(X1, X15)))) & (~ p101(X1) | p102(X1) | (? [X16] : (p102(X16) & ~ p103(X16) & p3(X16) & r1(X1, X16)) & ? [X17] : (p102(X17) & ~ p103(X17) & ~ p3(X17) & r1(X1, X17)))) & (~ p102(X1) | p103(X1) | (? [X18] : (p103(X18) & ~ p104(X18) & p4(X18) & r1(X1, X18)) & ? [X19] : (p103(X19) & ~ p104(X19) & ~ p4(X19) & r1(X1, X19)))) & (~ p103(X1) | p104(X1) | (? [X20] : (p104(X20) & ~ p105(X20) & p5(X20) & r1(X1, X20)) & ? [X21] : (p104(X21) & ~ p105(X21) & ~ p5(X21) & r1(X1, X21)))) & (~ p104(X1) | p105(X1) | (? [X22] : (p105(X22) & p6(X22) & r1(X1, X22)) & ? [X23] : (p105(X23) & ~ p6(X23) & r1(X1, X23))))) | ~ r1(X0, X1)) & ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(flattening, [], [f10])).
fof(f10, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & ((~ p100(X1) | p101(X1)) | (? [X14] : ((p101(X14) & ~ p102(X14) & p2(X14)) & r1(X1, X14)) & ? [X15] : ((p101(X15) & ~ p102(X15) & ~ p2(X15)) & r1(X1, X15)))) & ((~ p101(X1) | p102(X1)) | (? [X16] : ((p102(X16) & ~ p103(X16) & p3(X16)) & r1(X1, X16)) & ? [X17] : ((p102(X17) & ~ p103(X17) & ~ p3(X17)) & r1(X1, X17)))) & ((~ p102(X1) | p103(X1)) | (? [X18] : ((p103(X18) & ~ p104(X18) & p4(X18)) & r1(X1, X18)) & ? [X19] : ((p103(X19) & ~ p104(X19) & ~ p4(X19)) & r1(X1, X19)))) & ((~ p103(X1) | p104(X1)) | (? [X20] : ((p104(X20) & ~ p105(X20) & p5(X20)) & r1(X1, X20)) & ? [X21] : ((p104(X21) & ~ p105(X21) & ~ p5(X21)) & r1(X1, X21)))) & ((~ p104(X1) | p105(X1)) | (? [X22] : ((p105(X22) & p6(X22)) & r1(X1, X22)) & ? [X23] : ((p105(X23) & ~ p6(X23)) & r1(X1, X23))))) | ~ r1(X0, X1))) & ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & p2(X14)) | ~ r1(X1, X14)) & ~ ! [X15] : (~ (p101(X15) & ~ p102(X15) & ~ p2(X15)) | ~ r1(X1, X15)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & p3(X16)) | ~ r1(X1, X16)) & ~ ! [X17] : (~ (p102(X17) & ~ p103(X17) & ~ p3(X17)) | ~ r1(X1, X17)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & p4(X18)) | ~ r1(X1, X18)) & ~ ! [X19] : (~ (p103(X19) & ~ p104(X19) & ~ p4(X19)) | ~ r1(X1, X19)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & p5(X20)) | ~ r1(X1, X20)) & ~ ! [X21] : (~ (p104(X21) & ~ p105(X21) & ~ p5(X21)) | ~ r1(X1, X21)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X22] : (~ (p105(X22) & p6(X22)) | ~ r1(X1, X22)) & ~ ! [X23] : (~ (p105(X23) & ~ p6(X23)) | ~ r1(X1, X23))))) | ~ r1(X0, X1))) | ~ ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & p2(X14)) | ~ r1(X1, X14)) & ~ ! [X15] : (~ (p101(X15) & ~ p102(X15) & ~ p2(X15)) | ~ r1(X1, X15)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & p3(X16)) | ~ r1(X1, X16)) & ~ ! [X17] : (~ (p102(X17) & ~ p103(X17) & ~ p3(X17)) | ~ r1(X1, X17)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & p4(X18)) | ~ r1(X1, X18)) & ~ ! [X19] : (~ (p103(X19) & ~ p104(X19) & ~ p4(X19)) | ~ r1(X1, X19)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & p5(X20)) | ~ r1(X1, X20)) & ~ ! [X21] : (~ (p104(X21) & ~ p105(X21) & ~ p5(X21)) | ~ r1(X1, X21)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X22] : (~ (p105(X22) & ~ p106(X22) & p6(X22)) | ~ r1(X1, X22)) & ~ ! [X23] : (~ (p105(X23) & ~ p106(X23) & ~ p6(X23)) | ~ r1(X1, X23))))) | ~ r1(X0, X1))) | ~ ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X2] : (~ p100(X2) | p1(X2) | ~ r1(X1, X2))) & (p1(X1) | ! [X3] : (~ p100(X3) | ~ p1(X3) | ~ r1(X1, X3))))) & (~ p101(X1) | ((~ p2(X1) | ! [X4] : (~ p101(X4) | p2(X4) | ~ r1(X1, X4))) & (p2(X1) | ! [X5] : (~ p101(X5) | ~ p2(X5) | ~ r1(X1, X5))))) & (~ p102(X1) | ((~ p3(X1) | ! [X6] : (~ p102(X6) | p3(X6) | ~ r1(X1, X6))) & (p3(X1) | ! [X7] : (~ p102(X7) | ~ p3(X7) | ~ r1(X1, X7))))) & (~ p103(X1) | ((~ p4(X1) | ! [X8] : (~ p103(X8) | p4(X8) | ~ r1(X1, X8))) & (p4(X1) | ! [X9] : (~ p103(X9) | ~ p4(X9) | ~ r1(X1, X9))))) & (~ p104(X1) | ((~ p5(X1) | ! [X10] : (~ p104(X10) | p5(X10) | ~ r1(X1, X10))) & (p5(X1) | ! [X11] : (~ p104(X11) | ~ p5(X11) | ~ r1(X1, X11))))) & (~ p105(X1) | ((~ p6(X1) | ! [X12] : (~ p105(X12) | p6(X12) | ~ r1(X1, X12))) & (p6(X1) | ! [X13] : (~ p105(X13) | ~ p6(X13) | ~ r1(X1, X13))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X14] : (~ (p101(X14) & ~ p102(X14) & p2(X14)) | ~ r1(X1, X14)) & ~ ! [X15] : (~ (p101(X15) & ~ p102(X15) & ~ p2(X15)) | ~ r1(X1, X15)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X16] : (~ (p102(X16) & ~ p103(X16) & p3(X16)) | ~ r1(X1, X16)) & ~ ! [X17] : (~ (p102(X17) & ~ p103(X17) & ~ p3(X17)) | ~ r1(X1, X17)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X18] : (~ (p103(X18) & ~ p104(X18) & p4(X18)) | ~ r1(X1, X18)) & ~ ! [X19] : (~ (p103(X19) & ~ p104(X19) & ~ p4(X19)) | ~ r1(X1, X19)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X20] : (~ (p104(X20) & ~ p105(X20) & p5(X20)) | ~ r1(X1, X20)) & ~ ! [X21] : (~ (p104(X21) & ~ p105(X21) & ~ p5(X21)) | ~ r1(X1, X21)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X22] : (~ (p105(X22) & ~ p106(X22) & p6(X22)) | ~ r1(X1, X22)) & ~ ! [X23] : (~ (p105(X23) & ~ p106(X23) & ~ p6(X23)) | ~ r1(X1, X23))))) | ~ r1(X0, X1))) | ~ ! [X24] : (p3(X24) | ~ r1(X0, X24))), inference(rectify, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p3(X1) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (((~ p101(X1) | p100(X1)) & (~ p102(X1) | p101(X1)) & (~ p103(X1) | p102(X1)) & (~ p104(X1) | p103(X1)) & (~ p105(X1) | p104(X1)) & (~ p106(X1) | p105(X1)) & (~ p100(X1) | ((~ p1(X1) | ! [X0] : (~ p100(X0) | p1(X0) | ~ r1(X1, X0))) & (p1(X1) | ! [X0] : (~ p100(X0) | ~ p1(X0) | ~ r1(X1, X0))))) & (~ p101(X1) | ((~ p2(X1) | ! [X0] : (~ p101(X0) | p2(X0) | ~ r1(X1, X0))) & (p2(X1) | ! [X0] : (~ p101(X0) | ~ p2(X0) | ~ r1(X1, X0))))) & (~ p102(X1) | ((~ p3(X1) | ! [X0] : (~ p102(X0) | p3(X0) | ~ r1(X1, X0))) & (p3(X1) | ! [X0] : (~ p102(X0) | ~ p3(X0) | ~ r1(X1, X0))))) & (~ p103(X1) | ((~ p4(X1) | ! [X0] : (~ p103(X0) | p4(X0) | ~ r1(X1, X0))) & (p4(X1) | ! [X0] : (~ p103(X0) | ~ p4(X0) | ~ r1(X1, X0))))) & (~ p104(X1) | ((~ p5(X1) | ! [X0] : (~ p104(X0) | p5(X0) | ~ r1(X1, X0))) & (p5(X1) | ! [X0] : (~ p104(X0) | ~ p5(X0) | ~ r1(X1, X0))))) & (~ p105(X1) | ((~ p6(X1) | ! [X0] : (~ p105(X0) | p6(X0) | ~ r1(X1, X0))) & (p6(X1) | ! [X0] : (~ p105(X0) | ~ p6(X0) | ~ r1(X1, X0))))) & (~ (p100(X1) & ~ p101(X1)) | (~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & p2(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p101(X0) & ~ p102(X0) & ~ p2(X0)) | ~ r1(X1, X0)))) & (~ (p101(X1) & ~ p102(X1)) | (~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & p3(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p102(X0) & ~ p103(X0) & ~ p3(X0)) | ~ r1(X1, X0)))) & (~ (p102(X1) & ~ p103(X1)) | (~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & p4(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p103(X0) & ~ p104(X0) & ~ p4(X0)) | ~ r1(X1, X0)))) & (~ (p103(X1) & ~ p104(X1)) | (~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & p5(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p104(X0) & ~ p105(X0) & ~ p5(X0)) | ~ r1(X1, X0)))) & (~ (p104(X1) & ~ p105(X1)) | (~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & p6(X0)) | ~ r1(X1, X0)) & ~ ! [X0] : (~ (p105(X0) & ~ p106(X0) & ~ p6(X0)) | ~ r1(X1, X0))))) | ~ r1(X0, X1))) | ~ ! [X1] : (p3(X1) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.005.p', main)).
fof(f727, plain, (~ sP3(sK22) | ~ spl23_8), inference(subsumption_resolution, [], [f726, f138])).
fof(f138, plain, p100(sK22), inference(cnf_transformation, [], [f66])).
fof(f726, plain, (~ p100(sK22) | ~ sP3(sK22) | ~ spl23_8), inference(subsumption_resolution, [], [f725, f137])).
fof(f137, plain, ~ p101(sK22), inference(cnf_transformation, [], [f66])).
fof(f725, plain, (p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | ~ spl23_8), inference(resolution, [], [f227, f105])).
fof(f105, plain, ! [X0] : (~ p102(sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f48])).
fof(f48, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK14(X0)) & ~ p102(sK14(X0)) & p2(sK14(X0)) & r1(X0, sK14(X0))) & (p101(sK15(X0)) & ~ p102(sK15(X0)) & ~ p2(sK15(X0)) & r1(X0, sK15(X0)))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15])], [f45, f47, f46])).
fof(f46, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK14(X0)) & ~ p102(sK14(X0)) & p2(sK14(X0)) & r1(X0, sK14(X0)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK15(X0)) & ~ p102(sK15(X0)) & ~ p2(sK15(X0)) & r1(X0, sK15(X0)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP3(X0)), inference(rectify, [], [f44])).
fof(f44, plain, ! [X1] : (~ p100(X1) | p101(X1) | (? [X14] : (p101(X14) & ~ p102(X14) & p2(X14) & r1(X1, X14)) & ? [X15] : (p101(X15) & ~ p102(X15) & ~ p2(X15) & r1(X1, X15))) | ~ sP3(X1)), inference(nnf_transformation, [], [f15])).
fof(f227, plain, (p102(sK15(sK22)) | ~ spl23_8), inference(avatar_component_clause, [], [f225])).
fof(f225, plain, (spl23_8 <=> p102(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_8])])).
fof(f719, plain, (spl23_8 | ~ spl23_9), inference(avatar_split_clause, [], [f718, f230, f225])).
fof(f230, plain, (spl23_9 <=> p101(sK15(sK22))), introduced(avatar_definition, [new_symbols(naming, [spl23_9])])).
fof(f718, plain, (p102(sK15(sK22)) | ~ spl23_9), inference(subsumption_resolution, [], [f717, f145])).
fof(f717, plain, (p102(sK15(sK22)) | ~ sP3(sK22) | ~ spl23_9), inference(subsumption_resolution, [], [f716, f138])).
fof(f716, plain, (p102(sK15(sK22)) | ~ p100(sK22) | ~ sP3(sK22) | ~ spl23_9), inference(subsumption_resolution, [], [f715, f137])).
fof(f715, plain, (p102(sK15(sK22)) | p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | ~ spl23_9), inference(subsumption_resolution, [], [f706, f232])).
fof(f232, plain, (p101(sK15(sK22)) | ~ spl23_9), inference(avatar_component_clause, [], [f230])).
fof(f706, plain, (p102(sK15(sK22)) | ~ p101(sK15(sK22)) | p101(sK22) | ~ p100(sK22) | ~ sP3(sK22)), inference(subsumption_resolution, [], [f683, f197])).
fof(f197, plain, sP2(sK15(sK22)), inference(resolution, [], [f190, f72])).
fof(f72, plain, ! [X0] : (~ sP11(X0) | sP2(X0)), inference(cnf_transformation, [], [f26])).
fof(f190, plain, sP11(sK15(sK22)), inference(subsumption_resolution, [], [f189, f145])).
fof(f189, plain, (~ sP3(sK22) | sP11(sK15(sK22))), inference(subsumption_resolution, [], [f188, f138])).
fof(f188, plain, (~ p100(sK22) | ~ sP3(sK22) | sP11(sK15(sK22))), inference(subsumption_resolution, [], [f185, f137])).
fof(f185, plain, (p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | sP11(sK15(sK22))), inference(resolution, [], [f103, f136])).
fof(f103, plain, ! [X0] : (r1(X0, sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f48])).
fof(f683, plain, (~ sP2(sK15(sK22)) | p102(sK15(sK22)) | ~ p101(sK15(sK22)) | p101(sK22) | ~ p100(sK22) | ~ sP3(sK22)), inference(resolution, [], [f565, f103])).
fof(f565, plain, ! [X1] : (~ r1(sK22, X1) | ~ sP2(X1) | p102(X1) | ~ p101(X1)), inference(subsumption_resolution, [], [f551, f112])).
fof(f112, plain, ! [X0] : (~ p3(sK17(X0)) | p102(X0) | ~ p101(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f53])).
fof(f53, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK16(X0)) & ~ p103(sK16(X0)) & p3(sK16(X0)) & r1(X0, sK16(X0))) & (p102(sK17(X0)) & ~ p103(sK17(X0)) & ~ p3(sK17(X0)) & r1(X0, sK17(X0)))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17])], [f50, f52, f51])).
fof(f51, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK16(X0)) & ~ p103(sK16(X0)) & p3(sK16(X0)) & r1(X0, sK16(X0)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK17(X0)) & ~ p103(sK17(X0)) & ~ p3(sK17(X0)) & r1(X0, sK17(X0)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP2(X0)), inference(rectify, [], [f49])).
fof(f49, plain, ! [X1] : (~ p101(X1) | p102(X1) | (? [X16] : (p102(X16) & ~ p103(X16) & p3(X16) & r1(X1, X16)) & ? [X17] : (p102(X17) & ~ p103(X17) & ~ p3(X17) & r1(X1, X17))) | ~ sP2(X1)), inference(nnf_transformation, [], [f14])).
fof(f551, plain, ! [X1] : (~ p101(X1) | ~ sP2(X1) | p102(X1) | ~ r1(sK22, X1) | p3(sK17(X1))), inference(resolution, [], [f250, f135])).
fof(f135, plain, ! [X2] : (~ r1(sK22, X2) | p3(X2)), inference(cnf_transformation, [], [f66])).
fof(f250, plain, ! [X0, X1] : (r1(X1, sK17(X0)) | ~ p101(X0) | ~ sP2(X0) | p102(X0) | ~ r1(X1, X0)), inference(resolution, [], [f111, f68])).
fof(f68, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | r1(X0, X2) | ~ r1(X0, X1)), inference(cnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2] : (r1(X0, X2) | ~ r1(X1, X2) | ~ r1(X0, X1)), inference(flattening, [], [f8])).
fof(f8, plain, ! [X0, X1, X2] : (r1(X0, X2) | (~ r1(X1, X2) | ~ r1(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((r1(X1, X2) & r1(X0, X1)) => r1(X0, X2)), file('/home/ubuntu/library/tptp/Problems/LCL/LCL674+1.005.p', transitivity)).
fof(f111, plain, ! [X0] : (r1(X0, sK17(X0)) | p102(X0) | ~ p101(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f53])).
fof(f479, plain, spl23_9, inference(avatar_contradiction_clause, [], [f478])).
fof(f478, plain, ($false | spl23_9), inference(subsumption_resolution, [], [f477, f145])).
fof(f477, plain, (~ sP3(sK22) | spl23_9), inference(subsumption_resolution, [], [f476, f138])).
fof(f476, plain, (~ p100(sK22) | ~ sP3(sK22) | spl23_9), inference(subsumption_resolution, [], [f475, f137])).
fof(f475, plain, (p101(sK22) | ~ p100(sK22) | ~ sP3(sK22) | spl23_9), inference(resolution, [], [f231, f106])).
fof(f106, plain, ! [X0] : (p101(sK15(X0)) | p101(X0) | ~ p100(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f48])).
fof(f231, plain, (~ p101(sK15(sK22)) | spl23_9), inference(avatar_component_clause, [], [f230])).