fof(f207536, plain, $false, inference(avatar_sat_refutation, [], [f710, f715, f928, f947, f15862, f207513, f207517, f207525, f207531])).
fof(f207531, plain, (spl20_439 | ~ spl20_1 | ~ spl20_101 | spl20_1130), inference(avatar_split_clause, [], [f207530, f15854, f712, f88, f9083])).
fof(f9083, plain, (spl20_439 <=> p3(sK2(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl20_439])])).
fof(f88, plain, (spl20_1 <=> sP1(sK4)), introduced(avatar_definition, [new_symbols(naming, [spl20_1])])).
fof(f712, plain, (spl20_101 <=> r1(sK16, sK2(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl20_101])])).
fof(f15854, plain, (spl20_1130 <=> p4(sK2(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl20_1130])])).
fof(f207530, plain, (p3(sK2(sK16)) | (~ spl20_1 | ~ spl20_101 | spl20_1130)), inference(subsumption_resolution, [], [f95270, f15855])).
fof(f15855, plain, (~ p4(sK2(sK16)) | spl20_1130), inference(avatar_component_clause, [], [f15854])).
fof(f95270, plain, (p4(sK2(sK16)) | p3(sK2(sK16)) | (~ spl20_1 | ~ spl20_101)), inference(resolution, [], [f90107, f714])).
fof(f714, plain, (r1(sK16, sK2(sK16)) | ~ spl20_101), inference(avatar_component_clause, [], [f712])).
fof(f90107, plain, (! [X3] : (~ r1(sK16, X3) | p4(X3) | p3(X3)) | ~ spl20_1), inference(resolution, [], [f40768, f39])).
fof(f39, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL668+1.001.p', reflexivity)).
fof(f40768, plain, (! [X0, X1] : (~ r1(sK16, X0) | ~ r1(X0, X1) | p4(X1) | p3(X1)) | ~ spl20_1), inference(resolution, [], [f17915, f53])).
fof(f53, plain, r1(sK15, sK16), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ((((((r1(sK9, sK10) & r1(sK8, sK9)) & r1(sK7, sK8)) & r1(sK6, sK7)) & r1(sK5, sK6)) & r1(sK4, sK5)) & ! [X7] : ((! [X8] : ((sP1(X8) & (~ p5(sK11(X8)) & r1(X8, sK11(X8))) & ! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p4(X15) | ~ p5(X15)) & (p5(X15) | p4(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X8, X10))) | ~ r1(X7, X8)) & r1(X7, sK12(X7)) & ! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (((~ p5(X23) | ~ p1(X23)) & (p1(X23) | p5(X23))) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X7, X17))) | ~ r1(sK4, X7)) & r1(sK4, sK13) & (((((r1(sK18, sK19) & r1(sK17, sK18)) & r1(sK16, sK17)) & r1(sK15, sK16)) & r1(sK14, sK15)) & r1(sK4, sK14))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6, sK7, sK8, sK9, sK10, sK11, sK12, sK13, sK14, sK15, sK16, sK17, sK18, sK19])], [f21, f37, f36, f35, f34, f33, f32, f31, f30, f29, f28, f27, f26, f25, f24, f23, f22])).
fof(f22, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X7] : ((! [X8] : ((sP1(X8) & ? [X9] : (~ p5(X9) & r1(X8, X9)) & ! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p4(X15) | ~ p5(X15)) & (p5(X15) | p4(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X8, X10))) | ~ r1(X7, X8)) & ? [X16] : r1(X7, X16) & ! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (((~ p5(X23) | ~ p1(X23)) & (p1(X23) | p5(X23))) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X7, X17))) | ~ r1(X0, X7)) & ? [X24] : r1(X0, X24) & ? [X25] : (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(X25, X26)) & r1(X0, X25))) => (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(sK4, X1)) & ! [X7] : ((! [X8] : ((sP1(X8) & ? [X9] : (~ p5(X9) & r1(X8, X9)) & ! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p4(X15) | ~ p5(X15)) & (p5(X15) | p4(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X8, X10))) | ~ r1(X7, X8)) & ? [X16] : r1(X7, X16) & ! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (((~ p5(X23) | ~ p1(X23)) & (p1(X23) | p5(X23))) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X7, X17))) | ~ r1(sK4, X7)) & ? [X24] : r1(sK4, X24) & ? [X25] : (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(X25, X26)) & r1(sK4, X25)))), introduced(choice_axiom, [])).
fof(f23, plain, (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(sK4, X1)) => (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(sK5, X2)) & r1(sK4, sK5))), introduced(choice_axiom, [])).
fof(f24, plain, (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(sK5, X2)) => (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(sK6, X3)) & r1(sK5, sK6))), introduced(choice_axiom, [])).
fof(f25, plain, (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(sK6, X3)) => (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(sK7, X4)) & r1(sK6, sK7))), introduced(choice_axiom, [])).
fof(f26, plain, (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(sK7, X4)) => (? [X5] : (? [X6] : r1(X5, X6) & r1(sK8, X5)) & r1(sK7, sK8))), introduced(choice_axiom, [])).
fof(f27, plain, (? [X5] : (? [X6] : r1(X5, X6) & r1(sK8, X5)) => (? [X6] : r1(sK9, X6) & r1(sK8, sK9))), introduced(choice_axiom, [])).
fof(f28, plain, (? [X6] : r1(sK9, X6) => r1(sK9, sK10)), introduced(choice_axiom, [])).
fof(f29, plain, ! [X8] : (? [X9] : (~ p5(X9) & r1(X8, X9)) => (~ p5(sK11(X8)) & r1(X8, sK11(X8)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X7] : (? [X16] : r1(X7, X16) => r1(X7, sK12(X7))), introduced(choice_axiom, [])).
fof(f31, plain, (? [X24] : r1(sK4, X24) => r1(sK4, sK13)), introduced(choice_axiom, [])).
fof(f32, plain, (? [X25] : (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(X25, X26)) & r1(sK4, X25)) => (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(sK14, X26)) & r1(sK4, sK14))), introduced(choice_axiom, [])).
fof(f33, plain, (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(sK14, X26)) => (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(sK15, X27)) & r1(sK14, sK15))), introduced(choice_axiom, [])).
fof(f34, plain, (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(sK15, X27)) => (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(sK16, X28)) & r1(sK15, sK16))), introduced(choice_axiom, [])).
fof(f35, plain, (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(sK16, X28)) => (? [X29] : (? [X30] : r1(X29, X30) & r1(sK17, X29)) & r1(sK16, sK17))), introduced(choice_axiom, [])).
fof(f36, plain, (? [X29] : (? [X30] : r1(X29, X30) & r1(sK17, X29)) => (? [X30] : r1(sK18, X30) & r1(sK17, sK18))), introduced(choice_axiom, [])).
fof(f37, plain, (? [X30] : r1(sK18, X30) => r1(sK18, sK19)), introduced(choice_axiom, [])).
fof(f21, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X7] : ((! [X8] : ((sP1(X8) & ? [X9] : (~ p5(X9) & r1(X8, X9)) & ! [X10] : (! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (((~ p4(X15) | ~ p5(X15)) & (p5(X15) | p4(X15))) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ~ r1(X8, X10))) | ~ r1(X7, X8)) & ? [X16] : r1(X7, X16) & ! [X17] : (! [X18] : (! [X19] : (! [X20] : (! [X21] : (! [X22] : (! [X23] : (((~ p5(X23) | ~ p1(X23)) & (p1(X23) | p5(X23))) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X20, X21)) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X7, X17))) | ~ r1(X0, X7)) & ? [X24] : r1(X0, X24) & ? [X25] : (? [X26] : (? [X27] : (? [X28] : (? [X29] : (? [X30] : r1(X29, X30) & r1(X28, X29)) & r1(X27, X28)) & r1(X26, X27)) & r1(X25, X26)) & r1(X0, X25))), inference(rectify, [], [f12])).
fof(f12, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X7] : ((! [X8] : ((sP1(X8) & ? [X26] : (~ p5(X26) & r1(X8, X26)) & ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (((~ p4(X32) | ~ p5(X32)) & (p5(X32) | p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) & ? [X33] : r1(X7, X33) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (((~ p5(X40) | ~ p1(X40)) & (p1(X40) | p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) & ? [X41] : r1(X0, X41) & ? [X42] : (? [X43] : (? [X44] : (? [X45] : (? [X46] : (? [X47] : r1(X46, X47) & r1(X45, X46)) & r1(X44, X45)) & r1(X43, X44)) & r1(X42, X43)) & r1(X0, X42))), inference(definition_folding, [], [f9, e11, e10])).
fof(f10, plain, ! [X9] : (! [X10] : ((! [X11] : (! [X12] : (! [X13] : (! [X14] : (((~ p1(X14) | ~ p2(X14)) & (p2(X14) | p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) & ? [X15] : (~ p3(X15) & r1(X10, X15)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (((~ p2(X19) | ~ p3(X19)) & (p3(X19) | p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ~ sP0(X9)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X9] : (sP0(X9) <=> ! [X10] : ((! [X11] : (! [X12] : (! [X13] : (! [X14] : (((~ p1(X14) | ~ p2(X14)) & (p2(X14) | p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) & ? [X15] : (~ p3(X15) & r1(X10, X15)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (((~ p2(X19) | ~ p3(X19)) & (p3(X19) | p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f11, plain, ! [X8] : (! [X9] : ((sP0(X9) & ? [X20] : (~ p4(X20) & r1(X9, X20)) & ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (((~ p3(X25) | ~ p4(X25)) & (p4(X25) | p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ~ sP1(X8)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X8] : (sP1(X8) <=> ! [X9] : ((sP0(X9) & ? [X20] : (~ p4(X20) & r1(X9, X20)) & ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (((~ p3(X25) | ~ p4(X25)) & (p4(X25) | p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f9, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : r1(X5, X6) & r1(X4, X5)) & r1(X3, X4)) & r1(X2, X3)) & r1(X1, X2)) & r1(X0, X1)) & ! [X7] : ((! [X8] : ((! [X9] : ((! [X10] : ((! [X11] : (! [X12] : (! [X13] : (! [X14] : (((~ p1(X14) | ~ p2(X14)) & (p2(X14) | p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) & ? [X15] : (~ p3(X15) & r1(X10, X15)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (((~ p2(X19) | ~ p3(X19)) & (p3(X19) | p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) & ? [X20] : (~ p4(X20) & r1(X9, X20)) & ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (((~ p3(X25) | ~ p4(X25)) & (p4(X25) | p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) & ? [X26] : (~ p5(X26) & r1(X8, X26)) & ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (((~ p4(X32) | ~ p5(X32)) & (p5(X32) | p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) & ? [X33] : r1(X7, X33) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (((~ p5(X40) | ~ p1(X40)) & (p1(X40) | p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) & ? [X41] : r1(X0, X41) & ? [X42] : (? [X43] : (? [X44] : (? [X45] : (? [X46] : (? [X47] : r1(X46, X47) & r1(X45, X46)) & r1(X44, X45)) & r1(X43, X44)) & r1(X42, X43)) & r1(X0, X42))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : ~ r1(X5, X6) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X7] : (~ (~ ! [X8] : (~ (~ ! [X9] : (~ (~ ! [X10] : (~ (~ ! [X11] : (! [X12] : (! [X13] : (! [X14] : (~ ((p1(X14) & p2(X14)) | (~ p2(X14) & ~ p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ! [X15] : (p3(X15) | ~ r1(X10, X15)) | ~ ! [X16] : (! [X17] : (! [X18] : (! [X19] : (~ ((p2(X19) & p3(X19)) | (~ p3(X19) & ~ p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ! [X20] : (p4(X20) | ~ r1(X9, X20)) | ~ ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (~ ((p3(X25) & p4(X25)) | (~ p4(X25) & ~ p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ! [X26] : (p5(X26) | ~ r1(X8, X26)) | ~ ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (~ ((p4(X32) & p5(X32)) | (~ p5(X32) & ~ p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) | ! [X33] : ~ r1(X7, X33) | ~ ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (~ ((p5(X40) & p1(X40)) | (~ p1(X40) & ~ p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) | ! [X41] : ~ r1(X0, X41) | ! [X42] : (! [X43] : (! [X44] : (! [X45] : (! [X46] : (! [X47] : ~ r1(X46, X47) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X0, X42))), inference(pure_predicate_removal, [], [f7])).
fof(f7, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : ((p1(X6) & p2(X6) & p3(X6) & p4(X6) & p5(X6) & p6(X6)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X7] : (~ (~ ! [X8] : (~ (~ ! [X9] : (~ (~ ! [X10] : (~ (~ ! [X11] : (! [X12] : (! [X13] : (! [X14] : (~ ((p1(X14) & p2(X14)) | (~ p2(X14) & ~ p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ! [X15] : (p3(X15) | ~ r1(X10, X15)) | ~ ! [X16] : (! [X17] : (! [X18] : (! [X19] : (~ ((p2(X19) & p3(X19)) | (~ p3(X19) & ~ p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ! [X20] : (p4(X20) | ~ r1(X9, X20)) | ~ ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (~ ((p3(X25) & p4(X25)) | (~ p4(X25) & ~ p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ! [X26] : (p5(X26) | ~ r1(X8, X26)) | ~ ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (~ ((p4(X32) & p5(X32)) | (~ p5(X32) & ~ p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) | ! [X33] : (p6(X33) | ~ r1(X7, X33)) | ~ ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (~ ((p5(X40) & p1(X40)) | (~ p1(X40) & ~ p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) | ! [X41] : ~ r1(X0, X41) | ! [X42] : (! [X43] : (! [X44] : (! [X45] : (! [X46] : (! [X47] : ~ r1(X46, X47) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X0, X42))), inference(pure_predicate_removal, [], [f6])).
fof(f6, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : ((p1(X6) & p2(X6) & p3(X6) & p4(X6) & p5(X6) & p6(X6)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X7] : (~ (~ ! [X8] : (~ (~ ! [X9] : (~ (~ ! [X10] : (~ (~ ! [X11] : (! [X12] : (! [X13] : (! [X14] : (~ ((p1(X14) & p2(X14)) | (~ p2(X14) & ~ p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ! [X15] : (p3(X15) | ~ r1(X10, X15)) | ~ ! [X16] : (! [X17] : (! [X18] : (! [X19] : (~ ((p2(X19) & p3(X19)) | (~ p3(X19) & ~ p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ! [X20] : (p4(X20) | ~ r1(X9, X20)) | ~ ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (~ ((p3(X25) & p4(X25)) | (~ p4(X25) & ~ p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ! [X26] : (p5(X26) | ~ r1(X8, X26)) | ~ ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (~ ((p4(X32) & p5(X32)) | (~ p5(X32) & ~ p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) | ! [X33] : (p6(X33) | ~ r1(X7, X33)) | ~ ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (~ ((p5(X40) & p1(X40)) | (~ p1(X40) & ~ p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) | ! [X41] : ~ r1(X0, X41) | ! [X42] : (! [X43] : (! [X44] : (! [X45] : (! [X46] : (! [X47] : ((~ p2(X47) & ~ p4(X47) & ~ p6(X47) & ~ p8(X47) & ~ p10(X47) & ~ p12(X47)) | ~ r1(X46, X47)) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X0, X42))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : ((p1(X6) & p2(X6) & p3(X6) & p4(X6) & p5(X6) & p6(X6)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X7] : (~ (~ ! [X8] : (~ (~ ! [X9] : (~ (~ ! [X10] : (~ (~ ! [X11] : (! [X12] : (! [X13] : (! [X14] : (~ ((p1(X14) & p2(X14)) | (~ p2(X14) & ~ p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ! [X15] : (p3(X15) | ~ r1(X10, X15)) | ~ ! [X16] : (! [X17] : (! [X18] : (! [X19] : (~ ((p2(X19) & p3(X19)) | (~ p3(X19) & ~ p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ! [X20] : (p4(X20) | ~ r1(X9, X20)) | ~ ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (~ ((p3(X25) & p4(X25)) | (~ p4(X25) & ~ p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ! [X26] : (p5(X26) | ~ r1(X8, X26)) | ~ ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (~ ((p4(X32) & p5(X32)) | (~ p5(X32) & ~ p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) | ! [X33] : (p6(X33) | ~ r1(X7, X33)) | ~ ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (~ ((p5(X40) & p1(X40)) | (~ p1(X40) & ~ p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) | ! [X41] : (p7(X41) | ~ r1(X0, X41)) | ! [X42] : (! [X43] : (! [X44] : (! [X45] : (! [X46] : (! [X47] : ((~ p2(X47) & ~ p4(X47) & ~ p6(X47) & ~ p8(X47) & ~ p10(X47) & ~ p12(X47)) | ~ r1(X46, X47)) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X0, X42))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : ((p1(X6) & p2(X6) & p3(X6) & p4(X6) & p5(X6) & p6(X6)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ ! [X7] : (~ (~ ! [X8] : (~ (~ ! [X9] : (~ (~ ! [X10] : (~ (~ ! [X11] : (~ ~ ! [X12] : (! [X13] : (! [X14] : (~ ((p1(X14) & p2(X14)) | (~ p2(X14) & ~ p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) | ! [X15] : (p3(X15) | ~ r1(X10, X15)) | ~ ! [X16] : (! [X17] : (! [X18] : (! [X19] : (~ ((p2(X19) & p3(X19)) | (~ p3(X19) & ~ p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ! [X20] : (p4(X20) | ~ r1(X9, X20)) | ~ ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (~ ((p3(X25) & p4(X25)) | (~ p4(X25) & ~ p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ! [X26] : (p5(X26) | ~ r1(X8, X26)) | ~ ! [X27] : (! [X28] : (! [X29] : (! [X30] : (! [X31] : (! [X32] : (~ ((p4(X32) & p5(X32)) | (~ p5(X32) & ~ p4(X32))) | ~ r1(X31, X32)) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | ~ r1(X28, X29)) | ~ r1(X27, X28)) | ~ r1(X8, X27))) | ~ r1(X7, X8)) | ! [X33] : (p6(X33) | ~ r1(X7, X33)) | ~ ! [X34] : (! [X35] : (! [X36] : (! [X37] : (! [X38] : (! [X39] : (! [X40] : (~ ((p5(X40) & p1(X40)) | (~ p1(X40) & ~ p5(X40))) | ~ r1(X39, X40)) | ~ r1(X38, X39)) | ~ r1(X37, X38)) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X7, X34))) | ~ r1(X0, X7)) | ! [X41] : (p7(X41) | ~ r1(X0, X41)) | ! [X42] : (! [X43] : (! [X44] : (! [X45] : (! [X46] : (! [X47] : ((~ p2(X47) & ~ p4(X47) & ~ p6(X47) & ~ p8(X47) & ~ p10(X47) & ~ p12(X47)) | ~ r1(X46, X47)) | ~ r1(X45, X46)) | ~ r1(X44, X45)) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | ~ r1(X0, X42))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((p1(X0) & p2(X0) & p3(X0) & p4(X0) & p5(X0) & p6(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ ~ ! [X0] : (! [X1] : (! [X0] : (~ ((p1(X0) & p2(X0)) | (~ p2(X0) & ~ p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p2(X0) & p3(X0)) | (~ p3(X0) & ~ p2(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p4(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p3(X0) & p4(X0)) | (~ p4(X0) & ~ p3(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p5(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p4(X0) & p5(X0)) | (~ p5(X0) & ~ p4(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p6(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p5(X0) & p1(X0)) | (~ p1(X0) & ~ p5(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p7(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((~ p2(X0) & ~ p4(X0) & ~ p6(X0) & ~ p8(X0) & ~ p10(X0) & ~ p12(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((p1(X0) & p2(X0) & p3(X0) & p4(X0) & p5(X0) & p6(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ (~ ! [X0] : (~ (~ ! [X1] : (~ ~ ! [X0] : (! [X1] : (! [X0] : (~ ((p1(X0) & p2(X0)) | (~ p2(X0) & ~ p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p3(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p2(X0) & p3(X0)) | (~ p3(X0) & ~ p2(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p4(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p3(X0) & p4(X0)) | (~ p4(X0) & ~ p3(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p5(X1) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p4(X0) & p5(X0)) | (~ p5(X0) & ~ p4(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ r1(X1, X0)) | ! [X0] : (p6(X0) | ~ r1(X1, X0)) | ~ ! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ((p5(X0) & p1(X0)) | (~ p1(X0) & ~ p5(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0))) | ~ r1(X0, X1)) | ! [X1] : (p7(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((~ p2(X0) & ~ p4(X0) & ~ p6(X0) & ~ p8(X0) & ~ p10(X0) & ~ p12(X0)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL668+1.001.p', main)).
fof(f17915, plain, (! [X2, X0, X1] : (~ r1(sK15, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | p4(X1) | p3(X1)) | ~ spl20_1), inference(resolution, [], [f2892, f52])).
fof(f52, plain, r1(sK14, sK15), inference(cnf_transformation, [], [f38])).
fof(f2892, plain, (! [X54, X52, X55, X53] : (~ r1(sK14, X55) | ~ r1(X53, X52) | ~ r1(X54, X53) | ~ r1(X55, X54) | p4(X52) | p3(X52)) | ~ spl20_1), inference(subsumption_resolution, [], [f2853, f89])).
fof(f89, plain, (sP1(sK4) | ~ spl20_1), inference(avatar_component_clause, [], [f88])).
fof(f2853, plain, ! [X54, X52, X55, X53] : (p4(X52) | ~ r1(X53, X52) | ~ r1(X54, X53) | ~ r1(X55, X54) | ~ r1(sK14, X55) | p3(X52) | ~ sP1(sK4)), inference(resolution, [], [f518, f51])).
fof(f51, plain, r1(sK4, sK14), inference(cnf_transformation, [], [f38])).
fof(f518, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(X5, X4) | p4(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | p3(X0) | ~ sP1(X5)), inference(resolution, [], [f40, f39])).
fof(f40, plain, ! [X6, X4, X0, X7, X5, X3, X1] : (~ r1(X6, X7) | p3(X7) | p4(X7) | ~ r1(X5, X6) | ~ r1(X4, X5) | ~ r1(X3, X4) | ~ r1(X1, X3) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (! [X1] : ((sP0(X1) & (~ p4(sK2(X1)) & r1(X1, sK2(X1))) & ! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (((~ p3(X7) | ~ p4(X7)) & (p4(X7) | p3(X7))) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X1, X3))) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2])], [f14, f15])).
fof(f15, plain, ! [X1] : (? [X2] : (~ p4(X2) & r1(X1, X2)) => (~ p4(sK2(X1)) & r1(X1, sK2(X1)))), introduced(choice_axiom, [])).
fof(f14, plain, ! [X0] : (! [X1] : ((sP0(X1) & ? [X2] : (~ p4(X2) & r1(X1, X2)) & ! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (((~ p3(X7) | ~ p4(X7)) & (p4(X7) | p3(X7))) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X1, X3))) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f13])).
fof(f13, plain, ! [X8] : (! [X9] : ((sP0(X9) & ? [X20] : (~ p4(X20) & r1(X9, X20)) & ! [X21] : (! [X22] : (! [X23] : (! [X24] : (! [X25] : (((~ p3(X25) | ~ p4(X25)) & (p4(X25) | p3(X25))) | ~ r1(X24, X25)) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X21, X22)) | ~ r1(X9, X21))) | ~ r1(X8, X9)) | ~ sP1(X8)), inference(nnf_transformation, [], [f11])).
fof(f207525, plain, (~ spl20_439 | ~ spl20_438 | ~ spl20_25 | ~ spl20_101), inference(avatar_split_clause, [], [f38837, f712, f234, f9077, f9083])).
fof(f9077, plain, (spl20_438 <=> p2(sK2(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl20_438])])).
fof(f234, plain, (spl20_25 <=> sP0(sK4)), introduced(avatar_definition, [new_symbols(naming, [spl20_25])])).
fof(f38837, plain, (~ p2(sK2(sK16)) | ~ p3(sK2(sK16)) | (~ spl20_25 | ~ spl20_101)), inference(resolution, [], [f32560, f714])).
fof(f32560, plain, (! [X0] : (~ r1(sK16, X0) | ~ p2(X0) | ~ p3(X0)) | ~ spl20_25), inference(resolution, [], [f12301, f53])).
fof(f12301, plain, (! [X0, X1] : (~ r1(sK15, X0) | ~ r1(X0, X1) | ~ p2(X1) | ~ p3(X1)) | ~ spl20_25), inference(resolution, [], [f2710, f52])).
fof(f2710, plain, (! [X43, X41, X42] : (~ r1(sK14, X43) | ~ r1(X42, X41) | ~ r1(X43, X42) | ~ p2(X41) | ~ p3(X41)) | ~ spl20_25), inference(subsumption_resolution, [], [f2676, f235])).
fof(f235, plain, (sP0(sK4) | ~ spl20_25), inference(avatar_component_clause, [], [f234])).
fof(f2676, plain, ! [X43, X41, X42] : (~ p2(X41) | ~ r1(X42, X41) | ~ r1(X43, X42) | ~ r1(sK14, X43) | ~ p3(X41) | ~ sP0(sK4)), inference(resolution, [], [f473, f51])).
fof(f473, plain, ! [X4, X2, X0, X3, X1] : (~ r1(X4, X3) | ~ p2(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ p3(X0) | ~ sP0(X4)), inference(resolution, [], [f46, f39])).
fof(f46, plain, ! [X0, X10, X8, X7, X1, X9] : (~ r1(X9, X10) | ~ p3(X10) | ~ p2(X10) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X1, X7) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (! [X1] : ((! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p1(X5) | ~ p2(X5)) & (p2(X5) | p1(X5))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & (~ p3(sK3(X1)) & r1(X1, sK3(X1))) & ! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p2(X10) | ~ p3(X10)) & (p3(X10) | p2(X10))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X1, X7))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f18, f19])).
fof(f19, plain, ! [X1] : (? [X6] : (~ p3(X6) & r1(X1, X6)) => (~ p3(sK3(X1)) & r1(X1, sK3(X1)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (! [X1] : ((! [X2] : (! [X3] : (! [X4] : (! [X5] : (((~ p1(X5) | ~ p2(X5)) & (p2(X5) | p1(X5))) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ? [X6] : (~ p3(X6) & r1(X1, X6)) & ! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p2(X10) | ~ p3(X10)) & (p3(X10) | p2(X10))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X1, X7))) | ~ r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X9] : (! [X10] : ((! [X11] : (! [X12] : (! [X13] : (! [X14] : (((~ p1(X14) | ~ p2(X14)) & (p2(X14) | p1(X14))) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X10, X11)) & ? [X15] : (~ p3(X15) & r1(X10, X15)) & ! [X16] : (! [X17] : (! [X18] : (! [X19] : (((~ p2(X19) | ~ p3(X19)) & (p3(X19) | p2(X19))) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X10, X16))) | ~ r1(X9, X10)) | ~ sP0(X9)), inference(nnf_transformation, [], [f10])).
fof(f207517, plain, (spl20_437 | spl20_438 | ~ spl20_25 | ~ spl20_101), inference(avatar_split_clause, [], [f40207, f712, f234, f9077, f9073])).
fof(f9073, plain, (spl20_437 <=> p1(sK2(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl20_437])])).
fof(f40207, plain, (p2(sK2(sK16)) | p1(sK2(sK16)) | (~ spl20_25 | ~ spl20_101)), inference(resolution, [], [f32688, f714])).
fof(f32688, plain, (! [X0] : (~ r1(sK16, X0) | p2(X0) | p1(X0)) | ~ spl20_25), inference(resolution, [], [f12521, f53])).
fof(f12521, plain, (! [X0, X1] : (~ r1(sK15, X0) | ~ r1(X0, X1) | p2(X1) | p1(X1)) | ~ spl20_25), inference(resolution, [], [f2769, f52])).
fof(f2769, plain, (! [X43, X41, X42] : (~ r1(sK14, X43) | ~ r1(X42, X41) | ~ r1(X43, X42) | p2(X41) | p1(X41)) | ~ spl20_25), inference(subsumption_resolution, [], [f2735, f235])).
fof(f2735, plain, ! [X43, X41, X42] : (p2(X41) | ~ r1(X42, X41) | ~ r1(X43, X42) | ~ r1(sK14, X43) | p1(X41) | ~ sP0(sK4)), inference(resolution, [], [f488, f51])).
fof(f488, plain, ! [X4, X2, X0, X3, X1] : (~ r1(X4, X3) | p2(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | p1(X0) | ~ sP0(X4)), inference(resolution, [], [f49, f39])).
fof(f49, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(X4, X5) | p1(X5) | p2(X5) | ~ r1(X3, X4) | ~ r1(X2, X3) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f207513, plain, (~ spl20_437 | ~ spl20_101 | spl20_1130), inference(avatar_split_clause, [], [f207512, f15854, f712, f9073])).
fof(f207512, plain, (~ p1(sK2(sK16)) | (~ spl20_101 | spl20_1130)), inference(subsumption_resolution, [], [f195474, f207460])).
fof(f207460, plain, (p5(sK2(sK16)) | (~ spl20_101 | spl20_1130)), inference(subsumption_resolution, [], [f207451, f15855])).
fof(f207451, plain, (p4(sK2(sK16)) | p5(sK2(sK16)) | ~ spl20_101), inference(resolution, [], [f207421, f714])).
fof(f207421, plain, ! [X3] : (~ r1(sK16, X3) | p4(X3) | p5(X3)), inference(resolution, [], [f207398, f39])).
fof(f207398, plain, ! [X6, X7] : (~ r1(sK16, X6) | ~ r1(X6, X7) | p4(X7) | p5(X7)), inference(resolution, [], [f118160, f39])).
fof(f118160, plain, ! [X10, X11, X9] : (~ r1(sK16, X9) | ~ r1(X9, X10) | ~ r1(X10, X11) | p4(X11) | p5(X11)), inference(resolution, [], [f56222, f39])).
fof(f56222, plain, ! [X14, X12, X15, X13] : (~ r1(sK16, X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | ~ r1(X14, X15) | p4(X15) | p5(X15)), inference(resolution, [], [f20707, f39])).
fof(f20707, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK16, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | p4(X4) | p5(X4)), inference(resolution, [], [f3249, f53])).
fof(f3249, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK15, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | p4(X1) | p5(X1)), inference(resolution, [], [f684, f52])).
fof(f684, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK14, X6) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | p4(X0) | p5(X0)), inference(resolution, [], [f61, f51])).
fof(f61, plain, ! [X14, X12, X10, X8, X7, X15, X13, X11] : (~ r1(sK4, X7) | p4(X15) | ~ r1(X14, X15) | ~ r1(X13, X14) | ~ r1(X12, X13) | ~ r1(X11, X12) | ~ r1(X10, X11) | ~ r1(X8, X10) | ~ r1(X7, X8) | p5(X15)), inference(cnf_transformation, [], [f38])).
fof(f195474, plain, (~ p1(sK2(sK16)) | ~ p5(sK2(sK16)) | ~ spl20_101), inference(resolution, [], [f195452, f714])).
fof(f195452, plain, ! [X3] : (~ r1(sK16, X3) | ~ p1(X3) | ~ p5(X3)), inference(resolution, [], [f195421, f39])).
fof(f195421, plain, ! [X6, X7] : (~ r1(sK16, X6) | ~ r1(X6, X7) | ~ p1(X7) | ~ p5(X7)), inference(resolution, [], [f117962, f39])).
fof(f117962, plain, ! [X10, X11, X9] : (~ r1(sK16, X9) | ~ r1(X9, X10) | ~ r1(X10, X11) | ~ p1(X11) | ~ p5(X11)), inference(resolution, [], [f56012, f39])).
fof(f56012, plain, ! [X14, X12, X15, X13] : (~ r1(sK16, X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | ~ r1(X14, X15) | ~ p1(X15) | ~ p5(X15)), inference(resolution, [], [f20399, f39])).
fof(f20399, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK16, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | ~ p1(X4) | ~ p5(X4)), inference(resolution, [], [f3082, f53])).
fof(f3082, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK15, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | ~ p1(X1) | ~ p5(X1)), inference(resolution, [], [f678, f52])).
fof(f678, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK14, X6) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ p1(X0) | ~ p5(X0)), inference(resolution, [], [f59, f51])).
fof(f59, plain, ! [X23, X21, X19, X17, X7, X22, X20, X18] : (~ r1(sK4, X7) | ~ p1(X23) | ~ r1(X22, X23) | ~ r1(X21, X22) | ~ r1(X20, X21) | ~ r1(X19, X20) | ~ r1(X18, X19) | ~ r1(X17, X18) | ~ r1(X7, X17) | ~ p5(X23)), inference(cnf_transformation, [], [f38])).
fof(f15862, plain, (~ spl20_17 | ~ spl20_1130), inference(avatar_contradiction_clause, [], [f15861])).
fof(f15861, plain, ($false | (~ spl20_17 | ~ spl20_1130)), inference(subsumption_resolution, [], [f15860, f162])).
fof(f162, plain, (sP1(sK15) | ~ spl20_17), inference(avatar_component_clause, [], [f161])).
fof(f161, plain, (spl20_17 <=> sP1(sK15)), introduced(avatar_definition, [new_symbols(naming, [spl20_17])])).
fof(f15860, plain, (~ sP1(sK15) | ~ spl20_1130), inference(resolution, [], [f15858, f53])).
fof(f15858, plain, (! [X0] : (~ r1(X0, sK16) | ~ sP1(X0)) | ~ spl20_1130), inference(resolution, [], [f15856, f43])).
fof(f43, plain, ! [X0, X1] : (~ p4(sK2(X1)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f15856, plain, (p4(sK2(sK16)) | ~ spl20_1130), inference(avatar_component_clause, [], [f15854])).
fof(f947, plain, (spl20_25 | ~ spl20_1), inference(avatar_split_clause, [], [f944, f88, f234])).
fof(f944, plain, (sP0(sK4) | ~ spl20_1), inference(resolution, [], [f89, f72])).
fof(f72, plain, ! [X0] : (~ sP1(X0) | sP0(X0)), inference(resolution, [], [f44, f39])).
fof(f44, plain, ! [X0, X1] : (~ r1(X0, X1) | sP0(X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f928, plain, spl20_1, inference(avatar_split_clause, [], [f923, f88])).
fof(f923, plain, sP1(sK4), inference(resolution, [], [f199, f39])).
fof(f199, plain, ! [X3] : (~ r1(sK4, X3) | sP1(X3)), inference(resolution, [], [f65, f39])).
fof(f65, plain, ! [X8, X7] : (~ r1(sK4, X7) | ~ r1(X7, X8) | sP1(X8)), inference(cnf_transformation, [], [f38])).
fof(f715, plain, (~ spl20_17 | spl20_101), inference(avatar_split_clause, [], [f213, f712, f161])).
fof(f213, plain, (r1(sK16, sK2(sK16)) | ~ sP1(sK15)), inference(resolution, [], [f42, f53])).
fof(f42, plain, ! [X0, X1] : (~ r1(X0, X1) | r1(X1, sK2(X1)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f710, plain, spl20_17, inference(avatar_split_clause, [], [f696, f161])).
fof(f696, plain, sP1(sK15), inference(resolution, [], [f196, f52])).
fof(f196, plain, ! [X0] : (~ r1(sK14, X0) | sP1(X0)), inference(resolution, [], [f65, f51])).