fof(f49604, plain, $false, inference(avatar_sat_refutation, [], [f1548, f1555, f1568, f1594, f1654, f2876, f2952, f3072, f7330, f7372, f7479, f15360, f15826, f15942, f49603])).
fof(f49603, plain, (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031), inference(avatar_contradiction_clause, [], [f49602])).
fof(f49602, plain, ($false | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031)), inference(subsumption_resolution, [], [f49601, f12273])).
fof(f12273, plain, (sP5(sK28(sK26(sK25(sK42)))) | ~ spl43_2030), inference(avatar_component_clause, [], [f12272])).
fof(f12272, plain, (spl43_2030 <=> sP5(sK28(sK26(sK25(sK42))))), introduced(avatar_definition, [new_symbols(naming, [spl43_2030])])).
fof(f49601, plain, (~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031)), inference(subsumption_resolution, [], [f49600, f12262])).
fof(f12262, plain, (p103(sK28(sK26(sK25(sK42)))) | ~ spl43_2027), inference(avatar_component_clause, [], [f12260])).
fof(f12260, plain, (spl43_2027 <=> p103(sK28(sK26(sK25(sK42))))), introduced(avatar_definition, [new_symbols(naming, [spl43_2027])])).
fof(f49600, plain, (~ p103(sK28(sK26(sK25(sK42)))) | ~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031)), inference(subsumption_resolution, [], [f49599, f12277])).
fof(f12277, plain, (~ p104(sK28(sK26(sK25(sK42)))) | spl43_2031), inference(avatar_component_clause, [], [f12276])).
fof(f12276, plain, (spl43_2031 <=> p104(sK28(sK26(sK25(sK42))))), introduced(avatar_definition, [new_symbols(naming, [spl43_2031])])).
fof(f49599, plain, (p104(sK28(sK26(sK25(sK42)))) | ~ p103(sK28(sK26(sK25(sK42)))) | ~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031)), inference(resolution, [], [f49577, f194])).
fof(f194, plain, ! [X0] : (~ p5(sK31(X0)) | p104(X0) | ~ p103(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ! [X0] : (~ p103(X0) | p104(X0) | ((p104(sK30(X0)) & ~ p105(sK30(X0)) & p5(sK30(X0)) & r1(X0, sK30(X0))) & (p104(sK31(X0)) & ~ p105(sK31(X0)) & ~ p5(sK31(X0)) & r1(X0, sK31(X0)))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31])], [f77, f79, f78])).
fof(f78, plain, ! [X0] : (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) => (p104(sK30(X0)) & ~ p105(sK30(X0)) & p5(sK30(X0)) & r1(X0, sK30(X0)))), introduced(choice_axiom, [])).
fof(f79, plain, ! [X0] : (? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2)) => (p104(sK31(X0)) & ~ p105(sK31(X0)) & ~ p5(sK31(X0)) & r1(X0, sK31(X0)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X0] : (~ p103(X0) | p104(X0) | (? [X1] : (p104(X1) & ~ p105(X1) & p5(X1) & r1(X0, X1)) & ? [X2] : (p104(X2) & ~ p105(X2) & ~ p5(X2) & r1(X0, X2))) | ~ sP5(X0)), inference(rectify, [], [f76])).
fof(f76, plain, ! [X10] : (~ p103(X10) | p104(X10) | (? [X39] : (p104(X39) & ~ p105(X39) & p5(X39) & r1(X10, X39)) & ? [X40] : (p104(X40) & ~ p105(X40) & ~ p5(X40) & r1(X10, X40))) | ~ sP5(X10)), inference(nnf_transformation, [], [f14])).
fof(f14, plain, ! [X10] : (~ p103(X10) | p104(X10) | (? [X39] : (p104(X39) & ~ p105(X39) & p5(X39) & r1(X10, X39)) & ? [X40] : (p104(X40) & ~ p105(X40) & ~ p5(X40) & r1(X10, X40))) | ~ sP5(X10)), inference(usedef, [], [e14])).
fof(e14, plain, ! [X10] : (sP5(X10) <=> (~ p103(X10) | p104(X10) | (? [X39] : (p104(X39) & ~ p105(X39) & p5(X39) & r1(X10, X39)) & ? [X40] : (p104(X40) & ~ p105(X40) & ~ p5(X40) & r1(X10, X40))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f49577, plain, (p5(sK31(sK28(sK26(sK25(sK42))))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | ~ spl43_2030 | spl43_2031)), inference(subsumption_resolution, [], [f49576, f12273])).
fof(f49576, plain, (p5(sK31(sK28(sK26(sK25(sK42))))) | ~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2027 | spl43_2031)), inference(subsumption_resolution, [], [f49575, f12262])).
fof(f49575, plain, (p5(sK31(sK28(sK26(sK25(sK42))))) | ~ p103(sK28(sK26(sK25(sK42)))) | ~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810 | spl43_2031)), inference(subsumption_resolution, [], [f49564, f12277])).
fof(f49564, plain, (p5(sK31(sK28(sK26(sK25(sK42))))) | p104(sK28(sK26(sK25(sK42)))) | ~ p103(sK28(sK26(sK25(sK42)))) | ~ sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810)), inference(resolution, [], [f49523, f193])).
fof(f193, plain, ! [X0] : (r1(X0, sK31(X0)) | p104(X0) | ~ p103(X0) | ~ sP5(X0)), inference(cnf_transformation, [], [f80])).
fof(f49523, plain, (! [X7] : (~ r1(sK28(sK26(sK25(sK42))), X7) | p5(X7)) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810)), inference(subsumption_resolution, [], [f49522, f5319])).
fof(f5319, plain, (sP6(sK26(sK25(sK42))) | ~ spl43_809), inference(avatar_component_clause, [], [f5318])).
fof(f5318, plain, (spl43_809 <=> sP6(sK26(sK25(sK42)))), introduced(avatar_definition, [new_symbols(naming, [spl43_809])])).
fof(f49522, plain, (! [X7] : (~ r1(sK28(sK26(sK25(sK42))), X7) | p5(X7) | ~ sP6(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | spl43_810)), inference(subsumption_resolution, [], [f49521, f5308])).
fof(f5308, plain, (p102(sK26(sK25(sK42))) | ~ spl43_806), inference(avatar_component_clause, [], [f5306])).
fof(f5306, plain, (spl43_806 <=> p102(sK26(sK25(sK42)))), introduced(avatar_definition, [new_symbols(naming, [spl43_806])])).
fof(f49521, plain, (! [X7] : (~ r1(sK28(sK26(sK25(sK42))), X7) | p5(X7) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317 | spl43_810)), inference(subsumption_resolution, [], [f49507, f5323])).
fof(f5323, plain, (~ p103(sK26(sK25(sK42))) | spl43_810), inference(avatar_component_clause, [], [f5322])).
fof(f5322, plain, (spl43_810 <=> p103(sK26(sK25(sK42)))), introduced(avatar_definition, [new_symbols(naming, [spl43_810])])).
fof(f49507, plain, (! [X7] : (~ r1(sK28(sK26(sK25(sK42))), X7) | p5(X7) | p103(sK26(sK25(sK42))) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42)))) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(resolution, [], [f49469, f189])).
fof(f189, plain, ! [X0] : (r1(X0, sK28(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ! [X0] : (~ p102(X0) | p103(X0) | ((p103(sK28(X0)) & ~ p104(sK28(X0)) & p4(sK28(X0)) & r1(X0, sK28(X0))) & (p103(sK29(X0)) & ~ p104(sK29(X0)) & ~ p4(sK29(X0)) & r1(X0, sK29(X0)))) | ~ sP6(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29])], [f72, f74, f73])).
fof(f73, plain, ! [X0] : (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) => (p103(sK28(X0)) & ~ p104(sK28(X0)) & p4(sK28(X0)) & r1(X0, sK28(X0)))), introduced(choice_axiom, [])).
fof(f74, plain, ! [X0] : (? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2)) => (p103(sK29(X0)) & ~ p104(sK29(X0)) & ~ p4(sK29(X0)) & r1(X0, sK29(X0)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X0] : (~ p102(X0) | p103(X0) | (? [X1] : (p103(X1) & ~ p104(X1) & p4(X1) & r1(X0, X1)) & ? [X2] : (p103(X2) & ~ p104(X2) & ~ p4(X2) & r1(X0, X2))) | ~ sP6(X0)), inference(rectify, [], [f71])).
fof(f71, plain, ! [X10] : (~ p102(X10) | p103(X10) | (? [X37] : (p103(X37) & ~ p104(X37) & p4(X37) & r1(X10, X37)) & ? [X38] : (p103(X38) & ~ p104(X38) & ~ p4(X38) & r1(X10, X38))) | ~ sP6(X10)), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ! [X10] : (~ p102(X10) | p103(X10) | (? [X37] : (p103(X37) & ~ p104(X37) & p4(X37) & r1(X10, X37)) & ? [X38] : (p103(X38) & ~ p104(X38) & ~ p4(X38) & r1(X10, X38))) | ~ sP6(X10)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X10] : (sP6(X10) <=> (~ p102(X10) | p103(X10) | (? [X37] : (p103(X37) & ~ p104(X37) & p4(X37) & r1(X10, X37)) & ? [X38] : (p103(X38) & ~ p104(X38) & ~ p4(X38) & r1(X10, X38))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f49469, plain, (! [X10, X11] : (~ r1(sK26(sK25(sK42)), X11) | ~ r1(X11, X10) | p5(X10)) | (~ spl43_91 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(subsumption_resolution, [], [f49468, f2695])).
fof(f2695, plain, (sP7(sK25(sK42)) | ~ spl43_316), inference(avatar_component_clause, [], [f2694])).
fof(f2694, plain, (spl43_316 <=> sP7(sK25(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_316])])).
fof(f49468, plain, (! [X10, X11] : (p5(X10) | ~ r1(X11, X10) | ~ r1(sK26(sK25(sK42)), X11) | ~ sP7(sK25(sK42))) | (~ spl43_91 | ~ spl43_313 | spl43_317)), inference(subsumption_resolution, [], [f49467, f2684])).
fof(f2684, plain, (p101(sK25(sK42)) | ~ spl43_313), inference(avatar_component_clause, [], [f2682])).
fof(f2682, plain, (spl43_313 <=> p101(sK25(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_313])])).
fof(f49467, plain, (! [X10, X11] : (p5(X10) | ~ r1(X11, X10) | ~ r1(sK26(sK25(sK42)), X11) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42))) | (~ spl43_91 | spl43_317)), inference(subsumption_resolution, [], [f49451, f2699])).
fof(f2699, plain, (~ p102(sK25(sK42)) | spl43_317), inference(avatar_component_clause, [], [f2698])).
fof(f2698, plain, (spl43_317 <=> p102(sK25(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_317])])).
fof(f49451, plain, (! [X10, X11] : (p5(X10) | ~ r1(X11, X10) | ~ r1(sK26(sK25(sK42)), X11) | p102(sK25(sK42)) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42))) | ~ spl43_91), inference(resolution, [], [f1653, f181])).
fof(f181, plain, ! [X0] : (r1(X0, sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (~ p101(X0) | p102(X0) | ((p102(sK26(X0)) & ~ p103(sK26(X0)) & p3(sK26(X0)) & r1(X0, sK26(X0))) & (p102(sK27(X0)) & ~ p103(sK27(X0)) & ~ p3(sK27(X0)) & r1(X0, sK27(X0)))) | ~ sP7(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27])], [f67, f69, f68])).
fof(f68, plain, ! [X0] : (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) => (p102(sK26(X0)) & ~ p103(sK26(X0)) & p3(sK26(X0)) & r1(X0, sK26(X0)))), introduced(choice_axiom, [])).
fof(f69, plain, ! [X0] : (? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2)) => (p102(sK27(X0)) & ~ p103(sK27(X0)) & ~ p3(sK27(X0)) & r1(X0, sK27(X0)))), introduced(choice_axiom, [])).
fof(f67, plain, ! [X0] : (~ p101(X0) | p102(X0) | (? [X1] : (p102(X1) & ~ p103(X1) & p3(X1) & r1(X0, X1)) & ? [X2] : (p102(X2) & ~ p103(X2) & ~ p3(X2) & r1(X0, X2))) | ~ sP7(X0)), inference(rectify, [], [f66])).
fof(f66, plain, ! [X10] : (~ p101(X10) | p102(X10) | (? [X35] : (p102(X35) & ~ p103(X35) & p3(X35) & r1(X10, X35)) & ? [X36] : (p102(X36) & ~ p103(X36) & ~ p3(X36) & r1(X10, X36))) | ~ sP7(X10)), inference(nnf_transformation, [], [f16])).
fof(f16, plain, ! [X10] : (~ p101(X10) | p102(X10) | (? [X35] : (p102(X35) & ~ p103(X35) & p3(X35) & r1(X10, X35)) & ? [X36] : (p102(X36) & ~ p103(X36) & ~ p3(X36) & r1(X10, X36))) | ~ sP7(X10)), inference(usedef, [], [e16])).
fof(e16, plain, ! [X10] : (sP7(X10) <=> (~ p101(X10) | p102(X10) | (? [X35] : (p102(X35) & ~ p103(X35) & p3(X35) & r1(X10, X35)) & ? [X36] : (p102(X36) & ~ p103(X36) & ~ p3(X36) & r1(X10, X36))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f1653, plain, (! [X14, X12, X13] : (~ r1(sK25(sK42), X12) | p5(X14) | ~ r1(X13, X14) | ~ r1(X12, X13)) | ~ spl43_91), inference(avatar_component_clause, [], [f1652])).
fof(f1652, plain, (spl43_91 <=> ! [X13, X12, X14] : (~ r1(sK25(sK42), X12) | p5(X14) | ~ r1(X13, X14) | ~ r1(X12, X13))), introduced(avatar_definition, [new_symbols(naming, [spl43_91])])).
fof(f15942, plain, (~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2031), inference(avatar_contradiction_clause, [], [f15941])).
fof(f15941, plain, ($false | (~ spl43_806 | ~ spl43_809 | spl43_810 | ~ spl43_2031)), inference(subsumption_resolution, [], [f15940, f5319])).
fof(f15940, plain, (~ sP6(sK26(sK25(sK42))) | (~ spl43_806 | spl43_810 | ~ spl43_2031)), inference(subsumption_resolution, [], [f15939, f5308])).
fof(f15939, plain, (~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | (spl43_810 | ~ spl43_2031)), inference(subsumption_resolution, [], [f15938, f5323])).
fof(f15938, plain, (p103(sK26(sK25(sK42))) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | ~ spl43_2031), inference(resolution, [], [f12278, f191])).
fof(f191, plain, ! [X0] : (~ p104(sK28(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f75])).
fof(f12278, plain, (p104(sK28(sK26(sK25(sK42)))) | ~ spl43_2031), inference(avatar_component_clause, [], [f12276])).
fof(f15826, plain, (spl43_2030 | ~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810), inference(avatar_split_clause, [], [f15754, f5322, f5318, f5306, f2698, f2694, f2682, f1592, f12272])).
fof(f1592, plain, (spl43_81 <=> ! [X9, X8] : (~ r1(sK25(sK42), X8) | sP21(X9) | ~ r1(X8, X9))), introduced(avatar_definition, [new_symbols(naming, [spl43_81])])).
fof(f15754, plain, (sP5(sK28(sK26(sK25(sK42)))) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810)), inference(resolution, [], [f15717, f116])).
fof(f116, plain, ! [X0] : (~ sP21(X0) | sP5(X0)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & sP20(X0) & sP19(X0) & sP18(X0) & sP17(X0) & sP16(X0) & sP15(X0) & sP14(X0) & sP13(X0) & sP12(X0) & sP11(X0) & sP10(X0) & sP8(X0) & sP7(X0) & sP6(X0) & sP5(X0) & sP4(X0) & sP3(X0) & sP2(X0) & sP1(X0) & sP0(X0) & sP9(X0)) | ~ sP21(X0)), inference(rectify, [], [f32])).
fof(f32, plain, ! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & sP20(X10) & sP19(X10) & sP18(X10) & sP17(X10) & sP16(X10) & sP15(X10) & sP14(X10) & sP13(X10) & sP12(X10) & sP11(X10) & sP10(X10) & sP8(X10) & sP7(X10) & sP6(X10) & sP5(X10) & sP4(X10) & sP3(X10) & sP2(X10) & sP1(X10) & sP0(X10) & sP9(X10)) | ~ sP21(X10)), inference(nnf_transformation, [], [f30])).
fof(f30, plain, ! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & sP20(X10) & sP19(X10) & sP18(X10) & sP17(X10) & sP16(X10) & sP15(X10) & sP14(X10) & sP13(X10) & sP12(X10) & sP11(X10) & sP10(X10) & sP8(X10) & sP7(X10) & sP6(X10) & sP5(X10) & sP4(X10) & sP3(X10) & sP2(X10) & sP1(X10) & sP0(X10) & sP9(X10)) | ~ sP21(X10)), inference(usedef, [], [e30])).
fof(e30, plain, ! [X10] : (sP21(X10) <=> ((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & sP20(X10) & sP19(X10) & sP18(X10) & sP17(X10) & sP16(X10) & sP15(X10) & sP14(X10) & sP13(X10) & sP12(X10) & sP11(X10) & sP10(X10) & sP8(X10) & sP7(X10) & sP6(X10) & sP5(X10) & sP4(X10) & sP3(X10) & sP2(X10) & sP1(X10) & sP0(X10) & sP9(X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f15717, plain, (sP21(sK28(sK26(sK25(sK42)))) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | ~ spl43_809 | spl43_810)), inference(subsumption_resolution, [], [f15716, f5319])).
fof(f15716, plain, (sP21(sK28(sK26(sK25(sK42)))) | ~ sP6(sK26(sK25(sK42))) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_806 | spl43_810)), inference(subsumption_resolution, [], [f15715, f5308])).
fof(f15715, plain, (sP21(sK28(sK26(sK25(sK42)))) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317 | spl43_810)), inference(subsumption_resolution, [], [f15701, f5323])).
fof(f15701, plain, (sP21(sK28(sK26(sK25(sK42)))) | p103(sK26(sK25(sK42))) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(resolution, [], [f15663, f189])).
fof(f15663, plain, (! [X5] : (~ r1(sK26(sK25(sK42)), X5) | sP21(X5)) | (~ spl43_81 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(subsumption_resolution, [], [f15662, f2695])).
fof(f15662, plain, (! [X5] : (sP21(X5) | ~ r1(sK26(sK25(sK42)), X5) | ~ sP7(sK25(sK42))) | (~ spl43_81 | ~ spl43_313 | spl43_317)), inference(subsumption_resolution, [], [f15661, f2684])).
fof(f15661, plain, (! [X5] : (sP21(X5) | ~ r1(sK26(sK25(sK42)), X5) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42))) | (~ spl43_81 | spl43_317)), inference(subsumption_resolution, [], [f15645, f2699])).
fof(f15645, plain, (! [X5] : (sP21(X5) | ~ r1(sK26(sK25(sK42)), X5) | p102(sK25(sK42)) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42))) | ~ spl43_81), inference(resolution, [], [f1593, f181])).
fof(f1593, plain, (! [X8, X9] : (~ r1(sK25(sK42), X8) | sP21(X9) | ~ r1(X8, X9)) | ~ spl43_81), inference(avatar_component_clause, [], [f1592])).
fof(f15360, plain, (~ spl43_806 | ~ spl43_809 | spl43_810 | spl43_2027), inference(avatar_contradiction_clause, [], [f15359])).
fof(f15359, plain, ($false | (~ spl43_806 | ~ spl43_809 | spl43_810 | spl43_2027)), inference(subsumption_resolution, [], [f15358, f5319])).
fof(f15358, plain, (~ sP6(sK26(sK25(sK42))) | (~ spl43_806 | spl43_810 | spl43_2027)), inference(subsumption_resolution, [], [f15357, f5308])).
fof(f15357, plain, (~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | (spl43_810 | spl43_2027)), inference(subsumption_resolution, [], [f15356, f5323])).
fof(f15356, plain, (p103(sK26(sK25(sK42))) | ~ p102(sK26(sK25(sK42))) | ~ sP6(sK26(sK25(sK42))) | spl43_2027), inference(resolution, [], [f12261, f192])).
fof(f192, plain, ! [X0] : (p103(sK28(X0)) | p103(X0) | ~ p102(X0) | ~ sP6(X0)), inference(cnf_transformation, [], [f75])).
fof(f12261, plain, (~ p103(sK28(sK26(sK25(sK42)))) | spl43_2027), inference(avatar_component_clause, [], [f12260])).
fof(f7479, plain, (~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_810), inference(avatar_contradiction_clause, [], [f7478])).
fof(f7478, plain, ($false | (~ spl43_313 | ~ spl43_316 | spl43_317 | ~ spl43_810)), inference(subsumption_resolution, [], [f7477, f2695])).
fof(f7477, plain, (~ sP7(sK25(sK42)) | (~ spl43_313 | spl43_317 | ~ spl43_810)), inference(subsumption_resolution, [], [f7476, f2684])).
fof(f7476, plain, (~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | (spl43_317 | ~ spl43_810)), inference(subsumption_resolution, [], [f7475, f2699])).
fof(f7475, plain, (p102(sK25(sK42)) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | ~ spl43_810), inference(resolution, [], [f5324, f183])).
fof(f183, plain, ! [X0] : (~ p103(sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f70])).
fof(f5324, plain, (p103(sK26(sK25(sK42))) | ~ spl43_810), inference(avatar_component_clause, [], [f5322])).
fof(f7372, plain, (~ spl43_313 | ~ spl43_316 | spl43_317 | spl43_806), inference(avatar_contradiction_clause, [], [f7371])).
fof(f7371, plain, ($false | (~ spl43_313 | ~ spl43_316 | spl43_317 | spl43_806)), inference(subsumption_resolution, [], [f7370, f2695])).
fof(f7370, plain, (~ sP7(sK25(sK42)) | (~ spl43_313 | spl43_317 | spl43_806)), inference(subsumption_resolution, [], [f7369, f2684])).
fof(f7369, plain, (~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | (spl43_317 | spl43_806)), inference(subsumption_resolution, [], [f7368, f2699])).
fof(f7368, plain, (p102(sK25(sK42)) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | spl43_806), inference(resolution, [], [f5307, f184])).
fof(f184, plain, ! [X0] : (p102(sK26(X0)) | p102(X0) | ~ p101(X0) | ~ sP7(X0)), inference(cnf_transformation, [], [f70])).
fof(f5307, plain, (~ p102(sK26(sK25(sK42))) | spl43_806), inference(avatar_component_clause, [], [f5306])).
fof(f7330, plain, (spl43_809 | ~ spl43_77 | ~ spl43_313 | ~ spl43_316 | spl43_317), inference(avatar_split_clause, [], [f7274, f2698, f2694, f2682, f1566, f5318])).
fof(f1566, plain, (spl43_77 <=> ! [X4] : (~ r1(sK25(sK42), X4) | sP21(X4))), introduced(avatar_definition, [new_symbols(naming, [spl43_77])])).
fof(f7274, plain, (sP6(sK26(sK25(sK42))) | (~ spl43_77 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(resolution, [], [f7236, f117])).
fof(f117, plain, ! [X0] : (~ sP21(X0) | sP6(X0)), inference(cnf_transformation, [], [f33])).
fof(f7236, plain, (sP21(sK26(sK25(sK42))) | (~ spl43_77 | ~ spl43_313 | ~ spl43_316 | spl43_317)), inference(subsumption_resolution, [], [f7235, f2695])).
fof(f7235, plain, (sP21(sK26(sK25(sK42))) | ~ sP7(sK25(sK42)) | (~ spl43_77 | ~ spl43_313 | spl43_317)), inference(subsumption_resolution, [], [f7234, f2684])).
fof(f7234, plain, (sP21(sK26(sK25(sK42))) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | (~ spl43_77 | spl43_317)), inference(subsumption_resolution, [], [f7218, f2699])).
fof(f7218, plain, (sP21(sK26(sK25(sK42))) | p102(sK25(sK42)) | ~ p101(sK25(sK42)) | ~ sP7(sK25(sK42)) | ~ spl43_77), inference(resolution, [], [f1567, f181])).
fof(f1567, plain, (! [X4] : (~ r1(sK25(sK42), X4) | sP21(X4)) | ~ spl43_77), inference(avatar_component_clause, [], [f1566])).
fof(f3072, plain, (~ spl43_6 | ~ spl43_317), inference(avatar_contradiction_clause, [], [f3071])).
fof(f3071, plain, ($false | (~ spl43_6 | ~ spl43_317)), inference(subsumption_resolution, [], [f3070, f814])).
fof(f814, plain, (sP8(sK42) | ~ spl43_6), inference(avatar_component_clause, [], [f813])).
fof(f813, plain, (spl43_6 <=> sP8(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl43_6])])).
fof(f3070, plain, (~ sP8(sK42) | ~ spl43_317), inference(subsumption_resolution, [], [f3069, f244])).
fof(f244, plain, p100(sK42), inference(cnf_transformation, [], [f108])).
fof(f108, plain, (p100(sK42) & ~ p101(sK42) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (sP21(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK42, X1)) & ! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (p5(X20) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X15, X16)) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(sK42, X11))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f106, f107])).
fof(f107, plain, (? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (sP21(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (p5(X20) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X15, X16)) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X0, X11))) => (p100(sK42) & ~ p101(sK42) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (sP21(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(sK42, X1)) & ! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (p5(X20) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X15, X16)) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(sK42, X11)))), introduced(choice_axiom, [])).
fof(f106, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (sP21(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : (! [X13] : (! [X14] : (! [X15] : (! [X16] : (! [X17] : (! [X18] : (! [X19] : (! [X20] : (p5(X20) | ~ r1(X19, X20)) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | ~ r1(X16, X17)) | ~ r1(X15, X16)) | ~ r1(X14, X15)) | ~ r1(X13, X14)) | ~ r1(X12, X13)) | ~ r1(X11, X12)) | ~ r1(X0, X11))), inference(rectify, [], [f31])).
fof(f31, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (sP21(X10) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(definition_folding, [], [f8, e30, e29, e28, e27, e26, e25, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9])).
fof(f9, plain, ! [X10] : (~ p108(X10) | p109(X10) | (? [X49] : (p109(X49) & ~ p110(X49) & p10(X49) & r1(X10, X49)) & ? [X50] : (p109(X50) & ~ p110(X50) & ~ p10(X50) & r1(X10, X50))) | ~ sP0(X10)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X10] : (sP0(X10) <=> (~ p108(X10) | p109(X10) | (? [X49] : (p109(X49) & ~ p110(X49) & p10(X49) & r1(X10, X49)) & ? [X50] : (p109(X50) & ~ p110(X50) & ~ p10(X50) & r1(X10, X50))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f10, plain, ! [X10] : (~ p107(X10) | p108(X10) | (? [X47] : (p108(X47) & ~ p109(X47) & p9(X47) & r1(X10, X47)) & ? [X48] : (p108(X48) & ~ p109(X48) & ~ p9(X48) & r1(X10, X48))) | ~ sP1(X10)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X10] : (sP1(X10) <=> (~ p107(X10) | p108(X10) | (? [X47] : (p108(X47) & ~ p109(X47) & p9(X47) & r1(X10, X47)) & ? [X48] : (p108(X48) & ~ p109(X48) & ~ p9(X48) & r1(X10, X48))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f11, plain, ! [X10] : (~ p106(X10) | p107(X10) | (? [X45] : (p107(X45) & ~ p108(X45) & p8(X45) & r1(X10, X45)) & ? [X46] : (p107(X46) & ~ p108(X46) & ~ p8(X46) & r1(X10, X46))) | ~ sP2(X10)), inference(usedef, [], [e11])).
fof(e11, plain, ! [X10] : (sP2(X10) <=> (~ p106(X10) | p107(X10) | (? [X45] : (p107(X45) & ~ p108(X45) & p8(X45) & r1(X10, X45)) & ? [X46] : (p107(X46) & ~ p108(X46) & ~ p8(X46) & r1(X10, X46))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f12, plain, ! [X10] : (~ p105(X10) | p106(X10) | (? [X43] : (p106(X43) & ~ p107(X43) & p7(X43) & r1(X10, X43)) & ? [X44] : (p106(X44) & ~ p107(X44) & ~ p7(X44) & r1(X10, X44))) | ~ sP3(X10)), inference(usedef, [], [e12])).
fof(e12, plain, ! [X10] : (sP3(X10) <=> (~ p105(X10) | p106(X10) | (? [X43] : (p106(X43) & ~ p107(X43) & p7(X43) & r1(X10, X43)) & ? [X44] : (p106(X44) & ~ p107(X44) & ~ p7(X44) & r1(X10, X44))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f13, plain, ! [X10] : (~ p104(X10) | p105(X10) | (? [X41] : (p105(X41) & ~ p106(X41) & p6(X41) & r1(X10, X41)) & ? [X42] : (p105(X42) & ~ p106(X42) & ~ p6(X42) & r1(X10, X42))) | ~ sP4(X10)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X10] : (sP4(X10) <=> (~ p104(X10) | p105(X10) | (? [X41] : (p105(X41) & ~ p106(X41) & p6(X41) & r1(X10, X41)) & ? [X42] : (p105(X42) & ~ p106(X42) & ~ p6(X42) & r1(X10, X42))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f17, plain, ! [X10] : (~ p100(X10) | p101(X10) | (? [X33] : (p101(X33) & ~ p102(X33) & p2(X33) & r1(X10, X33)) & ? [X34] : (p101(X34) & ~ p102(X34) & ~ p2(X34) & r1(X10, X34))) | ~ sP8(X10)), inference(usedef, [], [e17])).
fof(e17, plain, ! [X10] : (sP8(X10) <=> (~ p100(X10) | p101(X10) | (? [X33] : (p101(X33) & ~ p102(X33) & p2(X33) & r1(X10, X33)) & ? [X34] : (p101(X34) & ~ p102(X34) & ~ p2(X34) & r1(X10, X34))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f18, plain, ! [X10] : (~ p109(X10) | p110(X10) | (? [X51] : (p110(X51) & p11(X51) & r1(X10, X51)) & ? [X52] : (p110(X52) & ~ p11(X52) & r1(X10, X52))) | ~ sP9(X10)), inference(usedef, [], [e18])).
fof(e18, plain, ! [X10] : (sP9(X10) <=> (~ p109(X10) | p110(X10) | (? [X51] : (p110(X51) & p11(X51) & r1(X10, X51)) & ? [X52] : (p110(X52) & ~ p11(X52) & r1(X10, X52))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f19, plain, ! [X10] : (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32)))) | ~ sP10(X10)), inference(usedef, [], [e19])).
fof(e19, plain, ! [X10] : (sP10(X10) <=> (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f20, plain, ! [X10] : (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30)))) | ~ sP11(X10)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X10] : (sP11(X10) <=> (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f21, plain, ! [X10] : (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28)))) | ~ sP12(X10)), inference(usedef, [], [e21])).
fof(e21, plain, ! [X10] : (sP12(X10) <=> (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f22, plain, ! [X10] : (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26)))) | ~ sP13(X10)), inference(usedef, [], [e22])).
fof(e22, plain, ! [X10] : (sP13(X10) <=> (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f23, plain, ! [X10] : (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24)))) | ~ sP14(X10)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X10] : (sP14(X10) <=> (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f24, plain, ! [X10] : (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22)))) | ~ sP15(X10)), inference(usedef, [], [e24])).
fof(e24, plain, ! [X10] : (sP15(X10) <=> (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f25, plain, ! [X10] : (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20)))) | ~ sP16(X10)), inference(usedef, [], [e25])).
fof(e25, plain, ! [X10] : (sP16(X10) <=> (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f26, plain, ! [X10] : (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18)))) | ~ sP17(X10)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X10] : (sP17(X10) <=> (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f27, plain, ! [X10] : (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16)))) | ~ sP18(X10)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X10] : (sP18(X10) <=> (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f28, plain, ! [X10] : (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14)))) | ~ sP19(X10)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X10] : (sP19(X10) <=> (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f29, plain, ! [X10] : (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12)))) | ~ sP20(X10)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X10] : (sP20(X10) <=> (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f8, plain, ? [X0] : (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12))))) & (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14))))) & (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16))))) & (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18))))) & (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20))))) & (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22))))) & (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24))))) & (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26))))) & (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28))))) & (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30))))) & (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32))))) & (~ p100(X10) | p101(X10) | (? [X33] : (p101(X33) & ~ p102(X33) & p2(X33) & r1(X10, X33)) & ? [X34] : (p101(X34) & ~ p102(X34) & ~ p2(X34) & r1(X10, X34)))) & (~ p101(X10) | p102(X10) | (? [X35] : (p102(X35) & ~ p103(X35) & p3(X35) & r1(X10, X35)) & ? [X36] : (p102(X36) & ~ p103(X36) & ~ p3(X36) & r1(X10, X36)))) & (~ p102(X10) | p103(X10) | (? [X37] : (p103(X37) & ~ p104(X37) & p4(X37) & r1(X10, X37)) & ? [X38] : (p103(X38) & ~ p104(X38) & ~ p4(X38) & r1(X10, X38)))) & (~ p103(X10) | p104(X10) | (? [X39] : (p104(X39) & ~ p105(X39) & p5(X39) & r1(X10, X39)) & ? [X40] : (p104(X40) & ~ p105(X40) & ~ p5(X40) & r1(X10, X40)))) & (~ p104(X10) | p105(X10) | (? [X41] : (p105(X41) & ~ p106(X41) & p6(X41) & r1(X10, X41)) & ? [X42] : (p105(X42) & ~ p106(X42) & ~ p6(X42) & r1(X10, X42)))) & (~ p105(X10) | p106(X10) | (? [X43] : (p106(X43) & ~ p107(X43) & p7(X43) & r1(X10, X43)) & ? [X44] : (p106(X44) & ~ p107(X44) & ~ p7(X44) & r1(X10, X44)))) & (~ p106(X10) | p107(X10) | (? [X45] : (p107(X45) & ~ p108(X45) & p8(X45) & r1(X10, X45)) & ? [X46] : (p107(X46) & ~ p108(X46) & ~ p8(X46) & r1(X10, X46)))) & (~ p107(X10) | p108(X10) | (? [X47] : (p108(X47) & ~ p109(X47) & p9(X47) & r1(X10, X47)) & ? [X48] : (p108(X48) & ~ p109(X48) & ~ p9(X48) & r1(X10, X48)))) & (~ p108(X10) | p109(X10) | (? [X49] : (p109(X49) & ~ p110(X49) & p10(X49) & r1(X10, X49)) & ? [X50] : (p109(X50) & ~ p110(X50) & ~ p10(X50) & r1(X10, X50)))) & (~ p109(X10) | p110(X10) | (? [X51] : (p110(X51) & p11(X51) & r1(X10, X51)) & ? [X52] : (p110(X52) & ~ p11(X52) & r1(X10, X52))))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1)) & ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(flattening, [], [f7])).
fof(f7, plain, ? [X0] : ((p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12))))) & (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14))))) & (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16))))) & (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18))))) & (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20))))) & (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22))))) & (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24))))) & (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26))))) & (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28))))) & (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30))))) & (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32))))) & ((~ p100(X10) | p101(X10)) | (? [X33] : ((p101(X33) & ~ p102(X33) & p2(X33)) & r1(X10, X33)) & ? [X34] : ((p101(X34) & ~ p102(X34) & ~ p2(X34)) & r1(X10, X34)))) & ((~ p101(X10) | p102(X10)) | (? [X35] : ((p102(X35) & ~ p103(X35) & p3(X35)) & r1(X10, X35)) & ? [X36] : ((p102(X36) & ~ p103(X36) & ~ p3(X36)) & r1(X10, X36)))) & ((~ p102(X10) | p103(X10)) | (? [X37] : ((p103(X37) & ~ p104(X37) & p4(X37)) & r1(X10, X37)) & ? [X38] : ((p103(X38) & ~ p104(X38) & ~ p4(X38)) & r1(X10, X38)))) & ((~ p103(X10) | p104(X10)) | (? [X39] : ((p104(X39) & ~ p105(X39) & p5(X39)) & r1(X10, X39)) & ? [X40] : ((p104(X40) & ~ p105(X40) & ~ p5(X40)) & r1(X10, X40)))) & ((~ p104(X10) | p105(X10)) | (? [X41] : ((p105(X41) & ~ p106(X41) & p6(X41)) & r1(X10, X41)) & ? [X42] : ((p105(X42) & ~ p106(X42) & ~ p6(X42)) & r1(X10, X42)))) & ((~ p105(X10) | p106(X10)) | (? [X43] : ((p106(X43) & ~ p107(X43) & p7(X43)) & r1(X10, X43)) & ? [X44] : ((p106(X44) & ~ p107(X44) & ~ p7(X44)) & r1(X10, X44)))) & ((~ p106(X10) | p107(X10)) | (? [X45] : ((p107(X45) & ~ p108(X45) & p8(X45)) & r1(X10, X45)) & ? [X46] : ((p107(X46) & ~ p108(X46) & ~ p8(X46)) & r1(X10, X46)))) & ((~ p107(X10) | p108(X10)) | (? [X47] : ((p108(X47) & ~ p109(X47) & p9(X47)) & r1(X10, X47)) & ? [X48] : ((p108(X48) & ~ p109(X48) & ~ p9(X48)) & r1(X10, X48)))) & ((~ p108(X10) | p109(X10)) | (? [X49] : ((p109(X49) & ~ p110(X49) & p10(X49)) & r1(X10, X49)) & ? [X50] : ((p109(X50) & ~ p110(X50) & ~ p10(X50)) & r1(X10, X50)))) & ((~ p109(X10) | p110(X10)) | (? [X51] : ((p110(X51) & p11(X51)) & r1(X10, X51)) & ? [X52] : ((p110(X52) & ~ p11(X52)) & r1(X10, X52))))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) & ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12))))) & (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14))))) & (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16))))) & (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18))))) & (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20))))) & (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22))))) & (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24))))) & (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26))))) & (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28))))) & (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30))))) & (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32))))) & (~ (p100(X10) & ~ p101(X10)) | (~ ! [X33] : (~ (p101(X33) & ~ p102(X33) & p2(X33)) | ~ r1(X10, X33)) & ~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & ~ p2(X34)) | ~ r1(X10, X34)))) & (~ (p101(X10) & ~ p102(X10)) | (~ ! [X35] : (~ (p102(X35) & ~ p103(X35) & p3(X35)) | ~ r1(X10, X35)) & ~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & ~ p3(X36)) | ~ r1(X10, X36)))) & (~ (p102(X10) & ~ p103(X10)) | (~ ! [X37] : (~ (p103(X37) & ~ p104(X37) & p4(X37)) | ~ r1(X10, X37)) & ~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & ~ p4(X38)) | ~ r1(X10, X38)))) & (~ (p103(X10) & ~ p104(X10)) | (~ ! [X39] : (~ (p104(X39) & ~ p105(X39) & p5(X39)) | ~ r1(X10, X39)) & ~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & ~ p5(X40)) | ~ r1(X10, X40)))) & (~ (p104(X10) & ~ p105(X10)) | (~ ! [X41] : (~ (p105(X41) & ~ p106(X41) & p6(X41)) | ~ r1(X10, X41)) & ~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & ~ p6(X42)) | ~ r1(X10, X42)))) & (~ (p105(X10) & ~ p106(X10)) | (~ ! [X43] : (~ (p106(X43) & ~ p107(X43) & p7(X43)) | ~ r1(X10, X43)) & ~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & ~ p7(X44)) | ~ r1(X10, X44)))) & (~ (p106(X10) & ~ p107(X10)) | (~ ! [X45] : (~ (p107(X45) & ~ p108(X45) & p8(X45)) | ~ r1(X10, X45)) & ~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & ~ p8(X46)) | ~ r1(X10, X46)))) & (~ (p107(X10) & ~ p108(X10)) | (~ ! [X47] : (~ (p108(X47) & ~ p109(X47) & p9(X47)) | ~ r1(X10, X47)) & ~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & ~ p9(X48)) | ~ r1(X10, X48)))) & (~ (p108(X10) & ~ p109(X10)) | (~ ! [X49] : (~ (p109(X49) & ~ p110(X49) & p10(X49)) | ~ r1(X10, X49)) & ~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & ~ p10(X50)) | ~ r1(X10, X50)))) & (~ (p109(X10) & ~ p110(X10)) | (~ ! [X51] : (~ (p110(X51) & p11(X51)) | ~ r1(X10, X51)) & ~ ! [X52] : (~ (p110(X52) & ~ p11(X52)) | ~ r1(X10, X52))))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(pure_predicate_removal, [], [f5])).
fof(f5, plain, ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & (~ p111(X10) | p110(X10)) & (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12))))) & (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14))))) & (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16))))) & (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18))))) & (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20))))) & (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22))))) & (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24))))) & (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26))))) & (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28))))) & (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30))))) & (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32))))) & (~ (p100(X10) & ~ p101(X10)) | (~ ! [X33] : (~ (p101(X33) & ~ p102(X33) & p2(X33)) | ~ r1(X10, X33)) & ~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & ~ p2(X34)) | ~ r1(X10, X34)))) & (~ (p101(X10) & ~ p102(X10)) | (~ ! [X35] : (~ (p102(X35) & ~ p103(X35) & p3(X35)) | ~ r1(X10, X35)) & ~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & ~ p3(X36)) | ~ r1(X10, X36)))) & (~ (p102(X10) & ~ p103(X10)) | (~ ! [X37] : (~ (p103(X37) & ~ p104(X37) & p4(X37)) | ~ r1(X10, X37)) & ~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & ~ p4(X38)) | ~ r1(X10, X38)))) & (~ (p103(X10) & ~ p104(X10)) | (~ ! [X39] : (~ (p104(X39) & ~ p105(X39) & p5(X39)) | ~ r1(X10, X39)) & ~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & ~ p5(X40)) | ~ r1(X10, X40)))) & (~ (p104(X10) & ~ p105(X10)) | (~ ! [X41] : (~ (p105(X41) & ~ p106(X41) & p6(X41)) | ~ r1(X10, X41)) & ~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & ~ p6(X42)) | ~ r1(X10, X42)))) & (~ (p105(X10) & ~ p106(X10)) | (~ ! [X43] : (~ (p106(X43) & ~ p107(X43) & p7(X43)) | ~ r1(X10, X43)) & ~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & ~ p7(X44)) | ~ r1(X10, X44)))) & (~ (p106(X10) & ~ p107(X10)) | (~ ! [X45] : (~ (p107(X45) & ~ p108(X45) & p8(X45)) | ~ r1(X10, X45)) & ~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & ~ p8(X46)) | ~ r1(X10, X46)))) & (~ (p107(X10) & ~ p108(X10)) | (~ ! [X47] : (~ (p108(X47) & ~ p109(X47) & p9(X47)) | ~ r1(X10, X47)) & ~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & ~ p9(X48)) | ~ r1(X10, X48)))) & (~ (p108(X10) & ~ p109(X10)) | (~ ! [X49] : (~ (p109(X49) & ~ p110(X49) & p10(X49)) | ~ r1(X10, X49)) & ~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & ~ p10(X50)) | ~ r1(X10, X50)))) & (~ (p109(X10) & ~ p110(X10)) | (~ ! [X51] : (~ (p110(X51) & ~ p111(X51) & p11(X51)) | ~ r1(X10, X51)) & ~ ! [X52] : (~ (p110(X52) & ~ p111(X52) & ~ p11(X52)) | ~ r1(X10, X52))))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(flattening, [], [f4])).
fof(f4, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X2] : (! [X3] : (! [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (! [X9] : (! [X10] : (((~ p101(X10) | p100(X10)) & (~ p102(X10) | p101(X10)) & (~ p103(X10) | p102(X10)) & (~ p104(X10) | p103(X10)) & (~ p105(X10) | p104(X10)) & (~ p106(X10) | p105(X10)) & (~ p107(X10) | p106(X10)) & (~ p108(X10) | p107(X10)) & (~ p109(X10) | p108(X10)) & (~ p110(X10) | p109(X10)) & (~ p111(X10) | p110(X10)) & (~ p100(X10) | ((~ p1(X10) | ! [X11] : (~ p100(X11) | p1(X11) | ~ r1(X10, X11))) & (p1(X10) | ! [X12] : (~ p100(X12) | ~ p1(X12) | ~ r1(X10, X12))))) & (~ p101(X10) | ((~ p2(X10) | ! [X13] : (~ p101(X13) | p2(X13) | ~ r1(X10, X13))) & (p2(X10) | ! [X14] : (~ p101(X14) | ~ p2(X14) | ~ r1(X10, X14))))) & (~ p102(X10) | ((~ p3(X10) | ! [X15] : (~ p102(X15) | p3(X15) | ~ r1(X10, X15))) & (p3(X10) | ! [X16] : (~ p102(X16) | ~ p3(X16) | ~ r1(X10, X16))))) & (~ p103(X10) | ((~ p4(X10) | ! [X17] : (~ p103(X17) | p4(X17) | ~ r1(X10, X17))) & (p4(X10) | ! [X18] : (~ p103(X18) | ~ p4(X18) | ~ r1(X10, X18))))) & (~ p104(X10) | ((~ p5(X10) | ! [X19] : (~ p104(X19) | p5(X19) | ~ r1(X10, X19))) & (p5(X10) | ! [X20] : (~ p104(X20) | ~ p5(X20) | ~ r1(X10, X20))))) & (~ p105(X10) | ((~ p6(X10) | ! [X21] : (~ p105(X21) | p6(X21) | ~ r1(X10, X21))) & (p6(X10) | ! [X22] : (~ p105(X22) | ~ p6(X22) | ~ r1(X10, X22))))) & (~ p106(X10) | ((~ p7(X10) | ! [X23] : (~ p106(X23) | p7(X23) | ~ r1(X10, X23))) & (p7(X10) | ! [X24] : (~ p106(X24) | ~ p7(X24) | ~ r1(X10, X24))))) & (~ p107(X10) | ((~ p8(X10) | ! [X25] : (~ p107(X25) | p8(X25) | ~ r1(X10, X25))) & (p8(X10) | ! [X26] : (~ p107(X26) | ~ p8(X26) | ~ r1(X10, X26))))) & (~ p108(X10) | ((~ p9(X10) | ! [X27] : (~ p108(X27) | p9(X27) | ~ r1(X10, X27))) & (p9(X10) | ! [X28] : (~ p108(X28) | ~ p9(X28) | ~ r1(X10, X28))))) & (~ p109(X10) | ((~ p10(X10) | ! [X29] : (~ p109(X29) | p10(X29) | ~ r1(X10, X29))) & (p10(X10) | ! [X30] : (~ p109(X30) | ~ p10(X30) | ~ r1(X10, X30))))) & (~ p110(X10) | ((~ p11(X10) | ! [X31] : (~ p110(X31) | p11(X31) | ~ r1(X10, X31))) & (p11(X10) | ! [X32] : (~ p110(X32) | ~ p11(X32) | ~ r1(X10, X32))))) & (~ (p100(X10) & ~ p101(X10)) | (~ ! [X33] : (~ (p101(X33) & ~ p102(X33) & p2(X33)) | ~ r1(X10, X33)) & ~ ! [X34] : (~ (p101(X34) & ~ p102(X34) & ~ p2(X34)) | ~ r1(X10, X34)))) & (~ (p101(X10) & ~ p102(X10)) | (~ ! [X35] : (~ (p102(X35) & ~ p103(X35) & p3(X35)) | ~ r1(X10, X35)) & ~ ! [X36] : (~ (p102(X36) & ~ p103(X36) & ~ p3(X36)) | ~ r1(X10, X36)))) & (~ (p102(X10) & ~ p103(X10)) | (~ ! [X37] : (~ (p103(X37) & ~ p104(X37) & p4(X37)) | ~ r1(X10, X37)) & ~ ! [X38] : (~ (p103(X38) & ~ p104(X38) & ~ p4(X38)) | ~ r1(X10, X38)))) & (~ (p103(X10) & ~ p104(X10)) | (~ ! [X39] : (~ (p104(X39) & ~ p105(X39) & p5(X39)) | ~ r1(X10, X39)) & ~ ! [X40] : (~ (p104(X40) & ~ p105(X40) & ~ p5(X40)) | ~ r1(X10, X40)))) & (~ (p104(X10) & ~ p105(X10)) | (~ ! [X41] : (~ (p105(X41) & ~ p106(X41) & p6(X41)) | ~ r1(X10, X41)) & ~ ! [X42] : (~ (p105(X42) & ~ p106(X42) & ~ p6(X42)) | ~ r1(X10, X42)))) & (~ (p105(X10) & ~ p106(X10)) | (~ ! [X43] : (~ (p106(X43) & ~ p107(X43) & p7(X43)) | ~ r1(X10, X43)) & ~ ! [X44] : (~ (p106(X44) & ~ p107(X44) & ~ p7(X44)) | ~ r1(X10, X44)))) & (~ (p106(X10) & ~ p107(X10)) | (~ ! [X45] : (~ (p107(X45) & ~ p108(X45) & p8(X45)) | ~ r1(X10, X45)) & ~ ! [X46] : (~ (p107(X46) & ~ p108(X46) & ~ p8(X46)) | ~ r1(X10, X46)))) & (~ (p107(X10) & ~ p108(X10)) | (~ ! [X47] : (~ (p108(X47) & ~ p109(X47) & p9(X47)) | ~ r1(X10, X47)) & ~ ! [X48] : (~ (p108(X48) & ~ p109(X48) & ~ p9(X48)) | ~ r1(X10, X48)))) & (~ (p108(X10) & ~ p109(X10)) | (~ ! [X49] : (~ (p109(X49) & ~ p110(X49) & p10(X49)) | ~ r1(X10, X49)) & ~ ! [X50] : (~ (p109(X50) & ~ p110(X50) & ~ p10(X50)) | ~ r1(X10, X50)))) & (~ (p109(X10) & ~ p110(X10)) | (~ ! [X51] : (~ (p110(X51) & ~ p111(X51) & p11(X51)) | ~ r1(X10, X51)) & ~ ! [X52] : (~ (p110(X52) & ~ p111(X52) & ~ p11(X52)) | ~ r1(X10, X52))))) | ~ r1(X9, X10)) | ~ r1(X8, X9)) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | ~ r1(X5, X6)) | ~ r1(X4, X5)) | ~ r1(X3, X4)) | ~ r1(X2, X3)) | ~ r1(X1, X2)) | ~ r1(X0, X1))) | ~ ! [X53] : (! [X54] : (! [X55] : (! [X56] : (! [X57] : (! [X58] : (! [X59] : (! [X60] : (! [X61] : (! [X62] : (p5(X62) | ~ r1(X61, X62)) | ~ r1(X60, X61)) | ~ r1(X59, X60)) | ~ r1(X58, X59)) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | ~ r1(X55, X56)) | ~ r1(X54, X55)) | ~ r1(X53, X54)) | ~ r1(X0, X53))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & (~ p111(X0) | p110(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ p106(X0) | ((~ p7(X0) | ! [X1] : (~ p106(X1) | p7(X1) | ~ r1(X0, X1))) & (p7(X0) | ! [X1] : (~ p106(X1) | ~ p7(X1) | ~ r1(X0, X1))))) & (~ p107(X0) | ((~ p8(X0) | ! [X1] : (~ p107(X1) | p8(X1) | ~ r1(X0, X1))) & (p8(X0) | ! [X1] : (~ p107(X1) | ~ p8(X1) | ~ r1(X0, X1))))) & (~ p108(X0) | ((~ p9(X0) | ! [X1] : (~ p108(X1) | p9(X1) | ~ r1(X0, X1))) & (p9(X0) | ! [X1] : (~ p108(X1) | ~ p9(X1) | ~ r1(X0, X1))))) & (~ p109(X0) | ((~ p10(X0) | ! [X1] : (~ p109(X1) | p10(X1) | ~ r1(X0, X1))) & (p10(X0) | ! [X1] : (~ p109(X1) | ~ p10(X1) | ~ r1(X0, X1))))) & (~ p110(X0) | ((~ p11(X0) | ! [X1] : (~ p110(X1) | p11(X1) | ~ r1(X0, X1))) & (p11(X0) | ! [X1] : (~ p110(X1) | ~ p11(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1)))) & (~ (p105(X0) & ~ p106(X0)) | (~ ! [X1] : (~ (p106(X1) & ~ p107(X1) & p7(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p106(X1) & ~ p107(X1) & ~ p7(X1)) | ~ r1(X0, X1)))) & (~ (p106(X0) & ~ p107(X0)) | (~ ! [X1] : (~ (p107(X1) & ~ p108(X1) & p8(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p107(X1) & ~ p108(X1) & ~ p8(X1)) | ~ r1(X0, X1)))) & (~ (p107(X0) & ~ p108(X0)) | (~ ! [X1] : (~ (p108(X1) & ~ p109(X1) & p9(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p108(X1) & ~ p109(X1) & ~ p9(X1)) | ~ r1(X0, X1)))) & (~ (p108(X0) & ~ p109(X0)) | (~ ! [X1] : (~ (p109(X1) & ~ p110(X1) & p10(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p109(X1) & ~ p110(X1) & ~ p10(X1)) | ~ r1(X0, X1)))) & (~ (p109(X0) & ~ p110(X0)) | (~ ! [X1] : (~ (p110(X1) & ~ p111(X1) & p11(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p110(X1) & ~ p111(X1) & ~ p11(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (p100(X0) & ~ p101(X0) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ p101(X0) | p100(X0)) & (~ p102(X0) | p101(X0)) & (~ p103(X0) | p102(X0)) & (~ p104(X0) | p103(X0)) & (~ p105(X0) | p104(X0)) & (~ p106(X0) | p105(X0)) & (~ p107(X0) | p106(X0)) & (~ p108(X0) | p107(X0)) & (~ p109(X0) | p108(X0)) & (~ p110(X0) | p109(X0)) & (~ p111(X0) | p110(X0)) & (~ p100(X0) | ((~ p1(X0) | ! [X1] : (~ p100(X1) | p1(X1) | ~ r1(X0, X1))) & (p1(X0) | ! [X1] : (~ p100(X1) | ~ p1(X1) | ~ r1(X0, X1))))) & (~ p101(X0) | ((~ p2(X0) | ! [X1] : (~ p101(X1) | p2(X1) | ~ r1(X0, X1))) & (p2(X0) | ! [X1] : (~ p101(X1) | ~ p2(X1) | ~ r1(X0, X1))))) & (~ p102(X0) | ((~ p3(X0) | ! [X1] : (~ p102(X1) | p3(X1) | ~ r1(X0, X1))) & (p3(X0) | ! [X1] : (~ p102(X1) | ~ p3(X1) | ~ r1(X0, X1))))) & (~ p103(X0) | ((~ p4(X0) | ! [X1] : (~ p103(X1) | p4(X1) | ~ r1(X0, X1))) & (p4(X0) | ! [X1] : (~ p103(X1) | ~ p4(X1) | ~ r1(X0, X1))))) & (~ p104(X0) | ((~ p5(X0) | ! [X1] : (~ p104(X1) | p5(X1) | ~ r1(X0, X1))) & (p5(X0) | ! [X1] : (~ p104(X1) | ~ p5(X1) | ~ r1(X0, X1))))) & (~ p105(X0) | ((~ p6(X0) | ! [X1] : (~ p105(X1) | p6(X1) | ~ r1(X0, X1))) & (p6(X0) | ! [X1] : (~ p105(X1) | ~ p6(X1) | ~ r1(X0, X1))))) & (~ p106(X0) | ((~ p7(X0) | ! [X1] : (~ p106(X1) | p7(X1) | ~ r1(X0, X1))) & (p7(X0) | ! [X1] : (~ p106(X1) | ~ p7(X1) | ~ r1(X0, X1))))) & (~ p107(X0) | ((~ p8(X0) | ! [X1] : (~ p107(X1) | p8(X1) | ~ r1(X0, X1))) & (p8(X0) | ! [X1] : (~ p107(X1) | ~ p8(X1) | ~ r1(X0, X1))))) & (~ p108(X0) | ((~ p9(X0) | ! [X1] : (~ p108(X1) | p9(X1) | ~ r1(X0, X1))) & (p9(X0) | ! [X1] : (~ p108(X1) | ~ p9(X1) | ~ r1(X0, X1))))) & (~ p109(X0) | ((~ p10(X0) | ! [X1] : (~ p109(X1) | p10(X1) | ~ r1(X0, X1))) & (p10(X0) | ! [X1] : (~ p109(X1) | ~ p10(X1) | ~ r1(X0, X1))))) & (~ p110(X0) | ((~ p11(X0) | ! [X1] : (~ p110(X1) | p11(X1) | ~ r1(X0, X1))) & (p11(X0) | ! [X1] : (~ p110(X1) | ~ p11(X1) | ~ r1(X0, X1))))) & (~ (p100(X0) & ~ p101(X0)) | (~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & p2(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p101(X1) & ~ p102(X1) & ~ p2(X1)) | ~ r1(X0, X1)))) & (~ (p101(X0) & ~ p102(X0)) | (~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & p3(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p102(X1) & ~ p103(X1) & ~ p3(X1)) | ~ r1(X0, X1)))) & (~ (p102(X0) & ~ p103(X0)) | (~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & p4(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p103(X1) & ~ p104(X1) & ~ p4(X1)) | ~ r1(X0, X1)))) & (~ (p103(X0) & ~ p104(X0)) | (~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & p5(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p104(X1) & ~ p105(X1) & ~ p5(X1)) | ~ r1(X0, X1)))) & (~ (p104(X0) & ~ p105(X0)) | (~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & p6(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p105(X1) & ~ p106(X1) & ~ p6(X1)) | ~ r1(X0, X1)))) & (~ (p105(X0) & ~ p106(X0)) | (~ ! [X1] : (~ (p106(X1) & ~ p107(X1) & p7(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p106(X1) & ~ p107(X1) & ~ p7(X1)) | ~ r1(X0, X1)))) & (~ (p106(X0) & ~ p107(X0)) | (~ ! [X1] : (~ (p107(X1) & ~ p108(X1) & p8(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p107(X1) & ~ p108(X1) & ~ p8(X1)) | ~ r1(X0, X1)))) & (~ (p107(X0) & ~ p108(X0)) | (~ ! [X1] : (~ (p108(X1) & ~ p109(X1) & p9(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p108(X1) & ~ p109(X1) & ~ p9(X1)) | ~ r1(X0, X1)))) & (~ (p108(X0) & ~ p109(X0)) | (~ ! [X1] : (~ (p109(X1) & ~ p110(X1) & p10(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p109(X1) & ~ p110(X1) & ~ p10(X1)) | ~ r1(X0, X1)))) & (~ (p109(X0) & ~ p110(X0)) | (~ ! [X1] : (~ (p110(X1) & ~ p111(X1) & p11(X1)) | ~ r1(X0, X1)) & ~ ! [X1] : (~ (p110(X1) & ~ p111(X1) & ~ p11(X1)) | ~ r1(X0, X1))))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (p5(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.010.p', main)).
fof(f3069, plain, (~ p100(sK42) | ~ sP8(sK42) | ~ spl43_317), inference(subsumption_resolution, [], [f3068, f243])).
fof(f243, plain, ~ p101(sK42), inference(cnf_transformation, [], [f108])).
fof(f3068, plain, (p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | ~ spl43_317), inference(resolution, [], [f2700, f171])).
fof(f171, plain, ! [X0] : (~ p102(sK25(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f65])).
fof(f65, plain, ! [X0] : (~ p100(X0) | p101(X0) | ((p101(sK24(X0)) & ~ p102(sK24(X0)) & p2(sK24(X0)) & r1(X0, sK24(X0))) & (p101(sK25(X0)) & ~ p102(sK25(X0)) & ~ p2(sK25(X0)) & r1(X0, sK25(X0)))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25])], [f62, f64, f63])).
fof(f63, plain, ! [X0] : (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) => (p101(sK24(X0)) & ~ p102(sK24(X0)) & p2(sK24(X0)) & r1(X0, sK24(X0)))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X0] : (? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2)) => (p101(sK25(X0)) & ~ p102(sK25(X0)) & ~ p2(sK25(X0)) & r1(X0, sK25(X0)))), introduced(choice_axiom, [])).
fof(f62, plain, ! [X0] : (~ p100(X0) | p101(X0) | (? [X1] : (p101(X1) & ~ p102(X1) & p2(X1) & r1(X0, X1)) & ? [X2] : (p101(X2) & ~ p102(X2) & ~ p2(X2) & r1(X0, X2))) | ~ sP8(X0)), inference(rectify, [], [f61])).
fof(f61, plain, ! [X10] : (~ p100(X10) | p101(X10) | (? [X33] : (p101(X33) & ~ p102(X33) & p2(X33) & r1(X10, X33)) & ? [X34] : (p101(X34) & ~ p102(X34) & ~ p2(X34) & r1(X10, X34))) | ~ sP8(X10)), inference(nnf_transformation, [], [f17])).
fof(f2700, plain, (p102(sK25(sK42)) | ~ spl43_317), inference(avatar_component_clause, [], [f2698])).
fof(f2952, plain, (spl43_316 | ~ spl43_75), inference(avatar_split_clause, [], [f2893, f1552, f2694])).
fof(f1552, plain, (spl43_75 <=> sP21(sK25(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl43_75])])).
fof(f2893, plain, (sP7(sK25(sK42)) | ~ spl43_75), inference(resolution, [], [f1554, f118])).
fof(f118, plain, ! [X0] : (~ sP21(X0) | sP7(X0)), inference(cnf_transformation, [], [f33])).
fof(f1554, plain, (sP21(sK25(sK42)) | ~ spl43_75), inference(avatar_component_clause, [], [f1552])).
fof(f2876, plain, (~ spl43_6 | spl43_313), inference(avatar_contradiction_clause, [], [f2875])).
fof(f2875, plain, ($false | (~ spl43_6 | spl43_313)), inference(subsumption_resolution, [], [f2874, f814])).
fof(f2874, plain, (~ sP8(sK42) | spl43_313), inference(subsumption_resolution, [], [f2873, f244])).
fof(f2873, plain, (~ p100(sK42) | ~ sP8(sK42) | spl43_313), inference(subsumption_resolution, [], [f2872, f243])).
fof(f2872, plain, (p101(sK42) | ~ p100(sK42) | ~ sP8(sK42) | spl43_313), inference(resolution, [], [f2683, f172])).
fof(f172, plain, ! [X0] : (p101(sK25(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f65])).
fof(f2683, plain, (~ p101(sK25(sK42)) | spl43_313), inference(avatar_component_clause, [], [f2682])).
fof(f1654, plain, (~ spl43_6 | spl43_91), inference(avatar_split_clause, [], [f1650, f1652, f813])).
fof(f1650, plain, ! [X14, X12, X13] : (~ r1(sK25(sK42), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | p5(X14) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1649, f244])).
fof(f1649, plain, ! [X14, X12, X13] : (~ r1(sK25(sK42), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | p5(X14) | ~ p100(sK42) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1152, f243])).
fof(f1152, plain, ! [X14, X12, X13] : (~ r1(sK25(sK42), X12) | ~ r1(X12, X13) | ~ r1(X13, X14) | p5(X14) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(resolution, [], [f1106, f169])).
fof(f169, plain, ! [X0] : (r1(X0, sK25(X0)) | p101(X0) | ~ p100(X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f65])).
fof(f1106, plain, ! [X2, X0, X3, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | p5(X3)), inference(resolution, [], [f1064, f109])).
fof(f109, plain, ! [X0] : r1(X0, X0), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : r1(X0, X0), file('/home/ubuntu/library/tptp/Problems/LCL/LCL656+1.010.p', reflexivity)).
fof(f1064, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | p5(X4)), inference(resolution, [], [f1022, f109])).
fof(f1022, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK42, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | p5(X5)), inference(resolution, [], [f980, f109])).
fof(f980, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK42, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | ~ r1(X5, X6) | p5(X6)), inference(resolution, [], [f959, f109])).
fof(f959, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK42, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | ~ r1(X1, X7) | p5(X7)), inference(resolution, [], [f769, f109])).
fof(f769, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK42, X8) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X0, X1) | p5(X1)), inference(resolution, [], [f241, f109])).
fof(f241, plain, ! [X14, X12, X19, X17, X15, X13, X11, X20, X18, X16] : (~ r1(sK42, X11) | ~ r1(X19, X20) | ~ r1(X18, X19) | ~ r1(X17, X18) | ~ r1(X16, X17) | ~ r1(X15, X16) | ~ r1(X14, X15) | ~ r1(X13, X14) | ~ r1(X12, X13) | ~ r1(X11, X12) | p5(X20)), inference(cnf_transformation, [], [f108])).
fof(f1594, plain, (~ spl43_6 | spl43_81), inference(avatar_split_clause, [], [f1590, f1592, f813])).
fof(f1590, plain, ! [X8, X9] : (~ r1(sK25(sK42), X8) | ~ r1(X8, X9) | sP21(X9) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1589, f244])).
fof(f1589, plain, ! [X8, X9] : (~ r1(sK25(sK42), X8) | ~ r1(X8, X9) | sP21(X9) | ~ p100(sK42) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1257, f243])).
fof(f1257, plain, ! [X8, X9] : (~ r1(sK25(sK42), X8) | ~ r1(X8, X9) | sP21(X9) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(resolution, [], [f1211, f169])).
fof(f1211, plain, ! [X2, X0, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP21(X2)), inference(resolution, [], [f1169, f109])).
fof(f1169, plain, ! [X2, X0, X3, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP21(X3)), inference(resolution, [], [f1127, f109])).
fof(f1127, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | sP21(X4)), inference(resolution, [], [f1085, f109])).
fof(f1085, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK42, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | sP21(X5)), inference(resolution, [], [f1043, f109])).
fof(f1043, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK42, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | ~ r1(X5, X6) | sP21(X6)), inference(resolution, [], [f1001, f109])).
fof(f1001, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK42, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | ~ r1(X1, X7) | sP21(X7)), inference(resolution, [], [f938, f109])).
fof(f938, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK42, X8) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X0, X1) | sP21(X1)), inference(resolution, [], [f242, f109])).
fof(f242, plain, ! [X6, X4, X2, X10, X8, X7, X5, X3, X1, X9] : (~ r1(sK42, X1) | ~ r1(X9, X10) | ~ r1(X8, X9) | ~ r1(X7, X8) | ~ r1(X6, X7) | ~ r1(X5, X6) | ~ r1(X4, X5) | ~ r1(X3, X4) | ~ r1(X2, X3) | ~ r1(X1, X2) | sP21(X10)), inference(cnf_transformation, [], [f108])).
fof(f1568, plain, (~ spl43_6 | spl43_77), inference(avatar_split_clause, [], [f1564, f1566, f813])).
fof(f1564, plain, ! [X4] : (~ r1(sK25(sK42), X4) | sP21(X4) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1563, f244])).
fof(f1563, plain, ! [X4] : (~ r1(sK25(sK42), X4) | sP21(X4) | ~ p100(sK42) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1299, f243])).
fof(f1299, plain, ! [X4] : (~ r1(sK25(sK42), X4) | sP21(X4) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(resolution, [], [f1253, f169])).
fof(f1253, plain, ! [X0, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | sP21(X1)), inference(resolution, [], [f1211, f109])).
fof(f1555, plain, (~ spl43_6 | spl43_75), inference(avatar_split_clause, [], [f1550, f1552, f813])).
fof(f1550, plain, (sP21(sK25(sK42)) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1549, f244])).
fof(f1549, plain, (sP21(sK25(sK42)) | ~ p100(sK42) | ~ sP8(sK42)), inference(subsumption_resolution, [], [f1320, f243])).
fof(f1320, plain, (sP21(sK25(sK42)) | p101(sK42) | ~ p100(sK42) | ~ sP8(sK42)), inference(resolution, [], [f1295, f169])).
fof(f1295, plain, ! [X0] : (~ r1(sK42, X0) | sP21(X0)), inference(resolution, [], [f1253, f109])).
fof(f1548, plain, spl43_6, inference(avatar_split_clause, [], [f1346, f813])).
fof(f1346, plain, sP8(sK42), inference(resolution, [], [f1316, f119])).
fof(f119, plain, ! [X0] : (~ sP21(X0) | sP8(X0)), inference(cnf_transformation, [], [f33])).
fof(f1316, plain, sP21(sK42), inference(resolution, [], [f1295, f109])).