fof(f117713, plain, $false, inference(avatar_sat_refutation, [], [f533, f2525, f2530, f2535, f2759, f3070, f3072, f26272, f26321, f85089, f85106, f101504, f109880, f110018, f115598, f115829, f115860, f115867, f115871, f115877, f115895, f117451, f117712])).
fof(f117712, plain, (spl67_6221 | spl67_26 | ~ spl67_54 | ~ spl67_243), inference(avatar_split_clause, [], [f117711, f2773, f517, f360, f115854])).
fof(f115854, plain, (spl67_6221 <=> ! [X2] : (r1(X2, sK7(X2)) | ~ r1(sK42, X2))), introduced(avatar_definition, [new_symbols(naming, [spl67_6221])])).
fof(f360, plain, (spl67_26 <=> p1(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl67_26])])).
fof(f517, plain, (spl67_54 <=> sP2(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl67_54])])).
fof(f2773, plain, (spl67_243 <=> p1(sK6(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl67_243])])).
fof(f117711, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK42, X0)) | (spl67_26 | ~ spl67_54 | ~ spl67_243)), inference(subsumption_resolution, [], [f117710, f518])).
fof(f518, plain, (sP2(sK42) | ~ spl67_54), inference(avatar_component_clause, [], [f517])).
fof(f117710, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK42, X0) | ~ sP2(sK42)) | (spl67_26 | ~ spl67_243)), inference(subsumption_resolution, [], [f115896, f361])).
fof(f361, plain, (~ p1(sK42) | spl67_26), inference(avatar_component_clause, [], [f360])).
fof(f115896, plain, (! [X0] : (r1(X0, sK7(X0)) | ~ r1(sK42, X0) | p1(sK42) | ~ sP2(sK42)) | ~ spl67_243), inference(resolution, [], [f2775, f95])).
fof(f95, plain, ! [X4, X0] : (~ p1(sK6(X0)) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ! [X0] : ((! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0))) | ! [X4] : ((~ p1(sK7(X4)) & r1(X4, sK7(X4))) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7])], [f18, f20, f19])).
fof(f19, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) => (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(sK6(X0), X2)) & ~ p1(sK6(X0)) & r1(X0, sK6(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) => (~ p1(sK7(X4)) & r1(X4, sK7(X4)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1] : (! [X2] : (~ p1(X2) | ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ~ r1(X1, X2)) & ~ p1(X1) & r1(X0, X1)) | ! [X4] : (? [X5] : (~ p1(X5) & r1(X4, X5)) | ~ r1(X0, X4)) | p1(X0) | ~ sP2(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X113] : (? [X125] : (! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) & ~ p1(X125) & r1(X113, X125)) | ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113) | ~ sP2(X113)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X113] : (? [X125] : (! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) & ~ p1(X125) & r1(X113, X125)) | ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113) | ~ sP2(X113)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X113] : (sP2(X113) <=> (? [X125] : (! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) & ~ p1(X125) & r1(X113, X125)) | ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f2775, plain, (p1(sK6(sK42)) | ~ spl67_243), inference(avatar_component_clause, [], [f2773])).
fof(f117451, plain, (spl67_57 | ~ spl67_6221), inference(avatar_contradiction_clause, [], [f117450])).
fof(f117450, plain, ($false | (spl67_57 | ~ spl67_6221)), inference(subsumption_resolution, [], [f115911, f175])).
fof(f175, plain, r1(sK42, sK45), inference(cnf_transformation, [], [f88])).
fof(f88, plain, (! [X1] : ((! [X2] : ((~ p1(sK14(X2)) & r1(X2, sK14(X2))) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & ((! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1))) | ! [X9] : ((~ p1(sK16(X9)) & r1(X9, sK16(X9))) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : ((~ p1(sK17(X13)) & r1(X13, sK17(X13))) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & ((! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12))) | ! [X20] : ((~ p1(sK19(X20)) & r1(X20, sK19(X20))) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : ((~ p1(sK20(X25)) & r1(X25, sK20(X25))) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & ((! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24))) | ! [X32] : ((~ p1(sK22(X32)) & r1(X32, sK22(X32))) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : ((~ p1(sK23(X38)) & r1(X38, sK23(X38))) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & ((! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37))) | ! [X45] : ((~ p1(sK25(X45)) & r1(X45, sK25(X45))) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : ((~ p1(sK26(X52)) & r1(X52, sK26(X52))) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & ((! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51))) | ! [X59] : ((~ p1(sK28(X59)) & r1(X59, sK28(X59))) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : ((~ p1(sK29(X67)) & r1(X67, sK29(X67))) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & ((! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(sK30(X66), X72)) & ~ p1(sK30(X66)) & r1(X66, sK30(X66))) | ! [X74] : ((~ p1(sK31(X74)) & r1(X74, sK31(X74))) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(sK13, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : ((~ p1(sK32(X83)) & r1(X83, sK32(X83))) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & ((! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(sK33(X82), X88)) & ~ p1(sK33(X82)) & r1(X82, sK33(X82))) | ! [X90] : ((~ p1(sK34(X90)) & r1(X90, sK34(X90))) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(sK13, X76)) & ((((((((! [X100] : ((p1(sK43(X100)) & (~ p1(sK44(X100)) & r1(sK43(X100), sK44(X100))) & r1(X100, sK43(X100))) | p1(X100) | ~ r1(sK42, X100)) & (! [X104] : (p1(X104) | ~ r1(sK45, X104)) & r1(sK42, sK45)) & (~ p1(sK46) & r1(sK42, sK46)) & r1(sK41, sK42)) & r1(sK40, sK41)) & r1(sK39, sK40)) & r1(sK38, sK39)) & r1(sK37, sK38)) & r1(sK36, sK37)) & r1(sK35, sK36)) & r1(sK13, sK35)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : ((((! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(sK47(X113), X115)) & ~ p1(sK47(X113)) & r1(X113, sK47(X113))) | sP1(X113)) & ! [X117] : ((~ p1(sK48(X117)) & r1(X117, sK48(X117))) | ! [X119] : (! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X117, X119)) | ~ r1(X113, X117)) & sP2(X113) & sP3(X113)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(sK13, X106)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : (! [X129] : ((! [X130] : ((~ p1(sK49(X130)) & r1(X130, sK49(X130))) | ! [X132] : (! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X130, X132)) | ~ r1(X129, X130)) & ((! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(sK50(X129), X135)) & ~ p1(sK50(X129)) & r1(X129, sK50(X129))) | ! [X137] : ((~ p1(sK51(X137)) & r1(X137, sK51(X137))) | ~ r1(X129, X137)) | p1(X129))) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(sK13, X121)) & ! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : ((~ p1(sK52(X149)) & r1(X149, sK52(X149))) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & ((! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(sK53(X148), X154)) & ~ p1(sK53(X148)) & r1(X148, sK53(X148))) | ! [X156] : ((~ p1(sK54(X156)) & r1(X156, sK54(X156))) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(sK13, X139)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : ((! [X169] : ((~ p1(sK55(X169)) & r1(X169, sK55(X169))) | ! [X171] : (! [X172] : (p1(X172) | ~ r1(X171, X172)) | ~ r1(X169, X171)) | ~ r1(X168, X169)) & ((! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(sK56(X168), X174)) & ~ p1(sK56(X168)) & r1(X168, sK56(X168))) | ! [X176] : ((~ p1(sK57(X176)) & r1(X176, sK57(X176))) | ~ r1(X168, X176)) | p1(X168))) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(sK13, X158)) & ! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : ((! [X190] : ((~ p1(sK58(X190)) & r1(X190, sK58(X190))) | ! [X192] : (! [X193] : (p1(X193) | ~ r1(X192, X193)) | ~ r1(X190, X192)) | ~ r1(X189, X190)) & ((! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(sK59(X189), X195)) & ~ p1(sK59(X189)) & r1(X189, sK59(X189))) | ! [X197] : ((~ p1(sK60(X197)) & r1(X197, sK60(X197))) | ~ r1(X189, X197)) | p1(X189))) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(sK13, X178)) & ! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : ((~ p1(sK61(X212)) & r1(X212, sK61(X212))) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & ((! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(sK62(X211), X217)) & ~ p1(sK62(X211)) & r1(X211, sK62(X211))) | ! [X219] : ((~ p1(sK63(X219)) & r1(X219, sK63(X219))) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(sK13, X199)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : ((! [X235] : ((~ p1(sK64(X235)) & r1(X235, sK64(X235))) | ! [X237] : (! [X238] : (p1(X238) | ~ r1(X237, X238)) | ~ r1(X235, X237)) | ~ r1(X234, X235)) & ((! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(sK65(X234), X240)) & ~ p1(sK65(X234)) & r1(X234, sK65(X234))) | ! [X242] : ((~ p1(sK66(X242)) & r1(X242, sK66(X242))) | ~ r1(X234, X242)) | p1(X234))) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(sK13, X221))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15, sK16, sK17, sK18, sK19, sK20, sK21, sK22, sK23, sK24, sK25, sK26, sK27, sK28, sK29, sK30, sK31, sK32, sK33, sK34, sK35, sK36, sK37, sK38, sK39, sK40, sK41, sK42, sK43, sK44, sK45, sK46, sK47, sK48, sK49, sK50, sK51, sK52, sK53, sK54, sK55, sK56, sK57, sK58, sK59, sK60, sK61, sK62, sK63, sK64, sK65, sK66])], [f33, f87, f86, f85, f84, f83, f82, f81, f80, f79, f78, f77, f76, f75, f74, f73, f72, f71, f70, f69, f68, f67, f66, f65, f64, f63, f62, f61, f60, f59, f58, f57, f56, f55, f54, f53, f52, f51, f50, f49, f48, f47, f46, f45, f44, f43, f42, f41, f40, f39, f38, f37, f36, f35, f34])).
fof(f34, plain, (? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(X0, X92)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | sP1(X113)) & ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ! [X119] : (! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X117, X119)) | ~ r1(X113, X117)) & sP2(X113) & sP3(X113)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : (! [X129] : ((! [X130] : (? [X131] : (~ p1(X131) & r1(X130, X131)) | ! [X132] : (! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X130, X132)) | ~ r1(X129, X130)) & (? [X134] : (! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(X134, X135)) & ~ p1(X134) & r1(X129, X134)) | ! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ~ r1(X129, X137)) | p1(X129))) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X0, X139)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : ((! [X169] : (? [X170] : (~ p1(X170) & r1(X169, X170)) | ! [X171] : (! [X172] : (p1(X172) | ~ r1(X171, X172)) | ~ r1(X169, X171)) | ~ r1(X168, X169)) & (? [X173] : (! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(X173, X174)) & ~ p1(X173) & r1(X168, X173)) | ! [X176] : (? [X177] : (~ p1(X177) & r1(X176, X177)) | ~ r1(X168, X176)) | p1(X168))) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) & ! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : ((! [X190] : (? [X191] : (~ p1(X191) & r1(X190, X191)) | ! [X192] : (! [X193] : (p1(X193) | ~ r1(X192, X193)) | ~ r1(X190, X192)) | ~ r1(X189, X190)) & (? [X194] : (! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X194, X195)) & ~ p1(X194) & r1(X189, X194)) | ! [X197] : (? [X198] : (~ p1(X198) & r1(X197, X198)) | ~ r1(X189, X197)) | p1(X189))) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X0, X178)) & ! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X0, X199)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : ((! [X235] : (? [X236] : (~ p1(X236) & r1(X235, X236)) | ! [X237] : (! [X238] : (p1(X238) | ~ r1(X237, X238)) | ~ r1(X235, X237)) | ~ r1(X234, X235)) & (? [X239] : (! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X239, X240)) & ~ p1(X239) & r1(X234, X239)) | ! [X242] : (? [X243] : (~ p1(X243) & r1(X242, X243)) | ~ r1(X234, X242)) | p1(X234))) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221))) => (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(sK13, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(sK13, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(sK13, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(sK13, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(sK13, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(sK13, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(sK13, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(sK13, X92)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | sP1(X113)) & ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ! [X119] : (! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X117, X119)) | ~ r1(X113, X117)) & sP2(X113) & sP3(X113)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(sK13, X106)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : (! [X129] : ((! [X130] : (? [X131] : (~ p1(X131) & r1(X130, X131)) | ! [X132] : (! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X130, X132)) | ~ r1(X129, X130)) & (? [X134] : (! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(X134, X135)) & ~ p1(X134) & r1(X129, X134)) | ! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ~ r1(X129, X137)) | p1(X129))) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(sK13, X121)) & ! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(sK13, X139)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : ((! [X169] : (? [X170] : (~ p1(X170) & r1(X169, X170)) | ! [X171] : (! [X172] : (p1(X172) | ~ r1(X171, X172)) | ~ r1(X169, X171)) | ~ r1(X168, X169)) & (? [X173] : (! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(X173, X174)) & ~ p1(X173) & r1(X168, X173)) | ! [X176] : (? [X177] : (~ p1(X177) & r1(X176, X177)) | ~ r1(X168, X176)) | p1(X168))) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(sK13, X158)) & ! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : ((! [X190] : (? [X191] : (~ p1(X191) & r1(X190, X191)) | ! [X192] : (! [X193] : (p1(X193) | ~ r1(X192, X193)) | ~ r1(X190, X192)) | ~ r1(X189, X190)) & (? [X194] : (! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X194, X195)) & ~ p1(X194) & r1(X189, X194)) | ! [X197] : (? [X198] : (~ p1(X198) & r1(X197, X198)) | ~ r1(X189, X197)) | p1(X189))) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(sK13, X178)) & ! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(sK13, X199)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : ((! [X235] : (? [X236] : (~ p1(X236) & r1(X235, X236)) | ! [X237] : (! [X238] : (p1(X238) | ~ r1(X237, X238)) | ~ r1(X235, X237)) | ~ r1(X234, X235)) & (? [X239] : (! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X239, X240)) & ~ p1(X239) & r1(X234, X239)) | ! [X242] : (? [X243] : (~ p1(X243) & r1(X242, X243)) | ~ r1(X234, X242)) | p1(X234))) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(sK13, X221)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK14(X2)) & r1(X2, sK14(X2)))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X1] : (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) => (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(sK15(X1), X7)) & ~ p1(sK15(X1)) & r1(X1, sK15(X1)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) => (~ p1(sK16(X9)) & r1(X9, sK16(X9)))), introduced(choice_axiom, [])).
fof(f38, plain, ! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) => (~ p1(sK17(X13)) & r1(X13, sK17(X13)))), introduced(choice_axiom, [])).
fof(f39, plain, ! [X12] : (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) => (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(sK18(X12), X18)) & ~ p1(sK18(X12)) & r1(X12, sK18(X12)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) => (~ p1(sK19(X20)) & r1(X20, sK19(X20)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) => (~ p1(sK20(X25)) & r1(X25, sK20(X25)))), introduced(choice_axiom, [])).
fof(f42, plain, ! [X24] : (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) => (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(sK21(X24), X30)) & ~ p1(sK21(X24)) & r1(X24, sK21(X24)))), introduced(choice_axiom, [])).
fof(f43, plain, ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) => (~ p1(sK22(X32)) & r1(X32, sK22(X32)))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) => (~ p1(sK23(X38)) & r1(X38, sK23(X38)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X37] : (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) => (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(sK24(X37), X43)) & ~ p1(sK24(X37)) & r1(X37, sK24(X37)))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) => (~ p1(sK25(X45)) & r1(X45, sK25(X45)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) => (~ p1(sK26(X52)) & r1(X52, sK26(X52)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X51] : (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) => (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(sK27(X51), X57)) & ~ p1(sK27(X51)) & r1(X51, sK27(X51)))), introduced(choice_axiom, [])).
fof(f49, plain, ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) => (~ p1(sK28(X59)) & r1(X59, sK28(X59)))), introduced(choice_axiom, [])).
fof(f50, plain, ! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) => (~ p1(sK29(X67)) & r1(X67, sK29(X67)))), introduced(choice_axiom, [])).
fof(f51, plain, ! [X66] : (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) => (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(sK30(X66), X72)) & ~ p1(sK30(X66)) & r1(X66, sK30(X66)))), introduced(choice_axiom, [])).
fof(f52, plain, ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) => (~ p1(sK31(X74)) & r1(X74, sK31(X74)))), introduced(choice_axiom, [])).
fof(f53, plain, ! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) => (~ p1(sK32(X83)) & r1(X83, sK32(X83)))), introduced(choice_axiom, [])).
fof(f54, plain, ! [X82] : (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) => (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(sK33(X82), X88)) & ~ p1(sK33(X82)) & r1(X82, sK33(X82)))), introduced(choice_axiom, [])).
fof(f55, plain, ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) => (~ p1(sK34(X90)) & r1(X90, sK34(X90)))), introduced(choice_axiom, [])).
fof(f56, plain, (? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(sK13, X92)) => (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(sK35, X93)) & r1(sK13, sK35))), introduced(choice_axiom, [])).
fof(f57, plain, (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(sK35, X93)) => (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(sK36, X94)) & r1(sK35, sK36))), introduced(choice_axiom, [])).
fof(f58, plain, (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(sK36, X94)) => (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(sK37, X95)) & r1(sK36, sK37))), introduced(choice_axiom, [])).
fof(f59, plain, (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(sK37, X95)) => (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(sK38, X96)) & r1(sK37, sK38))), introduced(choice_axiom, [])).
fof(f60, plain, (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(sK38, X96)) => (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(sK39, X97)) & r1(sK38, sK39))), introduced(choice_axiom, [])).
fof(f61, plain, (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(sK39, X97)) => (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(sK40, X98)) & r1(sK39, sK40))), introduced(choice_axiom, [])).
fof(f62, plain, (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(sK40, X98)) => (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(sK41, X99)) & r1(sK40, sK41))), introduced(choice_axiom, [])).
fof(f63, plain, (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(sK41, X99)) => (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(sK42, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(sK42, X103)) & ? [X105] : (~ p1(X105) & r1(sK42, X105)) & r1(sK41, sK42))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) => (p1(sK43(X100)) & ? [X102] : (~ p1(X102) & r1(sK43(X100), X102)) & r1(X100, sK43(X100)))), introduced(choice_axiom, [])).
fof(f65, plain, ! [X100] : (? [X102] : (~ p1(X102) & r1(sK43(X100), X102)) => (~ p1(sK44(X100)) & r1(sK43(X100), sK44(X100)))), introduced(choice_axiom, [])).
fof(f66, plain, (? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(sK42, X103)) => (! [X104] : (p1(X104) | ~ r1(sK45, X104)) & r1(sK42, sK45))), introduced(choice_axiom, [])).
fof(f67, plain, (? [X105] : (~ p1(X105) & r1(sK42, X105)) => (~ p1(sK46) & r1(sK42, sK46))), introduced(choice_axiom, [])).
fof(f68, plain, ! [X113] : (? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) => (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(sK47(X113), X115)) & ~ p1(sK47(X113)) & r1(X113, sK47(X113)))), introduced(choice_axiom, [])).
fof(f69, plain, ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) => (~ p1(sK48(X117)) & r1(X117, sK48(X117)))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X130] : (? [X131] : (~ p1(X131) & r1(X130, X131)) => (~ p1(sK49(X130)) & r1(X130, sK49(X130)))), introduced(choice_axiom, [])).
fof(f71, plain, ! [X129] : (? [X134] : (! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(X134, X135)) & ~ p1(X134) & r1(X129, X134)) => (! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(sK50(X129), X135)) & ~ p1(sK50(X129)) & r1(X129, sK50(X129)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) => (~ p1(sK51(X137)) & r1(X137, sK51(X137)))), introduced(choice_axiom, [])).
fof(f73, plain, ! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) => (~ p1(sK52(X149)) & r1(X149, sK52(X149)))), introduced(choice_axiom, [])).
fof(f74, plain, ! [X148] : (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) => (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(sK53(X148), X154)) & ~ p1(sK53(X148)) & r1(X148, sK53(X148)))), introduced(choice_axiom, [])).
fof(f75, plain, ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) => (~ p1(sK54(X156)) & r1(X156, sK54(X156)))), introduced(choice_axiom, [])).
fof(f76, plain, ! [X169] : (? [X170] : (~ p1(X170) & r1(X169, X170)) => (~ p1(sK55(X169)) & r1(X169, sK55(X169)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X168] : (? [X173] : (! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(X173, X174)) & ~ p1(X173) & r1(X168, X173)) => (! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(sK56(X168), X174)) & ~ p1(sK56(X168)) & r1(X168, sK56(X168)))), introduced(choice_axiom, [])).
fof(f78, plain, ! [X176] : (? [X177] : (~ p1(X177) & r1(X176, X177)) => (~ p1(sK57(X176)) & r1(X176, sK57(X176)))), introduced(choice_axiom, [])).
fof(f79, plain, ! [X190] : (? [X191] : (~ p1(X191) & r1(X190, X191)) => (~ p1(sK58(X190)) & r1(X190, sK58(X190)))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X189] : (? [X194] : (! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X194, X195)) & ~ p1(X194) & r1(X189, X194)) => (! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(sK59(X189), X195)) & ~ p1(sK59(X189)) & r1(X189, sK59(X189)))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X197] : (? [X198] : (~ p1(X198) & r1(X197, X198)) => (~ p1(sK60(X197)) & r1(X197, sK60(X197)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) => (~ p1(sK61(X212)) & r1(X212, sK61(X212)))), introduced(choice_axiom, [])).
fof(f83, plain, ! [X211] : (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) => (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(sK62(X211), X217)) & ~ p1(sK62(X211)) & r1(X211, sK62(X211)))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) => (~ p1(sK63(X219)) & r1(X219, sK63(X219)))), introduced(choice_axiom, [])).
fof(f85, plain, ! [X235] : (? [X236] : (~ p1(X236) & r1(X235, X236)) => (~ p1(sK64(X235)) & r1(X235, sK64(X235)))), introduced(choice_axiom, [])).
fof(f86, plain, ! [X234] : (? [X239] : (! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X239, X240)) & ~ p1(X239) & r1(X234, X239)) => (! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(sK65(X234), X240)) & ~ p1(sK65(X234)) & r1(X234, sK65(X234)))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X242] : (? [X243] : (~ p1(X243) & r1(X242, X243)) => (~ p1(sK66(X242)) & r1(X242, sK66(X242)))), introduced(choice_axiom, [])).
fof(f33, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(X0, X92)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | sP1(X113)) & ! [X117] : (? [X118] : (~ p1(X118) & r1(X117, X118)) | ! [X119] : (! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X117, X119)) | ~ r1(X113, X117)) & sP2(X113) & sP3(X113)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) & ! [X121] : (! [X122] : (! [X123] : (! [X124] : (! [X125] : (! [X126] : (! [X127] : (! [X128] : (! [X129] : ((! [X130] : (? [X131] : (~ p1(X131) & r1(X130, X131)) | ! [X132] : (! [X133] : (p1(X133) | ~ r1(X132, X133)) | ~ r1(X130, X132)) | ~ r1(X129, X130)) & (? [X134] : (! [X135] : (~ p1(X135) | ! [X136] : (p1(X136) | ~ r1(X135, X136)) | ~ r1(X134, X135)) & ~ p1(X134) & r1(X129, X134)) | ! [X137] : (? [X138] : (~ p1(X138) & r1(X137, X138)) | ~ r1(X129, X137)) | p1(X129))) | ~ r1(X128, X129)) | ~ r1(X127, X128)) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | ~ r1(X124, X125)) | ~ r1(X123, X124)) | ~ r1(X122, X123)) | ~ r1(X121, X122)) | ~ r1(X0, X121)) & ! [X139] : (! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X139, X140)) | ~ r1(X0, X139)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : (! [X168] : ((! [X169] : (? [X170] : (~ p1(X170) & r1(X169, X170)) | ! [X171] : (! [X172] : (p1(X172) | ~ r1(X171, X172)) | ~ r1(X169, X171)) | ~ r1(X168, X169)) & (? [X173] : (! [X174] : (~ p1(X174) | ! [X175] : (p1(X175) | ~ r1(X174, X175)) | ~ r1(X173, X174)) & ~ p1(X173) & r1(X168, X173)) | ! [X176] : (? [X177] : (~ p1(X177) & r1(X176, X177)) | ~ r1(X168, X176)) | p1(X168))) | ~ r1(X167, X168)) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) & ! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : (! [X188] : (! [X189] : ((! [X190] : (? [X191] : (~ p1(X191) & r1(X190, X191)) | ! [X192] : (! [X193] : (p1(X193) | ~ r1(X192, X193)) | ~ r1(X190, X192)) | ~ r1(X189, X190)) & (? [X194] : (! [X195] : (~ p1(X195) | ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X194, X195)) & ~ p1(X194) & r1(X189, X194)) | ! [X197] : (? [X198] : (~ p1(X198) & r1(X197, X198)) | ~ r1(X189, X197)) | p1(X189))) | ~ r1(X188, X189)) | ~ r1(X187, X188)) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X0, X178)) & ! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : (! [X209] : (! [X210] : (! [X211] : ((! [X212] : (? [X213] : (~ p1(X213) & r1(X212, X213)) | ! [X214] : (! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X212, X214)) | ~ r1(X211, X212)) & (? [X216] : (! [X217] : (~ p1(X217) | ! [X218] : (p1(X218) | ~ r1(X217, X218)) | ~ r1(X216, X217)) & ~ p1(X216) & r1(X211, X216)) | ! [X219] : (? [X220] : (~ p1(X220) & r1(X219, X220)) | ~ r1(X211, X219)) | p1(X211))) | ~ r1(X210, X211)) | ~ r1(X209, X210)) | ~ r1(X208, X209)) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X0, X199)) & ! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : (! [X231] : (! [X232] : (! [X233] : (! [X234] : ((! [X235] : (? [X236] : (~ p1(X236) & r1(X235, X236)) | ! [X237] : (! [X238] : (p1(X238) | ~ r1(X237, X238)) | ~ r1(X235, X237)) | ~ r1(X234, X235)) & (? [X239] : (! [X240] : (~ p1(X240) | ! [X241] : (p1(X241) | ~ r1(X240, X241)) | ~ r1(X239, X240)) & ~ p1(X239) & r1(X234, X239)) | ! [X242] : (? [X243] : (~ p1(X243) & r1(X242, X243)) | ~ r1(X234, X242)) | p1(X234))) | ~ r1(X233, X234)) | ~ r1(X232, X233)) | ~ r1(X231, X232)) | ~ r1(X230, X231)) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X0, X221))), inference(rectify, [], [f11])).
fof(f11, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(X0, X92)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | sP1(X113)) & ! [X121] : (? [X122] : (~ p1(X122) & r1(X121, X122)) | ! [X123] : (! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X121, X123)) | ~ r1(X113, X121)) & sP2(X113) & sP3(X113)) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) & ! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X0, X140)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : ((! [X168] : (? [X169] : (~ p1(X169) & r1(X168, X169)) | ! [X170] : (! [X171] : (p1(X171) | ~ r1(X170, X171)) | ~ r1(X168, X170)) | ~ r1(X167, X168)) & (? [X172] : (! [X173] : (~ p1(X173) | ! [X174] : (p1(X174) | ~ r1(X173, X174)) | ~ r1(X172, X173)) & ~ p1(X172) & r1(X167, X172)) | ! [X175] : (? [X176] : (~ p1(X176) & r1(X175, X176)) | ~ r1(X167, X175)) | p1(X167))) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) & ! [X177] : (! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : ((! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ! [X190] : (! [X191] : (p1(X191) | ~ r1(X190, X191)) | ~ r1(X188, X190)) | ~ r1(X187, X188)) & (? [X192] : (! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X192, X193)) & ~ p1(X192) & r1(X187, X192)) | ! [X195] : (? [X196] : (~ p1(X196) & r1(X195, X196)) | ~ r1(X187, X195)) | p1(X187))) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X177, X178)) | ~ r1(X0, X177)) & ! [X197] : (! [X198] : (! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : ((! [X209] : (? [X210] : (~ p1(X210) & r1(X209, X210)) | ! [X211] : (! [X212] : (p1(X212) | ~ r1(X211, X212)) | ~ r1(X209, X211)) | ~ r1(X208, X209)) & (? [X213] : (! [X214] : (~ p1(X214) | ! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X213, X214)) & ~ p1(X213) & r1(X208, X213)) | ! [X216] : (? [X217] : (~ p1(X217) & r1(X216, X217)) | ~ r1(X208, X216)) | p1(X208))) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X198, X199)) | ~ r1(X197, X198)) | ~ r1(X0, X197)) & ! [X218] : (! [X219] : (! [X220] : (! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : ((! [X231] : (? [X232] : (~ p1(X232) & r1(X231, X232)) | ! [X233] : (! [X234] : (p1(X234) | ~ r1(X233, X234)) | ~ r1(X231, X233)) | ~ r1(X230, X231)) & (? [X235] : (! [X236] : (~ p1(X236) | ! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X235, X236)) & ~ p1(X235) & r1(X230, X235)) | ! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ~ r1(X230, X238)) | p1(X230))) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X220, X221)) | ~ r1(X219, X220)) | ~ r1(X218, X219)) | ~ r1(X0, X218)) & ! [X240] : (! [X241] : (! [X242] : (! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : ((! [X254] : (? [X255] : (~ p1(X255) & r1(X254, X255)) | ! [X256] : (! [X257] : (p1(X257) | ~ r1(X256, X257)) | ~ r1(X254, X256)) | ~ r1(X253, X254)) & (? [X258] : (! [X259] : (~ p1(X259) | ! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X258, X259)) & ~ p1(X258) & r1(X253, X258)) | ! [X261] : (? [X262] : (~ p1(X262) & r1(X261, X262)) | ~ r1(X253, X261)) | p1(X253))) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X242, X243)) | ~ r1(X241, X242)) | ~ r1(X240, X241)) | ~ r1(X0, X240))), inference(definition_folding, [], [f6, e10, e9, e8, e7])).
fof(f7, plain, ! [X113] : (? [X130] : (! [X131] : ((p1(X131) & ? [X132] : (~ p1(X132) & r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) & p1(X130) & ? [X135] : (~ p1(X135) & r1(X130, X135)) & r1(X113, X130)) | ~ sP0(X113)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X113] : (sP0(X113) <=> ? [X130] : (! [X131] : ((p1(X131) & ? [X132] : (~ p1(X132) & r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) & p1(X130) & ? [X135] : (~ p1(X135) & r1(X130, X135)) & r1(X113, X130))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f8, plain, ! [X113] : (! [X117] : (! [X118] : (? [X119] : (p1(X119) & ? [X120] : (~ p1(X120) & r1(X119, X120)) & r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117)) | ~ sP1(X113)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X113] : (sP1(X113) <=> ! [X117] : (! [X118] : (? [X119] : (p1(X119) & ? [X120] : (~ p1(X120) & r1(X119, X120)) & r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f10, plain, ! [X113] : (sP0(X113) | ! [X136] : (? [X137] : (p1(X137) & ? [X138] : (~ p1(X138) & r1(X137, X138)) & r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)) | ~ sP3(X113)), inference(usedef, [], [e10])).
fof(e10, plain, ! [X113] : (sP3(X113) <=> (sP0(X113) | ! [X136] : (? [X137] : (p1(X137) & ? [X138] : (~ p1(X138) & r1(X137, X138)) & r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f6, plain, ? [X0] : (! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(X0, X92)) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | ! [X117] : (! [X118] : (? [X119] : (p1(X119) & ? [X120] : (~ p1(X120) & r1(X119, X120)) & r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117))) & ! [X121] : (? [X122] : (~ p1(X122) & r1(X121, X122)) | ! [X123] : (! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X121, X123)) | ~ r1(X113, X121)) & (? [X125] : (! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) & ~ p1(X125) & r1(X113, X125)) | ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113)) & (? [X130] : (! [X131] : ((p1(X131) & ? [X132] : (~ p1(X132) & r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) & p1(X130) & ? [X135] : (~ p1(X135) & r1(X130, X135)) & r1(X113, X130)) | ! [X136] : (? [X137] : (p1(X137) & ? [X138] : (~ p1(X138) & r1(X137, X138)) & r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)))) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) & ! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X0, X140)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : ((! [X168] : (? [X169] : (~ p1(X169) & r1(X168, X169)) | ! [X170] : (! [X171] : (p1(X171) | ~ r1(X170, X171)) | ~ r1(X168, X170)) | ~ r1(X167, X168)) & (? [X172] : (! [X173] : (~ p1(X173) | ! [X174] : (p1(X174) | ~ r1(X173, X174)) | ~ r1(X172, X173)) & ~ p1(X172) & r1(X167, X172)) | ! [X175] : (? [X176] : (~ p1(X176) & r1(X175, X176)) | ~ r1(X167, X175)) | p1(X167))) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) & ! [X177] : (! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : ((! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ! [X190] : (! [X191] : (p1(X191) | ~ r1(X190, X191)) | ~ r1(X188, X190)) | ~ r1(X187, X188)) & (? [X192] : (! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X192, X193)) & ~ p1(X192) & r1(X187, X192)) | ! [X195] : (? [X196] : (~ p1(X196) & r1(X195, X196)) | ~ r1(X187, X195)) | p1(X187))) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X177, X178)) | ~ r1(X0, X177)) & ! [X197] : (! [X198] : (! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : ((! [X209] : (? [X210] : (~ p1(X210) & r1(X209, X210)) | ! [X211] : (! [X212] : (p1(X212) | ~ r1(X211, X212)) | ~ r1(X209, X211)) | ~ r1(X208, X209)) & (? [X213] : (! [X214] : (~ p1(X214) | ! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X213, X214)) & ~ p1(X213) & r1(X208, X213)) | ! [X216] : (? [X217] : (~ p1(X217) & r1(X216, X217)) | ~ r1(X208, X216)) | p1(X208))) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X198, X199)) | ~ r1(X197, X198)) | ~ r1(X0, X197)) & ! [X218] : (! [X219] : (! [X220] : (! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : ((! [X231] : (? [X232] : (~ p1(X232) & r1(X231, X232)) | ! [X233] : (! [X234] : (p1(X234) | ~ r1(X233, X234)) | ~ r1(X231, X233)) | ~ r1(X230, X231)) & (? [X235] : (! [X236] : (~ p1(X236) | ! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X235, X236)) & ~ p1(X235) & r1(X230, X235)) | ! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ~ r1(X230, X238)) | p1(X230))) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X220, X221)) | ~ r1(X219, X220)) | ~ r1(X218, X219)) | ~ r1(X0, X218)) & ! [X240] : (! [X241] : (! [X242] : (! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : ((! [X254] : (? [X255] : (~ p1(X255) & r1(X254, X255)) | ! [X256] : (! [X257] : (p1(X257) | ~ r1(X256, X257)) | ~ r1(X254, X256)) | ~ r1(X253, X254)) & (? [X258] : (! [X259] : (~ p1(X259) | ! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X258, X259)) & ~ p1(X258) & r1(X253, X258)) | ! [X261] : (? [X262] : (~ p1(X262) & r1(X261, X262)) | ~ r1(X253, X261)) | p1(X253))) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X242, X243)) | ~ r1(X241, X242)) | ~ r1(X240, X241)) | ~ r1(X0, X240))), inference(flattening, [], [f5])).
fof(f5, plain, ? [X0] : ((! [X1] : ((! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (? [X6] : (! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) & ~ p1(X6) & r1(X1, X6)) | ! [X9] : (? [X10] : (~ p1(X10) & r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (? [X14] : (~ p1(X14) & r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (? [X17] : (! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) & ~ p1(X17) & r1(X12, X17)) | ! [X20] : (? [X21] : (~ p1(X21) & r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (? [X26] : (~ p1(X26) & r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (? [X29] : (! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) & ~ p1(X29) & r1(X24, X29)) | ! [X32] : (? [X33] : (~ p1(X33) & r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (? [X39] : (~ p1(X39) & r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (? [X42] : (! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) & ~ p1(X42) & r1(X37, X42)) | ! [X45] : (? [X46] : (~ p1(X46) & r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (? [X53] : (~ p1(X53) & r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (? [X56] : (! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) & ~ p1(X56) & r1(X51, X56)) | ! [X59] : (? [X60] : (~ p1(X60) & r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (? [X68] : (~ p1(X68) & r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (? [X71] : (! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) & ~ p1(X71) & r1(X66, X71)) | ! [X74] : (? [X75] : (~ p1(X75) & r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (? [X84] : (~ p1(X84) & r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (? [X87] : (! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) & ~ p1(X87) & r1(X82, X87)) | ! [X90] : (? [X91] : (~ p1(X91) & r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ? [X92] : (? [X93] : (? [X94] : (? [X95] : (? [X96] : (? [X97] : (? [X98] : (? [X99] : (! [X100] : (? [X101] : (p1(X101) & ? [X102] : (~ p1(X102) & r1(X101, X102)) & r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) & ? [X103] : (! [X104] : (p1(X104) | ~ r1(X103, X104)) & r1(X99, X103)) & ? [X105] : (~ p1(X105) & r1(X99, X105)) & r1(X98, X99)) & r1(X97, X98)) & r1(X96, X97)) & r1(X95, X96)) & r1(X94, X95)) & r1(X93, X94)) & r1(X92, X93)) & r1(X0, X92))) & ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((? [X114] : (! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) & ~ p1(X114) & r1(X113, X114)) | ! [X117] : (! [X118] : (? [X119] : (p1(X119) & ? [X120] : (~ p1(X120) & r1(X119, X120)) & r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117))) & ! [X121] : (? [X122] : (~ p1(X122) & r1(X121, X122)) | ! [X123] : (! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X121, X123)) | ~ r1(X113, X121)) & (? [X125] : (! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) & ~ p1(X125) & r1(X113, X125)) | ! [X128] : (? [X129] : (~ p1(X129) & r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113)) & (? [X130] : (! [X131] : ((p1(X131) & ? [X132] : (~ p1(X132) & r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) & p1(X130) & ? [X135] : (~ p1(X135) & r1(X130, X135)) & r1(X113, X130)) | ! [X136] : (? [X137] : (p1(X137) & ? [X138] : (~ p1(X138) & r1(X137, X138)) & r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)))) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) & ! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (? [X150] : (~ p1(X150) & r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (? [X153] : (! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) & ~ p1(X153) & r1(X148, X153)) | ! [X156] : (? [X157] : (~ p1(X157) & r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X0, X140)) & ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : ((! [X168] : (? [X169] : (~ p1(X169) & r1(X168, X169)) | ! [X170] : (! [X171] : (p1(X171) | ~ r1(X170, X171)) | ~ r1(X168, X170)) | ~ r1(X167, X168)) & (? [X172] : (! [X173] : (~ p1(X173) | ! [X174] : (p1(X174) | ~ r1(X173, X174)) | ~ r1(X172, X173)) & ~ p1(X172) & r1(X167, X172)) | ! [X175] : (? [X176] : (~ p1(X176) & r1(X175, X176)) | ~ r1(X167, X175)) | p1(X167))) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) & ! [X177] : (! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : ((! [X188] : (? [X189] : (~ p1(X189) & r1(X188, X189)) | ! [X190] : (! [X191] : (p1(X191) | ~ r1(X190, X191)) | ~ r1(X188, X190)) | ~ r1(X187, X188)) & (? [X192] : (! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X192, X193)) & ~ p1(X192) & r1(X187, X192)) | ! [X195] : (? [X196] : (~ p1(X196) & r1(X195, X196)) | ~ r1(X187, X195)) | p1(X187))) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X177, X178)) | ~ r1(X0, X177)) & ! [X197] : (! [X198] : (! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : ((! [X209] : (? [X210] : (~ p1(X210) & r1(X209, X210)) | ! [X211] : (! [X212] : (p1(X212) | ~ r1(X211, X212)) | ~ r1(X209, X211)) | ~ r1(X208, X209)) & (? [X213] : (! [X214] : (~ p1(X214) | ! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X213, X214)) & ~ p1(X213) & r1(X208, X213)) | ! [X216] : (? [X217] : (~ p1(X217) & r1(X216, X217)) | ~ r1(X208, X216)) | p1(X208))) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X198, X199)) | ~ r1(X197, X198)) | ~ r1(X0, X197)) & ! [X218] : (! [X219] : (! [X220] : (! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : ((! [X231] : (? [X232] : (~ p1(X232) & r1(X231, X232)) | ! [X233] : (! [X234] : (p1(X234) | ~ r1(X233, X234)) | ~ r1(X231, X233)) | ~ r1(X230, X231)) & (? [X235] : (! [X236] : (~ p1(X236) | ! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X235, X236)) & ~ p1(X235) & r1(X230, X235)) | ! [X238] : (? [X239] : (~ p1(X239) & r1(X238, X239)) | ~ r1(X230, X238)) | p1(X230))) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X220, X221)) | ~ r1(X219, X220)) | ~ r1(X218, X219)) | ~ r1(X0, X218)) & ! [X240] : (! [X241] : (! [X242] : (! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : ((! [X254] : (? [X255] : (~ p1(X255) & r1(X254, X255)) | ! [X256] : (! [X257] : (p1(X257) | ~ r1(X256, X257)) | ~ r1(X254, X256)) | ~ r1(X253, X254)) & (? [X258] : (! [X259] : (~ p1(X259) | ! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X258, X259)) & ~ p1(X258) & r1(X253, X258)) | ! [X261] : (? [X262] : (~ p1(X262) & r1(X261, X262)) | ~ r1(X253, X261)) | p1(X253))) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X242, X243)) | ~ r1(X241, X242)) | ~ r1(X240, X241)) | ~ r1(X0, X240))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (~ ! [X68] : (p1(X68) | ~ r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (~ ! [X71] : (~ ! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) | p1(X71) | ~ r1(X66, X71)) | ! [X74] : (~ ! [X75] : (p1(X75) | ~ r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (~ ! [X84] : (p1(X84) | ~ r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (~ ! [X87] : (~ ! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | p1(X87) | ~ r1(X82, X87)) | ! [X90] : (~ ! [X91] : (p1(X91) | ~ r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ~ ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : (~ ! [X100] : (~ ! [X101] : (~ p1(X101) | ! [X102] : (p1(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) | ! [X103] : (~ ! [X104] : (p1(X104) | ~ r1(X103, X104)) | ~ r1(X99, X103)) | ! [X105] : (p1(X105) | ~ r1(X99, X105)) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92))) | ~ ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((~ ! [X114] : (~ ! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | p1(X114) | ~ r1(X113, X114)) | ! [X117] : (! [X118] : (~ ! [X119] : (~ p1(X119) | ! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117))) & ! [X121] : (~ ! [X122] : (p1(X122) | ~ r1(X121, X122)) | ! [X123] : (! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X121, X123)) | ~ r1(X113, X121)) & (~ ! [X125] : (~ ! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | p1(X125) | ~ r1(X113, X125)) | ! [X128] : (~ ! [X129] : (p1(X129) | ~ r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113)) & (~ ! [X130] : (~ ! [X131] : (~ (~ p1(X131) | ! [X132] : (p1(X132) | ~ r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) | ~ p1(X130) | ! [X135] : (p1(X135) | ~ r1(X130, X135)) | ~ r1(X113, X130)) | ! [X136] : (~ ! [X137] : (~ p1(X137) | ! [X138] : (p1(X138) | ~ r1(X137, X138)) | ~ r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)))) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) | ~ ! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (~ ! [X150] : (p1(X150) | ~ r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (~ ! [X153] : (~ ! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) | p1(X153) | ~ r1(X148, X153)) | ! [X156] : (~ ! [X157] : (p1(X157) | ~ r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X0, X140)) | ~ ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : ((! [X168] : (~ ! [X169] : (p1(X169) | ~ r1(X168, X169)) | ! [X170] : (! [X171] : (p1(X171) | ~ r1(X170, X171)) | ~ r1(X168, X170)) | ~ r1(X167, X168)) & (~ ! [X172] : (~ ! [X173] : (~ p1(X173) | ! [X174] : (p1(X174) | ~ r1(X173, X174)) | ~ r1(X172, X173)) | p1(X172) | ~ r1(X167, X172)) | ! [X175] : (~ ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X167, X175)) | p1(X167))) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) | ~ ! [X177] : (! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : ((! [X188] : (~ ! [X189] : (p1(X189) | ~ r1(X188, X189)) | ! [X190] : (! [X191] : (p1(X191) | ~ r1(X190, X191)) | ~ r1(X188, X190)) | ~ r1(X187, X188)) & (~ ! [X192] : (~ ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X192, X193)) | p1(X192) | ~ r1(X187, X192)) | ! [X195] : (~ ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X187, X195)) | p1(X187))) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X177, X178)) | ~ r1(X0, X177)) | ~ ! [X197] : (! [X198] : (! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : ((! [X209] : (~ ! [X210] : (p1(X210) | ~ r1(X209, X210)) | ! [X211] : (! [X212] : (p1(X212) | ~ r1(X211, X212)) | ~ r1(X209, X211)) | ~ r1(X208, X209)) & (~ ! [X213] : (~ ! [X214] : (~ p1(X214) | ! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X213, X214)) | p1(X213) | ~ r1(X208, X213)) | ! [X216] : (~ ! [X217] : (p1(X217) | ~ r1(X216, X217)) | ~ r1(X208, X216)) | p1(X208))) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X198, X199)) | ~ r1(X197, X198)) | ~ r1(X0, X197)) | ~ ! [X218] : (! [X219] : (! [X220] : (! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : ((! [X231] : (~ ! [X232] : (p1(X232) | ~ r1(X231, X232)) | ! [X233] : (! [X234] : (p1(X234) | ~ r1(X233, X234)) | ~ r1(X231, X233)) | ~ r1(X230, X231)) & (~ ! [X235] : (~ ! [X236] : (~ p1(X236) | ! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | p1(X235) | ~ r1(X230, X235)) | ! [X238] : (~ ! [X239] : (p1(X239) | ~ r1(X238, X239)) | ~ r1(X230, X238)) | p1(X230))) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X220, X221)) | ~ r1(X219, X220)) | ~ r1(X218, X219)) | ~ r1(X0, X218)) | ~ ! [X240] : (! [X241] : (! [X242] : (! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : ((! [X254] : (~ ! [X255] : (p1(X255) | ~ r1(X254, X255)) | ! [X256] : (! [X257] : (p1(X257) | ~ r1(X256, X257)) | ~ r1(X254, X256)) | ~ r1(X253, X254)) & (~ ! [X258] : (~ ! [X259] : (~ p1(X259) | ! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | p1(X258) | ~ r1(X253, X258)) | ! [X261] : (~ ! [X262] : (p1(X262) | ~ r1(X261, X262)) | ~ r1(X253, X261)) | p1(X253))) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X242, X243)) | ~ r1(X241, X242)) | ~ r1(X240, X241)) | ~ r1(X0, X240))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X2] : (~ ! [X3] : (p1(X3) | ~ r1(X2, X3)) | ! [X4] : (! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & (~ ! [X6] : (~ ! [X7] : (~ p1(X7) | ! [X8] : (p1(X8) | ~ r1(X7, X8)) | ~ r1(X6, X7)) | p1(X6) | ~ r1(X1, X6)) | ! [X9] : (~ ! [X10] : (p1(X10) | ~ r1(X9, X10)) | ~ r1(X1, X9)) | p1(X1))) | ~ r1(X0, X1)) & ! [X11] : (! [X12] : ((! [X13] : (~ ! [X14] : (p1(X14) | ~ r1(X13, X14)) | ! [X15] : (! [X16] : (p1(X16) | ~ r1(X15, X16)) | ~ r1(X13, X15)) | ~ r1(X12, X13)) & (~ ! [X17] : (~ ! [X18] : (~ p1(X18) | ! [X19] : (p1(X19) | ~ r1(X18, X19)) | ~ r1(X17, X18)) | p1(X17) | ~ r1(X12, X17)) | ! [X20] : (~ ! [X21] : (p1(X21) | ~ r1(X20, X21)) | ~ r1(X12, X20)) | p1(X12))) | ~ r1(X11, X12)) | ~ r1(X0, X11)) & ! [X22] : (! [X23] : (! [X24] : ((! [X25] : (~ ! [X26] : (p1(X26) | ~ r1(X25, X26)) | ! [X27] : (! [X28] : (p1(X28) | ~ r1(X27, X28)) | ~ r1(X25, X27)) | ~ r1(X24, X25)) & (~ ! [X29] : (~ ! [X30] : (~ p1(X30) | ! [X31] : (p1(X31) | ~ r1(X30, X31)) | ~ r1(X29, X30)) | p1(X29) | ~ r1(X24, X29)) | ! [X32] : (~ ! [X33] : (p1(X33) | ~ r1(X32, X33)) | ~ r1(X24, X32)) | p1(X24))) | ~ r1(X23, X24)) | ~ r1(X22, X23)) | ~ r1(X0, X22)) & ! [X34] : (! [X35] : (! [X36] : (! [X37] : ((! [X38] : (~ ! [X39] : (p1(X39) | ~ r1(X38, X39)) | ! [X40] : (! [X41] : (p1(X41) | ~ r1(X40, X41)) | ~ r1(X38, X40)) | ~ r1(X37, X38)) & (~ ! [X42] : (~ ! [X43] : (~ p1(X43) | ! [X44] : (p1(X44) | ~ r1(X43, X44)) | ~ r1(X42, X43)) | p1(X42) | ~ r1(X37, X42)) | ! [X45] : (~ ! [X46] : (p1(X46) | ~ r1(X45, X46)) | ~ r1(X37, X45)) | p1(X37))) | ~ r1(X36, X37)) | ~ r1(X35, X36)) | ~ r1(X34, X35)) | ~ r1(X0, X34)) & ! [X47] : (! [X48] : (! [X49] : (! [X50] : (! [X51] : ((! [X52] : (~ ! [X53] : (p1(X53) | ~ r1(X52, X53)) | ! [X54] : (! [X55] : (p1(X55) | ~ r1(X54, X55)) | ~ r1(X52, X54)) | ~ r1(X51, X52)) & (~ ! [X56] : (~ ! [X57] : (~ p1(X57) | ! [X58] : (p1(X58) | ~ r1(X57, X58)) | ~ r1(X56, X57)) | p1(X56) | ~ r1(X51, X56)) | ! [X59] : (~ ! [X60] : (p1(X60) | ~ r1(X59, X60)) | ~ r1(X51, X59)) | p1(X51))) | ~ r1(X50, X51)) | ~ r1(X49, X50)) | ~ r1(X48, X49)) | ~ r1(X47, X48)) | ~ r1(X0, X47)) & ! [X61] : (! [X62] : (! [X63] : (! [X64] : (! [X65] : (! [X66] : ((! [X67] : (~ ! [X68] : (p1(X68) | ~ r1(X67, X68)) | ! [X69] : (! [X70] : (p1(X70) | ~ r1(X69, X70)) | ~ r1(X67, X69)) | ~ r1(X66, X67)) & (~ ! [X71] : (~ ! [X72] : (~ p1(X72) | ! [X73] : (p1(X73) | ~ r1(X72, X73)) | ~ r1(X71, X72)) | p1(X71) | ~ r1(X66, X71)) | ! [X74] : (~ ! [X75] : (p1(X75) | ~ r1(X74, X75)) | ~ r1(X66, X74)) | p1(X66))) | ~ r1(X65, X66)) | ~ r1(X64, X65)) | ~ r1(X63, X64)) | ~ r1(X62, X63)) | ~ r1(X61, X62)) | ~ r1(X0, X61)) & ! [X76] : (! [X77] : (! [X78] : (! [X79] : (! [X80] : (! [X81] : (! [X82] : ((! [X83] : (~ ! [X84] : (p1(X84) | ~ r1(X83, X84)) | ! [X85] : (! [X86] : (p1(X86) | ~ r1(X85, X86)) | ~ r1(X83, X85)) | ~ r1(X82, X83)) & (~ ! [X87] : (~ ! [X88] : (~ p1(X88) | ! [X89] : (p1(X89) | ~ r1(X88, X89)) | ~ r1(X87, X88)) | p1(X87) | ~ r1(X82, X87)) | ! [X90] : (~ ! [X91] : (p1(X91) | ~ r1(X90, X91)) | ~ r1(X82, X90)) | p1(X82))) | ~ r1(X81, X82)) | ~ r1(X80, X81)) | ~ r1(X79, X80)) | ~ r1(X78, X79)) | ~ r1(X77, X78)) | ~ r1(X76, X77)) | ~ r1(X0, X76)) & ~ ! [X92] : (! [X93] : (! [X94] : (! [X95] : (! [X96] : (! [X97] : (! [X98] : (! [X99] : (~ ! [X100] : (~ ! [X101] : (~ p1(X101) | ! [X102] : (p1(X102) | ~ r1(X101, X102)) | ~ r1(X100, X101)) | p1(X100) | ~ r1(X99, X100)) | ! [X103] : (~ ! [X104] : (p1(X104) | ~ r1(X103, X104)) | ~ r1(X99, X103)) | ! [X105] : (p1(X105) | ~ r1(X99, X105)) | ~ r1(X98, X99)) | ~ r1(X97, X98)) | ~ r1(X96, X97)) | ~ r1(X95, X96)) | ~ r1(X94, X95)) | ~ r1(X93, X94)) | ~ r1(X92, X93)) | ~ r1(X0, X92))) | ~ ! [X106] : (! [X107] : (! [X108] : (! [X109] : (! [X110] : (! [X111] : (! [X112] : (! [X113] : (((~ ! [X114] : (~ ! [X115] : (~ p1(X115) | ! [X116] : (p1(X116) | ~ r1(X115, X116)) | ~ r1(X114, X115)) | p1(X114) | ~ r1(X113, X114)) | ! [X117] : (! [X118] : (~ ! [X119] : (~ p1(X119) | ! [X120] : (p1(X120) | ~ r1(X119, X120)) | ~ r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117))) & ! [X121] : (~ ! [X122] : (p1(X122) | ~ r1(X121, X122)) | ! [X123] : (! [X124] : (p1(X124) | ~ r1(X123, X124)) | ~ r1(X121, X123)) | ~ r1(X113, X121)) & (~ ! [X125] : (~ ! [X126] : (~ p1(X126) | ! [X127] : (p1(X127) | ~ r1(X126, X127)) | ~ r1(X125, X126)) | p1(X125) | ~ r1(X113, X125)) | ! [X128] : (~ ! [X129] : (p1(X129) | ~ r1(X128, X129)) | ~ r1(X113, X128)) | p1(X113)) & (~ ! [X130] : (~ ! [X131] : (~ (~ p1(X131) | ! [X132] : (p1(X132) | ~ r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) | ~ p1(X130) | ! [X135] : (p1(X135) | ~ r1(X130, X135)) | ~ r1(X113, X130)) | ! [X136] : (~ ! [X137] : (~ p1(X137) | ! [X138] : (p1(X138) | ~ r1(X137, X138)) | ~ r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)))) | ~ r1(X112, X113)) | ~ r1(X111, X112)) | ~ r1(X110, X111)) | ~ r1(X109, X110)) | ~ r1(X108, X109)) | ~ r1(X107, X108)) | ~ r1(X106, X107)) | ~ r1(X0, X106)) | ~ ! [X140] : (! [X141] : (! [X142] : (! [X143] : (! [X144] : (! [X145] : (! [X146] : (! [X147] : (! [X148] : ((! [X149] : (~ ! [X150] : (p1(X150) | ~ r1(X149, X150)) | ! [X151] : (! [X152] : (p1(X152) | ~ r1(X151, X152)) | ~ r1(X149, X151)) | ~ r1(X148, X149)) & (~ ! [X153] : (~ ! [X154] : (~ p1(X154) | ! [X155] : (p1(X155) | ~ r1(X154, X155)) | ~ r1(X153, X154)) | p1(X153) | ~ r1(X148, X153)) | ! [X156] : (~ ! [X157] : (p1(X157) | ~ r1(X156, X157)) | ~ r1(X148, X156)) | p1(X148))) | ~ r1(X147, X148)) | ~ r1(X146, X147)) | ~ r1(X145, X146)) | ~ r1(X144, X145)) | ~ r1(X143, X144)) | ~ r1(X142, X143)) | ~ r1(X141, X142)) | ~ r1(X140, X141)) | ~ r1(X0, X140)) | ~ ! [X158] : (! [X159] : (! [X160] : (! [X161] : (! [X162] : (! [X163] : (! [X164] : (! [X165] : (! [X166] : (! [X167] : ((! [X168] : (~ ! [X169] : (p1(X169) | ~ r1(X168, X169)) | ! [X170] : (! [X171] : (p1(X171) | ~ r1(X170, X171)) | ~ r1(X168, X170)) | ~ r1(X167, X168)) & (~ ! [X172] : (~ ! [X173] : (~ p1(X173) | ! [X174] : (p1(X174) | ~ r1(X173, X174)) | ~ r1(X172, X173)) | p1(X172) | ~ r1(X167, X172)) | ! [X175] : (~ ! [X176] : (p1(X176) | ~ r1(X175, X176)) | ~ r1(X167, X175)) | p1(X167))) | ~ r1(X166, X167)) | ~ r1(X165, X166)) | ~ r1(X164, X165)) | ~ r1(X163, X164)) | ~ r1(X162, X163)) | ~ r1(X161, X162)) | ~ r1(X160, X161)) | ~ r1(X159, X160)) | ~ r1(X158, X159)) | ~ r1(X0, X158)) | ~ ! [X177] : (! [X178] : (! [X179] : (! [X180] : (! [X181] : (! [X182] : (! [X183] : (! [X184] : (! [X185] : (! [X186] : (! [X187] : ((! [X188] : (~ ! [X189] : (p1(X189) | ~ r1(X188, X189)) | ! [X190] : (! [X191] : (p1(X191) | ~ r1(X190, X191)) | ~ r1(X188, X190)) | ~ r1(X187, X188)) & (~ ! [X192] : (~ ! [X193] : (~ p1(X193) | ! [X194] : (p1(X194) | ~ r1(X193, X194)) | ~ r1(X192, X193)) | p1(X192) | ~ r1(X187, X192)) | ! [X195] : (~ ! [X196] : (p1(X196) | ~ r1(X195, X196)) | ~ r1(X187, X195)) | p1(X187))) | ~ r1(X186, X187)) | ~ r1(X185, X186)) | ~ r1(X184, X185)) | ~ r1(X183, X184)) | ~ r1(X182, X183)) | ~ r1(X181, X182)) | ~ r1(X180, X181)) | ~ r1(X179, X180)) | ~ r1(X178, X179)) | ~ r1(X177, X178)) | ~ r1(X0, X177)) | ~ ! [X197] : (! [X198] : (! [X199] : (! [X200] : (! [X201] : (! [X202] : (! [X203] : (! [X204] : (! [X205] : (! [X206] : (! [X207] : (! [X208] : ((! [X209] : (~ ! [X210] : (p1(X210) | ~ r1(X209, X210)) | ! [X211] : (! [X212] : (p1(X212) | ~ r1(X211, X212)) | ~ r1(X209, X211)) | ~ r1(X208, X209)) & (~ ! [X213] : (~ ! [X214] : (~ p1(X214) | ! [X215] : (p1(X215) | ~ r1(X214, X215)) | ~ r1(X213, X214)) | p1(X213) | ~ r1(X208, X213)) | ! [X216] : (~ ! [X217] : (p1(X217) | ~ r1(X216, X217)) | ~ r1(X208, X216)) | p1(X208))) | ~ r1(X207, X208)) | ~ r1(X206, X207)) | ~ r1(X205, X206)) | ~ r1(X204, X205)) | ~ r1(X203, X204)) | ~ r1(X202, X203)) | ~ r1(X201, X202)) | ~ r1(X200, X201)) | ~ r1(X199, X200)) | ~ r1(X198, X199)) | ~ r1(X197, X198)) | ~ r1(X0, X197)) | ~ ! [X218] : (! [X219] : (! [X220] : (! [X221] : (! [X222] : (! [X223] : (! [X224] : (! [X225] : (! [X226] : (! [X227] : (! [X228] : (! [X229] : (! [X230] : ((! [X231] : (~ ! [X232] : (p1(X232) | ~ r1(X231, X232)) | ! [X233] : (! [X234] : (p1(X234) | ~ r1(X233, X234)) | ~ r1(X231, X233)) | ~ r1(X230, X231)) & (~ ! [X235] : (~ ! [X236] : (~ p1(X236) | ! [X237] : (p1(X237) | ~ r1(X236, X237)) | ~ r1(X235, X236)) | p1(X235) | ~ r1(X230, X235)) | ! [X238] : (~ ! [X239] : (p1(X239) | ~ r1(X238, X239)) | ~ r1(X230, X238)) | p1(X230))) | ~ r1(X229, X230)) | ~ r1(X228, X229)) | ~ r1(X227, X228)) | ~ r1(X226, X227)) | ~ r1(X225, X226)) | ~ r1(X224, X225)) | ~ r1(X223, X224)) | ~ r1(X222, X223)) | ~ r1(X221, X222)) | ~ r1(X220, X221)) | ~ r1(X219, X220)) | ~ r1(X218, X219)) | ~ r1(X0, X218)) | ~ ! [X240] : (! [X241] : (! [X242] : (! [X243] : (! [X244] : (! [X245] : (! [X246] : (! [X247] : (! [X248] : (! [X249] : (! [X250] : (! [X251] : (! [X252] : (! [X253] : ((! [X254] : (~ ! [X255] : (p1(X255) | ~ r1(X254, X255)) | ! [X256] : (! [X257] : (p1(X257) | ~ r1(X256, X257)) | ~ r1(X254, X256)) | ~ r1(X253, X254)) & (~ ! [X258] : (~ ! [X259] : (~ p1(X259) | ! [X260] : (p1(X260) | ~ r1(X259, X260)) | ~ r1(X258, X259)) | p1(X258) | ~ r1(X253, X258)) | ! [X261] : (~ ! [X262] : (p1(X262) | ~ r1(X261, X262)) | ~ r1(X253, X261)) | p1(X253))) | ~ r1(X252, X253)) | ~ r1(X251, X252)) | ~ r1(X250, X251)) | ~ r1(X249, X250)) | ~ r1(X248, X249)) | ~ r1(X247, X248)) | ~ r1(X246, X247)) | ~ r1(X245, X246)) | ~ r1(X244, X245)) | ~ r1(X243, X244)) | ~ r1(X242, X243)) | ~ r1(X241, X242)) | ~ r1(X240, X241)) | ~ r1(X0, X240))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) & ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1))) | ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ? [X0] : ~ (~ (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (((~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1))) & ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0)) & (~ ! [X1] : (~ ! [X0] : (~ (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1))) | ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : ((! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ! [X1] : (! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) & (~ ! [X0] : (~ ! [X1] : (~ p1(X1) | ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0) | ~ r1(X1, X0)) | ! [X0] : (~ ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1))) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ ! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : (! [X1] : (! [X0] : ((! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ! [X0] : (! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) & (~ ! [X1] : (~ ! [X0] : (~ p1(X0) | ! [X1] : (p1(X1) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | p1(X1) | ~ r1(X0, X1)) | ! [X1] : (~ ! [X0] : (p1(X0) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | p1(X0))) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1)) | ~ r1(X1, X0)) | ~ r1(X0, X1))), file('/home/ubuntu/library/tptp/Problems/LCL/LCL640+1.015.p', main)).
fof(f115911, plain, (~ r1(sK42, sK45) | (spl67_57 | ~ spl67_6221)), inference(resolution, [], [f115855, f531])).
fof(f531, plain, (~ r1(sK45, sK7(sK45)) | spl67_57), inference(avatar_component_clause, [], [f530])).
fof(f530, plain, (spl67_57 <=> r1(sK45, sK7(sK45))), introduced(avatar_definition, [new_symbols(naming, [spl67_57])])).
fof(f115855, plain, (! [X2] : (r1(X2, sK7(X2)) | ~ r1(sK42, X2)) | ~ spl67_6221), inference(avatar_component_clause, [], [f115854])).
fof(f115895, plain, (spl67_242 | spl67_243 | ~ spl67_55), inference(avatar_split_clause, [], [f115802, f521, f2773, f2769])).
fof(f2769, plain, (spl67_242 <=> p1(sK43(sK6(sK42)))), introduced(avatar_definition, [new_symbols(naming, [spl67_242])])).
fof(f521, plain, (spl67_55 <=> r1(sK42, sK6(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl67_55])])).
fof(f115802, plain, (p1(sK6(sK42)) | p1(sK43(sK6(sK42))) | ~ spl67_55), inference(resolution, [], [f523, f180])).
fof(f180, plain, ! [X100] : (~ r1(sK42, X100) | p1(X100) | p1(sK43(X100))), inference(cnf_transformation, [], [f88])).
fof(f523, plain, (r1(sK42, sK6(sK42)) | ~ spl67_55), inference(avatar_component_clause, [], [f521])).
fof(f115877, plain, (~ spl67_57 | ~ spl67_6223), inference(avatar_contradiction_clause, [], [f115876])).
fof(f115876, plain, ($false | (~ spl67_57 | ~ spl67_6223)), inference(subsumption_resolution, [], [f115872, f175])).
fof(f115872, plain, (~ r1(sK42, sK45) | (~ spl67_57 | ~ spl67_6223)), inference(resolution, [], [f115866, f2820])).
fof(f2820, plain, (p1(sK7(sK45)) | ~ spl67_57), inference(resolution, [], [f532, f176])).
fof(f176, plain, ! [X104] : (~ r1(sK45, X104) | p1(X104)), inference(cnf_transformation, [], [f88])).
fof(f532, plain, (r1(sK45, sK7(sK45)) | ~ spl67_57), inference(avatar_component_clause, [], [f530])).
fof(f115866, plain, (! [X3] : (~ p1(sK7(X3)) | ~ r1(sK42, X3)) | ~ spl67_6223), inference(avatar_component_clause, [], [f115865])).
fof(f115865, plain, (spl67_6223 <=> ! [X3] : (~ p1(sK7(X3)) | ~ r1(sK42, X3))), introduced(avatar_definition, [new_symbols(naming, [spl67_6223])])).
fof(f115871, plain, (~ spl67_55 | spl67_243 | ~ spl67_6222), inference(avatar_contradiction_clause, [], [f115870])).
fof(f115870, plain, ($false | (~ spl67_55 | spl67_243 | ~ spl67_6222)), inference(subsumption_resolution, [], [f115869, f523])).
fof(f115869, plain, (~ r1(sK42, sK6(sK42)) | (spl67_243 | ~ spl67_6222)), inference(subsumption_resolution, [], [f115868, f2774])).
fof(f2774, plain, (~ p1(sK6(sK42)) | spl67_243), inference(avatar_component_clause, [], [f2773])).
fof(f115868, plain, (p1(sK6(sK42)) | ~ r1(sK42, sK6(sK42)) | ~ spl67_6222), inference(resolution, [], [f115859, f179])).
fof(f179, plain, ! [X100] : (~ p1(sK44(X100)) | p1(X100) | ~ r1(sK42, X100)), inference(cnf_transformation, [], [f88])).
fof(f115859, plain, (p1(sK44(sK6(sK42))) | ~ spl67_6222), inference(avatar_component_clause, [], [f115857])).
fof(f115857, plain, (spl67_6222 <=> p1(sK44(sK6(sK42)))), introduced(avatar_definition, [new_symbols(naming, [spl67_6222])])).
fof(f115867, plain, (spl67_6223 | spl67_6222 | spl67_26 | ~ spl67_54 | ~ spl67_55 | ~ spl67_242), inference(avatar_split_clause, [], [f115863, f2769, f521, f517, f360, f115857, f115865])).
fof(f115863, plain, (! [X3] : (p1(sK44(sK6(sK42))) | ~ p1(sK7(X3)) | ~ r1(sK42, X3)) | (spl67_26 | ~ spl67_54 | ~ spl67_55 | ~ spl67_242)), inference(subsumption_resolution, [], [f115862, f523])).
fof(f115862, plain, (! [X3] : (p1(sK44(sK6(sK42))) | ~ p1(sK7(X3)) | ~ r1(sK42, X3) | ~ r1(sK42, sK6(sK42))) | (spl67_26 | ~ spl67_54 | ~ spl67_242)), inference(subsumption_resolution, [], [f115861, f518])).
fof(f115861, plain, (! [X3] : (p1(sK44(sK6(sK42))) | ~ p1(sK7(X3)) | ~ r1(sK42, X3) | ~ sP2(sK42) | ~ r1(sK42, sK6(sK42))) | (spl67_26 | ~ spl67_242)), inference(subsumption_resolution, [], [f115849, f361])).
fof(f115849, plain, (! [X3] : (p1(sK44(sK6(sK42))) | ~ p1(sK7(X3)) | ~ r1(sK42, X3) | p1(sK42) | ~ sP2(sK42) | ~ r1(sK42, sK6(sK42))) | ~ spl67_242), inference(resolution, [], [f2771, f2091])).
fof(f2091, plain, ! [X0, X1] : (~ p1(sK43(sK6(X0))) | p1(sK44(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0))), inference(subsumption_resolution, [], [f2090, f96])).
fof(f96, plain, ! [X4, X0] : (~ p1(sK7(X4)) | ~ p1(sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f2090, plain, ! [X0, X1] : (p1(sK44(sK6(X0))) | ~ p1(sK43(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f2083])).
fof(f2083, plain, ! [X0, X1] : (p1(sK44(sK6(X0))) | ~ p1(sK43(sK6(X0))) | ~ p1(sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK42, sK6(X0))), inference(resolution, [], [f754, f178])).
fof(f178, plain, ! [X100] : (r1(sK43(X100), sK44(X100)) | p1(X100) | ~ r1(sK42, X100)), inference(cnf_transformation, [], [f88])).
fof(f754, plain, ! [X4, X5, X3] : (~ r1(sK43(sK6(X4)), X3) | p1(X3) | ~ p1(sK43(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK42, sK6(X4))), inference(subsumption_resolution, [], [f752, f96])).
fof(f752, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK43(sK6(X4)), X3) | ~ p1(sK43(sK6(X4))) | ~ p1(sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK42, sK6(X4))), inference(resolution, [], [f98, f177])).
fof(f177, plain, ! [X100] : (r1(X100, sK43(X100)) | p1(X100) | ~ r1(sK42, X100)), inference(cnf_transformation, [], [f88])).
fof(f98, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | ~ p1(sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f2771, plain, (p1(sK43(sK6(sK42))) | ~ spl67_242), inference(avatar_component_clause, [], [f2769])).
fof(f115860, plain, (spl67_6221 | spl67_6222 | spl67_26 | ~ spl67_54 | ~ spl67_55 | ~ spl67_242), inference(avatar_split_clause, [], [f115852, f2769, f521, f517, f360, f115857, f115854])).
fof(f115852, plain, (! [X2] : (p1(sK44(sK6(sK42))) | r1(X2, sK7(X2)) | ~ r1(sK42, X2)) | (spl67_26 | ~ spl67_54 | ~ spl67_55 | ~ spl67_242)), inference(subsumption_resolution, [], [f115851, f523])).
fof(f115851, plain, (! [X2] : (p1(sK44(sK6(sK42))) | r1(X2, sK7(X2)) | ~ r1(sK42, X2) | ~ r1(sK42, sK6(sK42))) | (spl67_26 | ~ spl67_54 | ~ spl67_242)), inference(subsumption_resolution, [], [f115850, f518])).
fof(f115850, plain, (! [X2] : (p1(sK44(sK6(sK42))) | r1(X2, sK7(X2)) | ~ r1(sK42, X2) | ~ sP2(sK42) | ~ r1(sK42, sK6(sK42))) | (spl67_26 | ~ spl67_242)), inference(subsumption_resolution, [], [f115848, f361])).
fof(f115848, plain, (! [X2] : (p1(sK44(sK6(sK42))) | r1(X2, sK7(X2)) | ~ r1(sK42, X2) | p1(sK42) | ~ sP2(sK42) | ~ r1(sK42, sK6(sK42))) | ~ spl67_242), inference(resolution, [], [f2771, f2173])).
fof(f2173, plain, ! [X0, X1] : (~ p1(sK43(sK6(X0))) | p1(sK44(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0))), inference(subsumption_resolution, [], [f2172, f95])).
fof(f2172, plain, ! [X0, X1] : (p1(sK44(sK6(X0))) | ~ p1(sK43(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0)) | p1(sK6(X0))), inference(duplicate_literal_removal, [], [f2165])).
fof(f2165, plain, ! [X0, X1] : (p1(sK44(sK6(X0))) | ~ p1(sK43(sK6(X0))) | r1(X1, sK7(X1)) | ~ r1(X0, X1) | p1(X0) | ~ sP2(X0) | ~ r1(sK42, sK6(X0)) | p1(sK6(X0)) | ~ r1(sK42, sK6(X0))), inference(resolution, [], [f767, f178])).
fof(f767, plain, ! [X4, X5, X3] : (~ r1(sK43(sK6(X4)), X3) | p1(X3) | ~ p1(sK43(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | ~ r1(sK42, sK6(X4))), inference(subsumption_resolution, [], [f765, f95])).
fof(f765, plain, ! [X4, X5, X3] : (p1(X3) | ~ r1(sK43(sK6(X4)), X3) | ~ p1(sK43(sK6(X4))) | r1(X5, sK7(X5)) | ~ r1(X4, X5) | p1(X4) | ~ sP2(X4) | p1(sK6(X4)) | ~ r1(sK42, sK6(X4))), inference(resolution, [], [f97, f177])).
fof(f97, plain, ! [X4, X2, X0, X3] : (~ r1(sK6(X0), X2) | p1(X3) | ~ r1(X2, X3) | ~ p1(X2) | r1(X4, sK7(X4)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f115829, plain, (spl67_26 | ~ spl67_54 | ~ spl67_57 | ~ spl67_243), inference(avatar_contradiction_clause, [], [f115828])).
fof(f115828, plain, ($false | (spl67_26 | ~ spl67_54 | ~ spl67_57 | ~ spl67_243)), inference(subsumption_resolution, [], [f115827, f518])).
fof(f115827, plain, (~ sP2(sK42) | (spl67_26 | ~ spl67_57 | ~ spl67_243)), inference(subsumption_resolution, [], [f115826, f361])).
fof(f115826, plain, (p1(sK42) | ~ sP2(sK42) | (~ spl67_57 | ~ spl67_243)), inference(subsumption_resolution, [], [f115824, f175])).
fof(f115824, plain, (~ r1(sK42, sK45) | p1(sK42) | ~ sP2(sK42) | (~ spl67_57 | ~ spl67_243)), inference(resolution, [], [f2775, f2838])).
fof(f2838, plain, (! [X1] : (~ p1(sK6(X1)) | ~ r1(X1, sK45) | p1(X1) | ~ sP2(X1)) | ~ spl67_57), inference(resolution, [], [f2820, f96])).
fof(f115598, plain, ~ spl67_5148, inference(avatar_contradiction_clause, [], [f115597])).
fof(f115597, plain, ($false | ~ spl67_5148), inference(subsumption_resolution, [], [f115583, f172])).
fof(f172, plain, r1(sK41, sK42), inference(cnf_transformation, [], [f88])).
fof(f115583, plain, (~ r1(sK41, sK42) | ~ spl67_5148), inference(resolution, [], [f115569, f171])).
fof(f171, plain, r1(sK40, sK41), inference(cnf_transformation, [], [f88])).
fof(f115569, plain, (! [X0] : (~ r1(sK40, X0) | ~ r1(X0, sK42)) | ~ spl67_5148), inference(resolution, [], [f115542, f170])).
fof(f170, plain, r1(sK39, sK40), inference(cnf_transformation, [], [f88])).
fof(f115542, plain, (! [X0, X1] : (~ r1(sK39, X1) | ~ r1(X0, sK42) | ~ r1(X1, X0)) | ~ spl67_5148), inference(resolution, [], [f115497, f169])).
fof(f169, plain, r1(sK38, sK39), inference(cnf_transformation, [], [f88])).
fof(f115497, plain, (! [X2, X0, X1] : (~ r1(sK38, X1) | ~ r1(X0, sK42) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl67_5148), inference(resolution, [], [f115466, f168])).
fof(f168, plain, r1(sK37, sK38), inference(cnf_transformation, [], [f88])).
fof(f115466, plain, (! [X2, X0, X3, X1] : (~ r1(sK37, X1) | ~ r1(X0, sK42) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl67_5148), inference(resolution, [], [f115422, f167])).
fof(f167, plain, r1(sK36, sK37), inference(cnf_transformation, [], [f88])).
fof(f115422, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK36, X1) | ~ r1(X0, sK42) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X0)) | ~ spl67_5148), inference(resolution, [], [f115320, f166])).
fof(f166, plain, r1(sK35, sK36), inference(cnf_transformation, [], [f88])).
fof(f115320, plain, (! [X4, X2, X0, X5, X3, X1] : (~ r1(sK35, X1) | ~ r1(X0, sK42) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X0)) | ~ spl67_5148), inference(resolution, [], [f85088, f165])).
fof(f165, plain, r1(sK13, sK35), inference(cnf_transformation, [], [f88])).
fof(f85088, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK13, X6) | ~ r1(X0, sK42) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0)) | ~ spl67_5148), inference(avatar_component_clause, [], [f85087])).
fof(f85087, plain, (spl67_5148 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ r1(X0, sK42) | ~ r1(sK13, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2) | ~ r1(X2, X1) | ~ r1(X1, X0))), introduced(avatar_definition, [new_symbols(naming, [spl67_5148])])).
fof(f110018, plain, (~ spl67_234 | ~ spl67_235 | ~ spl67_1624 | spl67_2532), inference(avatar_contradiction_clause, [], [f110017])).
fof(f110017, plain, ($false | (~ spl67_234 | ~ spl67_235 | ~ spl67_1624 | spl67_2532)), inference(subsumption_resolution, [], [f110016, f175])).
fof(f110016, plain, (~ r1(sK42, sK45) | (~ spl67_234 | ~ spl67_235 | ~ spl67_1624 | spl67_2532)), inference(subsumption_resolution, [], [f109995, f37733])).
fof(f37733, plain, (~ p1(sK5(sK45)) | spl67_2532), inference(avatar_component_clause, [], [f37732])).
fof(f37732, plain, (spl67_2532 <=> p1(sK5(sK45))), introduced(avatar_definition, [new_symbols(naming, [spl67_2532])])).
fof(f109995, plain, (p1(sK5(sK45)) | ~ r1(sK42, sK45) | (~ spl67_234 | ~ spl67_235 | ~ spl67_1624)), inference(resolution, [], [f109928, f2529])).
fof(f2529, plain, (! [X14] : (r1(sK4(X14), sK5(X14)) | ~ r1(sK42, X14)) | ~ spl67_234), inference(avatar_component_clause, [], [f2528])).
fof(f2528, plain, (spl67_234 <=> ! [X14] : (r1(sK4(X14), sK5(X14)) | ~ r1(sK42, X14))), introduced(avatar_definition, [new_symbols(naming, [spl67_234])])).
fof(f109928, plain, (! [X2] : (~ r1(sK4(sK45), X2) | p1(X2)) | (~ spl67_235 | ~ spl67_1624)), inference(subsumption_resolution, [], [f109909, f175])).
fof(f109909, plain, (! [X2] : (p1(X2) | ~ r1(sK4(sK45), X2) | ~ r1(sK42, sK45)) | (~ spl67_235 | ~ spl67_1624)), inference(resolution, [], [f26271, f2534])).
fof(f2534, plain, (! [X14] : (r1(X14, sK4(X14)) | ~ r1(sK42, X14)) | ~ spl67_235), inference(avatar_component_clause, [], [f2533])).
fof(f2533, plain, (spl67_235 <=> ! [X14] : (r1(X14, sK4(X14)) | ~ r1(sK42, X14))), introduced(avatar_definition, [new_symbols(naming, [spl67_235])])).
fof(f26271, plain, (! [X2, X3] : (~ r1(sK45, X2) | p1(X3) | ~ r1(X2, X3)) | ~ spl67_1624), inference(avatar_component_clause, [], [f26270])).
fof(f26270, plain, (spl67_1624 <=> ! [X3, X2] : (~ r1(sK45, X2) | p1(X3) | ~ r1(X2, X3))), introduced(avatar_definition, [new_symbols(naming, [spl67_1624])])).
fof(f109880, plain, ~ spl67_1627, inference(avatar_contradiction_clause, [], [f109879])).
fof(f109879, plain, ($false | ~ spl67_1627), inference(subsumption_resolution, [], [f109858, f175])).
fof(f109858, plain, (~ r1(sK42, sK45) | ~ spl67_1627), inference(resolution, [], [f109816, f172])).
fof(f109816, plain, (! [X0] : (~ r1(sK41, X0) | ~ r1(X0, sK45)) | ~ spl67_1627), inference(resolution, [], [f109795, f171])).
fof(f109795, plain, (! [X0, X1] : (~ r1(sK40, X1) | ~ r1(X0, sK45) | ~ r1(X1, X0)) | ~ spl67_1627), inference(resolution, [], [f109755, f170])).
fof(f109755, plain, (! [X2, X0, X1] : (~ r1(sK39, X1) | ~ r1(X0, sK45) | ~ r1(X1, X2) | ~ r1(X2, X0)) | ~ spl67_1627), inference(resolution, [], [f109711, f169])).
fof(f109711, plain, (! [X2, X0, X3, X1] : (~ r1(sK38, X1) | ~ r1(X0, sK45) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X0)) | ~ spl67_1627), inference(resolution, [], [f109670, f168])).
fof(f109670, plain, (! [X4, X2, X0, X3, X1] : (~ r1(sK37, X1) | ~ r1(X0, sK45) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X0)) | ~ spl67_1627), inference(resolution, [], [f109629, f167])).
fof(f109629, plain, (! [X4, X2, X0, X5, X3, X1] : (~ r1(sK36, X1) | ~ r1(X0, sK45) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X0)) | ~ spl67_1627), inference(resolution, [], [f109588, f166])).
fof(f109588, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK35, X1) | ~ r1(X0, sK45) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | ~ r1(X5, X6) | ~ r1(X6, X0)) | ~ spl67_1627), inference(resolution, [], [f26320, f165])).
fof(f26320, plain, (! [X6, X4, X2, X8, X7, X5, X3, X9] : (~ r1(sK13, X9) | ~ r1(X2, sK45) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2)) | ~ spl67_1627), inference(avatar_component_clause, [], [f26319])).
fof(f26319, plain, (spl67_1627 <=> ! [X9, X3, X5, X7, X8, X2, X4, X6] : (~ r1(X2, sK45) | ~ r1(sK13, X9) | ~ r1(X9, X8) | ~ r1(X8, X7) | ~ r1(X7, X6) | ~ r1(X6, X5) | ~ r1(X5, X4) | ~ r1(X4, X3) | ~ r1(X3, X2))), introduced(avatar_definition, [new_symbols(naming, [spl67_1627])])).
fof(f101504, plain, (~ spl67_26 | ~ spl67_81 | spl67_82 | ~ spl67_2532), inference(avatar_contradiction_clause, [], [f101503])).
fof(f101503, plain, ($false | (~ spl67_26 | ~ spl67_81 | spl67_82 | ~ spl67_2532)), inference(subsumption_resolution, [], [f101483, f174])).
fof(f174, plain, ~ p1(sK46), inference(cnf_transformation, [], [f88])).
fof(f101483, plain, (p1(sK46) | (~ spl67_26 | ~ spl67_81 | spl67_82 | ~ spl67_2532)), inference(resolution, [], [f101482, f173])).
fof(f173, plain, r1(sK42, sK46), inference(cnf_transformation, [], [f88])).
fof(f101482, plain, (! [X0] : (~ r1(sK42, X0) | p1(X0)) | (~ spl67_26 | ~ spl67_81 | spl67_82 | ~ spl67_2532)), inference(subsumption_resolution, [], [f101481, f665])).
fof(f665, plain, (sP3(sK42) | ~ spl67_81), inference(avatar_component_clause, [], [f664])).
fof(f664, plain, (spl67_81 <=> sP3(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl67_81])])).
fof(f101481, plain, (! [X0] : (p1(X0) | ~ r1(sK42, X0) | ~ sP3(sK42)) | (~ spl67_26 | spl67_82 | ~ spl67_2532)), inference(subsumption_resolution, [], [f101480, f362])).
fof(f362, plain, (p1(sK42) | ~ spl67_26), inference(avatar_component_clause, [], [f360])).
fof(f101480, plain, (! [X0] : (~ p1(sK42) | p1(X0) | ~ r1(sK42, X0) | ~ sP3(sK42)) | (spl67_82 | ~ spl67_2532)), inference(subsumption_resolution, [], [f101479, f669])).
fof(f669, plain, (~ sP0(sK42) | spl67_82), inference(avatar_component_clause, [], [f668])).
fof(f668, plain, (spl67_82 <=> sP0(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl67_82])])).
fof(f101479, plain, (! [X0] : (sP0(sK42) | ~ p1(sK42) | p1(X0) | ~ r1(sK42, X0) | ~ sP3(sK42)) | ~ spl67_2532), inference(resolution, [], [f101444, f175])).
fof(f101444, plain, (! [X0, X1] : (~ r1(X0, sK45) | sP0(X0) | ~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ sP3(X0)) | ~ spl67_2532), inference(resolution, [], [f37734, f91])).
fof(f91, plain, ! [X4, X0, X1] : (~ p1(sK5(X1)) | sP0(X0) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | ~ r1(X0, X4) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (sP0(X0) | ! [X1] : ((p1(sK4(X1)) & (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1))) & r1(X1, sK4(X1))) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f13, f15, f14])).
fof(f14, plain, ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) => (p1(sK4(X1)) & ? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) & r1(X1, sK4(X1)))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X1] : (? [X3] : (~ p1(X3) & r1(sK4(X1), X3)) => (~ p1(sK5(X1)) & r1(sK4(X1), sK5(X1)))), introduced(choice_axiom, [])).
fof(f13, plain, ! [X0] : (sP0(X0) | ! [X1] : (? [X2] : (p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3)) & r1(X1, X2)) | ~ r1(X0, X1)) | ~ p1(X0) | ! [X4] : (p1(X4) | ~ r1(X0, X4)) | ~ sP3(X0)), inference(rectify, [], [f12])).
fof(f12, plain, ! [X113] : (sP0(X113) | ! [X136] : (? [X137] : (p1(X137) & ? [X138] : (~ p1(X138) & r1(X137, X138)) & r1(X136, X137)) | ~ r1(X113, X136)) | ~ p1(X113) | ! [X139] : (p1(X139) | ~ r1(X113, X139)) | ~ sP3(X113)), inference(nnf_transformation, [], [f10])).
fof(f37734, plain, (p1(sK5(sK45)) | ~ spl67_2532), inference(avatar_component_clause, [], [f37732])).
fof(f85106, plain, (~ spl67_82 | ~ spl67_5143), inference(avatar_split_clause, [], [f85103, f85062, f668])).
fof(f85062, plain, (spl67_5143 <=> sP1(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl67_5143])])).
fof(f85103, plain, (~ sP0(sK42) | ~ spl67_5143), inference(resolution, [], [f85064, f892])).
fof(f892, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f891])).
fof(f891, plain, ! [X0] : (~ sP0(X0) | ~ sP1(X0) | ~ sP1(X0) | ~ sP0(X0)), inference(resolution, [], [f887, f103])).
fof(f103, plain, ! [X0] : (r1(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : ((! [X2] : ((p1(X2) & (~ p1(sK11(X2)) & r1(X2, sK11(X2)))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0))) & r1(X0, sK10(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12])], [f28, f31, f30, f29])).
fof(f29, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) => (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(sK10(X0), X2)) & p1(sK10(X0)) & ? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) & r1(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X2] : (? [X3] : (~ p1(X3) & r1(X2, X3)) => (~ p1(sK11(X2)) & r1(X2, sK11(X2)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (? [X6] : (~ p1(X6) & r1(sK10(X0), X6)) => (~ p1(sK12(X0)) & r1(sK10(X0), sK12(X0)))), introduced(choice_axiom, [])).
fof(f28, plain, ! [X0] : (? [X1] : (! [X2] : ((p1(X2) & ? [X3] : (~ p1(X3) & r1(X2, X3))) | ! [X4] : (~ p1(X4) | ! [X5] : (p1(X5) | ~ r1(X4, X5)) | ~ r1(X2, X4)) | ~ r1(X1, X2)) & p1(X1) & ? [X6] : (~ p1(X6) & r1(X1, X6)) & r1(X0, X1)) | ~ sP0(X0)), inference(rectify, [], [f27])).
fof(f27, plain, ! [X113] : (? [X130] : (! [X131] : ((p1(X131) & ? [X132] : (~ p1(X132) & r1(X131, X132))) | ! [X133] : (~ p1(X133) | ! [X134] : (p1(X134) | ~ r1(X133, X134)) | ~ r1(X131, X133)) | ~ r1(X130, X131)) & p1(X130) & ? [X135] : (~ p1(X135) & r1(X130, X135)) & r1(X113, X130)) | ~ sP0(X113)), inference(nnf_transformation, [], [f7])).
fof(f887, plain, ! [X0, X1] : (~ r1(X1, sK10(X0)) | ~ sP0(X0) | ~ sP1(X0) | ~ sP1(X1)), inference(duplicate_literal_removal, [], [f886])).
fof(f886, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | ~ r1(X1, sK10(X0)) | ~ sP1(X1) | ~ sP0(X0)), inference(resolution, [], [f885, f104])).
fof(f104, plain, ! [X0] : (r1(sK10(X0), sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f885, plain, ! [X2, X0, X1] : (~ r1(X1, sK12(X0)) | ~ sP1(X0) | ~ sP0(X0) | ~ r1(X2, X1) | ~ sP1(X2)), inference(subsumption_resolution, [], [f884, f105])).
fof(f105, plain, ! [X0] : (~ p1(sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f884, plain, ! [X2, X0, X1] : (~ sP0(X0) | ~ sP1(X0) | p1(sK12(X0)) | ~ r1(X1, sK12(X0)) | ~ r1(X2, X1) | ~ sP1(X2)), inference(resolution, [], [f883, f101])).
fof(f101, plain, ! [X2, X0, X1] : (~ p1(sK9(X2)) | p1(X2) | ~ r1(X1, X2) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : (! [X1] : (! [X2] : ((p1(sK8(X2)) & (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2))) & r1(X2, sK8(X2))) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9])], [f23, f25, f24])).
fof(f24, plain, ! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) => (p1(sK8(X2)) & ? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) & r1(X2, sK8(X2)))), introduced(choice_axiom, [])).
fof(f25, plain, ! [X2] : (? [X4] : (~ p1(X4) & r1(sK8(X2), X4)) => (~ p1(sK9(X2)) & r1(sK8(X2), sK9(X2)))), introduced(choice_axiom, [])).
fof(f23, plain, ! [X0] : (! [X1] : (! [X2] : (? [X3] : (p1(X3) & ? [X4] : (~ p1(X4) & r1(X3, X4)) & r1(X2, X3)) | p1(X2) | ~ r1(X1, X2)) | ~ r1(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f22])).
fof(f22, plain, ! [X113] : (! [X117] : (! [X118] : (? [X119] : (p1(X119) & ? [X120] : (~ p1(X120) & r1(X119, X120)) & r1(X118, X119)) | p1(X118) | ~ r1(X117, X118)) | ~ r1(X113, X117)) | ~ sP1(X113)), inference(nnf_transformation, [], [f8])).
fof(f883, plain, ! [X0] : (p1(sK9(sK12(X0))) | ~ sP0(X0) | ~ sP1(X0)), inference(duplicate_literal_removal, [], [f875])).
fof(f875, plain, ! [X0] : (~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0) | p1(sK9(sK12(X0))) | ~ sP1(X0)), inference(resolution, [], [f871, f824])).
fof(f824, plain, ! [X0, X1] : (~ r1(sK8(sK12(X0)), X1) | ~ sP0(X0) | p1(X1) | ~ sP1(X0)), inference(subsumption_resolution, [], [f823, f782])).
fof(f782, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f781])).
fof(f781, plain, ! [X0] : (p1(sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f267, f103])).
fof(f267, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | p1(sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f255, f105])).
fof(f255, plain, ! [X4, X5] : (p1(sK12(X4)) | p1(sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f102, f104])).
fof(f102, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | p1(sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f823, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0)))), inference(duplicate_literal_removal, [], [f815])).
fof(f815, plain, ! [X0, X1] : (~ sP1(X0) | ~ sP0(X0) | p1(X1) | ~ r1(sK8(sK12(X0)), X1) | ~ p1(sK8(sK12(X0))) | ~ sP0(X0)), inference(resolution, [], [f814, f718])).
fof(f718, plain, ! [X2, X0, X1] : (~ r1(sK12(X2), X0) | p1(X1) | ~ r1(X0, X1) | ~ p1(X0) | ~ sP0(X2)), inference(subsumption_resolution, [], [f717, f105])).
fof(f717, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2)), inference(duplicate_literal_removal, [], [f714])).
fof(f714, plain, ! [X2, X0, X1] : (~ p1(X0) | p1(X1) | ~ r1(X0, X1) | ~ r1(sK12(X2), X0) | p1(sK12(X2)) | ~ sP0(X2) | ~ sP0(X2)), inference(resolution, [], [f109, f104])).
fof(f109, plain, ! [X4, X2, X0, X5] : (~ r1(sK10(X0), X2) | ~ p1(X4) | p1(X5) | ~ r1(X4, X5) | ~ r1(X2, X4) | p1(X2) | ~ sP0(X0)), inference(cnf_transformation, [], [f32])).
fof(f814, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f813])).
fof(f813, plain, ! [X0] : (r1(sK12(X0), sK8(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f389, f103])).
fof(f389, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f377, f105])).
fof(f377, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK12(X4), sK8(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f99, f104])).
fof(f99, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(X2, sK8(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f871, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f870])).
fof(f870, plain, ! [X0] : (r1(sK8(sK12(X0)), sK9(sK12(X0))) | ~ sP1(X0) | ~ sP0(X0) | ~ sP0(X0)), inference(resolution, [], [f548, f103])).
fof(f548, plain, ! [X4, X5] : (~ r1(X5, sK10(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ sP1(X5) | ~ sP0(X4)), inference(subsumption_resolution, [], [f536, f105])).
fof(f536, plain, ! [X4, X5] : (p1(sK12(X4)) | r1(sK8(sK12(X4)), sK9(sK12(X4))) | ~ r1(X5, sK10(X4)) | ~ sP1(X5) | ~ sP0(X4)), inference(resolution, [], [f100, f104])).
fof(f100, plain, ! [X2, X0, X1] : (~ r1(X1, X2) | p1(X2) | r1(sK8(X2), sK9(X2)) | ~ r1(X0, X1) | ~ sP1(X0)), inference(cnf_transformation, [], [f26])).
fof(f85064, plain, (sP1(sK42) | ~ spl67_5143), inference(avatar_component_clause, [], [f85062])).
fof(f85089, plain, (spl67_5148 | spl67_5143), inference(avatar_split_clause, [], [f85085, f85062, f85087])).
fof(f85085, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (sP1(sK42) | ~ r1(X0, sK42) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(sK13, X6)), inference(subsumption_resolution, [], [f5745, f172])).
fof(f5745, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (sP1(sK42) | ~ r1(sK41, sK42) | ~ r1(X0, sK42) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(sK13, X6)), inference(duplicate_literal_removal, [], [f5470])).
fof(f5470, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (sP1(sK42) | ~ r1(sK41, sK42) | ~ r1(X0, sK42) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(sK13, X6) | sP1(sK42)), inference(resolution, [], [f5448, f2559])).
fof(f2559, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK42, sK47(X0)) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | sP1(X0)), inference(subsumption_resolution, [], [f2558, f163])).
fof(f163, plain, ! [X111, X109, X107, X112, X110, X108, X106, X113] : (~ p1(sK47(X113)) | sP1(X113) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | ~ r1(sK13, X106)), inference(cnf_transformation, [], [f88])).
fof(f2558, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | ~ r1(sK42, sK47(X0)) | p1(sK47(X0))), inference(subsumption_resolution, [], [f2557, f179])).
fof(f2557, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (p1(sK44(sK47(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | ~ r1(sK42, sK47(X0)) | p1(sK47(X0))), inference(duplicate_literal_removal, [], [f2550])).
fof(f2550, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (p1(sK44(sK47(X0))) | sP1(X0) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(sK13, X7) | ~ r1(sK42, sK47(X0)) | p1(sK47(X0)) | ~ r1(sK42, sK47(X0))), inference(resolution, [], [f1009, f178])).
fof(f1009, plain, ! [X14, X12, X10, X17, X15, X13, X11, X9, X16] : (~ r1(sK43(sK47(X10)), X9) | p1(X9) | sP1(X10) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK13, X17) | ~ r1(sK42, sK47(X10))), inference(subsumption_resolution, [], [f1008, f163])).
fof(f1008, plain, ! [X14, X12, X10, X17, X15, X13, X11, X9, X16] : (p1(X9) | ~ r1(sK43(sK47(X10)), X9) | sP1(X10) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK13, X17) | p1(sK47(X10)) | ~ r1(sK42, sK47(X10))), inference(subsumption_resolution, [], [f1006, f180])).
fof(f1006, plain, ! [X14, X12, X10, X17, X15, X13, X11, X9, X16] : (p1(X9) | ~ r1(sK43(sK47(X10)), X9) | ~ p1(sK43(sK47(X10))) | sP1(X10) | ~ r1(X11, X10) | ~ r1(X12, X11) | ~ r1(X13, X12) | ~ r1(X14, X13) | ~ r1(X15, X14) | ~ r1(X16, X15) | ~ r1(X17, X16) | ~ r1(sK13, X17) | p1(sK47(X10)) | ~ r1(sK42, sK47(X10))), inference(resolution, [], [f164, f177])).
fof(f164, plain, ! [X111, X109, X107, X116, X112, X110, X108, X106, X115, X113] : (~ r1(sK47(X113), X115) | p1(X116) | ~ r1(X115, X116) | ~ p1(X115) | sP1(X113) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | ~ r1(sK13, X106)), inference(cnf_transformation, [], [f88])).
fof(f5448, plain, ! [X0] : (r1(X0, sK47(X0)) | sP1(X0) | ~ r1(sK41, X0)), inference(resolution, [], [f5426, f171])).
fof(f5426, plain, ! [X0, X1] : (~ r1(sK40, X0) | ~ r1(X0, X1) | sP1(X1) | r1(X1, sK47(X1))), inference(resolution, [], [f5416, f170])).
fof(f5416, plain, ! [X2, X0, X1] : (~ r1(sK39, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP1(X2) | r1(X2, sK47(X2))), inference(resolution, [], [f5404, f169])).
fof(f5404, plain, ! [X2, X0, X3, X1] : (~ r1(sK38, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP1(X3) | r1(X3, sK47(X3))), inference(resolution, [], [f4887, f168])).
fof(f4887, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK37, X3) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X0, X1) | ~ r1(X1, X4) | sP1(X4) | r1(X4, sK47(X4))), inference(resolution, [], [f2206, f167])).
fof(f2206, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK36, X5) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X0, X1) | sP1(X1) | r1(X1, sK47(X1))), inference(resolution, [], [f872, f166])).
fof(f872, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK35, X6) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | sP1(X0) | r1(X0, sK47(X0))), inference(resolution, [], [f162, f165])).
fof(f162, plain, ! [X111, X109, X107, X112, X110, X108, X106, X113] : (~ r1(sK13, X106) | sP1(X113) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | r1(X113, sK47(X113))), inference(cnf_transformation, [], [f88])).
fof(f26321, plain, (spl67_1627 | spl67_1624 | ~ spl67_1623), inference(avatar_split_clause, [], [f26317, f26266, f26270, f26319])).
fof(f26266, plain, (spl67_1623 <=> r1(sK45, sK48(sK45))), introduced(avatar_definition, [new_symbols(naming, [spl67_1623])])).
fof(f26317, plain, (! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (p1(X0) | ~ r1(X1, X0) | ~ r1(sK45, X1) | ~ r1(X2, sK45) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | ~ r1(sK13, X9)) | ~ spl67_1623), inference(resolution, [], [f26285, f161])).
fof(f161, plain, ! [X111, X109, X107, X112, X120, X110, X108, X106, X119, X117, X113] : (~ p1(sK48(X117)) | p1(X120) | ~ r1(X119, X120) | ~ r1(X117, X119) | ~ r1(X113, X117) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | ~ r1(sK13, X106)), inference(cnf_transformation, [], [f88])).
fof(f26285, plain, (p1(sK48(sK45)) | ~ spl67_1623), inference(resolution, [], [f26268, f176])).
fof(f26268, plain, (r1(sK45, sK48(sK45)) | ~ spl67_1623), inference(avatar_component_clause, [], [f26266])).
fof(f26272, plain, (spl67_1623 | spl67_1624), inference(avatar_split_clause, [], [f26240, f26270, f26266])).
fof(f26240, plain, ! [X2, X3] : (~ r1(sK45, X2) | ~ r1(X2, X3) | p1(X3) | r1(sK45, sK48(sK45))), inference(resolution, [], [f26206, f175])).
fof(f26206, plain, ! [X2, X0, X1] : (~ r1(sK42, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | p1(X2) | r1(X0, sK48(X0))), inference(resolution, [], [f26172, f172])).
fof(f26172, plain, ! [X2, X0, X3, X1] : (~ r1(sK41, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | p1(X3) | r1(X1, sK48(X1))), inference(resolution, [], [f26154, f171])).
fof(f26154, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK40, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | ~ r1(X3, X4) | p1(X4) | r1(X2, sK48(X2))), inference(resolution, [], [f20226, f170])).
fof(f20226, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK39, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | ~ r1(X4, X5) | p1(X5) | r1(X3, sK48(X3))), inference(resolution, [], [f10896, f169])).
fof(f10896, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK38, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | ~ r1(X5, X6) | p1(X6) | r1(X1, sK48(X1))), inference(resolution, [], [f6667, f168])).
fof(f6667, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ r1(sK37, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | ~ r1(X1, X7) | p1(X7) | r1(X0, sK48(X0))), inference(resolution, [], [f2545, f167])).
fof(f2545, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1] : (~ r1(sK36, X8) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X0, X1) | p1(X1) | r1(X2, sK48(X2))), inference(resolution, [], [f1049, f166])).
fof(f1049, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ r1(sK35, X9) | ~ r1(X1, X0) | ~ r1(X2, X1) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X7, X6) | ~ r1(X8, X7) | ~ r1(X9, X8) | p1(X0) | r1(X2, sK48(X2))), inference(resolution, [], [f160, f165])).
fof(f160, plain, ! [X111, X109, X107, X112, X120, X110, X108, X106, X119, X117, X113] : (~ r1(sK13, X106) | p1(X120) | ~ r1(X119, X120) | ~ r1(X117, X119) | ~ r1(X113, X117) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | r1(X117, sK48(X117))), inference(cnf_transformation, [], [f88])).
fof(f3072, plain, (spl67_26 | ~ spl67_54 | spl67_55 | ~ spl67_57), inference(avatar_split_clause, [], [f3071, f530, f521, f517, f360])).
fof(f3071, plain, (p1(sK42) | (~ spl67_54 | spl67_55 | ~ spl67_57)), inference(subsumption_resolution, [], [f3058, f518])).
fof(f3058, plain, (p1(sK42) | ~ sP2(sK42) | (spl67_55 | ~ spl67_57)), inference(subsumption_resolution, [], [f2868, f175])).
fof(f2868, plain, (~ r1(sK42, sK45) | p1(sK42) | ~ sP2(sK42) | (spl67_55 | ~ spl67_57)), inference(resolution, [], [f2837, f522])).
fof(f522, plain, (~ r1(sK42, sK6(sK42)) | spl67_55), inference(avatar_component_clause, [], [f521])).
fof(f2837, plain, (! [X0] : (r1(X0, sK6(X0)) | ~ r1(X0, sK45) | p1(X0) | ~ sP2(X0)) | ~ spl67_57), inference(resolution, [], [f2820, f94])).
fof(f94, plain, ! [X4, X0] : (~ p1(sK7(X4)) | r1(X0, sK6(X0)) | ~ r1(X0, X4) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).
fof(f3070, plain, (spl67_242 | spl67_243 | spl67_26 | ~ spl67_54 | ~ spl67_57), inference(avatar_split_clause, [], [f3069, f530, f517, f360, f2773, f2769])).
fof(f3069, plain, (p1(sK42) | p1(sK6(sK42)) | p1(sK43(sK6(sK42))) | (~ spl67_54 | ~ spl67_57)), inference(subsumption_resolution, [], [f3062, f518])).
fof(f3062, plain, (p1(sK42) | ~ sP2(sK42) | p1(sK6(sK42)) | p1(sK43(sK6(sK42))) | ~ spl67_57), inference(subsumption_resolution, [], [f2966, f175])).
fof(f2966, plain, (~ r1(sK42, sK45) | p1(sK42) | ~ sP2(sK42) | p1(sK6(sK42)) | p1(sK43(sK6(sK42))) | ~ spl67_57), inference(resolution, [], [f2837, f180])).
fof(f2759, plain, spl67_54, inference(avatar_split_clause, [], [f2752, f517])).
fof(f2752, plain, sP2(sK42), inference(resolution, [], [f2737, f172])).
fof(f2737, plain, ! [X0] : (~ r1(sK41, X0) | sP2(X0)), inference(resolution, [], [f2722, f171])).
fof(f2722, plain, ! [X0, X1] : (~ r1(sK40, X0) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f2711, f170])).
fof(f2711, plain, ! [X2, X0, X1] : (~ r1(sK39, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP2(X2)), inference(resolution, [], [f2698, f169])).
fof(f2698, plain, ! [X2, X0, X3, X1] : (~ r1(sK38, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP2(X3)), inference(resolution, [], [f2683, f168])).
fof(f2683, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK37, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | sP2(X4)), inference(resolution, [], [f2046, f167])).
fof(f2046, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK36, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | sP2(X5)), inference(resolution, [], [f795, f166])).
fof(f795, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK35, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | sP2(X1)), inference(resolution, [], [f159, f165])).
fof(f159, plain, ! [X111, X109, X107, X112, X110, X108, X106, X113] : (~ r1(sK13, X106) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | sP2(X113)), inference(cnf_transformation, [], [f88])).
fof(f2535, plain, (~ spl67_81 | spl67_82 | ~ spl67_26 | spl67_235), inference(avatar_split_clause, [], [f2531, f2533, f360, f668, f664])).
fof(f2531, plain, ! [X14] : (r1(X14, sK4(X14)) | ~ r1(sK42, X14) | ~ p1(sK42) | sP0(sK42) | ~ sP3(sK42)), inference(subsumption_resolution, [], [f707, f174])).
fof(f707, plain, ! [X14] : (r1(X14, sK4(X14)) | ~ r1(sK42, X14) | ~ p1(sK42) | p1(sK46) | sP0(sK42) | ~ sP3(sK42)), inference(resolution, [], [f89, f173])).
fof(f89, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(X1, sK4(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f2530, plain, (~ spl67_81 | spl67_82 | ~ spl67_26 | spl67_234), inference(avatar_split_clause, [], [f2526, f2528, f360, f668, f664])).
fof(f2526, plain, ! [X14] : (r1(sK4(X14), sK5(X14)) | ~ r1(sK42, X14) | ~ p1(sK42) | sP0(sK42) | ~ sP3(sK42)), inference(subsumption_resolution, [], [f733, f174])).
fof(f733, plain, ! [X14] : (r1(sK4(X14), sK5(X14)) | ~ r1(sK42, X14) | ~ p1(sK42) | p1(sK46) | sP0(sK42) | ~ sP3(sK42)), inference(resolution, [], [f90, f173])).
fof(f90, plain, ! [X4, X0, X1] : (~ r1(X0, X4) | r1(sK4(X1), sK5(X1)) | ~ r1(X0, X1) | ~ p1(X0) | p1(X4) | sP0(X0) | ~ sP3(X0)), inference(cnf_transformation, [], [f16])).
fof(f2525, plain, spl67_81, inference(avatar_split_clause, [], [f2482, f664])).
fof(f2482, plain, sP3(sK42), inference(resolution, [], [f2456, f172])).
fof(f2456, plain, ! [X0] : (~ r1(sK41, X0) | sP3(X0)), inference(resolution, [], [f2447, f171])).
fof(f2447, plain, ! [X0, X1] : (~ r1(sK40, X0) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f2431, f170])).
fof(f2431, plain, ! [X2, X0, X1] : (~ r1(sK39, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | sP3(X2)), inference(resolution, [], [f2413, f169])).
fof(f2413, plain, ! [X2, X0, X3, X1] : (~ r1(sK38, X0) | ~ r1(X0, X1) | ~ r1(X1, X2) | ~ r1(X2, X3) | sP3(X3)), inference(resolution, [], [f2399, f168])).
fof(f2399, plain, ! [X4, X2, X0, X3, X1] : (~ r1(sK37, X2) | ~ r1(X2, X0) | ~ r1(X0, X1) | ~ r1(X1, X3) | ~ r1(X3, X4) | sP3(X4)), inference(resolution, [], [f2033, f167])).
fof(f2033, plain, ! [X4, X2, X0, X5, X3, X1] : (~ r1(sK36, X4) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X0, X1) | ~ r1(X1, X5) | sP3(X5)), inference(resolution, [], [f787, f166])).
fof(f787, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ r1(sK35, X6) | ~ r1(X2, X0) | ~ r1(X3, X2) | ~ r1(X4, X3) | ~ r1(X5, X4) | ~ r1(X6, X5) | ~ r1(X0, X1) | sP3(X1)), inference(resolution, [], [f158, f165])).
fof(f158, plain, ! [X111, X109, X107, X112, X110, X108, X106, X113] : (~ r1(sK13, X106) | ~ r1(X112, X113) | ~ r1(X111, X112) | ~ r1(X110, X111) | ~ r1(X109, X110) | ~ r1(X108, X109) | ~ r1(X107, X108) | ~ r1(X106, X107) | sP3(X113)), inference(cnf_transformation, [], [f88])).
fof(f533, plain, (~ spl67_54 | spl67_26 | spl67_55 | spl67_57), inference(avatar_split_clause, [], [f405, f530, f521, f360, f517])).
fof(f405, plain, (r1(sK45, sK7(sK45)) | r1(sK42, sK6(sK42)) | p1(sK42) | ~ sP2(sK42)), inference(resolution, [], [f93, f175])).
fof(f93, plain, ! [X4, X0] : (~ r1(X0, X4) | r1(X4, sK7(X4)) | r1(X0, sK6(X0)) | p1(X0) | ~ sP2(X0)), inference(cnf_transformation, [], [f21])).