fof(f118, plain, $false, inference(avatar_sat_refutation, [], [f80, f109, f117])).
fof(f117, plain, ~ spl6_1, inference(avatar_contradiction_clause, [], [f116])).
fof(f116, plain, ($false | ~ spl6_1), inference(subsumption_resolution, [], [f115, f41])).
fof(f41, plain, stable(sK5), inference(cnf_transformation, [], [f28])).
fof(f28, plain, (~ in_environment(sK5, critical_point(sK5)) & stable(sK5) & environment(sK5)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5])], [f14, f27])).
fof(f27, plain, (? [X0] : (~ in_environment(X0, critical_point(X0)) & stable(X0) & environment(X0)) => (~ in_environment(sK5, critical_point(sK5)) & stable(sK5) & environment(sK5))), introduced(choice_axiom, [])).
fof(f14, plain, ? [X0] : (~ in_environment(X0, critical_point(X0)) & stable(X0) & environment(X0)), inference(flattening, [], [f13])).
fof(f13, plain, ? [X0] : (~ in_environment(X0, critical_point(X0)) & (stable(X0) & environment(X0))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ~ ! [X0] : ((stable(X0) & environment(X0)) => in_environment(X0, critical_point(X0))), inference(negated_conjecture, [], [f4])).
fof(f4, plain, ~ ! [X0] : ((stable(X0) & environment(X0)) => in_environment(X0, critical_point(X0))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT023+2.p', prove_l5)).
fof(f115, plain, (~ stable(sK5) | ~ spl6_1), inference(subsumption_resolution, [], [f114, f40])).
fof(f40, plain, environment(sK5), inference(cnf_transformation, [], [f28])).
fof(f114, plain, (~ environment(sK5) | ~ stable(sK5) | ~ spl6_1), inference(resolution, [], [f75, f47])).
fof(f47, plain, ! [X0] : (~ sP0(X0) | ~ environment(X0) | ~ stable(X0)), inference(subsumption_resolution, [], [f46, f38])).
fof(f38, plain, ! [X0] : (in_environment(X0, sK4(X0)) | ~ stable(X0) | ~ environment(X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0] : ((! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater_or_equal(X2, sK4(X0)) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & in_environment(X0, sK4(X0))) | ~ stable(X0) | ~ environment(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4])], [f12, f25])).
fof(f25, plain, ! [X0] : (? [X1] : (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater_or_equal(X2, X1) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & in_environment(X0, X1)) => (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater_or_equal(X2, sK4(X0)) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & in_environment(X0, sK4(X0)))), introduced(choice_axiom, [])).
fof(f12, plain, ! [X0] : (? [X1] : (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater_or_equal(X2, X1) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & in_environment(X0, X1)) | ~ stable(X0) | ~ environment(X0)), inference(flattening, [], [f11])).
fof(f11, plain, ! [X0] : (? [X1] : (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | (~ greater_or_equal(X2, X1) | ~ subpopulations(first_movers, efficient_producers, X0, X2))) & in_environment(X0, X1)) | (~ stable(X0) | ~ environment(X0))), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : ((stable(X0) & environment(X0)) => ? [X1] : (! [X2] : ((greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2))) & in_environment(X0, X1))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT023+2.p', l1)).
fof(f46, plain, ! [X0] : (~ stable(X0) | ~ environment(X0) | ~ in_environment(X0, sK4(X0)) | ~ sP0(X0)), inference(duplicate_literal_removal, [], [f45])).
fof(f45, plain, ! [X0] : (~ stable(X0) | ~ environment(X0) | ~ in_environment(X0, sK4(X0)) | ~ sP0(X0) | ~ in_environment(X0, sK4(X0)) | ~ sP0(X0)), inference(resolution, [], [f44, f30])).
fof(f30, plain, ! [X0, X1] : (greater_or_equal(sK1(X0, X1), X1) | ~ in_environment(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ! [X0] : (! [X1] : ((~ greater(growth_rate(efficient_producers, sK1(X0, X1)), growth_rate(first_movers, sK1(X0, X1))) & greater_or_equal(sK1(X0, X1), X1) & subpopulations(first_movers, efficient_producers, X0, sK1(X0, X1))) | ~ in_environment(X0, X1)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1])], [f17, f18])).
fof(f18, plain, ! [X1, X0] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => (~ greater(growth_rate(efficient_producers, sK1(X0, X1)), growth_rate(first_movers, sK1(X0, X1))) & greater_or_equal(sK1(X0, X1), X1) & subpopulations(first_movers, efficient_producers, X0, sK1(X0, X1)))), introduced(choice_axiom, [])).
fof(f17, plain, ! [X0] : (! [X1] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) | ~ in_environment(X0, X1)) | ~ sP0(X0)), inference(nnf_transformation, [], [f15])).
fof(f15, plain, ! [X0] : (! [X1] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) | ~ in_environment(X0, X1)) | ~ sP0(X0)), inference(usedef, [], [e15])).
fof(e15, plain, ! [X0] : (sP0(X0) <=> ! [X1] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) | ~ in_environment(X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f44, plain, ! [X0, X1] : (~ greater_or_equal(sK1(X0, X1), sK4(X0)) | ~ stable(X0) | ~ environment(X0) | ~ in_environment(X0, X1) | ~ sP0(X0)), inference(subsumption_resolution, [], [f43, f31])).
fof(f31, plain, ! [X0, X1] : (~ greater(growth_rate(efficient_producers, sK1(X0, X1)), growth_rate(first_movers, sK1(X0, X1))) | ~ in_environment(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f19])).
fof(f43, plain, ! [X0, X1] : (~ greater_or_equal(sK1(X0, X1), sK4(X0)) | greater(growth_rate(efficient_producers, sK1(X0, X1)), growth_rate(first_movers, sK1(X0, X1))) | ~ stable(X0) | ~ environment(X0) | ~ in_environment(X0, X1) | ~ sP0(X0)), inference(resolution, [], [f39, f29])).
fof(f29, plain, ! [X0, X1] : (subpopulations(first_movers, efficient_producers, X0, sK1(X0, X1)) | ~ in_environment(X0, X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f19])).
fof(f39, plain, ! [X2, X0] : (~ subpopulations(first_movers, efficient_producers, X0, X2) | ~ greater_or_equal(X2, sK4(X0)) | greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ stable(X0) | ~ environment(X0)), inference(cnf_transformation, [], [f26])).
fof(f75, plain, (sP0(sK5) | ~ spl6_1), inference(avatar_component_clause, [], [f73])).
fof(f73, plain, (spl6_1 <=> sP0(sK5)), introduced(avatar_definition, [new_symbols(naming, [spl6_1])])).
fof(f109, plain, (spl6_1 | ~ spl6_2), inference(avatar_split_clause, [], [f108, f77, f73])).
fof(f77, plain, (spl6_2 <=> (critical_point(sK5) = sK2(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl6_2])])).
fof(f108, plain, (sP0(sK5) | ~ spl6_2), inference(subsumption_resolution, [], [f107, f40])).
fof(f107, plain, (sP0(sK5) | ~ environment(sK5) | ~ spl6_2), inference(subsumption_resolution, [], [f85, f42])).
fof(f42, plain, ~ in_environment(sK5, critical_point(sK5)), inference(cnf_transformation, [], [f28])).
fof(f85, plain, (in_environment(sK5, critical_point(sK5)) | sP0(sK5) | ~ environment(sK5) | ~ spl6_2), inference(superposition, [], [f32, f79])).
fof(f79, plain, ((critical_point(sK5) = sK2(sK5)) | ~ spl6_2), inference(avatar_component_clause, [], [f77])).
fof(f32, plain, ! [X0] : (in_environment(X0, sK2(X0)) | sP0(X0) | ~ environment(X0)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ! [X0] : ((! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater(X2, sK2(X0)) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & ~ greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) & in_environment(X0, sK2(X0))) | sP0(X0) | ~ environment(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2])], [f20, f21])).
fof(f21, plain, ! [X0] : (? [X1] : (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater(X2, X1) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & ~ greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) & in_environment(X0, X1)) => (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater(X2, sK2(X0)) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & ~ greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) & in_environment(X0, sK2(X0)))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X0] : (? [X1] : (! [X2] : (greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | ~ greater(X2, X1) | ~ subpopulations(first_movers, efficient_producers, X0, X2)) & ~ greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) & in_environment(X0, X1)) | sP0(X0) | ~ environment(X0)), inference(rectify, [], [f16])).
fof(f16, plain, ! [X0] : (? [X3] : (! [X4] : (greater(growth_rate(efficient_producers, X4), growth_rate(first_movers, X4)) | ~ greater(X4, X3) | ~ subpopulations(first_movers, efficient_producers, X0, X4)) & ~ greater(growth_rate(efficient_producers, X3), growth_rate(first_movers, X3)) & in_environment(X0, X3)) | sP0(X0) | ~ environment(X0)), inference(definition_folding, [], [f8, e15])).
fof(f8, plain, ! [X0] : (? [X3] : (! [X4] : (greater(growth_rate(efficient_producers, X4), growth_rate(first_movers, X4)) | ~ greater(X4, X3) | ~ subpopulations(first_movers, efficient_producers, X0, X4)) & ~ greater(growth_rate(efficient_producers, X3), growth_rate(first_movers, X3)) & in_environment(X0, X3)) | ! [X1] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) | ~ in_environment(X0, X1)) | ~ environment(X0)), inference(flattening, [], [f7])).
fof(f7, plain, ! [X0] : (? [X3] : (! [X4] : (greater(growth_rate(efficient_producers, X4), growth_rate(first_movers, X4)) | (~ greater(X4, X3) | ~ subpopulations(first_movers, efficient_producers, X0, X4))) & ~ greater(growth_rate(efficient_producers, X3), growth_rate(first_movers, X3)) & in_environment(X0, X3)) | (! [X1] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & (greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2))) | ~ in_environment(X0, X1)) | ~ environment(X0))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0] : ((? [X1] : (! [X2] : ((greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2))) & in_environment(X0, X1)) & environment(X0)) => ? [X3] : (! [X4] : ((greater(X4, X3) & subpopulations(first_movers, efficient_producers, X0, X4)) => greater(growth_rate(efficient_producers, X4), growth_rate(first_movers, X4))) & ~ greater(growth_rate(efficient_producers, X3), growth_rate(first_movers, X3)) & in_environment(X0, X3))), inference(rectify, [], [f1])).
fof(f1, plain, ! [X0] : ((? [X1] : (! [X2] : ((greater_or_equal(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2))) & in_environment(X0, X1)) & environment(X0)) => ? [X1] : (! [X2] : ((greater(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2))) & ~ greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) & in_environment(X0, X1))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT023+2.p', mp_earliest_time_growth_rate_exceeds)).
fof(f80, plain, (spl6_1 | spl6_2), inference(avatar_split_clause, [], [f71, f77, f73])).
fof(f71, plain, ((critical_point(sK5) = sK2(sK5)) | sP0(sK5)), inference(resolution, [], [f70, f40])).
fof(f70, plain, ! [X0] : (~ environment(X0) | (critical_point(X0) = sK2(X0)) | sP0(X0)), inference(subsumption_resolution, [], [f69, f33])).
fof(f33, plain, ! [X0] : (~ greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) | sP0(X0) | ~ environment(X0)), inference(cnf_transformation, [], [f22])).
fof(f69, plain, ! [X0] : ((critical_point(X0) = sK2(X0)) | ~ environment(X0) | greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) | sP0(X0)), inference(subsumption_resolution, [], [f68, f32])).
fof(f68, plain, ! [X0] : (~ in_environment(X0, sK2(X0)) | (critical_point(X0) = sK2(X0)) | ~ environment(X0) | greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) | sP0(X0)), inference(duplicate_literal_removal, [], [f65])).
fof(f65, plain, ! [X0] : (~ in_environment(X0, sK2(X0)) | (critical_point(X0) = sK2(X0)) | ~ environment(X0) | greater(growth_rate(efficient_producers, sK2(X0)), growth_rate(first_movers, sK2(X0))) | sP0(X0) | ~ in_environment(X0, sK2(X0)) | (critical_point(X0) = sK2(X0)) | ~ environment(X0) | sP0(X0) | ~ environment(X0)), inference(resolution, [], [f57, f49])).
fof(f49, plain, ! [X0, X1] : (greater(sK3(X0, sK2(X1)), sK2(X1)) | ~ in_environment(X0, sK2(X1)) | (critical_point(X0) = sK2(X1)) | ~ environment(X0) | sP0(X1) | ~ environment(X1)), inference(resolution, [], [f36, f33])).
fof(f36, plain, ! [X0, X1] : (greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | greater(sK3(X0, X1), X1) | ~ in_environment(X0, X1) | (critical_point(X0) = X1) | ~ environment(X0)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1] : ((critical_point(X0) = X1) | (~ greater(growth_rate(efficient_producers, sK3(X0, X1)), growth_rate(first_movers, sK3(X0, X1))) & greater(sK3(X0, X1), X1) & subpopulations(first_movers, efficient_producers, X0, sK3(X0, X1))) | ~ in_environment(X0, X1) | greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | ~ environment(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f10, f23])).
fof(f23, plain, ! [X1, X0] : (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => (~ greater(growth_rate(efficient_producers, sK3(X0, X1)), growth_rate(first_movers, sK3(X0, X1))) & greater(sK3(X0, X1), X1) & subpopulations(first_movers, efficient_producers, X0, sK3(X0, X1)))), introduced(choice_axiom, [])).
fof(f10, plain, ! [X0, X1] : ((critical_point(X0) = X1) | ? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & greater(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) | ~ in_environment(X0, X1) | greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | ~ environment(X0)), inference(flattening, [], [f9])).
fof(f9, plain, ! [X0, X1] : ((critical_point(X0) = X1) | (? [X2] : (~ greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) & (greater(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2))) | ~ in_environment(X0, X1) | greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | ~ environment(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1] : ((! [X2] : ((greater(X2, X1) & subpopulations(first_movers, efficient_producers, X0, X2)) => greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2))) & in_environment(X0, X1) & ~ greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) & environment(X0)) => (critical_point(X0) = X1)), file('/home/ubuntu/library/tptp/Problems/MGT/MGT023+2.p', d1)).
fof(f57, plain, ! [X0, X1] : (~ greater(sK3(X1, X0), sK2(X1)) | ~ in_environment(X1, X0) | (critical_point(X1) = X0) | ~ environment(X1) | greater(growth_rate(efficient_producers, X0), growth_rate(first_movers, X0)) | sP0(X1)), inference(subsumption_resolution, [], [f56, f37])).
fof(f37, plain, ! [X0, X1] : (~ greater(growth_rate(efficient_producers, sK3(X0, X1)), growth_rate(first_movers, sK3(X0, X1))) | (critical_point(X0) = X1) | ~ in_environment(X0, X1) | greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | ~ environment(X0)), inference(cnf_transformation, [], [f24])).
fof(f56, plain, ! [X0, X1] : (greater(growth_rate(efficient_producers, X0), growth_rate(first_movers, X0)) | ~ in_environment(X1, X0) | (critical_point(X1) = X0) | ~ environment(X1) | ~ greater(sK3(X1, X0), sK2(X1)) | greater(growth_rate(efficient_producers, sK3(X1, X0)), growth_rate(first_movers, sK3(X1, X0))) | sP0(X1)), inference(duplicate_literal_removal, [], [f53])).
fof(f53, plain, ! [X0, X1] : (greater(growth_rate(efficient_producers, X0), growth_rate(first_movers, X0)) | ~ in_environment(X1, X0) | (critical_point(X1) = X0) | ~ environment(X1) | ~ greater(sK3(X1, X0), sK2(X1)) | greater(growth_rate(efficient_producers, sK3(X1, X0)), growth_rate(first_movers, sK3(X1, X0))) | sP0(X1) | ~ environment(X1)), inference(resolution, [], [f35, f34])).
fof(f34, plain, ! [X2, X0] : (~ subpopulations(first_movers, efficient_producers, X0, X2) | ~ greater(X2, sK2(X0)) | greater(growth_rate(efficient_producers, X2), growth_rate(first_movers, X2)) | sP0(X0) | ~ environment(X0)), inference(cnf_transformation, [], [f22])).
fof(f35, plain, ! [X0, X1] : (greater(growth_rate(efficient_producers, X1), growth_rate(first_movers, X1)) | subpopulations(first_movers, efficient_producers, X0, sK3(X0, X1)) | ~ in_environment(X0, X1) | (critical_point(X0) = X1) | ~ environment(X0)), inference(cnf_transformation, [], [f24])).