fof(f2247, plain, $false, inference(avatar_sat_refutation, [], [f78, f84, f112, f171, f179, f1199, f1241, f1242, f1256, f1257, f1314, f1320, f1437, f1439, f1500, f1520, f2066, f2246])).
fof(f2246, plain, (~ spl4_13 | spl4_52), inference(avatar_contradiction_clause, [], [f2245])).
fof(f2245, plain, ($false | (~ spl4_13 | spl4_52)), inference(subsumption_resolution, [], [f2244, f1378])).
fof(f1378, plain, (~ (position(sK1, sK3) = position(sK1, sK2)) | spl4_52), inference(avatar_component_clause, [], [f1377])).
fof(f1377, plain, (spl4_52 <=> (position(sK1, sK3) = position(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_52])])).
fof(f2244, plain, ((position(sK1, sK3) = position(sK1, sK2)) | ~ spl4_13), inference(equality_resolution, [], [f1575])).
fof(f1575, plain, (! [X2] : (~ (external_ties(sK1, sK2) = external_ties(sK1, X2)) | (position(sK1, sK3) = position(sK1, X2))) | ~ spl4_13), inference(subsumption_resolution, [], [f1562, f60])).
fof(f60, plain, organization(sK1), inference(cnf_transformation, [], [f36])).
fof(f36, plain, (((~ smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) & greater(external_ties(sK1, sK3), external_ties(sK1, sK2)) & smaller_or_equal(internal_friction(sK1, sK3), internal_friction(sK1, sK2)) & greater_or_equal(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))) | sP0(sK2, sK1, sK3)) & ~ has_immunity(sK1, sK3) & ~ has_immunity(sK1, sK2) & organization(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3])], [f27, f35])).
fof(f35, plain, (? [X0, X1, X2] : (((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) | sP0(X1, X0, X2)) & ~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0)) => (((~ smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) & greater(external_ties(sK1, sK3), external_ties(sK1, sK2)) & smaller_or_equal(internal_friction(sK1, sK3), internal_friction(sK1, sK2)) & greater_or_equal(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))) | sP0(sK2, sK1, sK3)) & ~ has_immunity(sK1, sK3) & ~ has_immunity(sK1, sK2) & organization(sK1))), introduced(choice_axiom, [])).
fof(f27, plain, ? [X0, X1, X2] : (((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) | sP0(X1, X0, X2)) & ~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0)), inference(definition_folding, [], [f25, e26])).
fof(f26, plain, ! [X1, X0, X2] : ((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) | ~ sP0(X1, X0, X2)), inference(usedef, [], [e26])).
fof(e26, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> (~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f25, plain, ? [X0, X1, X2] : (((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) | (~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)))) & ~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0)), inference(flattening, [], [f24])).
fof(f24, plain, ? [X0, X1, X2] : (((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & (greater(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)))) | (~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & (greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))))) & (~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ~ ! [X0, X1, X2] : ((~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0)) => (((greater(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) => smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1))) & ((greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) => smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1))))), inference(rectify, [], [f11])).
fof(f11, plain, ~ ! [X0, X3, X4] : ((~ has_immunity(X0, X4) & ~ has_immunity(X0, X3) & organization(X0)) => (((greater(external_ties(X0, X4), external_ties(X0, X3)) & smaller_or_equal(internal_friction(X0, X4), internal_friction(X0, X3)) & greater_or_equal(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))) & ((greater_or_equal(external_ties(X0, X4), external_ties(X0, X3)) & smaller_or_equal(internal_friction(X0, X4), internal_friction(X0, X3)) & greater(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))))), inference(negated_conjecture, [], [f10])).
fof(f10, plain, ~ ! [X0, X3, X4] : ((~ has_immunity(X0, X4) & ~ has_immunity(X0, X3) & organization(X0)) => (((greater(external_ties(X0, X4), external_ties(X0, X3)) & smaller_or_equal(internal_friction(X0, X4), internal_friction(X0, X3)) & greater_or_equal(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))) & ((greater_or_equal(external_ties(X0, X4), external_ties(X0, X3)) & smaller_or_equal(internal_friction(X0, X4), internal_friction(X0, X3)) & greater(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', lemma_1)).
fof(f1562, plain, (! [X2] : (~ (external_ties(sK1, sK2) = external_ties(sK1, X2)) | (position(sK1, sK3) = position(sK1, X2)) | ~ organization(sK1)) | ~ spl4_13), inference(superposition, [], [f55, f310])).
fof(f310, plain, ((external_ties(sK1, sK3) = external_ties(sK1, sK2)) | ~ spl4_13), inference(avatar_component_clause, [], [f308])).
fof(f308, plain, (spl4_13 <=> (external_ties(sK1, sK3) = external_ties(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_13])])).
fof(f55, plain, ! [X2, X0, X1] : (~ (external_ties(X0, X2) = external_ties(X0, X1)) | (position(X0, X2) = position(X0, X1)) | ~ organization(X0)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1, X2] : ((((position(X0, X2) = position(X0, X1)) | ~ (external_ties(X0, X2) = external_ties(X0, X1))) & (greater(position(X0, X2), position(X0, X1)) | ~ greater(external_ties(X0, X2), external_ties(X0, X1)))) | ~ organization(X0)), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0, X1, X2] : (organization(X0) => (((external_ties(X0, X2) = external_ties(X0, X1)) => (position(X0, X2) = position(X0, X1))) & (greater(external_ties(X0, X2), external_ties(X0, X1)) => greater(position(X0, X2), position(X0, X1))))), inference(rectify, [], [f9])).
fof(f9, plain, ! [X0, X3, X4] : (organization(X0) => (((external_ties(X0, X4) = external_ties(X0, X3)) => (position(X0, X4) = position(X0, X3))) & (greater(external_ties(X0, X4), external_ties(X0, X3)) => greater(position(X0, X4), position(X0, X3))))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', assumption_6)).
fof(f2066, plain, (spl4_47 | ~ spl4_52 | ~ spl4_53), inference(avatar_contradiction_clause, [], [f2065])).
fof(f2065, plain, ($false | (spl4_47 | ~ spl4_52 | ~ spl4_53)), inference(subsumption_resolution, [], [f2047, f68])).
fof(f68, plain, ! [X1] : greater_or_equal(X1, X1), inference(equality_resolution, [], [f42])).
fof(f42, plain, ! [X0, X1] : (greater_or_equal(X0, X1) | ~ (X0 = X1)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ! [X0, X1] : ((greater_or_equal(X0, X1) | (~ (X0 = X1) & ~ greater(X0, X1))) & ((X0 = X1) | greater(X0, X1) | ~ greater_or_equal(X0, X1))), inference(flattening, [], [f30])).
fof(f30, plain, ! [X0, X1] : ((greater_or_equal(X0, X1) | (~ (X0 = X1) & ~ greater(X0, X1))) & (((X0 = X1) | greater(X0, X1)) | ~ greater_or_equal(X0, X1))), inference(nnf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1] : (greater_or_equal(X0, X1) <=> ((X0 = X1) | greater(X0, X1))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', definition_greater_or_equal)).
fof(f2047, plain, (~ greater_or_equal(position(sK1, sK2), position(sK1, sK2)) | (spl4_47 | ~ spl4_52 | ~ spl4_53)), inference(backward_demodulation, [], [f1509, f1379])).
fof(f1379, plain, ((position(sK1, sK3) = position(sK1, sK2)) | ~ spl4_52), inference(avatar_component_clause, [], [f1377])).
fof(f1509, plain, (~ greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | (spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1508, f60])).
fof(f1508, plain, (~ greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | ~ organization(sK1) | (spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1507, f61])).
fof(f61, plain, ~ has_immunity(sK1, sK2), inference(cnf_transformation, [], [f36])).
fof(f1507, plain, (~ greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | has_immunity(sK1, sK2) | ~ organization(sK1) | (spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1506, f62])).
fof(f62, plain, ~ has_immunity(sK1, sK3), inference(cnf_transformation, [], [f36])).
fof(f1506, plain, (~ greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | has_immunity(sK1, sK3) | has_immunity(sK1, sK2) | ~ organization(sK1) | (spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1487, f1254])).
fof(f1254, plain, (~ smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | spl4_47), inference(avatar_component_clause, [], [f1253])).
fof(f1253, plain, (spl4_47 <=> smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_47])])).
fof(f1487, plain, (~ greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | has_immunity(sK1, sK3) | has_immunity(sK1, sK2) | ~ organization(sK1) | ~ spl4_53), inference(resolution, [], [f1434, f48])).
fof(f48, plain, ! [X2, X0, X1] : (~ greater(capability(X0, X2), capability(X0, X1)) | ~ greater_or_equal(position(X0, X2), position(X0, X1)) | smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | has_immunity(X0, X2) | has_immunity(X0, X1) | ~ organization(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X1, X2] : ((((hazard_of_mortality(X0, X2) = hazard_of_mortality(X0, X1)) | ~ (position(X0, X2) = position(X0, X1)) | ~ (capability(X0, X2) = capability(X0, X1))) & (smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | ~ greater(position(X0, X2), position(X0, X1)) | ~ greater_or_equal(capability(X0, X2), capability(X0, X1))) & (smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | ~ greater_or_equal(position(X0, X2), position(X0, X1)) | ~ greater(capability(X0, X2), capability(X0, X1)))) | has_immunity(X0, X2) | has_immunity(X0, X1) | ~ organization(X0)), inference(flattening, [], [f19])).
fof(f19, plain, ! [X0, X1, X2] : ((((hazard_of_mortality(X0, X2) = hazard_of_mortality(X0, X1)) | (~ (position(X0, X2) = position(X0, X1)) | ~ (capability(X0, X2) = capability(X0, X1)))) & (smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | (~ greater(position(X0, X2), position(X0, X1)) | ~ greater_or_equal(capability(X0, X2), capability(X0, X1)))) & (smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | (~ greater_or_equal(position(X0, X2), position(X0, X1)) | ~ greater(capability(X0, X2), capability(X0, X1))))) | (has_immunity(X0, X2) | has_immunity(X0, X1) | ~ organization(X0))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0, X1, X2] : ((~ has_immunity(X0, X2) & ~ has_immunity(X0, X1) & organization(X0)) => ((((position(X0, X2) = position(X0, X1)) & (capability(X0, X2) = capability(X0, X1))) => (hazard_of_mortality(X0, X2) = hazard_of_mortality(X0, X1))) & ((greater(position(X0, X2), position(X0, X1)) & greater_or_equal(capability(X0, X2), capability(X0, X1))) => smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1))) & ((greater_or_equal(position(X0, X2), position(X0, X1)) & greater(capability(X0, X2), capability(X0, X1))) => smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1))))), inference(rectify, [], [f7])).
fof(f7, plain, ! [X0, X3, X4] : ((~ has_immunity(X0, X4) & ~ has_immunity(X0, X3) & organization(X0)) => ((((position(X0, X4) = position(X0, X3)) & (capability(X0, X4) = capability(X0, X3))) => (hazard_of_mortality(X0, X4) = hazard_of_mortality(X0, X3))) & ((greater(position(X0, X4), position(X0, X3)) & greater_or_equal(capability(X0, X4), capability(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))) & ((greater_or_equal(position(X0, X4), position(X0, X3)) & greater(capability(X0, X4), capability(X0, X3))) => smaller(hazard_of_mortality(X0, X4), hazard_of_mortality(X0, X3))))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', assumption_4)).
fof(f1434, plain, (greater(capability(sK1, sK3), capability(sK1, sK2)) | ~ spl4_53), inference(avatar_component_clause, [], [f1432])).
fof(f1432, plain, (spl4_53 <=> greater(capability(sK1, sK3), capability(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_53])])).
fof(f1520, plain, (spl4_13 | ~ spl4_1 | spl4_2), inference(avatar_split_clause, [], [f1519, f75, f71, f308])).
fof(f71, plain, (spl4_1 <=> sP0(sK2, sK1, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl4_1])])).
fof(f75, plain, (spl4_2 <=> greater(external_ties(sK1, sK3), external_ties(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_2])])).
fof(f1519, plain, ((external_ties(sK1, sK3) = external_ties(sK1, sK2)) | (~ spl4_1 | spl4_2)), inference(subsumption_resolution, [], [f1515, f73])).
fof(f73, plain, (sP0(sK2, sK1, sK3) | ~ spl4_1), inference(avatar_component_clause, [], [f71])).
fof(f1515, plain, (~ sP0(sK2, sK1, sK3) | (external_ties(sK1, sK3) = external_ties(sK1, sK2)) | spl4_2), inference(resolution, [], [f76, f160])).
fof(f160, plain, ! [X2, X0, X1] : (greater(external_ties(X1, X2), external_ties(X1, X0)) | ~ sP0(X0, X1, X2) | (external_ties(X1, X2) = external_ties(X1, X0))), inference(resolution, [], [f58, f40])).
fof(f40, plain, ! [X0, X1] : (~ greater_or_equal(X0, X1) | greater(X0, X1) | (X0 = X1)), inference(cnf_transformation, [], [f31])).
fof(f58, plain, ! [X2, X0, X1] : (greater_or_equal(external_ties(X1, X2), external_ties(X1, X0)) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ! [X0, X1, X2] : ((~ smaller(hazard_of_mortality(X1, X2), hazard_of_mortality(X1, X0)) & greater_or_equal(external_ties(X1, X2), external_ties(X1, X0)) & smaller_or_equal(internal_friction(X1, X2), internal_friction(X1, X0)) & greater(stock_of_knowledge(X1, X2), stock_of_knowledge(X1, X0))) | ~ sP0(X0, X1, X2)), inference(rectify, [], [f33])).
fof(f33, plain, ! [X1, X0, X2] : ((~ smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) & greater_or_equal(external_ties(X0, X2), external_ties(X0, X1)) & smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) | ~ sP0(X1, X0, X2)), inference(nnf_transformation, [], [f26])).
fof(f76, plain, (~ greater(external_ties(sK1, sK3), external_ties(sK1, sK2)) | spl4_2), inference(avatar_component_clause, [], [f75])).
fof(f1500, plain, (~ spl4_2 | spl4_47 | ~ spl4_53), inference(avatar_contradiction_clause, [], [f1499])).
fof(f1499, plain, ($false | (~ spl4_2 | spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1498, f60])).
fof(f1498, plain, (~ organization(sK1) | (~ spl4_2 | spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1497, f61])).
fof(f1497, plain, (has_immunity(sK1, sK2) | ~ organization(sK1) | (~ spl4_2 | spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1496, f62])).
fof(f1496, plain, (has_immunity(sK1, sK3) | has_immunity(sK1, sK2) | ~ organization(sK1) | (~ spl4_2 | spl4_47 | ~ spl4_53)), inference(subsumption_resolution, [], [f1495, f1254])).
fof(f1495, plain, (smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | has_immunity(sK1, sK3) | has_immunity(sK1, sK2) | ~ organization(sK1) | (~ spl4_2 | ~ spl4_53)), inference(subsumption_resolution, [], [f1487, f1371])).
fof(f1371, plain, (greater_or_equal(position(sK1, sK3), position(sK1, sK2)) | ~ spl4_2), inference(resolution, [], [f1348, f41])).
fof(f41, plain, ! [X0, X1] : (~ greater(X0, X1) | greater_or_equal(X0, X1)), inference(cnf_transformation, [], [f31])).
fof(f1348, plain, (greater(position(sK1, sK3), position(sK1, sK2)) | ~ spl4_2), inference(subsumption_resolution, [], [f1341, f60])).
fof(f1341, plain, (greater(position(sK1, sK3), position(sK1, sK2)) | ~ organization(sK1) | ~ spl4_2), inference(resolution, [], [f77, f54])).
fof(f54, plain, ! [X2, X0, X1] : (~ greater(external_ties(X0, X2), external_ties(X0, X1)) | greater(position(X0, X2), position(X0, X1)) | ~ organization(X0)), inference(cnf_transformation, [], [f23])).
fof(f77, plain, (greater(external_ties(sK1, sK3), external_ties(sK1, sK2)) | ~ spl4_2), inference(avatar_component_clause, [], [f75])).
fof(f1439, plain, (spl4_12 | spl4_11 | spl4_46), inference(avatar_split_clause, [], [f1438, f1249, f241, f245])).
fof(f245, plain, (spl4_12 <=> smaller(capability(sK1, sK3), capability(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_12])])).
fof(f241, plain, (spl4_11 <=> (capability(sK1, sK3) = capability(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_11])])).
fof(f1249, plain, (spl4_46 <=> greater_or_equal(capability(sK1, sK3), capability(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_46])])).
fof(f1438, plain, (smaller(capability(sK1, sK3), capability(sK1, sK2)) | (spl4_11 | spl4_46)), inference(subsumption_resolution, [], [f1423, f242])).
fof(f242, plain, (~ (capability(sK1, sK3) = capability(sK1, sK2)) | spl4_11), inference(avatar_component_clause, [], [f241])).
fof(f1423, plain, (smaller(capability(sK1, sK3), capability(sK1, sK2)) | (capability(sK1, sK3) = capability(sK1, sK2)) | spl4_46), inference(resolution, [], [f1251, f118])).
fof(f118, plain, ! [X8, X7] : (greater_or_equal(X7, X8) | smaller(X7, X8) | (X7 = X8)), inference(resolution, [], [f47, f41])).
fof(f47, plain, ! [X0, X1] : (greater(X0, X1) | (X0 = X1) | smaller(X0, X1)), inference(cnf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1] : (greater(X0, X1) | (X0 = X1) | smaller(X0, X1)), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', meaning_postulate_greater_comparable)).
fof(f1251, plain, (~ greater_or_equal(capability(sK1, sK3), capability(sK1, sK2)) | spl4_46), inference(avatar_component_clause, [], [f1249])).
fof(f1437, plain, (spl4_10 | spl4_6 | spl4_53), inference(avatar_split_clause, [], [f1436, f1432, f105, f152])).
fof(f152, plain, (spl4_10 <=> smaller(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_10])])).
fof(f105, plain, (spl4_6 <=> (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_6])])).
fof(f1436, plain, (greater(capability(sK1, sK3), capability(sK1, sK2)) | (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, sK2)) | smaller(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))), inference(subsumption_resolution, [], [f1203, f60])).
fof(f1203, plain, (greater(capability(sK1, sK3), capability(sK1, sK2)) | ~ organization(sK1) | (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, sK2)) | smaller(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))), inference(resolution, [], [f79, f198])).
fof(f198, plain, ! [X4, X5, X3] : (~ smaller_or_equal(internal_friction(X3, X4), internal_friction(X3, X5)) | greater(capability(X3, X4), capability(X3, X5)) | ~ organization(X3) | (stock_of_knowledge(X3, X4) = stock_of_knowledge(X3, X5)) | smaller(stock_of_knowledge(X3, X4), stock_of_knowledge(X3, X5))), inference(resolution, [], [f51, f47])).
fof(f51, plain, ! [X2, X0, X1] : (~ greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)) | ~ smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) | greater(capability(X0, X2), capability(X0, X1)) | ~ organization(X0)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ! [X0, X1, X2] : ((((capability(X0, X2) = capability(X0, X1)) | ~ (internal_friction(X0, X2) = internal_friction(X0, X1)) | ~ (stock_of_knowledge(X0, X2) = stock_of_knowledge(X0, X1))) & (smaller(capability(X0, X2), capability(X0, X1)) | ~ greater(internal_friction(X0, X2), internal_friction(X0, X1)) | ~ smaller_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) & (greater(capability(X0, X2), capability(X0, X1)) | ~ smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) | ~ greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)))) | ~ organization(X0)), inference(flattening, [], [f21])).
fof(f21, plain, ! [X0, X1, X2] : ((((capability(X0, X2) = capability(X0, X1)) | (~ (internal_friction(X0, X2) = internal_friction(X0, X1)) | ~ (stock_of_knowledge(X0, X2) = stock_of_knowledge(X0, X1)))) & (smaller(capability(X0, X2), capability(X0, X1)) | (~ greater(internal_friction(X0, X2), internal_friction(X0, X1)) | ~ smaller_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)))) & (greater(capability(X0, X2), capability(X0, X1)) | (~ smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) | ~ greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))))) | ~ organization(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X1, X2] : (organization(X0) => ((((internal_friction(X0, X2) = internal_friction(X0, X1)) & (stock_of_knowledge(X0, X2) = stock_of_knowledge(X0, X1))) => (capability(X0, X2) = capability(X0, X1))) & ((greater(internal_friction(X0, X2), internal_friction(X0, X1)) & smaller_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) => smaller(capability(X0, X2), capability(X0, X1))) & ((smaller_or_equal(internal_friction(X0, X2), internal_friction(X0, X1)) & greater(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1))) => greater(capability(X0, X2), capability(X0, X1))))), inference(rectify, [], [f8])).
fof(f8, plain, ! [X0, X3, X4] : (organization(X0) => ((((internal_friction(X0, X4) = internal_friction(X0, X3)) & (stock_of_knowledge(X0, X4) = stock_of_knowledge(X0, X3))) => (capability(X0, X4) = capability(X0, X3))) & ((greater(internal_friction(X0, X4), internal_friction(X0, X3)) & smaller_or_equal(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => smaller(capability(X0, X4), capability(X0, X3))) & ((smaller_or_equal(internal_friction(X0, X4), internal_friction(X0, X3)) & greater(stock_of_knowledge(X0, X4), stock_of_knowledge(X0, X3))) => greater(capability(X0, X4), capability(X0, X3))))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', assumption_5)).
fof(f79, plain, smaller_or_equal(internal_friction(sK1, sK3), internal_friction(sK1, sK2)), inference(subsumption_resolution, [], [f64, f57])).
fof(f57, plain, ! [X2, X0, X1] : (smaller_or_equal(internal_friction(X1, X2), internal_friction(X1, X0)) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f34])).
fof(f64, plain, (smaller_or_equal(internal_friction(sK1, sK3), internal_friction(sK1, sK2)) | sP0(sK2, sK1, sK3)), inference(cnf_transformation, [], [f36])).
fof(f1320, plain, (~ spl4_1 | ~ spl4_6), inference(avatar_split_clause, [], [f457, f105, f71])).
fof(f457, plain, (~ sP0(sK2, sK1, sK3) | ~ spl4_6), inference(duplicate_literal_removal, [], [f452])).
fof(f452, plain, (~ sP0(sK2, sK1, sK3) | ~ sP0(sK2, sK1, sK3) | ~ spl4_6), inference(resolution, [], [f252, f191])).
fof(f191, plain, (! [X1] : (greater(stock_of_knowledge(sK1, sK2), stock_of_knowledge(sK1, X1)) | ~ sP0(X1, sK1, sK3)) | ~ spl4_6), inference(superposition, [], [f56, f107])).
fof(f107, plain, ((stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, sK2)) | ~ spl4_6), inference(avatar_component_clause, [], [f105])).
fof(f56, plain, ! [X2, X0, X1] : (greater(stock_of_knowledge(X1, X2), stock_of_knowledge(X1, X0)) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f34])).
fof(f252, plain, (! [X0] : (~ greater(stock_of_knowledge(sK1, X0), stock_of_knowledge(sK1, sK2)) | ~ sP0(X0, sK1, sK3)) | ~ spl4_6), inference(superposition, [], [f122, f107])).
fof(f122, plain, ! [X6, X4, X5] : (~ greater(stock_of_knowledge(X5, X4), stock_of_knowledge(X5, X6)) | ~ sP0(X4, X5, X6)), inference(resolution, [], [f56, f45])).
fof(f45, plain, ! [X0, X1] : (~ greater(X1, X0) | ~ greater(X0, X1)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0, X1] : (~ greater(X1, X0) | ~ greater(X0, X1)), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0, X1] : ~ (greater(X1, X0) & greater(X0, X1)), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', meaning_postulate_greater_strict)).
fof(f1314, plain, (~ spl4_11 | spl4_46), inference(avatar_contradiction_clause, [], [f1313])).
fof(f1313, plain, ($false | (~ spl4_11 | spl4_46)), inference(subsumption_resolution, [], [f1312, f68])).
fof(f1312, plain, (~ greater_or_equal(capability(sK1, sK2), capability(sK1, sK2)) | (~ spl4_11 | spl4_46)), inference(forward_demodulation, [], [f1251, f243])).
fof(f243, plain, ((capability(sK1, sK3) = capability(sK1, sK2)) | ~ spl4_11), inference(avatar_component_clause, [], [f241])).
fof(f1257, plain, ~ spl4_47, inference(avatar_split_clause, [], [f69, f1253])).
fof(f69, plain, ~ smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)), inference(subsumption_resolution, [], [f66, f59])).
fof(f59, plain, ! [X2, X0, X1] : (~ smaller(hazard_of_mortality(X1, X2), hazard_of_mortality(X1, X0)) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f34])).
fof(f66, plain, (~ smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | sP0(sK2, sK1, sK3)), inference(cnf_transformation, [], [f36])).
fof(f1256, plain, (~ spl4_46 | spl4_47 | ~ spl4_2), inference(avatar_split_clause, [], [f230, f75, f1253, f1249])).
fof(f230, plain, (smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | ~ greater_or_equal(capability(sK1, sK3), capability(sK1, sK2)) | ~ spl4_2), inference(subsumption_resolution, [], [f229, f60])).
fof(f229, plain, (smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | ~ greater_or_equal(capability(sK1, sK3), capability(sK1, sK2)) | ~ organization(sK1) | ~ spl4_2), inference(subsumption_resolution, [], [f228, f61])).
fof(f228, plain, (smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | ~ greater_or_equal(capability(sK1, sK3), capability(sK1, sK2)) | has_immunity(sK1, sK2) | ~ organization(sK1) | ~ spl4_2), inference(subsumption_resolution, [], [f226, f62])).
fof(f226, plain, (smaller(hazard_of_mortality(sK1, sK3), hazard_of_mortality(sK1, sK2)) | ~ greater_or_equal(capability(sK1, sK3), capability(sK1, sK2)) | has_immunity(sK1, sK3) | has_immunity(sK1, sK2) | ~ organization(sK1) | ~ spl4_2), inference(resolution, [], [f49, f174])).
fof(f174, plain, (greater(position(sK1, sK3), position(sK1, sK2)) | ~ spl4_2), inference(subsumption_resolution, [], [f172, f60])).
fof(f172, plain, (greater(position(sK1, sK3), position(sK1, sK2)) | ~ organization(sK1) | ~ spl4_2), inference(resolution, [], [f54, f77])).
fof(f49, plain, ! [X2, X0, X1] : (~ greater(position(X0, X2), position(X0, X1)) | smaller(hazard_of_mortality(X0, X2), hazard_of_mortality(X0, X1)) | ~ greater_or_equal(capability(X0, X2), capability(X0, X1)) | has_immunity(X0, X2) | has_immunity(X0, X1) | ~ organization(X0)), inference(cnf_transformation, [], [f20])).
fof(f1242, plain, (spl4_4 | spl4_5), inference(avatar_split_clause, [], [f1204, f95, f91])).
fof(f91, plain, (spl4_4 <=> (internal_friction(sK1, sK3) = internal_friction(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_4])])).
fof(f95, plain, (spl4_5 <=> smaller(internal_friction(sK1, sK3), internal_friction(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_5])])).
fof(f1204, plain, (smaller(internal_friction(sK1, sK3), internal_friction(sK1, sK2)) | (internal_friction(sK1, sK3) = internal_friction(sK1, sK2))), inference(resolution, [], [f79, f37])).
fof(f37, plain, ! [X0, X1] : (~ smaller_or_equal(X0, X1) | smaller(X0, X1) | (X0 = X1)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1] : ((smaller_or_equal(X0, X1) | (~ (X0 = X1) & ~ smaller(X0, X1))) & ((X0 = X1) | smaller(X0, X1) | ~ smaller_or_equal(X0, X1))), inference(flattening, [], [f28])).
fof(f28, plain, ! [X0, X1] : ((smaller_or_equal(X0, X1) | (~ (X0 = X1) & ~ smaller(X0, X1))) & (((X0 = X1) | smaller(X0, X1)) | ~ smaller_or_equal(X0, X1))), inference(nnf_transformation, [], [f1])).
fof(f1, plain, ! [X0, X1] : (smaller_or_equal(X0, X1) <=> ((X0 = X1) | smaller(X0, X1))), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', definition_smaller_or_equal)).
fof(f1241, plain, (~ spl4_5 | ~ spl4_6 | ~ spl4_12), inference(avatar_contradiction_clause, [], [f1240])).
fof(f1240, plain, ($false | (~ spl4_5 | ~ spl4_6 | ~ spl4_12)), inference(subsumption_resolution, [], [f1238, f1055])).
fof(f1055, plain, (~ greater(capability(sK1, sK3), capability(sK1, sK2)) | ~ spl4_12), inference(resolution, [], [f719, f45])).
fof(f719, plain, (greater(capability(sK1, sK2), capability(sK1, sK3)) | ~ spl4_12), inference(resolution, [], [f247, f43])).
fof(f43, plain, ! [X0, X1] : (~ smaller(X0, X1) | greater(X1, X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0, X1] : ((smaller(X0, X1) | ~ greater(X1, X0)) & (greater(X1, X0) | ~ smaller(X0, X1))), inference(nnf_transformation, [], [f3])).
fof(f3, plain, ! [X0, X1] : (smaller(X0, X1) <=> greater(X1, X0)), file('/home/ubuntu/library/tptp/Problems/MGT/MGT042+1.p', definition_smaller)).
fof(f247, plain, (smaller(capability(sK1, sK3), capability(sK1, sK2)) | ~ spl4_12), inference(avatar_component_clause, [], [f245])).
fof(f1238, plain, (greater(capability(sK1, sK3), capability(sK1, sK2)) | (~ spl4_5 | ~ spl4_6)), inference(resolution, [], [f1216, f43])).
fof(f1216, plain, (smaller(capability(sK1, sK2), capability(sK1, sK3)) | (~ spl4_5 | ~ spl4_6)), inference(subsumption_resolution, [], [f1215, f67])).
fof(f67, plain, ! [X1] : smaller_or_equal(X1, X1), inference(equality_resolution, [], [f39])).
fof(f39, plain, ! [X0, X1] : (smaller_or_equal(X0, X1) | ~ (X0 = X1)), inference(cnf_transformation, [], [f29])).
fof(f1215, plain, (~ smaller_or_equal(stock_of_knowledge(sK1, sK2), stock_of_knowledge(sK1, sK2)) | smaller(capability(sK1, sK2), capability(sK1, sK3)) | (~ spl4_5 | ~ spl4_6)), inference(forward_demodulation, [], [f1214, f107])).
fof(f1214, plain, (smaller(capability(sK1, sK2), capability(sK1, sK3)) | ~ smaller_or_equal(stock_of_knowledge(sK1, sK2), stock_of_knowledge(sK1, sK3)) | ~ spl4_5), inference(subsumption_resolution, [], [f1207, f60])).
fof(f1207, plain, (smaller(capability(sK1, sK2), capability(sK1, sK3)) | ~ smaller_or_equal(stock_of_knowledge(sK1, sK2), stock_of_knowledge(sK1, sK3)) | ~ organization(sK1) | ~ spl4_5), inference(resolution, [], [f1205, f52])).
fof(f52, plain, ! [X2, X0, X1] : (~ greater(internal_friction(X0, X2), internal_friction(X0, X1)) | smaller(capability(X0, X2), capability(X0, X1)) | ~ smaller_or_equal(stock_of_knowledge(X0, X2), stock_of_knowledge(X0, X1)) | ~ organization(X0)), inference(cnf_transformation, [], [f22])).
fof(f1205, plain, (greater(internal_friction(sK1, sK2), internal_friction(sK1, sK3)) | ~ spl4_5), inference(resolution, [], [f97, f43])).
fof(f97, plain, (smaller(internal_friction(sK1, sK3), internal_friction(sK1, sK2)) | ~ spl4_5), inference(avatar_component_clause, [], [f95])).
fof(f1199, plain, (~ spl4_4 | ~ spl4_6 | spl4_11), inference(avatar_contradiction_clause, [], [f1198])).
fof(f1198, plain, ($false | (~ spl4_4 | ~ spl4_6 | spl4_11)), inference(subsumption_resolution, [], [f1197, f242])).
fof(f1197, plain, ((capability(sK1, sK3) = capability(sK1, sK2)) | (~ spl4_4 | ~ spl4_6)), inference(trivial_inequality_removal, [], [f1196])).
fof(f1196, plain, (~ (stock_of_knowledge(sK1, sK2) = stock_of_knowledge(sK1, sK2)) | (capability(sK1, sK3) = capability(sK1, sK2)) | (~ spl4_4 | ~ spl4_6)), inference(equality_resolution, [], [f219])).
fof(f219, plain, (! [X0] : (~ (internal_friction(sK1, sK2) = internal_friction(sK1, X0)) | ~ (stock_of_knowledge(sK1, sK2) = stock_of_knowledge(sK1, X0)) | (capability(sK1, sK3) = capability(sK1, X0))) | (~ spl4_4 | ~ spl4_6)), inference(forward_demodulation, [], [f218, f107])).
fof(f218, plain, (! [X0] : (~ (internal_friction(sK1, sK2) = internal_friction(sK1, X0)) | (capability(sK1, sK3) = capability(sK1, X0)) | ~ (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, X0))) | ~ spl4_4), inference(subsumption_resolution, [], [f215, f60])).
fof(f215, plain, (! [X0] : (~ (internal_friction(sK1, sK2) = internal_friction(sK1, X0)) | (capability(sK1, sK3) = capability(sK1, X0)) | ~ (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, X0)) | ~ organization(sK1)) | ~ spl4_4), inference(superposition, [], [f53, f93])).
fof(f93, plain, ((internal_friction(sK1, sK3) = internal_friction(sK1, sK2)) | ~ spl4_4), inference(avatar_component_clause, [], [f91])).
fof(f53, plain, ! [X2, X0, X1] : (~ (internal_friction(X0, X2) = internal_friction(X0, X1)) | (capability(X0, X2) = capability(X0, X1)) | ~ (stock_of_knowledge(X0, X2) = stock_of_knowledge(X0, X1)) | ~ organization(X0)), inference(cnf_transformation, [], [f22])).
fof(f179, plain, (~ spl4_1 | spl4_7), inference(avatar_split_clause, [], [f175, f109, f71])).
fof(f109, plain, (spl4_7 <=> greater(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_7])])).
fof(f175, plain, (~ sP0(sK2, sK1, sK3) | spl4_7), inference(resolution, [], [f110, f56])).
fof(f110, plain, (~ greater(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | spl4_7), inference(avatar_component_clause, [], [f109])).
fof(f171, plain, (~ spl4_7 | ~ spl4_10), inference(avatar_split_clause, [], [f168, f152, f109])).
fof(f168, plain, (~ greater(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | ~ spl4_10), inference(resolution, [], [f161, f45])).
fof(f161, plain, (greater(stock_of_knowledge(sK1, sK2), stock_of_knowledge(sK1, sK3)) | ~ spl4_10), inference(resolution, [], [f154, f43])).
fof(f154, plain, (smaller(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | ~ spl4_10), inference(avatar_component_clause, [], [f152])).
fof(f112, plain, (spl4_6 | spl4_7 | ~ spl4_3), inference(avatar_split_clause, [], [f103, f81, f109, f105])).
fof(f81, plain, (spl4_3 <=> greater_or_equal(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl4_3])])).
fof(f103, plain, (greater(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | (stock_of_knowledge(sK1, sK3) = stock_of_knowledge(sK1, sK2)) | ~ spl4_3), inference(resolution, [], [f40, f83])).
fof(f83, plain, (greater_or_equal(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | ~ spl4_3), inference(avatar_component_clause, [], [f81])).
fof(f84, plain, (spl4_1 | spl4_3), inference(avatar_split_clause, [], [f63, f81, f71])).
fof(f63, plain, (greater_or_equal(stock_of_knowledge(sK1, sK3), stock_of_knowledge(sK1, sK2)) | sP0(sK2, sK1, sK3)), inference(cnf_transformation, [], [f36])).
fof(f78, plain, (spl4_1 | spl4_2), inference(avatar_split_clause, [], [f65, f75, f71])).
fof(f65, plain, (greater(external_ties(sK1, sK3), external_ties(sK1, sK2)) | sP0(sK2, sK1, sK3)), inference(cnf_transformation, [], [f36])).