fof(f4163, plain, $false, inference(avatar_sat_refutation, [], [f161, f166, f171, f176, f181, f190, f195, f200, f205, f210, f219, f224, f229, f234, f239, f248, f253, f258, f263, f268, f277, f282, f287, f292, f297, f306, f311, f316, f321, f326, f335, f340, f345, f350, f355, f364, f369, f374, f379, f384, f389, f398, f403, f408, f413, f423, f464, f469, f474, f479, f484, f493, f494, f495, f496, f497, f498, f499, f500, f792, f1052, f1334, f1417, f1737, f2146, f2623, f3053, f3459, f4149, f4162])).
fof(f4162, plain, (spl31_51 | ~ spl31_52 | ~ spl31_53 | ~ spl31_54 | ~ spl31_56 | ~ spl31_71), inference(avatar_contradiction_clause, [], [f4161])).
fof(f4161, plain, ($false | (spl31_51 | ~ spl31_52 | ~ spl31_53 | ~ spl31_54 | ~ spl31_56 | ~ spl31_71)), inference(subsumption_resolution, [], [f4160, f4157])).
fof(f4157, plain, (~ (init = a_select3(id_ds1_filter_init, pv63, sK26)) | (spl31_51 | ~ spl31_53)), inference(forward_demodulation, [], [f397, f407])).
fof(f407, plain, ((pv63 = sK25) | ~ spl31_53), inference(avatar_component_clause, [], [f405])).
fof(f405, plain, (spl31_53 <=> (pv63 = sK25)), introduced(avatar_definition, [new_symbols(naming, [spl31_53])])).
fof(f397, plain, (~ (init = a_select3(id_ds1_filter_init, sK25, sK26)) | spl31_51), inference(avatar_component_clause, [], [f395])).
fof(f395, plain, (spl31_51 <=> (init = a_select3(id_ds1_filter_init, sK25, sK26))), introduced(avatar_definition, [new_symbols(naming, [spl31_51])])).
fof(f4160, plain, ((init = a_select3(id_ds1_filter_init, pv63, sK26)) | (~ spl31_52 | ~ spl31_54 | ~ spl31_56 | ~ spl31_71)), inference(subsumption_resolution, [], [f4159, f412])).
fof(f412, plain, (leq(sK26, n5) | ~ spl31_54), inference(avatar_component_clause, [], [f410])).
fof(f410, plain, (spl31_54 <=> leq(sK26, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_54])])).
fof(f4159, plain, (~ leq(sK26, n5) | (init = a_select3(id_ds1_filter_init, pv63, sK26)) | (~ spl31_52 | ~ spl31_56 | ~ spl31_71)), inference(subsumption_resolution, [], [f4158, f422])).
fof(f422, plain, (leq(n0, sK26) | ~ spl31_56), inference(avatar_component_clause, [], [f420])).
fof(f420, plain, (spl31_56 <=> leq(n0, sK26)), introduced(avatar_definition, [new_symbols(naming, [spl31_56])])).
fof(f4158, plain, (~ leq(n0, sK26) | ~ leq(sK26, n5) | (init = a_select3(id_ds1_filter_init, pv63, sK26)) | (~ spl31_52 | ~ spl31_71)), inference(resolution, [], [f402, f492])).
fof(f492, plain, (! [X7] : (~ gt(pv64, X7) | ~ leq(n0, X7) | ~ leq(X7, n5) | (init = a_select3(id_ds1_filter_init, pv63, X7))) | ~ spl31_71), inference(avatar_component_clause, [], [f491])).
fof(f491, plain, (spl31_71 <=> ! [X7] : ((init = a_select3(id_ds1_filter_init, pv63, X7)) | ~ leq(n0, X7) | ~ leq(X7, n5) | ~ gt(pv64, X7))), introduced(avatar_definition, [new_symbols(naming, [spl31_71])])).
fof(f402, plain, (gt(pv64, sK26) | ~ spl31_52), inference(avatar_component_clause, [], [f400])).
fof(f400, plain, (spl31_52 <=> gt(pv64, sK26)), introduced(avatar_definition, [new_symbols(naming, [spl31_52])])).
fof(f4149, plain, (spl31_63 | ~ spl31_67 | ~ spl31_68 | ~ spl31_69 | ~ spl31_70), inference(avatar_contradiction_clause, [], [f4148])).
fof(f4148, plain, ($false | (spl31_63 | ~ spl31_67 | ~ spl31_68 | ~ spl31_69 | ~ spl31_70)), inference(subsumption_resolution, [], [f4147, f452])).
fof(f452, plain, (~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) | spl31_63), inference(avatar_component_clause, [], [f450])).
fof(f450, plain, (spl31_63 <=> (init = a_select3(id_ds1_filter_init, sK27, sK28))), introduced(avatar_definition, [new_symbols(naming, [spl31_63])])).
fof(f4147, plain, ((init = a_select3(id_ds1_filter_init, sK27, sK28)) | (~ spl31_67 | ~ spl31_68 | ~ spl31_69 | ~ spl31_70)), inference(subsumption_resolution, [], [f4136, f468])).
fof(f468, plain, (leq(sK28, n5) | ~ spl31_67), inference(avatar_component_clause, [], [f466])).
fof(f466, plain, (spl31_67 <=> leq(sK28, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_67])])).
fof(f4136, plain, (~ leq(sK28, n5) | (init = a_select3(id_ds1_filter_init, sK27, sK28)) | (~ spl31_68 | ~ spl31_69 | ~ spl31_70)), inference(resolution, [], [f3916, f473])).
fof(f473, plain, (leq(n0, sK28) | ~ spl31_68), inference(avatar_component_clause, [], [f471])).
fof(f471, plain, (spl31_68 <=> leq(n0, sK28)), introduced(avatar_definition, [new_symbols(naming, [spl31_68])])).
fof(f3916, plain, (! [X0] : (~ leq(n0, X0) | ~ leq(X0, n5) | (init = a_select3(id_ds1_filter_init, sK27, X0))) | (~ spl31_69 | ~ spl31_70)), inference(subsumption_resolution, [], [f3915, f483])).
fof(f483, plain, (leq(n0, sK27) | ~ spl31_70), inference(avatar_component_clause, [], [f481])).
fof(f481, plain, (spl31_70 <=> leq(n0, sK27)), introduced(avatar_definition, [new_symbols(naming, [spl31_70])])).
fof(f3915, plain, (! [X0] : (~ leq(X0, n5) | ~ leq(n0, X0) | (init = a_select3(id_ds1_filter_init, sK27, X0)) | ~ leq(n0, sK27)) | ~ spl31_69), inference(resolution, [], [f478, f133])).
fof(f133, plain, ! [X2, X3] : (~ leq(X2, pred(pv63)) | ~ leq(X3, n5) | ~ leq(n0, X3) | (init = a_select3(id_ds1_filter_init, X2, X3)) | ~ leq(n0, X2)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ((~ (init = a_select3(id_ds1_filter_init, sK29, sK30)) & leq(sK30, pv64) & (pv63 = sK29) & ~ (pv64 = sK30) & leq(sK30, n5) & leq(sK29, n5) & leq(n0, sK30) & leq(n0, sK29)) & ! [X2] : (! [X3] : ((init = a_select3(id_ds1_filter_init, X2, X3)) | ~ leq(X3, n5) | ~ leq(n0, X3)) | ~ leq(X2, pred(pv63)) | ~ leq(n0, X2)) & ! [X4, X5] : ((init = a_select3(id_ds1_filter_init, X4, X5)) | ~ gt(pv63, X4) | ~ leq(X5, n5) | ~ leq(X4, n5) | ~ leq(n0, X5) | ~ leq(n0, X4)) & ! [X6, X7] : ((init = a_select3(id_ds1_filter_init, X6, X7)) | ~ gt(pv64, X7) | ~ (pv63 = X6) | ~ leq(X7, n5) | ~ leq(X6, n5) | ~ leq(n0, X7) | ~ leq(n0, X6)) & ! [X8, X9] : ((init = a_select3(pminus_ds1_filter_init, X8, X9)) | ~ leq(X9, n5) | ~ leq(X8, n5) | ~ leq(n0, X9) | ~ leq(n0, X8)) & ! [X10, X11] : ((init = a_select3(xhatmin_ds1_filter_init, X10, X11)) | ~ leq(X11, n0) | ~ leq(X10, n5) | ~ leq(n0, X11) | ~ leq(n0, X10)) & ! [X12, X13] : ((init = a_select3(r_ds1_filter_init, X12, X13)) | ~ leq(X13, n2) | ~ leq(X12, n2) | ~ leq(n0, X13) | ~ leq(n0, X12)) & ! [X14, X15] : ((init = a_select3(q_ds1_filter_init, X14, X15)) | ~ leq(X15, n5) | ~ leq(X14, n5) | ~ leq(n0, X15) | ~ leq(n0, X14)) & ! [X16, X17] : ((init = a_select3(dv_ds1_filter_init, X16, X17)) | ~ leq(X17, n0) | ~ leq(X16, n5) | ~ leq(n0, X17) | ~ leq(n0, X16)) & ! [X18, X19] : ((init = a_select3(phi_ds1_filter_init, X18, X19)) | ~ leq(X19, n5) | ~ leq(X18, n5) | ~ leq(n0, X19) | ~ leq(n0, X18)) & ! [X20, X21] : ((init = a_select3(h_ds1_filter_init, X20, X21)) | ~ leq(X21, n5) | ~ leq(X20, n2) | ~ leq(n0, X21) | ~ leq(n0, X20)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30])], [f61, f62])).
fof(f62, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & leq(X1, pv64) & (pv63 = X0) & ~ (pv64 = X1) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(id_ds1_filter_init, sK29, sK30)) & leq(sK30, pv64) & (pv63 = sK29) & ~ (pv64 = sK30) & leq(sK30, n5) & leq(sK29, n5) & leq(n0, sK30) & leq(n0, sK29))), introduced(choice_axiom, [])).
fof(f61, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & leq(X1, pv64) & (pv63 = X0) & ~ (pv64 = X1) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) & ! [X2] : (! [X3] : ((init = a_select3(id_ds1_filter_init, X2, X3)) | ~ leq(X3, n5) | ~ leq(n0, X3)) | ~ leq(X2, pred(pv63)) | ~ leq(n0, X2)) & ! [X4, X5] : ((init = a_select3(id_ds1_filter_init, X4, X5)) | ~ gt(pv63, X4) | ~ leq(X5, n5) | ~ leq(X4, n5) | ~ leq(n0, X5) | ~ leq(n0, X4)) & ! [X6, X7] : ((init = a_select3(id_ds1_filter_init, X6, X7)) | ~ gt(pv64, X7) | ~ (pv63 = X6) | ~ leq(X7, n5) | ~ leq(X6, n5) | ~ leq(n0, X7) | ~ leq(n0, X6)) & ! [X8, X9] : ((init = a_select3(pminus_ds1_filter_init, X8, X9)) | ~ leq(X9, n5) | ~ leq(X8, n5) | ~ leq(n0, X9) | ~ leq(n0, X8)) & ! [X10, X11] : ((init = a_select3(xhatmin_ds1_filter_init, X10, X11)) | ~ leq(X11, n0) | ~ leq(X10, n5) | ~ leq(n0, X11) | ~ leq(n0, X10)) & ! [X12, X13] : ((init = a_select3(r_ds1_filter_init, X12, X13)) | ~ leq(X13, n2) | ~ leq(X12, n2) | ~ leq(n0, X13) | ~ leq(n0, X12)) & ! [X14, X15] : ((init = a_select3(q_ds1_filter_init, X14, X15)) | ~ leq(X15, n5) | ~ leq(X14, n5) | ~ leq(n0, X15) | ~ leq(n0, X14)) & ! [X16, X17] : ((init = a_select3(dv_ds1_filter_init, X16, X17)) | ~ leq(X17, n0) | ~ leq(X16, n5) | ~ leq(n0, X17) | ~ leq(n0, X16)) & ! [X18, X19] : ((init = a_select3(phi_ds1_filter_init, X18, X19)) | ~ leq(X19, n5) | ~ leq(X18, n5) | ~ leq(n0, X19) | ~ leq(n0, X18)) & ! [X20, X21] : ((init = a_select3(h_ds1_filter_init, X20, X21)) | ~ leq(X21, n5) | ~ leq(X20, n2) | ~ leq(n0, X21) | ~ leq(n0, X20)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)), inference(rectify, [], [f10])).
fof(f10, plain, (? [X20, X21] : (~ (init = a_select3(id_ds1_filter_init, X20, X21)) & leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21) & leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) & ! [X0] : (! [X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, n5) | ~ leq(n0, X1)) | ~ leq(X0, pred(pv63)) | ~ leq(n0, X0)) & ! [X2, X3] : ((init = a_select3(id_ds1_filter_init, X2, X3)) | ~ gt(pv63, X2) | ~ leq(X3, n5) | ~ leq(X2, n5) | ~ leq(n0, X3) | ~ leq(n0, X2)) & ! [X4, X5] : ((init = a_select3(id_ds1_filter_init, X4, X5)) | ~ gt(pv64, X5) | ~ (pv63 = X4) | ~ leq(X5, n5) | ~ leq(X4, n5) | ~ leq(n0, X5) | ~ leq(n0, X4)) & ! [X6, X7] : ((init = a_select3(pminus_ds1_filter_init, X6, X7)) | ~ leq(X7, n5) | ~ leq(X6, n5) | ~ leq(n0, X7) | ~ leq(n0, X6)) & ! [X8, X9] : ((init = a_select3(xhatmin_ds1_filter_init, X8, X9)) | ~ leq(X9, n0) | ~ leq(X8, n5) | ~ leq(n0, X9) | ~ leq(n0, X8)) & ! [X10, X11] : ((init = a_select3(r_ds1_filter_init, X10, X11)) | ~ leq(X11, n2) | ~ leq(X10, n2) | ~ leq(n0, X11) | ~ leq(n0, X10)) & ! [X12, X13] : ((init = a_select3(q_ds1_filter_init, X12, X13)) | ~ leq(X13, n5) | ~ leq(X12, n5) | ~ leq(n0, X13) | ~ leq(n0, X12)) & ! [X14, X15] : ((init = a_select3(dv_ds1_filter_init, X14, X15)) | ~ leq(X15, n0) | ~ leq(X14, n5) | ~ leq(n0, X15) | ~ leq(n0, X14)) & ! [X16, X17] : ((init = a_select3(phi_ds1_filter_init, X16, X17)) | ~ leq(X17, n5) | ~ leq(X16, n5) | ~ leq(n0, X17) | ~ leq(n0, X16)) & ! [X18, X19] : ((init = a_select3(h_ds1_filter_init, X18, X19)) | ~ leq(X19, n5) | ~ leq(X18, n2) | ~ leq(n0, X19) | ~ leq(n0, X18)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)), inference(flattening, [], [f9])).
fof(f9, plain, (? [X20, X21] : ((~ (init = a_select3(id_ds1_filter_init, X20, X21)) & (leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21))) & (leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20))) & (! [X0] : (! [X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | (~ leq(X1, n5) | ~ leq(n0, X1))) | (~ leq(X0, pred(pv63)) | ~ leq(n0, X0))) & ! [X2, X3] : (((init = a_select3(id_ds1_filter_init, X2, X3)) | ~ gt(pv63, X2)) | (~ leq(X3, n5) | ~ leq(X2, n5) | ~ leq(n0, X3) | ~ leq(n0, X2))) & ! [X4, X5] : (((init = a_select3(id_ds1_filter_init, X4, X5)) | (~ gt(pv64, X5) | ~ (pv63 = X4))) | (~ leq(X5, n5) | ~ leq(X4, n5) | ~ leq(n0, X5) | ~ leq(n0, X4))) & ! [X6, X7] : ((init = a_select3(pminus_ds1_filter_init, X6, X7)) | (~ leq(X7, n5) | ~ leq(X6, n5) | ~ leq(n0, X7) | ~ leq(n0, X6))) & ! [X8, X9] : ((init = a_select3(xhatmin_ds1_filter_init, X8, X9)) | (~ leq(X9, n0) | ~ leq(X8, n5) | ~ leq(n0, X9) | ~ leq(n0, X8))) & ! [X10, X11] : ((init = a_select3(r_ds1_filter_init, X10, X11)) | (~ leq(X11, n2) | ~ leq(X10, n2) | ~ leq(n0, X11) | ~ leq(n0, X10))) & ! [X12, X13] : ((init = a_select3(q_ds1_filter_init, X12, X13)) | (~ leq(X13, n5) | ~ leq(X12, n5) | ~ leq(n0, X13) | ~ leq(n0, X12))) & ! [X14, X15] : ((init = a_select3(dv_ds1_filter_init, X14, X15)) | (~ leq(X15, n0) | ~ leq(X14, n5) | ~ leq(n0, X15) | ~ leq(n0, X14))) & ! [X16, X17] : ((init = a_select3(phi_ds1_filter_init, X16, X17)) | (~ leq(X17, n5) | ~ leq(X16, n5) | ~ leq(n0, X17) | ~ leq(n0, X16))) & ! [X18, X19] : ((init = a_select3(h_ds1_filter_init, X18, X19)) | (~ leq(X19, n5) | ~ leq(X18, n2) | ~ leq(n0, X19) | ~ leq(n0, X18))) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ ((! [X0] : ((leq(X0, pred(pv63)) & leq(n0, X0)) => ! [X1] : ((leq(X1, n5) & leq(n0, X1)) => (init = a_select3(id_ds1_filter_init, X0, X1)))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (gt(pv63, X2) => (init = a_select3(id_ds1_filter_init, X2, X3)))) & ! [X4, X5] : ((leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => ((gt(pv64, X5) & (pv63 = X4)) => (init = a_select3(id_ds1_filter_init, X4, X5)))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(pminus_ds1_filter_init, X6, X7))) & ! [X8, X9] : ((leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(xhatmin_ds1_filter_init, X8, X9))) & ! [X10, X11] : ((leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(r_ds1_filter_init, X10, X11))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(q_ds1_filter_init, X12, X13))) & ! [X14, X15] : ((leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => (init = a_select3(dv_ds1_filter_init, X14, X15))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (init = a_select3(phi_ds1_filter_init, X16, X17))) & ! [X18, X19] : ((leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) => (init = a_select3(h_ds1_filter_init, X18, X19))) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), inference(flattening, [], [f5])).
fof(f5, plain, ~ ~ ~ ((! [X0] : ((leq(X0, pred(pv63)) & leq(n0, X0)) => ! [X1] : ((leq(X1, n5) & leq(n0, X1)) => (init = a_select3(id_ds1_filter_init, X0, X1)))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (gt(pv63, X2) => (init = a_select3(id_ds1_filter_init, X2, X3)))) & ! [X4, X5] : ((leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => ((gt(pv64, X5) & (pv63 = X4)) => (init = a_select3(id_ds1_filter_init, X4, X5)))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(pminus_ds1_filter_init, X6, X7))) & ! [X8, X9] : ((leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(xhatmin_ds1_filter_init, X8, X9))) & ! [X10, X11] : ((leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(r_ds1_filter_init, X10, X11))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(q_ds1_filter_init, X12, X13))) & ! [X14, X15] : ((leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => (init = a_select3(dv_ds1_filter_init, X14, X15))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (init = a_select3(phi_ds1_filter_init, X16, X17))) & ! [X18, X19] : ((leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) => (init = a_select3(h_ds1_filter_init, X18, X19))) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), inference(rectify, [], [f3])).
fof(f3, plain, ~ ~ ~ ((! [X18] : ((leq(X18, pred(pv63)) & leq(n0, X18)) => ! [X19] : ((leq(X19, n5) & leq(n0, X19)) => (init = a_select3(id_ds1_filter_init, X18, X19)))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (gt(pv63, X16) => (init = a_select3(id_ds1_filter_init, X16, X17)))) & ! [X14, X15] : ((leq(X15, n5) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => ((gt(pv64, X15) & (pv63 = X14)) => (init = a_select3(id_ds1_filter_init, X14, X15)))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(pminus_ds1_filter_init, X12, X13))) & ! [X10, X11] : ((leq(X11, n0) & leq(X10, n5) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(xhatmin_ds1_filter_init, X10, X11))) & ! [X8, X9] : ((leq(X9, n2) & leq(X8, n2) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(r_ds1_filter_init, X8, X9))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(q_ds1_filter_init, X6, X7))) & ! [X4, X5] : ((leq(X5, n0) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => (init = a_select3(dv_ds1_filter_init, X4, X5))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (init = a_select3(phi_ds1_filter_init, X2, X3))) & ! [X0, X1] : ((leq(X1, n5) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) => (a_select3(h_ds1_filter_init, X0, X1) = init)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), inference(negated_conjecture, [], [f2])).
fof(f2, plain, ~ ~ ~ ((! [X18] : ((leq(X18, pred(pv63)) & leq(n0, X18)) => ! [X19] : ((leq(X19, n5) & leq(n0, X19)) => (init = a_select3(id_ds1_filter_init, X18, X19)))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (gt(pv63, X16) => (init = a_select3(id_ds1_filter_init, X16, X17)))) & ! [X14, X15] : ((leq(X15, n5) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => ((gt(pv64, X15) & (pv63 = X14)) => (init = a_select3(id_ds1_filter_init, X14, X15)))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(pminus_ds1_filter_init, X12, X13))) & ! [X10, X11] : ((leq(X11, n0) & leq(X10, n5) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(xhatmin_ds1_filter_init, X10, X11))) & ! [X8, X9] : ((leq(X9, n2) & leq(X8, n2) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(r_ds1_filter_init, X8, X9))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(q_ds1_filter_init, X6, X7))) & ! [X4, X5] : ((leq(X5, n0) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => (init = a_select3(dv_ds1_filter_init, X4, X5))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (init = a_select3(phi_ds1_filter_init, X2, X3))) & ! [X0, X1] : ((leq(X1, n5) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) => (a_select3(h_ds1_filter_init, X0, X1) = init)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), file('/home/ubuntu/library/tptp/Problems/MSC/MSC010+1.p', n91)).
fof(f478, plain, (leq(sK27, pred(pv63)) | ~ spl31_69), inference(avatar_component_clause, [], [f476])).
fof(f476, plain, (spl31_69 <=> leq(sK27, pred(pv63))), introduced(avatar_definition, [new_symbols(naming, [spl31_69])])).
fof(f3459, plain, (spl31_44 | ~ spl31_45 | ~ spl31_46 | ~ spl31_47 | ~ spl31_48 | ~ spl31_49), inference(avatar_contradiction_clause, [], [f3458])).
fof(f3458, plain, ($false | (spl31_44 | ~ spl31_45 | ~ spl31_46 | ~ spl31_47 | ~ spl31_48 | ~ spl31_49)), inference(subsumption_resolution, [], [f3457, f363])).
fof(f363, plain, (~ (init = a_select3(id_ds1_filter_init, sK23, sK24)) | spl31_44), inference(avatar_component_clause, [], [f361])).
fof(f361, plain, (spl31_44 <=> (init = a_select3(id_ds1_filter_init, sK23, sK24))), introduced(avatar_definition, [new_symbols(naming, [spl31_44])])).
fof(f3457, plain, ((init = a_select3(id_ds1_filter_init, sK23, sK24)) | (~ spl31_45 | ~ spl31_46 | ~ spl31_47 | ~ spl31_48 | ~ spl31_49)), inference(subsumption_resolution, [], [f3446, f373])).
fof(f373, plain, (leq(sK24, n5) | ~ spl31_46), inference(avatar_component_clause, [], [f371])).
fof(f371, plain, (spl31_46 <=> leq(sK24, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_46])])).
fof(f3446, plain, (~ leq(sK24, n5) | (init = a_select3(id_ds1_filter_init, sK23, sK24)) | (~ spl31_45 | ~ spl31_47 | ~ spl31_48 | ~ spl31_49)), inference(resolution, [], [f3056, f383])).
fof(f383, plain, (leq(n0, sK24) | ~ spl31_48), inference(avatar_component_clause, [], [f381])).
fof(f381, plain, (spl31_48 <=> leq(n0, sK24)), introduced(avatar_definition, [new_symbols(naming, [spl31_48])])).
fof(f3056, plain, (! [X0] : (~ leq(n0, X0) | ~ leq(X0, n5) | (init = a_select3(id_ds1_filter_init, sK23, X0))) | (~ spl31_45 | ~ spl31_47 | ~ spl31_49)), inference(subsumption_resolution, [], [f3055, f388])).
fof(f388, plain, (leq(n0, sK23) | ~ spl31_49), inference(avatar_component_clause, [], [f386])).
fof(f386, plain, (spl31_49 <=> leq(n0, sK23)), introduced(avatar_definition, [new_symbols(naming, [spl31_49])])).
fof(f3055, plain, (! [X0] : ((init = a_select3(id_ds1_filter_init, sK23, X0)) | ~ leq(X0, n5) | ~ leq(n0, X0) | ~ leq(n0, sK23)) | (~ spl31_45 | ~ spl31_47)), inference(subsumption_resolution, [], [f3054, f378])).
fof(f378, plain, (leq(sK23, n5) | ~ spl31_47), inference(avatar_component_clause, [], [f376])).
fof(f376, plain, (spl31_47 <=> leq(sK23, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_47])])).
fof(f3054, plain, (! [X0] : ((init = a_select3(id_ds1_filter_init, sK23, X0)) | ~ leq(X0, n5) | ~ leq(sK23, n5) | ~ leq(n0, X0) | ~ leq(n0, sK23)) | ~ spl31_45), inference(resolution, [], [f368, f132])).
fof(f132, plain, ! [X4, X5] : (~ gt(pv63, X4) | (init = a_select3(id_ds1_filter_init, X4, X5)) | ~ leq(X5, n5) | ~ leq(X4, n5) | ~ leq(n0, X5) | ~ leq(n0, X4)), inference(cnf_transformation, [], [f63])).
fof(f368, plain, (gt(pv63, sK23) | ~ spl31_45), inference(avatar_component_clause, [], [f366])).
fof(f366, plain, (spl31_45 <=> gt(pv63, sK23)), introduced(avatar_definition, [new_symbols(naming, [spl31_45])])).
fof(f3053, plain, (spl31_38 | ~ spl31_39 | ~ spl31_40 | ~ spl31_41 | ~ spl31_42), inference(avatar_contradiction_clause, [], [f3052])).
fof(f3052, plain, ($false | (spl31_38 | ~ spl31_39 | ~ spl31_40 | ~ spl31_41 | ~ spl31_42)), inference(subsumption_resolution, [], [f3051, f344])).
fof(f344, plain, (leq(sK21, n2) | ~ spl31_40), inference(avatar_component_clause, [], [f342])).
fof(f342, plain, (spl31_40 <=> leq(sK21, n2)), introduced(avatar_definition, [new_symbols(naming, [spl31_40])])).
fof(f3051, plain, (~ leq(sK21, n2) | (spl31_38 | ~ spl31_39 | ~ spl31_41 | ~ spl31_42)), inference(subsumption_resolution, [], [f3047, f334])).
fof(f334, plain, (~ (init = a_select3(h_ds1_filter_init, sK21, sK22)) | spl31_38), inference(avatar_component_clause, [], [f332])).
fof(f332, plain, (spl31_38 <=> (init = a_select3(h_ds1_filter_init, sK21, sK22))), introduced(avatar_definition, [new_symbols(naming, [spl31_38])])).
fof(f3047, plain, ((init = a_select3(h_ds1_filter_init, sK21, sK22)) | ~ leq(sK21, n2) | (~ spl31_39 | ~ spl31_41 | ~ spl31_42)), inference(resolution, [], [f2712, f354])).
fof(f354, plain, (leq(n0, sK21) | ~ spl31_42), inference(avatar_component_clause, [], [f352])).
fof(f352, plain, (spl31_42 <=> leq(n0, sK21)), introduced(avatar_definition, [new_symbols(naming, [spl31_42])])).
fof(f2712, plain, (! [X0] : (~ leq(n0, X0) | (init = a_select3(h_ds1_filter_init, X0, sK22)) | ~ leq(X0, n2)) | (~ spl31_39 | ~ spl31_41)), inference(subsumption_resolution, [], [f2668, f339])).
fof(f339, plain, (leq(sK22, n5) | ~ spl31_39), inference(avatar_component_clause, [], [f337])).
fof(f337, plain, (spl31_39 <=> leq(sK22, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_39])])).
fof(f2668, plain, (! [X0] : (~ leq(sK22, n5) | ~ leq(X0, n2) | (init = a_select3(h_ds1_filter_init, X0, sK22)) | ~ leq(n0, X0)) | ~ spl31_41), inference(resolution, [], [f349, f124])).
fof(f124, plain, ! [X21, X20] : (~ leq(n0, X21) | ~ leq(X21, n5) | ~ leq(X20, n2) | (init = a_select3(h_ds1_filter_init, X20, X21)) | ~ leq(n0, X20)), inference(cnf_transformation, [], [f63])).
fof(f349, plain, (leq(n0, sK22) | ~ spl31_41), inference(avatar_component_clause, [], [f347])).
fof(f347, plain, (spl31_41 <=> leq(n0, sK22)), introduced(avatar_definition, [new_symbols(naming, [spl31_41])])).
fof(f2623, plain, (~ spl31_16 | spl31_14 | ~ spl31_15 | ~ spl31_17 | ~ spl31_18), inference(avatar_split_clause, [], [f2619, f236, f231, f221, f216, f226])).
fof(f226, plain, (spl31_16 <=> leq(sK13, n2)), introduced(avatar_definition, [new_symbols(naming, [spl31_16])])).
fof(f216, plain, (spl31_14 <=> (init = a_select3(r_ds1_filter_init, sK13, sK14))), introduced(avatar_definition, [new_symbols(naming, [spl31_14])])).
fof(f221, plain, (spl31_15 <=> leq(sK14, n2)), introduced(avatar_definition, [new_symbols(naming, [spl31_15])])).
fof(f231, plain, (spl31_17 <=> leq(n0, sK14)), introduced(avatar_definition, [new_symbols(naming, [spl31_17])])).
fof(f236, plain, (spl31_18 <=> leq(n0, sK13)), introduced(avatar_definition, [new_symbols(naming, [spl31_18])])).
fof(f2619, plain, (~ leq(sK13, n2) | (spl31_14 | ~ spl31_15 | ~ spl31_17 | ~ spl31_18)), inference(subsumption_resolution, [], [f2613, f218])).
fof(f218, plain, (~ (init = a_select3(r_ds1_filter_init, sK13, sK14)) | spl31_14), inference(avatar_component_clause, [], [f216])).
fof(f2613, plain, ((init = a_select3(r_ds1_filter_init, sK13, sK14)) | ~ leq(sK13, n2) | (~ spl31_15 | ~ spl31_17 | ~ spl31_18)), inference(resolution, [], [f2215, f238])).
fof(f238, plain, (leq(n0, sK13) | ~ spl31_18), inference(avatar_component_clause, [], [f236])).
fof(f2215, plain, (! [X4] : (~ leq(n0, X4) | (init = a_select3(r_ds1_filter_init, X4, sK14)) | ~ leq(X4, n2)) | (~ spl31_15 | ~ spl31_17)), inference(subsumption_resolution, [], [f2153, f223])).
fof(f223, plain, (leq(sK14, n2) | ~ spl31_15), inference(avatar_component_clause, [], [f221])).
fof(f2153, plain, (! [X4] : (~ leq(sK14, n2) | ~ leq(X4, n2) | (init = a_select3(r_ds1_filter_init, X4, sK14)) | ~ leq(n0, X4)) | ~ spl31_17), inference(resolution, [], [f233, f128])).
fof(f128, plain, ! [X12, X13] : (~ leq(n0, X13) | ~ leq(X13, n2) | ~ leq(X12, n2) | (init = a_select3(r_ds1_filter_init, X12, X13)) | ~ leq(n0, X12)), inference(cnf_transformation, [], [f63])).
fof(f233, plain, (leq(n0, sK14) | ~ spl31_17), inference(avatar_component_clause, [], [f231])).
fof(f2146, plain, (spl31_26 | ~ spl31_27 | ~ spl31_28 | ~ spl31_29 | ~ spl31_30), inference(avatar_contradiction_clause, [], [f2145])).
fof(f2145, plain, ($false | (spl31_26 | ~ spl31_27 | ~ spl31_28 | ~ spl31_29 | ~ spl31_30)), inference(subsumption_resolution, [], [f2144, f286])).
fof(f286, plain, (leq(sK17, n5) | ~ spl31_28), inference(avatar_component_clause, [], [f284])).
fof(f284, plain, (spl31_28 <=> leq(sK17, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_28])])).
fof(f2144, plain, (~ leq(sK17, n5) | (spl31_26 | ~ spl31_27 | ~ spl31_29 | ~ spl31_30)), inference(subsumption_resolution, [], [f2134, f276])).
fof(f276, plain, (~ (init = a_select3(dv_ds1_filter_init, sK17, sK18)) | spl31_26), inference(avatar_component_clause, [], [f274])).
fof(f274, plain, (spl31_26 <=> (init = a_select3(dv_ds1_filter_init, sK17, sK18))), introduced(avatar_definition, [new_symbols(naming, [spl31_26])])).
fof(f2134, plain, ((init = a_select3(dv_ds1_filter_init, sK17, sK18)) | ~ leq(sK17, n5) | (~ spl31_27 | ~ spl31_29 | ~ spl31_30)), inference(resolution, [], [f1793, f296])).
fof(f296, plain, (leq(n0, sK17) | ~ spl31_30), inference(avatar_component_clause, [], [f294])).
fof(f294, plain, (spl31_30 <=> leq(n0, sK17)), introduced(avatar_definition, [new_symbols(naming, [spl31_30])])).
fof(f1793, plain, (! [X2] : (~ leq(n0, X2) | (init = a_select3(dv_ds1_filter_init, X2, sK18)) | ~ leq(X2, n5)) | (~ spl31_27 | ~ spl31_29)), inference(subsumption_resolution, [], [f1743, f281])).
fof(f281, plain, (leq(sK18, n0) | ~ spl31_27), inference(avatar_component_clause, [], [f279])).
fof(f279, plain, (spl31_27 <=> leq(sK18, n0)), introduced(avatar_definition, [new_symbols(naming, [spl31_27])])).
fof(f1743, plain, (! [X2] : (~ leq(sK18, n0) | ~ leq(X2, n5) | (init = a_select3(dv_ds1_filter_init, X2, sK18)) | ~ leq(n0, X2)) | ~ spl31_29), inference(resolution, [], [f291, f126])).
fof(f126, plain, ! [X17, X16] : (~ leq(n0, X17) | ~ leq(X17, n0) | ~ leq(X16, n5) | (init = a_select3(dv_ds1_filter_init, X16, X17)) | ~ leq(n0, X16)), inference(cnf_transformation, [], [f63])).
fof(f291, plain, (leq(n0, sK18) | ~ spl31_29), inference(avatar_component_clause, [], [f289])).
fof(f289, plain, (spl31_29 <=> leq(n0, sK18)), introduced(avatar_definition, [new_symbols(naming, [spl31_29])])).
fof(f1737, plain, (spl31_20 | ~ spl31_21 | ~ spl31_22 | ~ spl31_23 | ~ spl31_24), inference(avatar_contradiction_clause, [], [f1736])).
fof(f1736, plain, ($false | (spl31_20 | ~ spl31_21 | ~ spl31_22 | ~ spl31_23 | ~ spl31_24)), inference(subsumption_resolution, [], [f1735, f257])).
fof(f257, plain, (leq(sK15, n5) | ~ spl31_22), inference(avatar_component_clause, [], [f255])).
fof(f255, plain, (spl31_22 <=> leq(sK15, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_22])])).
fof(f1735, plain, (~ leq(sK15, n5) | (spl31_20 | ~ spl31_21 | ~ spl31_23 | ~ spl31_24)), inference(subsumption_resolution, [], [f1726, f247])).
fof(f247, plain, (~ (init = a_select3(q_ds1_filter_init, sK15, sK16)) | spl31_20), inference(avatar_component_clause, [], [f245])).
fof(f245, plain, (spl31_20 <=> (init = a_select3(q_ds1_filter_init, sK15, sK16))), introduced(avatar_definition, [new_symbols(naming, [spl31_20])])).
fof(f1726, plain, ((init = a_select3(q_ds1_filter_init, sK15, sK16)) | ~ leq(sK15, n5) | (~ spl31_21 | ~ spl31_23 | ~ spl31_24)), inference(resolution, [], [f1463, f267])).
fof(f267, plain, (leq(n0, sK15) | ~ spl31_24), inference(avatar_component_clause, [], [f265])).
fof(f265, plain, (spl31_24 <=> leq(n0, sK15)), introduced(avatar_definition, [new_symbols(naming, [spl31_24])])).
fof(f1463, plain, (! [X3] : (~ leq(n0, X3) | (init = a_select3(q_ds1_filter_init, X3, sK16)) | ~ leq(X3, n5)) | (~ spl31_21 | ~ spl31_23)), inference(subsumption_resolution, [], [f1421, f252])).
fof(f252, plain, (leq(sK16, n5) | ~ spl31_21), inference(avatar_component_clause, [], [f250])).
fof(f250, plain, (spl31_21 <=> leq(sK16, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_21])])).
fof(f1421, plain, (! [X3] : (~ leq(sK16, n5) | ~ leq(X3, n5) | (init = a_select3(q_ds1_filter_init, X3, sK16)) | ~ leq(n0, X3)) | ~ spl31_23), inference(resolution, [], [f262, f127])).
fof(f127, plain, ! [X14, X15] : (~ leq(n0, X15) | ~ leq(X15, n5) | ~ leq(X14, n5) | (init = a_select3(q_ds1_filter_init, X14, X15)) | ~ leq(n0, X14)), inference(cnf_transformation, [], [f63])).
fof(f262, plain, (leq(n0, sK16) | ~ spl31_23), inference(avatar_component_clause, [], [f260])).
fof(f260, plain, (spl31_23 <=> leq(n0, sK16)), introduced(avatar_definition, [new_symbols(naming, [spl31_23])])).
fof(f1417, plain, ~ spl31_66, inference(avatar_contradiction_clause, [], [f1416])).
fof(f1416, plain, ($false | ~ spl31_66), inference(subsumption_resolution, [], [f1415, f140])).
fof(f140, plain, leq(sK30, pv64), inference(cnf_transformation, [], [f63])).
fof(f1415, plain, (~ leq(sK30, pv64) | ~ spl31_66), inference(subsumption_resolution, [], [f1414, f138])).
fof(f138, plain, ~ (pv64 = sK30), inference(cnf_transformation, [], [f63])).
fof(f1414, plain, ((pv64 = sK30) | ~ leq(sK30, pv64) | ~ spl31_66), inference(subsumption_resolution, [], [f1413, f137])).
fof(f137, plain, leq(sK30, n5), inference(cnf_transformation, [], [f63])).
fof(f1413, plain, (~ leq(sK30, n5) | (pv64 = sK30) | ~ leq(sK30, pv64) | ~ spl31_66), inference(subsumption_resolution, [], [f1345, f485])).
fof(f485, plain, ~ (init = a_select3(id_ds1_filter_init, pv63, sK30)), inference(backward_demodulation, [], [f141, f139])).
fof(f139, plain, (pv63 = sK29), inference(cnf_transformation, [], [f63])).
fof(f141, plain, ~ (init = a_select3(id_ds1_filter_init, sK29, sK30)), inference(cnf_transformation, [], [f63])).
fof(f1345, plain, ((init = a_select3(id_ds1_filter_init, pv63, sK30)) | ~ leq(sK30, n5) | (pv64 = sK30) | ~ leq(sK30, pv64) | ~ spl31_66), inference(resolution, [], [f463, f135])).
fof(f135, plain, leq(n0, sK30), inference(cnf_transformation, [], [f63])).
fof(f463, plain, (! [X1] : (~ leq(n0, X1) | (init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, n5) | (pv64 = X1) | ~ leq(X1, pv64)) | ~ spl31_66), inference(avatar_component_clause, [], [f462])).
fof(f462, plain, (spl31_66 <=> ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(n0, X1) | ~ leq(X1, n5) | (pv64 = X1) | ~ leq(X1, pv64))), introduced(avatar_definition, [new_symbols(naming, [spl31_66])])).
fof(f1334, plain, (spl31_32 | ~ spl31_33 | ~ spl31_34 | ~ spl31_35 | ~ spl31_36), inference(avatar_contradiction_clause, [], [f1333])).
fof(f1333, plain, ($false | (spl31_32 | ~ spl31_33 | ~ spl31_34 | ~ spl31_35 | ~ spl31_36)), inference(subsumption_resolution, [], [f1332, f315])).
fof(f315, plain, (leq(sK19, n5) | ~ spl31_34), inference(avatar_component_clause, [], [f313])).
fof(f313, plain, (spl31_34 <=> leq(sK19, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_34])])).
fof(f1332, plain, (~ leq(sK19, n5) | (spl31_32 | ~ spl31_33 | ~ spl31_35 | ~ spl31_36)), inference(subsumption_resolution, [], [f1324, f305])).
fof(f305, plain, (~ (init = a_select3(phi_ds1_filter_init, sK19, sK20)) | spl31_32), inference(avatar_component_clause, [], [f303])).
fof(f303, plain, (spl31_32 <=> (init = a_select3(phi_ds1_filter_init, sK19, sK20))), introduced(avatar_definition, [new_symbols(naming, [spl31_32])])).
fof(f1324, plain, ((init = a_select3(phi_ds1_filter_init, sK19, sK20)) | ~ leq(sK19, n5) | (~ spl31_33 | ~ spl31_35 | ~ spl31_36)), inference(resolution, [], [f1083, f325])).
fof(f325, plain, (leq(n0, sK19) | ~ spl31_36), inference(avatar_component_clause, [], [f323])).
fof(f323, plain, (spl31_36 <=> leq(n0, sK19)), introduced(avatar_definition, [new_symbols(naming, [spl31_36])])).
fof(f1083, plain, (! [X1] : (~ leq(n0, X1) | (init = a_select3(phi_ds1_filter_init, X1, sK20)) | ~ leq(X1, n5)) | (~ spl31_33 | ~ spl31_35)), inference(subsumption_resolution, [], [f1054, f310])).
fof(f310, plain, (leq(sK20, n5) | ~ spl31_33), inference(avatar_component_clause, [], [f308])).
fof(f308, plain, (spl31_33 <=> leq(sK20, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_33])])).
fof(f1054, plain, (! [X1] : (~ leq(sK20, n5) | ~ leq(X1, n5) | (init = a_select3(phi_ds1_filter_init, X1, sK20)) | ~ leq(n0, X1)) | ~ spl31_35), inference(resolution, [], [f320, f125])).
fof(f125, plain, ! [X19, X18] : (~ leq(n0, X19) | ~ leq(X19, n5) | ~ leq(X18, n5) | (init = a_select3(phi_ds1_filter_init, X18, X19)) | ~ leq(n0, X18)), inference(cnf_transformation, [], [f63])).
fof(f320, plain, (leq(n0, sK20) | ~ spl31_35), inference(avatar_component_clause, [], [f318])).
fof(f318, plain, (spl31_35 <=> leq(n0, sK20)), introduced(avatar_definition, [new_symbols(naming, [spl31_35])])).
fof(f1052, plain, (spl31_2 | ~ spl31_3 | ~ spl31_4 | ~ spl31_5 | ~ spl31_6), inference(avatar_split_clause, [], [f1051, f178, f173, f168, f163, f158])).
fof(f158, plain, (spl31_2 <=> (init = a_select3(pminus_ds1_filter_init, sK9, sK10))), introduced(avatar_definition, [new_symbols(naming, [spl31_2])])).
fof(f163, plain, (spl31_3 <=> leq(sK10, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_3])])).
fof(f168, plain, (spl31_4 <=> leq(sK9, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_4])])).
fof(f173, plain, (spl31_5 <=> leq(n0, sK10)), introduced(avatar_definition, [new_symbols(naming, [spl31_5])])).
fof(f178, plain, (spl31_6 <=> leq(n0, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl31_6])])).
fof(f1051, plain, ((init = a_select3(pminus_ds1_filter_init, sK9, sK10)) | (~ spl31_3 | ~ spl31_4 | ~ spl31_5 | ~ spl31_6)), inference(subsumption_resolution, [], [f1038, f170])).
fof(f170, plain, (leq(sK9, n5) | ~ spl31_4), inference(avatar_component_clause, [], [f168])).
fof(f1038, plain, ((init = a_select3(pminus_ds1_filter_init, sK9, sK10)) | ~ leq(sK9, n5) | (~ spl31_3 | ~ spl31_5 | ~ spl31_6)), inference(resolution, [], [f837, f180])).
fof(f180, plain, (leq(n0, sK9) | ~ spl31_6), inference(avatar_component_clause, [], [f178])).
fof(f837, plain, (! [X6] : (~ leq(n0, X6) | (init = a_select3(pminus_ds1_filter_init, X6, sK10)) | ~ leq(X6, n5)) | (~ spl31_3 | ~ spl31_5)), inference(subsumption_resolution, [], [f799, f165])).
fof(f165, plain, (leq(sK10, n5) | ~ spl31_3), inference(avatar_component_clause, [], [f163])).
fof(f799, plain, (! [X6] : (~ leq(sK10, n5) | ~ leq(X6, n5) | (init = a_select3(pminus_ds1_filter_init, X6, sK10)) | ~ leq(n0, X6)) | ~ spl31_5), inference(resolution, [], [f175, f130])).
fof(f130, plain, ! [X8, X9] : (~ leq(n0, X9) | ~ leq(X9, n5) | ~ leq(X8, n5) | (init = a_select3(pminus_ds1_filter_init, X8, X9)) | ~ leq(n0, X8)), inference(cnf_transformation, [], [f63])).
fof(f175, plain, (leq(n0, sK10) | ~ spl31_5), inference(avatar_component_clause, [], [f173])).
fof(f792, plain, (spl31_8 | ~ spl31_9 | ~ spl31_10 | ~ spl31_11 | ~ spl31_12), inference(avatar_split_clause, [], [f791, f207, f202, f197, f192, f187])).
fof(f187, plain, (spl31_8 <=> (init = a_select3(xhatmin_ds1_filter_init, sK11, sK12))), introduced(avatar_definition, [new_symbols(naming, [spl31_8])])).
fof(f192, plain, (spl31_9 <=> leq(sK12, n0)), introduced(avatar_definition, [new_symbols(naming, [spl31_9])])).
fof(f197, plain, (spl31_10 <=> leq(sK11, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_10])])).
fof(f202, plain, (spl31_11 <=> leq(n0, sK12)), introduced(avatar_definition, [new_symbols(naming, [spl31_11])])).
fof(f207, plain, (spl31_12 <=> leq(n0, sK11)), introduced(avatar_definition, [new_symbols(naming, [spl31_12])])).
fof(f791, plain, ((init = a_select3(xhatmin_ds1_filter_init, sK11, sK12)) | (~ spl31_9 | ~ spl31_10 | ~ spl31_11 | ~ spl31_12)), inference(subsumption_resolution, [], [f782, f199])).
fof(f199, plain, (leq(sK11, n5) | ~ spl31_10), inference(avatar_component_clause, [], [f197])).
fof(f782, plain, ((init = a_select3(xhatmin_ds1_filter_init, sK11, sK12)) | ~ leq(sK11, n5) | (~ spl31_9 | ~ spl31_11 | ~ spl31_12)), inference(resolution, [], [f654, f209])).
fof(f209, plain, (leq(n0, sK11) | ~ spl31_12), inference(avatar_component_clause, [], [f207])).
fof(f654, plain, (! [X4] : (~ leq(n0, X4) | (init = a_select3(xhatmin_ds1_filter_init, X4, sK12)) | ~ leq(X4, n5)) | (~ spl31_9 | ~ spl31_11)), inference(subsumption_resolution, [], [f652, f194])).
fof(f194, plain, (leq(sK12, n0) | ~ spl31_9), inference(avatar_component_clause, [], [f192])).
fof(f652, plain, (! [X4] : (~ leq(sK12, n0) | ~ leq(X4, n5) | (init = a_select3(xhatmin_ds1_filter_init, X4, sK12)) | ~ leq(n0, X4)) | ~ spl31_11), inference(resolution, [], [f129, f204])).
fof(f204, plain, (leq(n0, sK12) | ~ spl31_11), inference(avatar_component_clause, [], [f202])).
fof(f129, plain, ! [X10, X11] : (~ leq(n0, X11) | ~ leq(X11, n0) | ~ leq(X10, n5) | (init = a_select3(xhatmin_ds1_filter_init, X10, X11)) | ~ leq(n0, X10)), inference(cnf_transformation, [], [f63])).
fof(f500, plain, ~ spl31_58, inference(avatar_split_clause, [], [f117, f430])).
fof(f430, plain, (spl31_58 <=> (pv63 = pv64)), introduced(avatar_definition, [new_symbols(naming, [spl31_58])])).
fof(f117, plain, ~ (pv63 = pv64), inference(cnf_transformation, [], [f63])).
fof(f499, plain, spl31_59, inference(avatar_split_clause, [], [f118, f434])).
fof(f434, plain, (spl31_59 <=> leq(n0, pv5)), introduced(avatar_definition, [new_symbols(naming, [spl31_59])])).
fof(f118, plain, leq(n0, pv5), inference(cnf_transformation, [], [f63])).
fof(f498, plain, spl31_64, inference(avatar_split_clause, [], [f119, f454])).
fof(f454, plain, (spl31_64 <=> leq(n0, pv63)), introduced(avatar_definition, [new_symbols(naming, [spl31_64])])).
fof(f119, plain, leq(n0, pv63), inference(cnf_transformation, [], [f63])).
fof(f497, plain, spl31_60, inference(avatar_split_clause, [], [f120, f438])).
fof(f438, plain, (spl31_60 <=> leq(n0, pv64)), introduced(avatar_definition, [new_symbols(naming, [spl31_60])])).
fof(f120, plain, leq(n0, pv64), inference(cnf_transformation, [], [f63])).
fof(f496, plain, spl31_61, inference(avatar_split_clause, [], [f121, f442])).
fof(f442, plain, (spl31_61 <=> leq(pv5, n998)), introduced(avatar_definition, [new_symbols(naming, [spl31_61])])).
fof(f121, plain, leq(pv5, n998), inference(cnf_transformation, [], [f63])).
fof(f495, plain, spl31_65, inference(avatar_split_clause, [], [f122, f458])).
fof(f458, plain, (spl31_65 <=> leq(pv63, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_65])])).
fof(f122, plain, leq(pv63, n5), inference(cnf_transformation, [], [f63])).
fof(f494, plain, spl31_62, inference(avatar_split_clause, [], [f123, f446])).
fof(f446, plain, (spl31_62 <=> leq(pv64, n5)), introduced(avatar_definition, [new_symbols(naming, [spl31_62])])).
fof(f123, plain, leq(pv64, n5), inference(cnf_transformation, [], [f63])).
fof(f493, plain, (~ spl31_64 | ~ spl31_65 | spl31_71), inference(avatar_split_clause, [], [f147, f491, f458, f454])).
fof(f147, plain, ! [X7] : ((init = a_select3(id_ds1_filter_init, pv63, X7)) | ~ gt(pv64, X7) | ~ leq(X7, n5) | ~ leq(pv63, n5) | ~ leq(n0, X7) | ~ leq(n0, pv63)), inference(equality_resolution, [], [f131])).
fof(f131, plain, ! [X6, X7] : ((init = a_select3(id_ds1_filter_init, X6, X7)) | ~ gt(pv64, X7) | ~ (pv63 = X6) | ~ leq(X7, n5) | ~ leq(X6, n5) | ~ leq(n0, X7) | ~ leq(n0, X6)), inference(cnf_transformation, [], [f63])).
fof(f484, plain, (spl31_58 | ~ spl31_59 | ~ spl31_60 | ~ spl31_61 | ~ spl31_62 | spl31_37 | spl31_31 | spl31_25 | spl31_19 | spl31_13 | spl31_7 | spl31_1 | spl31_50 | spl31_43 | spl31_70 | ~ spl31_64 | ~ spl31_65 | spl31_66), inference(avatar_split_clause, [], [f148, f462, f458, f454, f481, f357, f391, f154, f183, f212, f241, f270, f299, f328, f446, f442, f438, f434, f430])).
fof(f328, plain, (spl31_37 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl31_37])])).
fof(f299, plain, (spl31_31 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl31_31])])).
fof(f270, plain, (spl31_25 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl31_25])])).
fof(f241, plain, (spl31_19 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl31_19])])).
fof(f212, plain, (spl31_13 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl31_13])])).
fof(f183, plain, (spl31_7 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl31_7])])).
fof(f154, plain, (spl31_1 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl31_1])])).
fof(f391, plain, (spl31_50 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl31_50])])).
fof(f357, plain, (spl31_43 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl31_43])])).
fof(f148, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(n0, sK27) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(duplicate_literal_removal, [], [f146])).
fof(f146, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(n0, sK27) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(equality_resolution, [], [f112])).
fof(f112, plain, ! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0) | leq(n0, sK27) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(cnf_transformation, [], [f60])).
fof(f60, plain, (! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0)) | ((~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) & leq(sK28, n5) & leq(n0, sK28)) & leq(sK27, pred(pv63)) & leq(n0, sK27)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK27, sK28])], [f57, f59, f58])).
fof(f58, plain, (? [X2] : (? [X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & leq(X3, n5) & leq(n0, X3)) & leq(X2, pred(pv63)) & leq(n0, X2)) => (? [X3] : (~ (init = a_select3(id_ds1_filter_init, sK27, X3)) & leq(X3, n5) & leq(n0, X3)) & leq(sK27, pred(pv63)) & leq(n0, sK27))), introduced(choice_axiom, [])).
fof(f59, plain, (? [X3] : (~ (init = a_select3(id_ds1_filter_init, sK27, X3)) & leq(X3, n5) & leq(n0, X3)) => (~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) & leq(sK28, n5) & leq(n0, sK28))), introduced(choice_axiom, [])).
fof(f57, plain, (! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0)) | ? [X2] : (? [X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & leq(X3, n5) & leq(n0, X3)) & leq(X2, pred(pv63)) & leq(n0, X2)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(rectify, [], [f20])).
fof(f20, plain, (! [X20, X21] : ((init = a_select3(id_ds1_filter_init, X20, X21)) | ~ leq(X21, pv64) | ~ (pv63 = X20) | (pv64 = X21) | ~ leq(X21, n5) | ~ leq(X20, n5) | ~ leq(n0, X21) | ~ leq(n0, X20)) | ? [X0] : (? [X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(n0, X1)) & leq(X0, pred(pv63)) & leq(n0, X0)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(definition_folding, [], [f8, e19, e18, e17, e16, e15, e14, e13, e12, e11])).
fof(f11, plain, (? [X4, X5] : (~ (init = a_select3(id_ds1_filter_init, X4, X5)) & gt(pv64, X5) & (pv63 = X4) & leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) | ~ sP0), inference(usedef, [], [e11])).
fof(e11, plain, (sP0 <=> ? [X4, X5] : (~ (init = a_select3(id_ds1_filter_init, X4, X5)) & gt(pv64, X5) & (pv63 = X4) & leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f12, plain, (? [X2, X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & gt(pv63, X2) & leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) | ~ sP1), inference(usedef, [], [e12])).
fof(e12, plain, (sP1 <=> ? [X2, X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & gt(pv63, X2) & leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f13, plain, (? [X18, X19] : (~ (init = a_select3(h_ds1_filter_init, X18, X19)) & leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) | ~ sP2), inference(usedef, [], [e13])).
fof(e13, plain, (sP2 <=> ? [X18, X19] : (~ (init = a_select3(h_ds1_filter_init, X18, X19)) & leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f14, plain, (? [X16, X17] : (~ (init = a_select3(phi_ds1_filter_init, X16, X17)) & leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) | ~ sP3), inference(usedef, [], [e14])).
fof(e14, plain, (sP3 <=> ? [X16, X17] : (~ (init = a_select3(phi_ds1_filter_init, X16, X17)) & leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f15, plain, (? [X14, X15] : (~ (init = a_select3(dv_ds1_filter_init, X14, X15)) & leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) | ~ sP4), inference(usedef, [], [e15])).
fof(e15, plain, (sP4 <=> ? [X14, X15] : (~ (init = a_select3(dv_ds1_filter_init, X14, X15)) & leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f16, plain, (? [X12, X13] : (~ (init = a_select3(q_ds1_filter_init, X12, X13)) & leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) | ~ sP5), inference(usedef, [], [e16])).
fof(e16, plain, (sP5 <=> ? [X12, X13] : (~ (init = a_select3(q_ds1_filter_init, X12, X13)) & leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f17, plain, (? [X10, X11] : (~ (init = a_select3(r_ds1_filter_init, X10, X11)) & leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) | ~ sP6), inference(usedef, [], [e17])).
fof(e17, plain, (sP6 <=> ? [X10, X11] : (~ (init = a_select3(r_ds1_filter_init, X10, X11)) & leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f18, plain, (? [X8, X9] : (~ (init = a_select3(xhatmin_ds1_filter_init, X8, X9)) & leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) | ~ sP7), inference(usedef, [], [e18])).
fof(e18, plain, (sP7 <=> ? [X8, X9] : (~ (init = a_select3(xhatmin_ds1_filter_init, X8, X9)) & leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f19, plain, (? [X6, X7] : (~ (init = a_select3(pminus_ds1_filter_init, X6, X7)) & leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) | ~ sP8), inference(usedef, [], [e19])).
fof(e19, plain, (sP8 <=> ? [X6, X7] : (~ (init = a_select3(pminus_ds1_filter_init, X6, X7)) & leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f8, plain, (! [X20, X21] : ((init = a_select3(id_ds1_filter_init, X20, X21)) | ~ leq(X21, pv64) | ~ (pv63 = X20) | (pv64 = X21) | ~ leq(X21, n5) | ~ leq(X20, n5) | ~ leq(n0, X21) | ~ leq(n0, X20)) | ? [X0] : (? [X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(n0, X1)) & leq(X0, pred(pv63)) & leq(n0, X0)) | ? [X2, X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & gt(pv63, X2) & leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) | ? [X4, X5] : (~ (init = a_select3(id_ds1_filter_init, X4, X5)) & gt(pv64, X5) & (pv63 = X4) & leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) | ? [X6, X7] : (~ (init = a_select3(pminus_ds1_filter_init, X6, X7)) & leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) | ? [X8, X9] : (~ (init = a_select3(xhatmin_ds1_filter_init, X8, X9)) & leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) | ? [X10, X11] : (~ (init = a_select3(r_ds1_filter_init, X10, X11)) & leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) | ? [X12, X13] : (~ (init = a_select3(q_ds1_filter_init, X12, X13)) & leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) | ? [X14, X15] : (~ (init = a_select3(dv_ds1_filter_init, X14, X15)) & leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) | ? [X16, X17] : (~ (init = a_select3(phi_ds1_filter_init, X16, X17)) & leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) | ? [X18, X19] : (~ (init = a_select3(h_ds1_filter_init, X18, X19)) & leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(flattening, [], [f7])).
fof(f7, plain, (! [X20, X21] : (((init = a_select3(id_ds1_filter_init, X20, X21)) | (~ leq(X21, pv64) | ~ (pv63 = X20) | (pv64 = X21))) | (~ leq(X21, n5) | ~ leq(X20, n5) | ~ leq(n0, X21) | ~ leq(n0, X20))) | (? [X0] : (? [X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & (leq(X1, n5) & leq(n0, X1))) & (leq(X0, pred(pv63)) & leq(n0, X0))) | ? [X2, X3] : ((~ (init = a_select3(id_ds1_filter_init, X2, X3)) & gt(pv63, X2)) & (leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2))) | ? [X4, X5] : ((~ (init = a_select3(id_ds1_filter_init, X4, X5)) & (gt(pv64, X5) & (pv63 = X4))) & (leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4))) | ? [X6, X7] : (~ (init = a_select3(pminus_ds1_filter_init, X6, X7)) & (leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6))) | ? [X8, X9] : (~ (init = a_select3(xhatmin_ds1_filter_init, X8, X9)) & (leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8))) | ? [X10, X11] : (~ (init = a_select3(r_ds1_filter_init, X10, X11)) & (leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10))) | ? [X12, X13] : (~ (init = a_select3(q_ds1_filter_init, X12, X13)) & (leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12))) | ? [X14, X15] : (~ (init = a_select3(dv_ds1_filter_init, X14, X15)) & (leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14))) | ? [X16, X17] : (~ (init = a_select3(phi_ds1_filter_init, X16, X17)) & (leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16))) | ? [X18, X19] : (~ (init = a_select3(h_ds1_filter_init, X18, X19)) & (leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18))) | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ((! [X0] : ((leq(X0, pred(pv63)) & leq(n0, X0)) => ! [X1] : ((leq(X1, n5) & leq(n0, X1)) => (init = a_select3(id_ds1_filter_init, X0, X1)))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (gt(pv63, X2) => (init = a_select3(id_ds1_filter_init, X2, X3)))) & ! [X4, X5] : ((leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => ((gt(pv64, X5) & (pv63 = X4)) => (init = a_select3(id_ds1_filter_init, X4, X5)))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(pminus_ds1_filter_init, X6, X7))) & ! [X8, X9] : ((leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(xhatmin_ds1_filter_init, X8, X9))) & ! [X10, X11] : ((leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(r_ds1_filter_init, X10, X11))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(q_ds1_filter_init, X12, X13))) & ! [X14, X15] : ((leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => (init = a_select3(dv_ds1_filter_init, X14, X15))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (init = a_select3(phi_ds1_filter_init, X16, X17))) & ! [X18, X19] : ((leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) => (init = a_select3(h_ds1_filter_init, X18, X19))) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), inference(rectify, [], [f1])).
fof(f1, plain, ((! [X18] : ((leq(X18, pred(pv63)) & leq(n0, X18)) => ! [X19] : ((leq(X19, n5) & leq(n0, X19)) => (init = a_select3(id_ds1_filter_init, X18, X19)))) & ! [X16, X17] : ((leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) => (gt(pv63, X16) => (init = a_select3(id_ds1_filter_init, X16, X17)))) & ! [X14, X15] : ((leq(X15, n5) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) => ((gt(pv64, X15) & (pv63 = X14)) => (init = a_select3(id_ds1_filter_init, X14, X15)))) & ! [X12, X13] : ((leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) => (init = a_select3(pminus_ds1_filter_init, X12, X13))) & ! [X10, X11] : ((leq(X11, n0) & leq(X10, n5) & leq(n0, X11) & leq(n0, X10)) => (init = a_select3(xhatmin_ds1_filter_init, X10, X11))) & ! [X8, X9] : ((leq(X9, n2) & leq(X8, n2) & leq(n0, X9) & leq(n0, X8)) => (init = a_select3(r_ds1_filter_init, X8, X9))) & ! [X6, X7] : ((leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) => (init = a_select3(q_ds1_filter_init, X6, X7))) & ! [X4, X5] : ((leq(X5, n0) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) => (init = a_select3(dv_ds1_filter_init, X4, X5))) & ! [X2, X3] : ((leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) => (init = a_select3(phi_ds1_filter_init, X2, X3))) & ! [X0, X1] : ((leq(X1, n5) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) => (a_select3(h_ds1_filter_init, X0, X1) = init)) & leq(pv64, n5) & leq(pv63, n5) & leq(pv5, n998) & leq(n0, pv64) & leq(n0, pv63) & leq(n0, pv5) & ~ (pv63 = pv64)) => ! [X20, X21] : ((leq(X21, n5) & leq(X20, n5) & leq(n0, X21) & leq(n0, X20)) => ((leq(X21, pv64) & (pv63 = X20) & ~ (pv64 = X21)) => (init = a_select3(id_ds1_filter_init, X20, X21))))), file('/home/ubuntu/library/tptp/Problems/MSC/MSC010+1.p', n1)).
fof(f479, plain, (spl31_58 | ~ spl31_59 | ~ spl31_60 | ~ spl31_61 | ~ spl31_62 | spl31_37 | spl31_31 | spl31_25 | spl31_19 | spl31_13 | spl31_7 | spl31_1 | spl31_50 | spl31_43 | spl31_69 | ~ spl31_64 | ~ spl31_65 | spl31_66), inference(avatar_split_clause, [], [f149, f462, f458, f454, f476, f357, f391, f154, f183, f212, f241, f270, f299, f328, f446, f442, f438, f434, f430])).
fof(f149, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(sK27, pred(pv63)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(duplicate_literal_removal, [], [f145])).
fof(f145, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(sK27, pred(pv63)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(equality_resolution, [], [f113])).
fof(f113, plain, ! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0) | leq(sK27, pred(pv63)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(cnf_transformation, [], [f60])).
fof(f474, plain, (spl31_58 | ~ spl31_59 | ~ spl31_60 | ~ spl31_61 | ~ spl31_62 | spl31_37 | spl31_31 | spl31_25 | spl31_19 | spl31_13 | spl31_7 | spl31_1 | spl31_50 | spl31_43 | spl31_68 | ~ spl31_64 | ~ spl31_65 | spl31_66), inference(avatar_split_clause, [], [f150, f462, f458, f454, f471, f357, f391, f154, f183, f212, f241, f270, f299, f328, f446, f442, f438, f434, f430])).
fof(f150, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(n0, sK28) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(duplicate_literal_removal, [], [f144])).
fof(f144, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(n0, sK28) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(equality_resolution, [], [f114])).
fof(f114, plain, ! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0) | leq(n0, sK28) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(cnf_transformation, [], [f60])).
fof(f469, plain, (spl31_58 | ~ spl31_59 | ~ spl31_60 | ~ spl31_61 | ~ spl31_62 | spl31_37 | spl31_31 | spl31_25 | spl31_19 | spl31_13 | spl31_7 | spl31_1 | spl31_50 | spl31_43 | spl31_67 | ~ spl31_64 | ~ spl31_65 | spl31_66), inference(avatar_split_clause, [], [f151, f462, f458, f454, f466, f357, f391, f154, f183, f212, f241, f270, f299, f328, f446, f442, f438, f434, f430])).
fof(f151, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(sK28, n5) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(duplicate_literal_removal, [], [f143])).
fof(f143, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | leq(sK28, n5) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(equality_resolution, [], [f115])).
fof(f115, plain, ! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0) | leq(sK28, n5) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(cnf_transformation, [], [f60])).
fof(f464, plain, (spl31_58 | ~ spl31_59 | ~ spl31_60 | ~ spl31_61 | ~ spl31_62 | spl31_37 | spl31_31 | spl31_25 | spl31_19 | spl31_13 | spl31_7 | spl31_1 | spl31_50 | spl31_43 | ~ spl31_63 | ~ spl31_64 | ~ spl31_65 | spl31_66), inference(avatar_split_clause, [], [f152, f462, f458, f454, f450, f357, f391, f154, f183, f212, f241, f270, f299, f328, f446, f442, f438, f434, f430])).
fof(f152, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | ~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(duplicate_literal_removal, [], [f142])).
fof(f142, plain, ! [X1] : ((init = a_select3(id_ds1_filter_init, pv63, X1)) | ~ leq(X1, pv64) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(pv63, n5) | ~ leq(n0, X1) | ~ leq(n0, pv63) | ~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(equality_resolution, [], [f116])).
fof(f116, plain, ! [X0, X1] : ((init = a_select3(id_ds1_filter_init, X0, X1)) | ~ leq(X1, pv64) | ~ (pv63 = X0) | (pv64 = X1) | ~ leq(X1, n5) | ~ leq(X0, n5) | ~ leq(n0, X1) | ~ leq(n0, X0) | ~ (init = a_select3(id_ds1_filter_init, sK27, sK28)) | sP1 | sP0 | sP8 | sP7 | sP6 | sP5 | sP4 | sP3 | sP2 | ~ leq(pv64, n5) | ~ leq(pv63, n5) | ~ leq(pv5, n998) | ~ leq(n0, pv64) | ~ leq(n0, pv63) | ~ leq(n0, pv5) | (pv63 = pv64)), inference(cnf_transformation, [], [f60])).
fof(f423, plain, (~ spl31_50 | spl31_56), inference(avatar_split_clause, [], [f106, f420, f391])).
fof(f106, plain, (leq(n0, sK26) | ~ sP0), inference(cnf_transformation, [], [f56])).
fof(f56, plain, ((~ (init = a_select3(id_ds1_filter_init, sK25, sK26)) & gt(pv64, sK26) & (pv63 = sK25) & leq(sK26, n5) & leq(sK25, n5) & leq(n0, sK26) & leq(n0, sK25)) | ~ sP0), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25, sK26])], [f54, f55])).
fof(f55, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & gt(pv64, X1) & (pv63 = X0) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(id_ds1_filter_init, sK25, sK26)) & gt(pv64, sK26) & (pv63 = sK25) & leq(sK26, n5) & leq(sK25, n5) & leq(n0, sK26) & leq(n0, sK25))), introduced(choice_axiom, [])).
fof(f54, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & gt(pv64, X1) & (pv63 = X0) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP0), inference(rectify, [], [f53])).
fof(f53, plain, (? [X4, X5] : (~ (init = a_select3(id_ds1_filter_init, X4, X5)) & gt(pv64, X5) & (pv63 = X4) & leq(X5, n5) & leq(X4, n5) & leq(n0, X5) & leq(n0, X4)) | ~ sP0), inference(nnf_transformation, [], [f11])).
fof(f413, plain, (~ spl31_50 | spl31_54), inference(avatar_split_clause, [], [f108, f410, f391])).
fof(f108, plain, (leq(sK26, n5) | ~ sP0), inference(cnf_transformation, [], [f56])).
fof(f408, plain, (~ spl31_50 | spl31_53), inference(avatar_split_clause, [], [f109, f405, f391])).
fof(f109, plain, ((pv63 = sK25) | ~ sP0), inference(cnf_transformation, [], [f56])).
fof(f403, plain, (~ spl31_50 | spl31_52), inference(avatar_split_clause, [], [f110, f400, f391])).
fof(f110, plain, (gt(pv64, sK26) | ~ sP0), inference(cnf_transformation, [], [f56])).
fof(f398, plain, (~ spl31_50 | ~ spl31_51), inference(avatar_split_clause, [], [f111, f395, f391])).
fof(f111, plain, (~ (init = a_select3(id_ds1_filter_init, sK25, sK26)) | ~ sP0), inference(cnf_transformation, [], [f56])).
fof(f389, plain, (~ spl31_43 | spl31_49), inference(avatar_split_clause, [], [f99, f386, f357])).
fof(f99, plain, (leq(n0, sK23) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f52, plain, ((~ (init = a_select3(id_ds1_filter_init, sK23, sK24)) & gt(pv63, sK23) & leq(sK24, n5) & leq(sK23, n5) & leq(n0, sK24) & leq(n0, sK23)) | ~ sP1), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23, sK24])], [f50, f51])).
fof(f51, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & gt(pv63, X0) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(id_ds1_filter_init, sK23, sK24)) & gt(pv63, sK23) & leq(sK24, n5) & leq(sK23, n5) & leq(n0, sK24) & leq(n0, sK23))), introduced(choice_axiom, [])).
fof(f50, plain, (? [X0, X1] : (~ (init = a_select3(id_ds1_filter_init, X0, X1)) & gt(pv63, X0) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP1), inference(rectify, [], [f49])).
fof(f49, plain, (? [X2, X3] : (~ (init = a_select3(id_ds1_filter_init, X2, X3)) & gt(pv63, X2) & leq(X3, n5) & leq(X2, n5) & leq(n0, X3) & leq(n0, X2)) | ~ sP1), inference(nnf_transformation, [], [f12])).
fof(f384, plain, (~ spl31_43 | spl31_48), inference(avatar_split_clause, [], [f100, f381, f357])).
fof(f100, plain, (leq(n0, sK24) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f379, plain, (~ spl31_43 | spl31_47), inference(avatar_split_clause, [], [f101, f376, f357])).
fof(f101, plain, (leq(sK23, n5) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f374, plain, (~ spl31_43 | spl31_46), inference(avatar_split_clause, [], [f102, f371, f357])).
fof(f102, plain, (leq(sK24, n5) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f369, plain, (~ spl31_43 | spl31_45), inference(avatar_split_clause, [], [f103, f366, f357])).
fof(f103, plain, (gt(pv63, sK23) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f364, plain, (~ spl31_43 | ~ spl31_44), inference(avatar_split_clause, [], [f104, f361, f357])).
fof(f104, plain, (~ (init = a_select3(id_ds1_filter_init, sK23, sK24)) | ~ sP1), inference(cnf_transformation, [], [f52])).
fof(f355, plain, (~ spl31_37 | spl31_42), inference(avatar_split_clause, [], [f94, f352, f328])).
fof(f94, plain, (leq(n0, sK21) | ~ sP2), inference(cnf_transformation, [], [f48])).
fof(f48, plain, ((~ (init = a_select3(h_ds1_filter_init, sK21, sK22)) & leq(sK22, n5) & leq(sK21, n2) & leq(n0, sK22) & leq(n0, sK21)) | ~ sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22])], [f46, f47])).
fof(f47, plain, (? [X0, X1] : (~ (a_select3(h_ds1_filter_init, X0, X1) = init) & leq(X1, n5) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(h_ds1_filter_init, sK21, sK22)) & leq(sK22, n5) & leq(sK21, n2) & leq(n0, sK22) & leq(n0, sK21))), introduced(choice_axiom, [])).
fof(f46, plain, (? [X0, X1] : (~ (a_select3(h_ds1_filter_init, X0, X1) = init) & leq(X1, n5) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) | ~ sP2), inference(rectify, [], [f45])).
fof(f45, plain, (? [X18, X19] : (~ (init = a_select3(h_ds1_filter_init, X18, X19)) & leq(X19, n5) & leq(X18, n2) & leq(n0, X19) & leq(n0, X18)) | ~ sP2), inference(nnf_transformation, [], [f13])).
fof(f350, plain, (~ spl31_37 | spl31_41), inference(avatar_split_clause, [], [f95, f347, f328])).
fof(f95, plain, (leq(n0, sK22) | ~ sP2), inference(cnf_transformation, [], [f48])).
fof(f345, plain, (~ spl31_37 | spl31_40), inference(avatar_split_clause, [], [f96, f342, f328])).
fof(f96, plain, (leq(sK21, n2) | ~ sP2), inference(cnf_transformation, [], [f48])).
fof(f340, plain, (~ spl31_37 | spl31_39), inference(avatar_split_clause, [], [f97, f337, f328])).
fof(f97, plain, (leq(sK22, n5) | ~ sP2), inference(cnf_transformation, [], [f48])).
fof(f335, plain, (~ spl31_37 | ~ spl31_38), inference(avatar_split_clause, [], [f98, f332, f328])).
fof(f98, plain, (~ (init = a_select3(h_ds1_filter_init, sK21, sK22)) | ~ sP2), inference(cnf_transformation, [], [f48])).
fof(f326, plain, (~ spl31_31 | spl31_36), inference(avatar_split_clause, [], [f89, f323, f299])).
fof(f89, plain, (leq(n0, sK19) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f44, plain, ((~ (init = a_select3(phi_ds1_filter_init, sK19, sK20)) & leq(sK20, n5) & leq(sK19, n5) & leq(n0, sK20) & leq(n0, sK19)) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19, sK20])], [f42, f43])).
fof(f43, plain, (? [X0, X1] : (~ (init = a_select3(phi_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(phi_ds1_filter_init, sK19, sK20)) & leq(sK20, n5) & leq(sK19, n5) & leq(n0, sK20) & leq(n0, sK19))), introduced(choice_axiom, [])).
fof(f42, plain, (? [X0, X1] : (~ (init = a_select3(phi_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP3), inference(rectify, [], [f41])).
fof(f41, plain, (? [X16, X17] : (~ (init = a_select3(phi_ds1_filter_init, X16, X17)) & leq(X17, n5) & leq(X16, n5) & leq(n0, X17) & leq(n0, X16)) | ~ sP3), inference(nnf_transformation, [], [f14])).
fof(f321, plain, (~ spl31_31 | spl31_35), inference(avatar_split_clause, [], [f90, f318, f299])).
fof(f90, plain, (leq(n0, sK20) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f316, plain, (~ spl31_31 | spl31_34), inference(avatar_split_clause, [], [f91, f313, f299])).
fof(f91, plain, (leq(sK19, n5) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f311, plain, (~ spl31_31 | spl31_33), inference(avatar_split_clause, [], [f92, f308, f299])).
fof(f92, plain, (leq(sK20, n5) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f306, plain, (~ spl31_31 | ~ spl31_32), inference(avatar_split_clause, [], [f93, f303, f299])).
fof(f93, plain, (~ (init = a_select3(phi_ds1_filter_init, sK19, sK20)) | ~ sP3), inference(cnf_transformation, [], [f44])).
fof(f297, plain, (~ spl31_25 | spl31_30), inference(avatar_split_clause, [], [f84, f294, f270])).
fof(f84, plain, (leq(n0, sK17) | ~ sP4), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ((~ (init = a_select3(dv_ds1_filter_init, sK17, sK18)) & leq(sK18, n0) & leq(sK17, n5) & leq(n0, sK18) & leq(n0, sK17)) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18])], [f38, f39])).
fof(f39, plain, (? [X0, X1] : (~ (init = a_select3(dv_ds1_filter_init, X0, X1)) & leq(X1, n0) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(dv_ds1_filter_init, sK17, sK18)) & leq(sK18, n0) & leq(sK17, n5) & leq(n0, sK18) & leq(n0, sK17))), introduced(choice_axiom, [])).
fof(f38, plain, (? [X0, X1] : (~ (init = a_select3(dv_ds1_filter_init, X0, X1)) & leq(X1, n0) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP4), inference(rectify, [], [f37])).
fof(f37, plain, (? [X14, X15] : (~ (init = a_select3(dv_ds1_filter_init, X14, X15)) & leq(X15, n0) & leq(X14, n5) & leq(n0, X15) & leq(n0, X14)) | ~ sP4), inference(nnf_transformation, [], [f15])).
fof(f292, plain, (~ spl31_25 | spl31_29), inference(avatar_split_clause, [], [f85, f289, f270])).
fof(f85, plain, (leq(n0, sK18) | ~ sP4), inference(cnf_transformation, [], [f40])).
fof(f287, plain, (~ spl31_25 | spl31_28), inference(avatar_split_clause, [], [f86, f284, f270])).
fof(f86, plain, (leq(sK17, n5) | ~ sP4), inference(cnf_transformation, [], [f40])).
fof(f282, plain, (~ spl31_25 | spl31_27), inference(avatar_split_clause, [], [f87, f279, f270])).
fof(f87, plain, (leq(sK18, n0) | ~ sP4), inference(cnf_transformation, [], [f40])).
fof(f277, plain, (~ spl31_25 | ~ spl31_26), inference(avatar_split_clause, [], [f88, f274, f270])).
fof(f88, plain, (~ (init = a_select3(dv_ds1_filter_init, sK17, sK18)) | ~ sP4), inference(cnf_transformation, [], [f40])).
fof(f268, plain, (~ spl31_19 | spl31_24), inference(avatar_split_clause, [], [f79, f265, f241])).
fof(f79, plain, (leq(n0, sK15) | ~ sP5), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ((~ (init = a_select3(q_ds1_filter_init, sK15, sK16)) & leq(sK16, n5) & leq(sK15, n5) & leq(n0, sK16) & leq(n0, sK15)) | ~ sP5), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16])], [f34, f35])).
fof(f35, plain, (? [X0, X1] : (~ (init = a_select3(q_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(q_ds1_filter_init, sK15, sK16)) & leq(sK16, n5) & leq(sK15, n5) & leq(n0, sK16) & leq(n0, sK15))), introduced(choice_axiom, [])).
fof(f34, plain, (? [X0, X1] : (~ (init = a_select3(q_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP5), inference(rectify, [], [f33])).
fof(f33, plain, (? [X12, X13] : (~ (init = a_select3(q_ds1_filter_init, X12, X13)) & leq(X13, n5) & leq(X12, n5) & leq(n0, X13) & leq(n0, X12)) | ~ sP5), inference(nnf_transformation, [], [f16])).
fof(f263, plain, (~ spl31_19 | spl31_23), inference(avatar_split_clause, [], [f80, f260, f241])).
fof(f80, plain, (leq(n0, sK16) | ~ sP5), inference(cnf_transformation, [], [f36])).
fof(f258, plain, (~ spl31_19 | spl31_22), inference(avatar_split_clause, [], [f81, f255, f241])).
fof(f81, plain, (leq(sK15, n5) | ~ sP5), inference(cnf_transformation, [], [f36])).
fof(f253, plain, (~ spl31_19 | spl31_21), inference(avatar_split_clause, [], [f82, f250, f241])).
fof(f82, plain, (leq(sK16, n5) | ~ sP5), inference(cnf_transformation, [], [f36])).
fof(f248, plain, (~ spl31_19 | ~ spl31_20), inference(avatar_split_clause, [], [f83, f245, f241])).
fof(f83, plain, (~ (init = a_select3(q_ds1_filter_init, sK15, sK16)) | ~ sP5), inference(cnf_transformation, [], [f36])).
fof(f239, plain, (~ spl31_13 | spl31_18), inference(avatar_split_clause, [], [f74, f236, f212])).
fof(f74, plain, (leq(n0, sK13) | ~ sP6), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ((~ (init = a_select3(r_ds1_filter_init, sK13, sK14)) & leq(sK14, n2) & leq(sK13, n2) & leq(n0, sK14) & leq(n0, sK13)) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14])], [f30, f31])).
fof(f31, plain, (? [X0, X1] : (~ (init = a_select3(r_ds1_filter_init, X0, X1)) & leq(X1, n2) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(r_ds1_filter_init, sK13, sK14)) & leq(sK14, n2) & leq(sK13, n2) & leq(n0, sK14) & leq(n0, sK13))), introduced(choice_axiom, [])).
fof(f30, plain, (? [X0, X1] : (~ (init = a_select3(r_ds1_filter_init, X0, X1)) & leq(X1, n2) & leq(X0, n2) & leq(n0, X1) & leq(n0, X0)) | ~ sP6), inference(rectify, [], [f29])).
fof(f29, plain, (? [X10, X11] : (~ (init = a_select3(r_ds1_filter_init, X10, X11)) & leq(X11, n2) & leq(X10, n2) & leq(n0, X11) & leq(n0, X10)) | ~ sP6), inference(nnf_transformation, [], [f17])).
fof(f234, plain, (~ spl31_13 | spl31_17), inference(avatar_split_clause, [], [f75, f231, f212])).
fof(f75, plain, (leq(n0, sK14) | ~ sP6), inference(cnf_transformation, [], [f32])).
fof(f229, plain, (~ spl31_13 | spl31_16), inference(avatar_split_clause, [], [f76, f226, f212])).
fof(f76, plain, (leq(sK13, n2) | ~ sP6), inference(cnf_transformation, [], [f32])).
fof(f224, plain, (~ spl31_13 | spl31_15), inference(avatar_split_clause, [], [f77, f221, f212])).
fof(f77, plain, (leq(sK14, n2) | ~ sP6), inference(cnf_transformation, [], [f32])).
fof(f219, plain, (~ spl31_13 | ~ spl31_14), inference(avatar_split_clause, [], [f78, f216, f212])).
fof(f78, plain, (~ (init = a_select3(r_ds1_filter_init, sK13, sK14)) | ~ sP6), inference(cnf_transformation, [], [f32])).
fof(f210, plain, (~ spl31_7 | spl31_12), inference(avatar_split_clause, [], [f69, f207, f183])).
fof(f69, plain, (leq(n0, sK11) | ~ sP7), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ((~ (init = a_select3(xhatmin_ds1_filter_init, sK11, sK12)) & leq(sK12, n0) & leq(sK11, n5) & leq(n0, sK12) & leq(n0, sK11)) | ~ sP7), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12])], [f26, f27])).
fof(f27, plain, (? [X0, X1] : (~ (init = a_select3(xhatmin_ds1_filter_init, X0, X1)) & leq(X1, n0) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(xhatmin_ds1_filter_init, sK11, sK12)) & leq(sK12, n0) & leq(sK11, n5) & leq(n0, sK12) & leq(n0, sK11))), introduced(choice_axiom, [])).
fof(f26, plain, (? [X0, X1] : (~ (init = a_select3(xhatmin_ds1_filter_init, X0, X1)) & leq(X1, n0) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP7), inference(rectify, [], [f25])).
fof(f25, plain, (? [X8, X9] : (~ (init = a_select3(xhatmin_ds1_filter_init, X8, X9)) & leq(X9, n0) & leq(X8, n5) & leq(n0, X9) & leq(n0, X8)) | ~ sP7), inference(nnf_transformation, [], [f18])).
fof(f205, plain, (~ spl31_7 | spl31_11), inference(avatar_split_clause, [], [f70, f202, f183])).
fof(f70, plain, (leq(n0, sK12) | ~ sP7), inference(cnf_transformation, [], [f28])).
fof(f200, plain, (~ spl31_7 | spl31_10), inference(avatar_split_clause, [], [f71, f197, f183])).
fof(f71, plain, (leq(sK11, n5) | ~ sP7), inference(cnf_transformation, [], [f28])).
fof(f195, plain, (~ spl31_7 | spl31_9), inference(avatar_split_clause, [], [f72, f192, f183])).
fof(f72, plain, (leq(sK12, n0) | ~ sP7), inference(cnf_transformation, [], [f28])).
fof(f190, plain, (~ spl31_7 | ~ spl31_8), inference(avatar_split_clause, [], [f73, f187, f183])).
fof(f73, plain, (~ (init = a_select3(xhatmin_ds1_filter_init, sK11, sK12)) | ~ sP7), inference(cnf_transformation, [], [f28])).
fof(f181, plain, (~ spl31_1 | spl31_6), inference(avatar_split_clause, [], [f64, f178, f154])).
fof(f64, plain, (leq(n0, sK9) | ~ sP8), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ((~ (init = a_select3(pminus_ds1_filter_init, sK9, sK10)) & leq(sK10, n5) & leq(sK9, n5) & leq(n0, sK10) & leq(n0, sK9)) | ~ sP8), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10])], [f22, f23])).
fof(f23, plain, (? [X0, X1] : (~ (init = a_select3(pminus_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) => (~ (init = a_select3(pminus_ds1_filter_init, sK9, sK10)) & leq(sK10, n5) & leq(sK9, n5) & leq(n0, sK10) & leq(n0, sK9))), introduced(choice_axiom, [])).
fof(f22, plain, (? [X0, X1] : (~ (init = a_select3(pminus_ds1_filter_init, X0, X1)) & leq(X1, n5) & leq(X0, n5) & leq(n0, X1) & leq(n0, X0)) | ~ sP8), inference(rectify, [], [f21])).
fof(f21, plain, (? [X6, X7] : (~ (init = a_select3(pminus_ds1_filter_init, X6, X7)) & leq(X7, n5) & leq(X6, n5) & leq(n0, X7) & leq(n0, X6)) | ~ sP8), inference(nnf_transformation, [], [f19])).
fof(f176, plain, (~ spl31_1 | spl31_5), inference(avatar_split_clause, [], [f65, f173, f154])).
fof(f65, plain, (leq(n0, sK10) | ~ sP8), inference(cnf_transformation, [], [f24])).
fof(f171, plain, (~ spl31_1 | spl31_4), inference(avatar_split_clause, [], [f66, f168, f154])).
fof(f66, plain, (leq(sK9, n5) | ~ sP8), inference(cnf_transformation, [], [f24])).
fof(f166, plain, (~ spl31_1 | spl31_3), inference(avatar_split_clause, [], [f67, f163, f154])).
fof(f67, plain, (leq(sK10, n5) | ~ sP8), inference(cnf_transformation, [], [f24])).
fof(f161, plain, (~ spl31_1 | ~ spl31_2), inference(avatar_split_clause, [], [f68, f158, f154])).
fof(f68, plain, (~ (init = a_select3(pminus_ds1_filter_init, sK9, sK10)) | ~ sP8), inference(cnf_transformation, [], [f24])).