fof(f693, plain, $false, inference(avatar_sat_refutation, [], [f107, f111, f115, f119, f124, f129, f133, f137, f141, f145, f150, f155, f426, f692])).
fof(f692, plain, (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6 | ~ spl25_7), inference(avatar_contradiction_clause, [], [f691])).
fof(f691, plain, ($false | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6 | ~ spl25_7)), inference(subsumption_resolution, [], [f690, f128])).
fof(f128, plain, (actual_world(sK7) | ~ spl25_7), inference(avatar_component_clause, [], [f126])).
fof(f126, plain, (spl25_7 <=> actual_world(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl25_7])])).
fof(f690, plain, (~ actual_world(sK7) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6)), inference(resolution, [], [f689, f123])).
fof(f123, plain, (sP1(sK7) | ~ spl25_6), inference(avatar_component_clause, [], [f121])).
fof(f121, plain, (spl25_6 <=> sP1(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl25_6])])).
fof(f689, plain, (! [X0] : (~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f688, f78])).
fof(f78, plain, ! [X0] : (revenge(X0, sK18(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : ((of(X0, sK20(X0), sK18(X0)) & scream(X0, sK20(X0)) & nonreflexive(X0, sK20(X0)) & present(X0, sK20(X0)) & patient(X0, sK20(X0), sK19(X0)) & agent(X0, sK20(X0), sK17(X0)) & event(X0, sK20(X0)) & cry(X0, sK19(X0)) & revenge(X0, sK18(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK17(X0))) & group(X0, sK17(X0)) & six(X0, sK17(X0)) & sP0(sK16(X0), X0, sK15(X0), sK17(X0)) & cannon(X0, sK16(X0)) & of(X0, sK16(X0), sK15(X0)) & man(X0, sK15(X0)) & male(X0, sK15(X0)) & male(X0, sK17(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17, sK18, sK19, sK20])], [f26, f27])).
fof(f27, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP0(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) => (of(X0, sK20(X0), sK18(X0)) & scream(X0, sK20(X0)) & nonreflexive(X0, sK20(X0)) & present(X0, sK20(X0)) & patient(X0, sK20(X0), sK19(X0)) & agent(X0, sK20(X0), sK17(X0)) & event(X0, sK20(X0)) & cry(X0, sK19(X0)) & revenge(X0, sK18(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK17(X0))) & group(X0, sK17(X0)) & six(X0, sK17(X0)) & sP0(sK16(X0), X0, sK15(X0), sK17(X0)) & cannon(X0, sK16(X0)) & of(X0, sK16(X0), sK15(X0)) & man(X0, sK15(X0)) & male(X0, sK15(X0)) & male(X0, sK17(X0)))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP0(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) | ~ sP1(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23)) | ~ sP1(X20)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23)) | ~ sP1(X20)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X20] : (sP1(X20) <=> ? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f688, plain, (! [X0] : (~ revenge(X0, sK18(X0)) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f687])).
fof(f687, plain, (! [X0] : (~ revenge(X0, sK18(X0)) | ~ sP1(X0) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f685, f86])).
fof(f86, plain, ! [X0] : (of(X0, sK20(X0), sK18(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f685, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f684, f79])).
fof(f79, plain, ! [X0] : (cry(X0, sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f684, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f683])).
fof(f683, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f679, f82])).
fof(f82, plain, ! [X0] : (patient(X0, sK20(X0), sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f679, plain, (! [X2, X0, X1] : (~ patient(X0, sK20(X0), X2) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f678, f85])).
fof(f85, plain, ! [X0] : (scream(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f678, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X2) | ~ scream(X0, sK20(X0))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f677, f84])).
fof(f84, plain, ! [X0] : (nonreflexive(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f677, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X2) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f676, f83])).
fof(f83, plain, ! [X0] : (present(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f676, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X2) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f675, f80])).
fof(f80, plain, ! [X0] : (event(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f675, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ event(X0, sK20(X0)) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X2) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f674])).
fof(f674, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK20(X0), X1) | ~ cry(X0, X2) | ~ revenge(X0, X1) | ~ event(X0, sK20(X0)) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X2) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f664, f81])).
fof(f81, plain, ! [X0] : (agent(X0, sK20(X0), sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f664, plain, (! [X6, X4, X7, X5] : (~ agent(X4, X5, sK17(X4)) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ sP1(X4) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f663, f69])).
fof(f69, plain, ! [X0] : (male(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f663, plain, (! [X6, X4, X7, X5] : (~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f662, f605])).
fof(f605, plain, (! [X0] : (shot(X0, sK5(X0, sK17(X0))) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f604])).
fof(f604, plain, (! [X0] : (~ actual_world(X0) | ~ sP1(X0) | shot(X0, sK5(X0, sK17(X0))) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f596, f77])).
fof(f77, plain, ! [X0, X7] : (~ member(X0, X7, sK17(X0)) | shot(X0, X7) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f596, plain, (! [X0] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f595, f78])).
fof(f595, plain, (! [X0] : (~ revenge(X0, sK18(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f594])).
fof(f594, plain, (! [X0] : (~ revenge(X0, sK18(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f591, f86])).
fof(f591, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ revenge(X0, X1) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f590, f79])).
fof(f590, plain, (! [X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ of(X0, sK20(X0), X1) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f589])).
fof(f589, plain, (! [X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ of(X0, sK20(X0), X1) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f585, f82])).
fof(f585, plain, (! [X2, X0, X1] : (~ patient(X0, sK20(X0), X1) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ of(X0, sK20(X0), X2) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f584, f85])).
fof(f584, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ of(X0, sK20(X0), X2) | ~ patient(X0, sK20(X0), X1) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f583, f84])).
fof(f583, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ of(X0, sK20(X0), X2) | ~ patient(X0, sK20(X0), X1) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f582, f83])).
fof(f582, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ of(X0, sK20(X0), X2) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f581, f80])).
fof(f581, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f580])).
fof(f580, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f576, f81])).
fof(f576, plain, (! [X2, X0, X3, X1] : (~ agent(X0, X1, sK17(X0)) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ of(X0, X1, X2) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f575, f70])).
fof(f70, plain, ! [X0] : (male(X0, sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f575, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP1(X0) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f574, f71])).
fof(f71, plain, ! [X0] : (man(X0, sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f574, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP1(X0) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f573, f73])).
fof(f73, plain, ! [X0] : (cannon(X0, sK16(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f573, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP1(X0) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f572, f72])).
fof(f72, plain, ! [X0] : (of(X0, sK16(X0), sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f572, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f571, f69])).
fof(f571, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK17(X0)) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f570, f76])).
fof(f76, plain, ! [X0] : (group(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f570, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK17(X0)) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f569, f75])).
fof(f75, plain, ! [X0] : (six(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f569, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK17(X0)) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | ~ male(X0, sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f568])).
fof(f568, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X3) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK17(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK17(X0)) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ cannon(X0, sK16(X0)) | ~ man(X0, sK15(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ male(X0, sK15(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f552, f480])).
fof(f480, plain, (! [X2, X0, X1] : (member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f479, f78])).
fof(f479, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl25_5), inference(duplicate_literal_removal, [], [f478])).
fof(f478, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl25_5), inference(resolution, [], [f467, f86])).
fof(f467, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK20(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f466, f76])).
fof(f466, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f465, f75])).
fof(f465, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f464, f69])).
fof(f464, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK17(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0)) | ~ spl25_5), inference(duplicate_literal_removal, [], [f463])).
fof(f463, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK17(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl25_5), inference(resolution, [], [f444, f81])).
fof(f444, plain, (! [X14, X12, X15, X13, X11] : (~ agent(X11, sK20(X11), X12) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ revenge(X11, X15) | ~ actual_world(X11) | ~ of(X11, sK20(X11), X15) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f443, f85])).
fof(f443, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f442, f84])).
fof(f442, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f441, f83])).
fof(f441, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f440, f80])).
fof(f440, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ revenge(X11, X15) | ~ event(X11, sK20(X11)) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f432, f79])).
fof(f432, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, X12), X12) | ~ six(X11, X12) | ~ group(X11, X12) | member(X11, sK5(X11, X12), X12) | ~ cry(X11, sK19(X11)) | ~ revenge(X11, X15) | ~ event(X11, sK20(X11)) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(resolution, [], [f118, f82])).
fof(f118, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ patient(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ of(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_5), inference(avatar_component_clause, [], [f117])).
fof(f117, plain, (spl25_5 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_5])])).
fof(f552, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f551, f72])).
fof(f551, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f550, f71])).
fof(f550, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f549, f70])).
fof(f549, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(duplicate_literal_removal, [], [f548])).
fof(f548, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(resolution, [], [f514, f167])).
fof(f167, plain, ! [X0, X1] : (agent(X0, sK21(sK16(X0), X0, sK15(X0), X1), sK15(X0)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f88, f74])).
fof(f74, plain, ! [X0] : (sP0(sK16(X0), X0, sK15(X0), sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f88, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK21(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK21(X0, X1, X2, X4), X0) & fire(X1, sK21(X0, X1, X2, X4)) & nonreflexive(X1, sK21(X0, X1, X2, X4)) & present(X1, sK21(X0, X1, X2, X4)) & patient(X1, sK21(X0, X1, X2, X4), X4) & agent(X1, sK21(X0, X1, X2, X4), X2) & event(X1, sK21(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21])], [f30, f31])).
fof(f31, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK21(X0, X1, X2, X4), X0) & fire(X1, sK21(X0, X1, X2, X4)) & nonreflexive(X1, sK21(X0, X1, X2, X4)) & present(X1, sK21(X0, X1, X2, X4)) & patient(X1, sK21(X0, X1, X2, X4), X4) & agent(X1, sK21(X0, X1, X2, X4), X2) & event(X1, sK21(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(rectify, [], [f29])).
fof(f29, plain, ! [X22, X20, X21, X23] : (! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) | ~ sP0(X22, X20, X21, X23)), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ! [X22, X20, X21, X23] : (! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) | ~ sP0(X22, X20, X21, X23)), inference(usedef, [], [e6])).
fof(e6, plain, ! [X22, X20, X21, X23] : (sP0(X22, X20, X21, X23) <=> ! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f514, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f513, f73])).
fof(f513, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(duplicate_literal_removal, [], [f512])).
fof(f512, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK5(X0, X1), X1) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(resolution, [], [f477, f169])).
fof(f169, plain, ! [X0, X1] : (from_loc(X0, sK21(sK16(X0), X0, sK15(X0), X1), sK16(X0)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f93, f74])).
fof(f93, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK21(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f32])).
fof(f477, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ actual_world(X7) | ~ six(X7, X8) | ~ group(X7, X8) | member(X7, sK5(X7, X8), X8) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f476, f163])).
fof(f163, plain, ! [X0, X1] : (fire(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f92, f74])).
fof(f92, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f476, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | member(X7, sK5(X7, X8), X8) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f475, f162])).
fof(f162, plain, ! [X0, X1] : (nonreflexive(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f91, f74])).
fof(f91, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f475, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | member(X7, sK5(X7, X8), X8) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f474, f161])).
fof(f161, plain, ! [X0, X1] : (present(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f90, f74])).
fof(f90, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f474, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | member(X7, sK5(X7, X8), X8) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f469, f160])).
fof(f160, plain, ! [X0, X1] : (event(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f87, f74])).
fof(f87, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f469, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ event(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | member(X7, sK5(X7, X8), X8) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(resolution, [], [f114, f168])).
fof(f168, plain, ! [X0, X1] : (patient(X0, sK21(sK16(X0), X0, sK15(X0), X1), X1) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f89, f74])).
fof(f89, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK21(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f32])).
fof(f114, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X5) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_4), inference(avatar_component_clause, [], [f113])).
fof(f113, plain, (spl25_4 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_4])])).
fof(f662, plain, (! [X6, X4, X7, X5] : (~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f661, f76])).
fof(f661, plain, (! [X6, X4, X7, X5] : (~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f660, f75])).
fof(f660, plain, (! [X6, X4, X7, X5] : (~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f659, f70])).
fof(f659, plain, (! [X6, X4, X7, X5] : (~ male(X4, sK15(X4)) | ~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f658, f71])).
fof(f658, plain, (! [X6, X4, X7, X5] : (~ man(X4, sK15(X4)) | ~ male(X4, sK15(X4)) | ~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f657, f73])).
fof(f657, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK16(X4)) | ~ man(X4, sK15(X4)) | ~ male(X4, sK15(X4)) | ~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f655, f72])).
fof(f655, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK16(X4), sK15(X4)) | ~ cannon(X4, sK16(X4)) | ~ man(X4, sK15(X4)) | ~ male(X4, sK15(X4)) | ~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f652])).
fof(f652, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK16(X4), sK15(X4)) | ~ cannon(X4, sK16(X4)) | ~ man(X4, sK15(X4)) | ~ male(X4, sK15(X4)) | ~ sP1(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ actual_world(X4) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X7) | ~ revenge(X4, X6) | ~ event(X4, X5) | ~ agent(X4, X5, sK17(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5) | ~ male(X4, sK17(X4)) | ~ sP1(X4)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f650, f533])).
fof(f533, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f532, f72])).
fof(f532, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f531, f71])).
fof(f531, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f530, f70])).
fof(f530, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(duplicate_literal_removal, [], [f529])).
fof(f529, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(resolution, [], [f508, f167])).
fof(f508, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f507, f73])).
fof(f507, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(duplicate_literal_removal, [], [f506])).
fof(f506, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK16(X0), X2) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X2, sK16(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK5(X0, X1)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X2, sK16(X0), X1), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(resolution, [], [f462, f169])).
fof(f462, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ actual_world(X7) | ~ six(X7, X8) | ~ group(X7, X8) | ~ shot(X7, sK5(X7, X8)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f461, f163])).
fof(f461, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | ~ shot(X7, sK5(X7, X8)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f460, f162])).
fof(f460, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | ~ shot(X7, sK5(X7, X8)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f459, f161])).
fof(f459, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | ~ shot(X7, sK5(X7, X8)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f454, f160])).
fof(f454, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ male(X7, X9) | ~ man(X7, X9) | ~ of(X7, X10, X9) | ~ cannon(X7, X10) | ~ event(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X9) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X9, X10, X8)), X10) | ~ six(X7, X8) | ~ group(X7, X8) | ~ shot(X7, sK5(X7, X8)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X9, X10, X8), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(resolution, [], [f106, f168])).
fof(f106, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X5) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_2), inference(avatar_component_clause, [], [f105])).
fof(f105, plain, (spl25_2 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_2])])).
fof(f650, plain, (! [X2, X0, X1] : (member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ male(X0, X1) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f649, f78])).
fof(f649, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK18(X0)) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f648])).
fof(f648, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK18(X0)) | ~ sP1(X0) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f633, f86])).
fof(f633, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK20(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ male(X0, X1) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f632, f79])).
fof(f632, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f631])).
fof(f631, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, sK19(X0)) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f627, f82])).
fof(f627, plain, (! [X4, X2, X0, X3, X1] : (~ patient(X0, sK20(X0), X4) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f626, f85])).
fof(f626, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X4) | ~ scream(X0, sK20(X0))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f625, f84])).
fof(f625, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X4) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f624, f83])).
fof(f624, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X4) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f623, f80])).
fof(f623, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, sK20(X0)) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X4) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f622])).
fof(f622, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X3) | ~ cry(X0, X4) | ~ revenge(X0, X3) | ~ event(X0, sK20(X0)) | ~ sP1(X0) | ~ patient(X0, sK20(X0), X4) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f612, f81])).
fof(f612, plain, (! [X6, X4, X8, X7, X5, X9] : (~ agent(X4, X7, sK17(X4)) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ sP1(X4) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f611, f76])).
fof(f611, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, sK17(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f610, f75])).
fof(f610, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, sK17(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f608, f69])).
fof(f608, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ male(X4, sK17(X4)) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, sK17(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f607])).
fof(f607, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ actual_world(X4) | ~ male(X4, sK17(X4)) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, sK17(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f605, f110])).
fof(f110, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ shot(X0, sK5(X0, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ of(X0, X6, X5) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_3), inference(avatar_component_clause, [], [f109])).
fof(f109, plain, (spl25_3 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_3])])).
fof(f426, plain, (~ spl25_13 | ~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11 | ~ spl25_12), inference(avatar_split_clause, [], [f423, f147, f143, f139, f135, f131, f152])).
fof(f152, plain, (spl25_13 <=> actual_world(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl25_13])])).
fof(f131, plain, (spl25_8 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_8])])).
fof(f135, plain, (spl25_9 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_9])])).
fof(f139, plain, (spl25_10 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_10])])).
fof(f143, plain, (spl25_11 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_11])])).
fof(f147, plain, (spl25_12 <=> sP3(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl25_12])])).
fof(f423, plain, (~ actual_world(sK24) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11 | ~ spl25_12)), inference(resolution, [], [f419, f149])).
fof(f149, plain, (sP3(sK24) | ~ spl25_12), inference(avatar_component_clause, [], [f147])).
fof(f419, plain, (! [X0] : (~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f418, f54])).
fof(f54, plain, ! [X0] : (revenge(X0, sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((of(X0, sK13(X0), sK12(X0)) & scream(X0, sK13(X0)) & nonreflexive(X0, sK13(X0)) & present(X0, sK13(X0)) & patient(X0, sK13(X0), sK11(X0)) & agent(X0, sK13(X0), sK10(X0)) & event(X0, sK13(X0)) & revenge(X0, sK12(X0)) & cry(X0, sK11(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK10(X0))) & group(X0, sK10(X0)) & six(X0, sK10(X0)) & sP2(sK9(X0), X0, sK8(X0), sK10(X0)) & cannon(X0, sK9(X0)) & of(X0, sK9(X0), sK8(X0)) & man(X0, sK8(X0)) & male(X0, sK8(X0)) & male(X0, sK10(X0))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10, sK11, sK12, sK13])], [f18, f19])).
fof(f19, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) => (of(X0, sK13(X0), sK12(X0)) & scream(X0, sK13(X0)) & nonreflexive(X0, sK13(X0)) & present(X0, sK13(X0)) & patient(X0, sK13(X0), sK11(X0)) & agent(X0, sK13(X0), sK10(X0)) & event(X0, sK13(X0)) & revenge(X0, sK12(X0)) & cry(X0, sK11(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK10(X0))) & group(X0, sK10(X0)) & six(X0, sK10(X0)) & sP2(sK9(X0), X0, sK8(X0), sK10(X0)) & cannon(X0, sK9(X0)) & of(X0, sK9(X0), sK8(X0)) & man(X0, sK8(X0)) & male(X0, sK8(X0)) & male(X0, sK10(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) | ~ sP3(X0)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) | ~ sP3(X0)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X0] : (sP3(X0) <=> ? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f418, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f417])).
fof(f417, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f410, f61])).
fof(f61, plain, ! [X0] : (of(X0, sK13(X0), sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f410, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f409, f53])).
fof(f53, plain, ! [X0] : (cry(X0, sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f409, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f408])).
fof(f408, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f407, f57])).
fof(f57, plain, ! [X0] : (patient(X0, sK13(X0), sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f407, plain, (! [X2, X0, X1] : (~ patient(X0, sK13(X0), X2) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f406, f60])).
fof(f60, plain, ! [X0] : (scream(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f406, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X2) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f405, f59])).
fof(f59, plain, ! [X0] : (nonreflexive(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f405, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X2) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f404, f58])).
fof(f58, plain, ! [X0] : (present(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f404, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f403, f55])).
fof(f55, plain, ! [X0] : (event(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f403, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f402])).
fof(f402, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f384, f56])).
fof(f56, plain, ! [X0] : (agent(X0, sK13(X0), sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f384, plain, (! [X6, X4, X7, X5] : (~ agent(X4, X5, sK10(X4)) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ sP3(X4) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f383, f330])).
fof(f330, plain, (! [X0] : (shot(X0, sK22(X0, sK10(X0))) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f329])).
fof(f329, plain, (! [X0] : (~ actual_world(X0) | ~ sP3(X0) | shot(X0, sK22(X0, sK10(X0))) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f325, f52])).
fof(f52, plain, ! [X0, X7] : (~ member(X0, X7, sK10(X0)) | shot(X0, X7) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f325, plain, (! [X0] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f324, f54])).
fof(f324, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f323])).
fof(f323, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f316, f61])).
fof(f316, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f315, f53])).
fof(f315, plain, (! [X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ of(X0, sK13(X0), X1) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f314])).
fof(f314, plain, (! [X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ of(X0, sK13(X0), X1) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f304, f57])).
fof(f304, plain, (! [X2, X0, X1] : (~ patient(X0, sK13(X0), X2) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ of(X0, sK13(X0), X1) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f303, f60])).
fof(f303, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ of(X0, sK13(X0), X1) | ~ patient(X0, sK13(X0), X2) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f302, f59])).
fof(f302, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ of(X0, sK13(X0), X1) | ~ patient(X0, sK13(X0), X2) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f301, f58])).
fof(f301, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ of(X0, sK13(X0), X1) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f300, f55])).
fof(f300, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f299])).
fof(f299, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f298, f56])).
fof(f298, plain, (! [X2, X0, X3, X1] : (~ agent(X0, X1, sK10(X0)) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ of(X0, X1, X2) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f297, f45])).
fof(f45, plain, ! [X0] : (male(X0, sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f297, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP3(X0) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f296, f46])).
fof(f46, plain, ! [X0] : (man(X0, sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f296, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP3(X0) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f295, f48])).
fof(f48, plain, ! [X0] : (cannon(X0, sK9(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f295, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP3(X0) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f294, f47])).
fof(f47, plain, ! [X0] : (of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f294, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f293, f44])).
fof(f44, plain, ! [X0] : (male(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f293, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK10(X0)) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f292, f51])).
fof(f51, plain, ! [X0] : (group(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f292, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK10(X0)) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f291, f50])).
fof(f50, plain, ! [X0] : (six(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f291, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK10(X0)) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | ~ male(X0, sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f290])).
fof(f290, plain, (! [X2, X0, X3, X1] : (~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X2) | ~ cry(X0, X3) | ~ event(X0, X1) | ~ agent(X0, X1, sK10(X0)) | ~ patient(X0, X1, X3) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ male(X0, sK10(X0)) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ cannon(X0, sK9(X0)) | ~ man(X0, sK8(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ male(X0, sK8(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f275, f217])).
fof(f217, plain, (! [X2, X0, X1] : (member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f216, f54])).
fof(f216, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl25_11), inference(duplicate_literal_removal, [], [f215])).
fof(f215, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f214, f61])).
fof(f214, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK13(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f213, f51])).
fof(f213, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f212, f50])).
fof(f212, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f211, f44])).
fof(f211, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK10(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0)) | ~ spl25_11), inference(duplicate_literal_removal, [], [f210])).
fof(f210, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK10(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f176, f56])).
fof(f176, plain, (! [X4, X2, X0, X3, X1] : (~ agent(X0, sK13(X0), X1) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X4) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f175, f60])).
fof(f175, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f174, f59])).
fof(f174, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f173, f58])).
fof(f173, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f172, f55])).
fof(f172, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f170, f53])).
fof(f170, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK23(X0, X2, X3, X1), X1) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ cry(X0, sK11(X0)) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f144, f57])).
fof(f144, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ patient(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ of(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_11), inference(avatar_component_clause, [], [f143])).
fof(f275, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f274, f47])).
fof(f274, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f273, f46])).
fof(f273, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f272, f45])).
fof(f272, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(duplicate_literal_removal, [], [f271])).
fof(f271, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f248, f164])).
fof(f164, plain, ! [X0, X1] : (agent(X0, sK14(sK9(X0), X0, sK8(X0), X1), sK8(X0)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f63, f49])).
fof(f49, plain, ! [X0] : (sP2(sK9(X0), X0, sK8(X0), sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f63, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK14(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK14(X0, X1, X2, X4), X0) & fire(X1, sK14(X0, X1, X2, X4)) & nonreflexive(X1, sK14(X0, X1, X2, X4)) & present(X1, sK14(X0, X1, X2, X4)) & patient(X1, sK14(X0, X1, X2, X4), X4) & agent(X1, sK14(X0, X1, X2, X4), X2) & event(X1, sK14(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f22, f23])).
fof(f23, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK14(X0, X1, X2, X4), X0) & fire(X1, sK14(X0, X1, X2, X4)) & nonreflexive(X1, sK14(X0, X1, X2, X4)) & present(X1, sK14(X0, X1, X2, X4)) & patient(X1, sK14(X0, X1, X2, X4), X4) & agent(X1, sK14(X0, X1, X2, X4), X2) & event(X1, sK14(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(rectify, [], [f21])).
fof(f21, plain, ! [X2, X0, X1, X3] : (! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) | ~ sP2(X2, X0, X1, X3)), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ! [X2, X0, X1, X3] : (! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) | ~ sP2(X2, X0, X1, X3)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X2, X0, X1, X3] : (sP2(X2, X0, X1, X3) <=> ! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f248, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f247, f48])).
fof(f247, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(duplicate_literal_removal, [], [f246])).
fof(f246, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ sP3(X0) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f195, f166])).
fof(f166, plain, ! [X0, X1] : (from_loc(X0, sK14(sK9(X0), X0, sK8(X0), X1), sK9(X0)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f68, f49])).
fof(f68, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK14(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f24])).
fof(f195, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X3) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X2, X3, X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f194, f159])).
fof(f159, plain, ! [X0, X1] : (fire(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f67, f49])).
fof(f67, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f194, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X2) | ~ of(X0, X4, X5) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X3) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X2, X3, X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f193, f158])).
fof(f158, plain, ! [X0, X1] : (nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f66, f49])).
fof(f66, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f193, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X2) | ~ of(X0, X4, X5) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X3) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X2, X3, X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f192, f157])).
fof(f157, plain, ! [X0, X1] : (present(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f65, f49])).
fof(f65, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f192, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X2) | ~ of(X0, X4, X5) | ~ present(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X3) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X2, X3, X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f191, f156])).
fof(f156, plain, ! [X0, X1] : (event(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f62, f49])).
fof(f62, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f191, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X2) | ~ of(X0, X4, X5) | ~ present(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, X3, X1)), X3) | ~ six(X0, X1) | ~ group(X0, X1) | member(X0, sK22(X0, X1), X1) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X2, X3, X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f140, f165])).
fof(f165, plain, ! [X0, X1] : (patient(X0, sK14(sK9(X0), X0, sK8(X0), X1), X1) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f64, f49])).
fof(f64, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK14(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f24])).
fof(f140, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X4) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_10), inference(avatar_component_clause, [], [f139])).
fof(f383, plain, (! [X6, X4, X7, X5] : (~ sP3(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f382, f51])).
fof(f382, plain, (! [X6, X4, X7, X5] : (~ sP3(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f381, f50])).
fof(f381, plain, (! [X6, X4, X7, X5] : (~ sP3(X4) | ~ actual_world(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f380, f44])).
fof(f380, plain, (! [X6, X4, X7, X5] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f379, f45])).
fof(f379, plain, (! [X6, X4, X7, X5] : (~ male(X4, sK8(X4)) | ~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f378, f46])).
fof(f378, plain, (! [X6, X4, X7, X5] : (~ man(X4, sK8(X4)) | ~ male(X4, sK8(X4)) | ~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f377, f48])).
fof(f377, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK9(X4)) | ~ man(X4, sK8(X4)) | ~ male(X4, sK8(X4)) | ~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f375, f47])).
fof(f375, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK9(X4), sK8(X4)) | ~ cannon(X4, sK9(X4)) | ~ man(X4, sK8(X4)) | ~ male(X4, sK8(X4)) | ~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f372])).
fof(f372, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK9(X4), sK8(X4)) | ~ cannon(X4, sK9(X4)) | ~ man(X4, sK8(X4)) | ~ male(X4, sK8(X4)) | ~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ actual_world(X4) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK10(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f367, f256])).
fof(f256, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ male(X0, X1) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl25_8), inference(subsumption_resolution, [], [f255, f47])).
fof(f255, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl25_8), inference(subsumption_resolution, [], [f254, f46])).
fof(f254, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl25_8), inference(subsumption_resolution, [], [f253, f45])).
fof(f253, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl25_8), inference(duplicate_literal_removal, [], [f252])).
fof(f252, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X4) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_8), inference(resolution, [], [f240, f164])).
fof(f240, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ sP3(X0) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3)) | ~ spl25_8), inference(subsumption_resolution, [], [f239, f48])).
fof(f239, plain, (! [X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3)) | ~ spl25_8), inference(duplicate_literal_removal, [], [f238])).
fof(f238, plain, (! [X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK9(X0), X2) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X2, sK9(X0), X1)), X2) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ six(X0, X1) | ~ group(X0, X1) | ~ shot(X0, sK22(X0, X1)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X2, sK9(X0), X1), sK10(X0)) | ~ sP3(X0)) | ~ spl25_8), inference(resolution, [], [f190, f166])).
fof(f190, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X9) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X9) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f189, f159])).
fof(f189, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X9) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X9) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f188, f158])).
fof(f188, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X9) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X9) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f187, f157])).
fof(f187, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X9) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ present(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X9) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f183, f156])).
fof(f183, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X9) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ event(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ present(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X9) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(resolution, [], [f165, f132])).
fof(f132, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X4) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_8), inference(avatar_component_clause, [], [f131])).
fof(f367, plain, (! [X2, X0, X1] : (member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ male(X0, X1) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f366, f54])).
fof(f366, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK12(X0)) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f365])).
fof(f365, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK12(X0)) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f358, f61])).
fof(f358, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK13(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ male(X0, X1) | ~ revenge(X0, X3) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f357, f53])).
fof(f357, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, sK11(X0)) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f356])).
fof(f356, plain, (! [X2, X0, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, sK11(X0)) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f355, f57])).
fof(f355, plain, (! [X4, X2, X0, X3, X1] : (~ patient(X0, sK13(X0), X4) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f354, f60])).
fof(f354, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X4) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f353, f59])).
fof(f353, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X4) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f352, f58])).
fof(f352, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f351, f55])).
fof(f351, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, sK13(X0)) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f350])).
fof(f350, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, X4) | ~ event(X0, sK13(X0)) | ~ sP3(X0) | ~ patient(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f340, f56])).
fof(f340, plain, (! [X6, X4, X8, X7, X5, X9] : (~ agent(X4, X7, sK10(X4)) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ sP3(X4) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f339, f51])).
fof(f339, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, sK10(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f338, f50])).
fof(f338, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, sK10(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f336, f44])).
fof(f336, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, sK10(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f335])).
fof(f335, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ actual_world(X4) | ~ male(X4, sK10(X4)) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, sK10(X4)) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f330, f136])).
fof(f136, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ shot(X0, sK22(X0, X3)) | ~ actual_world(X0) | ~ male(X0, X3) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ of(X0, X6, X4) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X3) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_9), inference(avatar_component_clause, [], [f135])).
fof(f155, plain, (spl25_1 | spl25_13), inference(avatar_split_clause, [], [f94, f152, f101])).
fof(f101, plain, (spl25_1 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl25_1])])).
fof(f94, plain, (actual_world(sK24) | sP4), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | (~ shot(X0, sK22(X0, X3)) & member(X0, sK22(X0, X3), X3)) | ~ group(X0, X3) | ~ six(X0, X3) | (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK23(X0, X1, X2, X3), X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3)) | ~ actual_world(X0)) & (sP3(sK24) & actual_world(sK24))) | sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22, sK23, sK24])], [f33, f36, f35, f34])).
fof(f34, plain, ! [X3, X0] : (? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) => (~ shot(X0, sK22(X0, X3)) & member(X0, sK22(X0, X3), X3))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X3, X2, X1, X0] : (? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) => (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK23(X0, X1, X2, X3), X3))), introduced(choice_axiom, [])).
fof(f36, plain, (? [X10] : (sP3(X10) & actual_world(X10)) => (sP3(sK24) & actual_world(sK24))), introduced(choice_axiom, [])).
fof(f33, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3)) | ~ actual_world(X0)) & ? [X10] : (sP3(X10) & actual_world(X10))) | sP4), inference(rectify, [], [f11])).
fof(f11, plain, ((! [X10] : (! [X11, X12, X13, X14, X15, X16] : (~ of(X10, X16, X14) | ~ scream(X10, X16) | ~ nonreflexive(X10, X16) | ~ present(X10, X16) | ~ patient(X10, X16, X15) | ~ agent(X10, X16, X13) | ~ event(X10, X16) | ~ cry(X10, X15) | ~ revenge(X10, X14) | ? [X17] : (~ shot(X10, X17) & member(X10, X17, X13)) | ~ group(X10, X13) | ~ six(X10, X13) | ? [X18] : (! [X19] : (~ from_loc(X10, X19, X12) | ~ fire(X10, X19) | ~ nonreflexive(X10, X19) | ~ present(X10, X19) | ~ patient(X10, X19, X18) | ~ agent(X10, X19, X11) | ~ event(X10, X19)) & member(X10, X18, X13)) | ~ cannon(X10, X12) | ~ of(X10, X12, X11) | ~ man(X10, X11) | ~ male(X10, X11) | ~ male(X10, X13)) | ~ actual_world(X10)) & ? [X0] : (sP3(X0) & actual_world(X0))) | sP4), inference(definition_folding, [], [f5, e10, e9, e8, e7, e6])).
fof(f10, plain, ((! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X33) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31) | ~ male(X30, X33)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X33) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31) | ~ male(X30, X33)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f5, plain, ((! [X10] : (! [X11, X12, X13, X14, X15, X16] : (~ of(X10, X16, X14) | ~ scream(X10, X16) | ~ nonreflexive(X10, X16) | ~ present(X10, X16) | ~ patient(X10, X16, X15) | ~ agent(X10, X16, X13) | ~ event(X10, X16) | ~ cry(X10, X15) | ~ revenge(X10, X14) | ? [X17] : (~ shot(X10, X17) & member(X10, X17, X13)) | ~ group(X10, X13) | ~ six(X10, X13) | ? [X18] : (! [X19] : (~ from_loc(X10, X19, X12) | ~ fire(X10, X19) | ~ nonreflexive(X10, X19) | ~ present(X10, X19) | ~ patient(X10, X19, X18) | ~ agent(X10, X19, X11) | ~ event(X10, X19)) & member(X10, X18, X13)) | ~ cannon(X10, X12) | ~ of(X10, X12, X11) | ~ man(X10, X11) | ~ male(X10, X11) | ~ male(X10, X13)) | ~ actual_world(X10)) & ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & ! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0))) | (! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X33) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31) | ~ male(X30, X33)) | ~ actual_world(X30)) & ? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & ! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23)) & actual_world(X20)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (member(X0, X7, X3) => shot(X0, X7)) & group(X0, X3) & six(X0, X3) & ! [X8] : (member(X0, X8, X3) => ? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X14) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X15) & agent(X10, X16, X13) & event(X10, X16) & cry(X10, X15) & revenge(X10, X14) & ! [X17] : (member(X10, X17, X13) => shot(X10, X17)) & group(X10, X13) & six(X10, X13) & ! [X18] : (member(X10, X18, X13) => ? [X19] : (from_loc(X10, X19, X12) & fire(X10, X19) & nonreflexive(X10, X19) & present(X10, X19) & patient(X10, X19, X18) & agent(X10, X19, X11) & event(X10, X19))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10))) & (? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (member(X20, X27, X23) => shot(X20, X27)) & group(X20, X23) & six(X20, X23) & ! [X28] : (member(X20, X28, X23) => ? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29))) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23)) & actual_world(X20)) => ? [X30] : (? [X31, X32, X33, X34, X35, X36] : (of(X30, X36, X35) & scream(X30, X36) & nonreflexive(X30, X36) & present(X30, X36) & patient(X30, X36, X34) & agent(X30, X36, X33) & event(X30, X36) & revenge(X30, X35) & cry(X30, X34) & ! [X37] : (member(X30, X37, X33) => shot(X30, X37)) & group(X30, X33) & six(X30, X33) & ! [X38] : (member(X30, X38, X33) => ? [X39] : (from_loc(X30, X39, X32) & fire(X30, X39) & nonreflexive(X30, X39) & present(X30, X39) & patient(X30, X39, X38) & agent(X30, X39, X31) & event(X30, X39))) & cannon(X30, X32) & of(X30, X32, X31) & man(X30, X31) & male(X30, X31) & male(X30, X33)) & actual_world(X30)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X3) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (member(X0, X7, X3) => shot(X0, X7)) & group(X0, X3) & six(X0, X3) & ! [X8] : (member(X0, X8, X3) => ? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X14) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X15) & agent(X10, X16, X13) & event(X10, X16) & cry(X10, X15) & revenge(X10, X14) & ! [X17] : (member(X10, X17, X13) => shot(X10, X17)) & group(X10, X13) & six(X10, X13) & ! [X18] : (member(X10, X18, X13) => ? [X19] : (from_loc(X10, X19, X12) & fire(X10, X19) & nonreflexive(X10, X19) & present(X10, X19) & patient(X10, X19, X18) & agent(X10, X19, X11) & event(X10, X19))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10))) & (? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X23) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (member(X20, X27, X23) => shot(X20, X27)) & group(X20, X23) & six(X20, X23) & ! [X28] : (member(X20, X28, X23) => ? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29))) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21) & male(X20, X23)) & actual_world(X20)) => ? [X30] : (? [X31, X32, X33, X34, X35, X36] : (of(X30, X36, X35) & scream(X30, X36) & nonreflexive(X30, X36) & present(X30, X36) & patient(X30, X36, X34) & agent(X30, X36, X33) & event(X30, X36) & revenge(X30, X35) & cry(X30, X34) & ! [X37] : (member(X30, X37, X33) => shot(X30, X37)) & group(X30, X33) & six(X30, X33) & ! [X38] : (member(X30, X38, X33) => ? [X39] : (from_loc(X30, X39, X32) & fire(X30, X39) & nonreflexive(X30, X39) & present(X30, X39) & patient(X30, X39, X38) & agent(X30, X39, X31) & event(X30, X39))) & cannon(X30, X32) & of(X30, X32, X31) & man(X30, X31) & male(X30, X31) & male(X30, X33)) & actual_world(X30)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X13) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X13) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X13) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X3) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1) & male(X0, X3)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X13) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11) & male(X10, X13)) & actual_world(X10)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP080+1.p', co1)).
fof(f150, plain, (spl25_1 | spl25_12), inference(avatar_split_clause, [], [f95, f147, f101])).
fof(f95, plain, (sP3(sK24) | sP4), inference(cnf_transformation, [], [f37])).
fof(f145, plain, (spl25_1 | spl25_11), inference(avatar_split_clause, [], [f96, f143, f101])).
fof(f96, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | member(X0, sK22(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f141, plain, (spl25_1 | spl25_10), inference(avatar_split_clause, [], [f97, f139, f101])).
fof(f97, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | member(X0, sK22(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f137, plain, (spl25_1 | spl25_9), inference(avatar_split_clause, [], [f98, f135, f101])).
fof(f98, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ shot(X0, sK22(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f133, plain, (spl25_1 | spl25_8), inference(avatar_split_clause, [], [f99, f131, f101])).
fof(f99, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ shot(X0, sK22(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f129, plain, (~ spl25_1 | spl25_7), inference(avatar_split_clause, [], [f38, f126, f101])).
fof(f38, plain, (actual_world(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | (~ shot(X0, sK5(X0, X3)) & member(X0, sK5(X0, X3), X3)) | ~ group(X0, X3) | ~ six(X0, X3) | (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK6(X0, X1, X2, X3), X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3)) | ~ actual_world(X0)) & (sP1(sK7) & actual_world(sK7))) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6, sK7])], [f13, f16, f15, f14])).
fof(f14, plain, ! [X3, X0] : (? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) => (~ shot(X0, sK5(X0, X3)) & member(X0, sK5(X0, X3), X3))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X3, X2, X1, X0] : (? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) => (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK6(X0, X1, X2, X3), X3))), introduced(choice_axiom, [])).
fof(f16, plain, (? [X10] : (sP1(X10) & actual_world(X10)) => (sP1(sK7) & actual_world(sK7))), introduced(choice_axiom, [])).
fof(f13, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3)) | ~ actual_world(X0)) & ? [X10] : (sP1(X10) & actual_world(X10))) | ~ sP4), inference(rectify, [], [f12])).
fof(f12, plain, ((! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X33) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31) | ~ male(X30, X33)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f124, plain, (~ spl25_1 | spl25_6), inference(avatar_split_clause, [], [f39, f121, f101])).
fof(f39, plain, (sP1(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f119, plain, (~ spl25_1 | spl25_5), inference(avatar_split_clause, [], [f40, f117, f101])).
fof(f40, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | member(X0, sK5(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f115, plain, (~ spl25_1 | spl25_4), inference(avatar_split_clause, [], [f41, f113, f101])).
fof(f41, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | member(X0, sK5(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f111, plain, (~ spl25_1 | spl25_3), inference(avatar_split_clause, [], [f42, f109, f101])).
fof(f42, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ~ shot(X0, sK5(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f107, plain, (~ spl25_1 | spl25_2), inference(avatar_split_clause, [], [f43, f105, f101])).
fof(f43, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X3) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ~ shot(X0, sK5(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ male(X0, X3) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).