fof(f302, plain, $false, inference(avatar_sat_refutation, [], [f75, f79, f83, f87, f91, f96, f100, f104, f108, f112, f116, f121, f136, f139, f152, f195, f228, f244, f247, f301])).
fof(f301, plain, (~ spl17_2 | ~ spl17_7 | ~ spl17_20), inference(avatar_contradiction_clause, [], [f300])).
fof(f300, plain, ($false | (~ spl17_2 | ~ spl17_7 | ~ spl17_20)), inference(subsumption_resolution, [], [f299, f95])).
fof(f95, plain, (actual_world(sK5) | ~ spl17_7), inference(avatar_component_clause, [], [f93])).
fof(f93, plain, (spl17_7 <=> actual_world(sK5)), introduced(avatar_definition, [new_symbols(naming, [spl17_7])])).
fof(f299, plain, (~ actual_world(sK5) | (~ spl17_2 | ~ spl17_20)), inference(resolution, [], [f296, f219])).
fof(f219, plain, (sP0(sK5, sK3(sK5)) | ~ spl17_20), inference(avatar_component_clause, [], [f217])).
fof(f217, plain, (spl17_20 <=> sP0(sK5, sK3(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl17_20])])).
fof(f296, plain, (! [X0] : (~ sP0(X0, sK3(X0)) | ~ actual_world(X0)) | ~ spl17_2), inference(subsumption_resolution, [], [f295, f49])).
fof(f49, plain, ! [X0, X1] : (coffee(X0, sK11(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1] : ((see(X0, sK13(X0, X1)) & nonreflexive(X0, sK13(X0, X1)) & past(X0, sK13(X0, X1)) & patient(X0, sK13(X0, X1), sK10(X0, X1)) & agent(X0, sK13(X0, X1), X1) & event(X0, sK13(X0, X1)) & drink(X0, sK12(X0, X1)) & nonreflexive(X0, sK12(X0, X1)) & past(X0, sK12(X0, X1)) & patient(X0, sK12(X0, X1), sK11(X0, X1)) & agent(X0, sK12(X0, X1), sK10(X0, X1)) & event(X0, sK12(X0, X1)) & coffee(X0, sK11(X0, X1)) & human_person(X0, sK10(X0, X1))) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10, sK11, sK12, sK13])], [f21, f22])).
fof(f22, plain, ! [X1, X0] : (? [X2, X3, X4, X5] : (see(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X2) & agent(X0, X5, X1) & event(X0, X5) & drink(X0, X4) & nonreflexive(X0, X4) & past(X0, X4) & patient(X0, X4, X3) & agent(X0, X4, X2) & event(X0, X4) & coffee(X0, X3) & human_person(X0, X2)) => (see(X0, sK13(X0, X1)) & nonreflexive(X0, sK13(X0, X1)) & past(X0, sK13(X0, X1)) & patient(X0, sK13(X0, X1), sK10(X0, X1)) & agent(X0, sK13(X0, X1), X1) & event(X0, sK13(X0, X1)) & drink(X0, sK12(X0, X1)) & nonreflexive(X0, sK12(X0, X1)) & past(X0, sK12(X0, X1)) & patient(X0, sK12(X0, X1), sK11(X0, X1)) & agent(X0, sK12(X0, X1), sK10(X0, X1)) & event(X0, sK12(X0, X1)) & coffee(X0, sK11(X0, X1)) & human_person(X0, sK10(X0, X1)))), introduced(choice_axiom, [])).
fof(f21, plain, ! [X0, X1] : (? [X2, X3, X4, X5] : (see(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X2) & agent(X0, X5, X1) & event(X0, X5) & drink(X0, X4) & nonreflexive(X0, X4) & past(X0, X4) & patient(X0, X4, X3) & agent(X0, X4, X2) & event(X0, X4) & coffee(X0, X3) & human_person(X0, X2)) | ~ sP0(X0, X1)), inference(rectify, [], [f20])).
fof(f20, plain, ! [X14, X15] : (? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17)) | ~ sP0(X14, X15)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X14, X15] : (? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17)) | ~ sP0(X14, X15)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X14, X15] : (sP0(X14, X15) <=> ? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f295, plain, (! [X0] : (~ coffee(X0, sK11(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(duplicate_literal_removal, [], [f294])).
fof(f294, plain, (! [X0] : (~ coffee(X0, sK11(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0)) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(resolution, [], [f280, f52])).
fof(f52, plain, ! [X0, X1] : (patient(X0, sK12(X0, X1), sK11(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f280, plain, (! [X0, X1] : (~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ coffee(X0, X1) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(subsumption_resolution, [], [f279, f55])).
fof(f55, plain, ! [X0, X1] : (drink(X0, sK12(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f279, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ drink(X0, sK12(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(subsumption_resolution, [], [f278, f54])).
fof(f54, plain, ! [X0, X1] : (nonreflexive(X0, sK12(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f278, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ nonreflexive(X0, sK12(X0, sK3(X0))) | ~ drink(X0, sK12(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(subsumption_resolution, [], [f277, f53])).
fof(f53, plain, ! [X0, X1] : (past(X0, sK12(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f277, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ past(X0, sK12(X0, sK3(X0))) | ~ nonreflexive(X0, sK12(X0, sK3(X0))) | ~ drink(X0, sK12(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(subsumption_resolution, [], [f276, f50])).
fof(f50, plain, ! [X0, X1] : (event(X0, sK12(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f276, plain, (! [X0, X1] : (~ event(X0, sK12(X0, sK3(X0))) | ~ coffee(X0, X1) | ~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ past(X0, sK12(X0, sK3(X0))) | ~ nonreflexive(X0, sK12(X0, sK3(X0))) | ~ drink(X0, sK12(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(duplicate_literal_removal, [], [f273])).
fof(f273, plain, (! [X0, X1] : (~ event(X0, sK12(X0, sK3(X0))) | ~ coffee(X0, X1) | ~ patient(X0, sK12(X0, sK3(X0)), X1) | ~ past(X0, sK12(X0, sK3(X0))) | ~ nonreflexive(X0, sK12(X0, sK3(X0))) | ~ drink(X0, sK12(X0, sK3(X0))) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0)) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(resolution, [], [f269, f51])).
fof(f51, plain, ! [X0, X1] : (agent(X0, sK12(X0, X1), sK10(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f269, plain, (! [X2, X0, X1] : (~ agent(X0, X2, sK10(X0, sK3(X0))) | ~ event(X0, X2) | ~ coffee(X0, X1) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(subsumption_resolution, [], [f268, f48])).
fof(f48, plain, ! [X0, X1] : (human_person(X0, sK10(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f268, plain, (! [X2, X0, X1] : (~ coffee(X0, X1) | ~ human_person(X0, sK10(X0, sK3(X0))) | ~ event(X0, X2) | ~ agent(X0, X2, sK10(X0, sK3(X0))) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(duplicate_literal_removal, [], [f267])).
fof(f267, plain, (! [X2, X0, X1] : (~ coffee(X0, X1) | ~ human_person(X0, sK10(X0, sK3(X0))) | ~ event(X0, X2) | ~ agent(X0, X2, sK10(X0, sK3(X0))) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP0(X0, sK3(X0)) | ~ sP0(X0, sK3(X0))) | ~ spl17_2), inference(resolution, [], [f238, f58])).
fof(f58, plain, ! [X0, X1] : (patient(X0, sK13(X0, X1), sK10(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f238, plain, (! [X6, X4, X7, X5] : (~ patient(X4, sK13(X4, sK3(X4)), X6) | ~ coffee(X4, X5) | ~ human_person(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, X6) | ~ patient(X4, X7, X5) | ~ past(X4, X7) | ~ nonreflexive(X4, X7) | ~ drink(X4, X7) | ~ actual_world(X4) | ~ sP0(X4, sK3(X4))) | ~ spl17_2), inference(subsumption_resolution, [], [f237, f60])).
fof(f60, plain, ! [X0, X1] : (nonreflexive(X0, sK13(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f237, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ coffee(X4, X5) | ~ human_person(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, X6) | ~ patient(X4, X7, X5) | ~ past(X4, X7) | ~ nonreflexive(X4, X7) | ~ drink(X4, X7) | ~ patient(X4, sK13(X4, sK3(X4)), X6) | ~ nonreflexive(X4, sK13(X4, sK3(X4))) | ~ sP0(X4, sK3(X4))) | ~ spl17_2), inference(subsumption_resolution, [], [f236, f59])).
fof(f59, plain, ! [X0, X1] : (past(X0, sK13(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f236, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ coffee(X4, X5) | ~ human_person(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, X6) | ~ patient(X4, X7, X5) | ~ past(X4, X7) | ~ nonreflexive(X4, X7) | ~ drink(X4, X7) | ~ patient(X4, sK13(X4, sK3(X4)), X6) | ~ past(X4, sK13(X4, sK3(X4))) | ~ nonreflexive(X4, sK13(X4, sK3(X4))) | ~ sP0(X4, sK3(X4))) | ~ spl17_2), inference(subsumption_resolution, [], [f235, f61])).
fof(f61, plain, ! [X0, X1] : (see(X0, sK13(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f235, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ coffee(X4, X5) | ~ human_person(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, X6) | ~ patient(X4, X7, X5) | ~ past(X4, X7) | ~ nonreflexive(X4, X7) | ~ drink(X4, X7) | ~ see(X4, sK13(X4, sK3(X4))) | ~ patient(X4, sK13(X4, sK3(X4)), X6) | ~ past(X4, sK13(X4, sK3(X4))) | ~ nonreflexive(X4, sK13(X4, sK3(X4))) | ~ sP0(X4, sK3(X4))) | ~ spl17_2), inference(subsumption_resolution, [], [f230, f56])).
fof(f56, plain, ! [X0, X1] : (event(X0, sK13(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f230, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ coffee(X4, X5) | ~ human_person(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, X6) | ~ patient(X4, X7, X5) | ~ past(X4, X7) | ~ nonreflexive(X4, X7) | ~ drink(X4, X7) | ~ event(X4, sK13(X4, sK3(X4))) | ~ see(X4, sK13(X4, sK3(X4))) | ~ patient(X4, sK13(X4, sK3(X4)), X6) | ~ past(X4, sK13(X4, sK3(X4))) | ~ nonreflexive(X4, sK13(X4, sK3(X4))) | ~ sP0(X4, sK3(X4))) | ~ spl17_2), inference(resolution, [], [f74, f57])).
fof(f57, plain, ! [X0, X1] : (agent(X0, sK13(X0, X1), X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f23])).
fof(f74, plain, (! [X6, X4, X0, X5, X3] : (~ agent(X0, X6, sK3(X0)) | ~ actual_world(X0) | ~ coffee(X0, X3) | ~ human_person(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X4) | ~ patient(X0, X5, X3) | ~ past(X0, X5) | ~ nonreflexive(X0, X5) | ~ drink(X0, X5) | ~ event(X0, X6) | ~ see(X0, X6) | ~ patient(X0, X6, X4) | ~ past(X0, X6) | ~ nonreflexive(X0, X6)) | ~ spl17_2), inference(avatar_component_clause, [], [f73])).
fof(f73, plain, (spl17_2 <=> ! [X3, X5, X0, X4, X6] : (~ see(X0, X6) | ~ actual_world(X0) | ~ coffee(X0, X3) | ~ human_person(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X4) | ~ patient(X0, X5, X3) | ~ past(X0, X5) | ~ nonreflexive(X0, X5) | ~ drink(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, sK3(X0)) | ~ patient(X0, X6, X4) | ~ past(X0, X6) | ~ nonreflexive(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl17_2])])).
fof(f247, plain, (~ spl17_5 | ~ spl17_7 | spl17_22), inference(avatar_contradiction_clause, [], [f246])).
fof(f246, plain, ($false | (~ spl17_5 | ~ spl17_7 | spl17_22)), inference(subsumption_resolution, [], [f245, f95])).
fof(f245, plain, (~ actual_world(sK5) | (~ spl17_5 | spl17_22)), inference(resolution, [], [f227, f86])).
fof(f86, plain, (! [X0] : (customer(X0, sK3(X0)) | ~ actual_world(X0)) | ~ spl17_5), inference(avatar_component_clause, [], [f85])).
fof(f85, plain, (spl17_5 <=> ! [X0] : (customer(X0, sK3(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_5])])).
fof(f227, plain, (~ customer(sK5, sK3(sK5)) | spl17_22), inference(avatar_component_clause, [], [f225])).
fof(f225, plain, (spl17_22 <=> customer(sK5, sK3(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl17_22])])).
fof(f244, plain, (~ spl17_4 | ~ spl17_7 | spl17_21), inference(avatar_contradiction_clause, [], [f243])).
fof(f243, plain, ($false | (~ spl17_4 | ~ spl17_7 | spl17_21)), inference(subsumption_resolution, [], [f242, f95])).
fof(f242, plain, (~ actual_world(sK5) | (~ spl17_4 | spl17_21)), inference(resolution, [], [f223, f82])).
fof(f82, plain, (! [X0] : (restaurant(X0, sK4(X0)) | ~ actual_world(X0)) | ~ spl17_4), inference(avatar_component_clause, [], [f81])).
fof(f81, plain, (spl17_4 <=> ! [X0] : (restaurant(X0, sK4(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_4])])).
fof(f223, plain, (~ restaurant(sK5, sK4(sK5)) | spl17_21), inference(avatar_component_clause, [], [f221])).
fof(f221, plain, (spl17_21 <=> restaurant(sK5, sK4(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl17_21])])).
fof(f228, plain, (spl17_20 | ~ spl17_21 | ~ spl17_22 | ~ spl17_3 | ~ spl17_6 | ~ spl17_7), inference(avatar_split_clause, [], [f215, f93, f89, f77, f225, f221, f217])).
fof(f77, plain, (spl17_3 <=> ! [X0] : (in(X0, sK3(X0), sK4(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_3])])).
fof(f89, plain, (spl17_6 <=> ! [X9, X8] : (sP0(sK5, X8) | ~ customer(sK5, X8) | ~ restaurant(sK5, X9) | ~ in(sK5, X8, X9))), introduced(avatar_definition, [new_symbols(naming, [spl17_6])])).
fof(f215, plain, (~ customer(sK5, sK3(sK5)) | ~ restaurant(sK5, sK4(sK5)) | sP0(sK5, sK3(sK5)) | (~ spl17_3 | ~ spl17_6 | ~ spl17_7)), inference(subsumption_resolution, [], [f214, f95])).
fof(f214, plain, (~ customer(sK5, sK3(sK5)) | ~ restaurant(sK5, sK4(sK5)) | sP0(sK5, sK3(sK5)) | ~ actual_world(sK5) | (~ spl17_3 | ~ spl17_6)), inference(resolution, [], [f90, f78])).
fof(f78, plain, (! [X0] : (in(X0, sK3(X0), sK4(X0)) | ~ actual_world(X0)) | ~ spl17_3), inference(avatar_component_clause, [], [f77])).
fof(f90, plain, (! [X8, X9] : (~ in(sK5, X8, X9) | ~ customer(sK5, X8) | ~ restaurant(sK5, X9) | sP0(sK5, X8)) | ~ spl17_6), inference(avatar_component_clause, [], [f89])).
fof(f195, plain, (~ spl17_8 | ~ spl17_13 | ~ spl17_14), inference(avatar_contradiction_clause, [], [f194])).
fof(f194, plain, ($false | (~ spl17_8 | ~ spl17_13 | ~ spl17_14)), inference(subsumption_resolution, [], [f193, f120])).
fof(f120, plain, (actual_world(sK16) | ~ spl17_13), inference(avatar_component_clause, [], [f118])).
fof(f118, plain, (spl17_13 <=> actual_world(sK16)), introduced(avatar_definition, [new_symbols(naming, [spl17_13])])).
fof(f193, plain, (~ actual_world(sK16) | (~ spl17_8 | ~ spl17_14)), inference(resolution, [], [f189, f127])).
fof(f127, plain, (sP1(sK16, sK14(sK16)) | ~ spl17_14), inference(avatar_component_clause, [], [f125])).
fof(f125, plain, (spl17_14 <=> sP1(sK16, sK14(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl17_14])])).
fof(f189, plain, (! [X0] : (~ sP1(X0, sK14(X0)) | ~ actual_world(X0)) | ~ spl17_8), inference(subsumption_resolution, [], [f188, f34])).
fof(f34, plain, ! [X0, X1] : (coffee(X0, sK6(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1] : ((see(X0, sK9(X0, X1)) & nonreflexive(X0, sK9(X0, X1)) & past(X0, sK9(X0, X1)) & patient(X0, sK9(X0, X1), sK7(X0, X1)) & agent(X0, sK9(X0, X1), X1) & event(X0, sK9(X0, X1)) & drink(X0, sK8(X0, X1)) & nonreflexive(X0, sK8(X0, X1)) & past(X0, sK8(X0, X1)) & patient(X0, sK8(X0, X1), sK6(X0, X1)) & agent(X0, sK8(X0, X1), sK7(X0, X1)) & event(X0, sK8(X0, X1)) & human_person(X0, sK7(X0, X1)) & coffee(X0, sK6(X0, X1))) | ~ sP1(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7, sK8, sK9])], [f17, f18])).
fof(f18, plain, ! [X1, X0] : (? [X2, X3, X4, X5] : (see(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X1) & event(X0, X5) & drink(X0, X4) & nonreflexive(X0, X4) & past(X0, X4) & patient(X0, X4, X2) & agent(X0, X4, X3) & event(X0, X4) & human_person(X0, X3) & coffee(X0, X2)) => (see(X0, sK9(X0, X1)) & nonreflexive(X0, sK9(X0, X1)) & past(X0, sK9(X0, X1)) & patient(X0, sK9(X0, X1), sK7(X0, X1)) & agent(X0, sK9(X0, X1), X1) & event(X0, sK9(X0, X1)) & drink(X0, sK8(X0, X1)) & nonreflexive(X0, sK8(X0, X1)) & past(X0, sK8(X0, X1)) & patient(X0, sK8(X0, X1), sK6(X0, X1)) & agent(X0, sK8(X0, X1), sK7(X0, X1)) & event(X0, sK8(X0, X1)) & human_person(X0, sK7(X0, X1)) & coffee(X0, sK6(X0, X1)))), introduced(choice_axiom, [])).
fof(f17, plain, ! [X0, X1] : (? [X2, X3, X4, X5] : (see(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X1) & event(X0, X5) & drink(X0, X4) & nonreflexive(X0, X4) & past(X0, X4) & patient(X0, X4, X2) & agent(X0, X4, X3) & event(X0, X4) & human_person(X0, X3) & coffee(X0, X2)) | ~ sP1(X0, X1)), inference(rectify, [], [f16])).
fof(f16, plain, ! [X0, X1] : (? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3)) | ~ sP1(X0, X1)), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3)) | ~ sP1(X0, X1)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X0, X1] : (sP1(X0, X1) <=> ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f188, plain, (! [X0] : (~ coffee(X0, sK6(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(duplicate_literal_removal, [], [f187])).
fof(f187, plain, (! [X0] : (~ coffee(X0, sK6(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0)) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(resolution, [], [f166, f38])).
fof(f38, plain, ! [X0, X1] : (patient(X0, sK8(X0, X1), sK6(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f166, plain, (! [X0, X1] : (~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ coffee(X0, X1) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f165, f41])).
fof(f41, plain, ! [X0, X1] : (drink(X0, sK8(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f165, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ drink(X0, sK8(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f164, f40])).
fof(f40, plain, ! [X0, X1] : (nonreflexive(X0, sK8(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f164, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ nonreflexive(X0, sK8(X0, sK14(X0))) | ~ drink(X0, sK8(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f163, f39])).
fof(f39, plain, ! [X0, X1] : (past(X0, sK8(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f163, plain, (! [X0, X1] : (~ coffee(X0, X1) | ~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ past(X0, sK8(X0, sK14(X0))) | ~ nonreflexive(X0, sK8(X0, sK14(X0))) | ~ drink(X0, sK8(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f162, f36])).
fof(f36, plain, ! [X0, X1] : (event(X0, sK8(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f162, plain, (! [X0, X1] : (~ event(X0, sK8(X0, sK14(X0))) | ~ coffee(X0, X1) | ~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ past(X0, sK8(X0, sK14(X0))) | ~ nonreflexive(X0, sK8(X0, sK14(X0))) | ~ drink(X0, sK8(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(duplicate_literal_removal, [], [f159])).
fof(f159, plain, (! [X0, X1] : (~ event(X0, sK8(X0, sK14(X0))) | ~ coffee(X0, X1) | ~ patient(X0, sK8(X0, sK14(X0)), X1) | ~ past(X0, sK8(X0, sK14(X0))) | ~ nonreflexive(X0, sK8(X0, sK14(X0))) | ~ drink(X0, sK8(X0, sK14(X0))) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0)) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(resolution, [], [f155, f37])).
fof(f37, plain, ! [X0, X1] : (agent(X0, sK8(X0, X1), sK7(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f155, plain, (! [X2, X0, X1] : (~ agent(X0, X2, sK7(X0, sK14(X0))) | ~ event(X0, X2) | ~ coffee(X0, X1) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f154, f35])).
fof(f35, plain, ! [X0, X1] : (human_person(X0, sK7(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f154, plain, (! [X2, X0, X1] : (~ human_person(X0, sK7(X0, sK14(X0))) | ~ coffee(X0, X1) | ~ event(X0, X2) | ~ agent(X0, X2, sK7(X0, sK14(X0))) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(duplicate_literal_removal, [], [f153])).
fof(f153, plain, (! [X2, X0, X1] : (~ human_person(X0, sK7(X0, sK14(X0))) | ~ coffee(X0, X1) | ~ event(X0, X2) | ~ agent(X0, X2, sK7(X0, sK14(X0))) | ~ patient(X0, X2, X1) | ~ past(X0, X2) | ~ nonreflexive(X0, X2) | ~ drink(X0, X2) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0)) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(resolution, [], [f145, f44])).
fof(f44, plain, ! [X0, X1] : (patient(X0, sK9(X0, X1), sK7(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f145, plain, (! [X2, X0, X3, X1] : (~ patient(X0, sK9(X0, sK14(X0)), X1) | ~ human_person(X0, X1) | ~ coffee(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X2) | ~ past(X0, X3) | ~ nonreflexive(X0, X3) | ~ drink(X0, X3) | ~ actual_world(X0) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f144, f46])).
fof(f46, plain, ! [X0, X1] : (nonreflexive(X0, sK9(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f144, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ human_person(X0, X1) | ~ coffee(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X2) | ~ past(X0, X3) | ~ nonreflexive(X0, X3) | ~ drink(X0, X3) | ~ patient(X0, sK9(X0, sK14(X0)), X1) | ~ nonreflexive(X0, sK9(X0, sK14(X0))) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f143, f45])).
fof(f45, plain, ! [X0, X1] : (past(X0, sK9(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f143, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ human_person(X0, X1) | ~ coffee(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X2) | ~ past(X0, X3) | ~ nonreflexive(X0, X3) | ~ drink(X0, X3) | ~ patient(X0, sK9(X0, sK14(X0)), X1) | ~ past(X0, sK9(X0, sK14(X0))) | ~ nonreflexive(X0, sK9(X0, sK14(X0))) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f142, f47])).
fof(f47, plain, ! [X0, X1] : (see(X0, sK9(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f142, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ human_person(X0, X1) | ~ coffee(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X2) | ~ past(X0, X3) | ~ nonreflexive(X0, X3) | ~ drink(X0, X3) | ~ see(X0, sK9(X0, sK14(X0))) | ~ patient(X0, sK9(X0, sK14(X0)), X1) | ~ past(X0, sK9(X0, sK14(X0))) | ~ nonreflexive(X0, sK9(X0, sK14(X0))) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(subsumption_resolution, [], [f140, f42])).
fof(f42, plain, ! [X0, X1] : (event(X0, sK9(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f140, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ human_person(X0, X1) | ~ coffee(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X2) | ~ past(X0, X3) | ~ nonreflexive(X0, X3) | ~ drink(X0, X3) | ~ event(X0, sK9(X0, sK14(X0))) | ~ see(X0, sK9(X0, sK14(X0))) | ~ patient(X0, sK9(X0, sK14(X0)), X1) | ~ past(X0, sK9(X0, sK14(X0))) | ~ nonreflexive(X0, sK9(X0, sK14(X0))) | ~ sP1(X0, sK14(X0))) | ~ spl17_8), inference(resolution, [], [f99, f43])).
fof(f43, plain, ! [X0, X1] : (agent(X0, sK9(X0, X1), X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f19])).
fof(f99, plain, (! [X6, X4, X0, X5, X3] : (~ agent(X0, X6, sK14(X0)) | ~ actual_world(X0) | ~ human_person(X0, X3) | ~ coffee(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X3) | ~ patient(X0, X5, X4) | ~ past(X0, X5) | ~ nonreflexive(X0, X5) | ~ drink(X0, X5) | ~ event(X0, X6) | ~ see(X0, X6) | ~ patient(X0, X6, X3) | ~ past(X0, X6) | ~ nonreflexive(X0, X6)) | ~ spl17_8), inference(avatar_component_clause, [], [f98])).
fof(f98, plain, (spl17_8 <=> ! [X3, X5, X0, X4, X6] : (~ see(X0, X6) | ~ actual_world(X0) | ~ human_person(X0, X3) | ~ coffee(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X3) | ~ patient(X0, X5, X4) | ~ past(X0, X5) | ~ nonreflexive(X0, X5) | ~ drink(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, sK14(X0)) | ~ patient(X0, X6, X3) | ~ past(X0, X6) | ~ nonreflexive(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl17_8])])).
fof(f152, plain, (~ spl17_11 | ~ spl17_13 | spl17_16), inference(avatar_contradiction_clause, [], [f151])).
fof(f151, plain, ($false | (~ spl17_11 | ~ spl17_13 | spl17_16)), inference(subsumption_resolution, [], [f150, f120])).
fof(f150, plain, (~ actual_world(sK16) | (~ spl17_11 | spl17_16)), inference(resolution, [], [f135, f111])).
fof(f111, plain, (! [X0] : (customer(X0, sK14(X0)) | ~ actual_world(X0)) | ~ spl17_11), inference(avatar_component_clause, [], [f110])).
fof(f110, plain, (spl17_11 <=> ! [X0] : (customer(X0, sK14(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_11])])).
fof(f135, plain, (~ customer(sK16, sK14(sK16)) | spl17_16), inference(avatar_component_clause, [], [f133])).
fof(f133, plain, (spl17_16 <=> customer(sK16, sK14(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl17_16])])).
fof(f139, plain, (~ spl17_10 | ~ spl17_13 | spl17_15), inference(avatar_contradiction_clause, [], [f138])).
fof(f138, plain, ($false | (~ spl17_10 | ~ spl17_13 | spl17_15)), inference(subsumption_resolution, [], [f137, f120])).
fof(f137, plain, (~ actual_world(sK16) | (~ spl17_10 | spl17_15)), inference(resolution, [], [f131, f107])).
fof(f107, plain, (! [X0] : (restaurant(X0, sK15(X0)) | ~ actual_world(X0)) | ~ spl17_10), inference(avatar_component_clause, [], [f106])).
fof(f106, plain, (spl17_10 <=> ! [X0] : (restaurant(X0, sK15(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_10])])).
fof(f131, plain, (~ restaurant(sK16, sK15(sK16)) | spl17_15), inference(avatar_component_clause, [], [f129])).
fof(f129, plain, (spl17_15 <=> restaurant(sK16, sK15(sK16))), introduced(avatar_definition, [new_symbols(naming, [spl17_15])])).
fof(f136, plain, (spl17_14 | ~ spl17_15 | ~ spl17_16 | ~ spl17_9 | ~ spl17_12 | ~ spl17_13), inference(avatar_split_clause, [], [f123, f118, f114, f102, f133, f129, f125])).
fof(f102, plain, (spl17_9 <=> ! [X0] : (in(X0, sK14(X0), sK15(X0)) | ~ actual_world(X0))), introduced(avatar_definition, [new_symbols(naming, [spl17_9])])).
fof(f114, plain, (spl17_12 <=> ! [X9, X8] : (sP1(sK16, X8) | ~ customer(sK16, X8) | ~ restaurant(sK16, X9) | ~ in(sK16, X8, X9))), introduced(avatar_definition, [new_symbols(naming, [spl17_12])])).
fof(f123, plain, (~ customer(sK16, sK14(sK16)) | ~ restaurant(sK16, sK15(sK16)) | sP1(sK16, sK14(sK16)) | (~ spl17_9 | ~ spl17_12 | ~ spl17_13)), inference(subsumption_resolution, [], [f122, f120])).
fof(f122, plain, (~ customer(sK16, sK14(sK16)) | ~ restaurant(sK16, sK15(sK16)) | sP1(sK16, sK14(sK16)) | ~ actual_world(sK16) | (~ spl17_9 | ~ spl17_12)), inference(resolution, [], [f115, f103])).
fof(f103, plain, (! [X0] : (in(X0, sK14(X0), sK15(X0)) | ~ actual_world(X0)) | ~ spl17_9), inference(avatar_component_clause, [], [f102])).
fof(f115, plain, (! [X8, X9] : (~ in(sK16, X8, X9) | ~ customer(sK16, X8) | ~ restaurant(sK16, X9) | sP1(sK16, X8)) | ~ spl17_12), inference(avatar_component_clause, [], [f114])).
fof(f121, plain, (spl17_1 | spl17_13), inference(avatar_split_clause, [], [f62, f118, f69])).
fof(f69, plain, (spl17_1 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl17_1])])).
fof(f62, plain, (actual_world(sK16) | sP2), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ((! [X0] : ((! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X3) | ~ agent(X0, X6, sK14(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X4) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ coffee(X0, X4) | ~ human_person(X0, X3)) & in(X0, sK14(X0), sK15(X0)) & restaurant(X0, sK15(X0)) & customer(X0, sK14(X0))) | ~ actual_world(X0)) & (! [X8, X9] : (sP1(sK16, X8) | ~ in(sK16, X8, X9) | ~ restaurant(sK16, X9) | ~ customer(sK16, X8)) & actual_world(sK16))) | sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15, sK16])], [f24, f26, f25])).
fof(f25, plain, ! [X0] : (? [X1, X2] : (! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X3) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X4) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ coffee(X0, X4) | ~ human_person(X0, X3)) & in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => (! [X6, X5, X4, X3] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X3) | ~ agent(X0, X6, sK14(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X4) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ coffee(X0, X4) | ~ human_person(X0, X3)) & in(X0, sK14(X0), sK15(X0)) & restaurant(X0, sK15(X0)) & customer(X0, sK14(X0)))), introduced(choice_axiom, [])).
fof(f26, plain, (? [X7] : (! [X8, X9] : (sP1(X7, X8) | ~ in(X7, X8, X9) | ~ restaurant(X7, X9) | ~ customer(X7, X8)) & actual_world(X7)) => (! [X9, X8] : (sP1(sK16, X8) | ~ in(sK16, X8, X9) | ~ restaurant(sK16, X9) | ~ customer(sK16, X8)) & actual_world(sK16))), introduced(choice_axiom, [])).
fof(f24, plain, ((! [X0] : (? [X1, X2] : (! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X3) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X4) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ coffee(X0, X4) | ~ human_person(X0, X3)) & in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) | ~ actual_world(X0)) & ? [X7] : (! [X8, X9] : (sP1(X7, X8) | ~ in(X7, X8, X9) | ~ restaurant(X7, X9) | ~ customer(X7, X8)) & actual_world(X7))) | sP2), inference(rectify, [], [f10])).
fof(f10, plain, ((! [X7] : (? [X8, X9] : (! [X10, X11, X12, X13] : (~ see(X7, X13) | ~ nonreflexive(X7, X13) | ~ past(X7, X13) | ~ patient(X7, X13, X10) | ~ agent(X7, X13, X8) | ~ event(X7, X13) | ~ drink(X7, X12) | ~ nonreflexive(X7, X12) | ~ past(X7, X12) | ~ patient(X7, X12, X11) | ~ agent(X7, X12, X10) | ~ event(X7, X12) | ~ coffee(X7, X11) | ~ human_person(X7, X10)) & in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) | ~ actual_world(X7)) & ? [X0] : (! [X1, X2] : (sP1(X0, X1) | ~ in(X0, X1, X2) | ~ restaurant(X0, X2) | ~ customer(X0, X1)) & actual_world(X0))) | sP2), inference(definition_folding, [], [f6, e9, e8, e7])).
fof(f9, plain, ((! [X21] : (? [X22, X23] : (! [X24, X25, X26, X27] : (~ see(X21, X27) | ~ nonreflexive(X21, X27) | ~ past(X21, X27) | ~ patient(X21, X27, X25) | ~ agent(X21, X27, X22) | ~ event(X21, X27) | ~ drink(X21, X26) | ~ nonreflexive(X21, X26) | ~ past(X21, X26) | ~ patient(X21, X26, X24) | ~ agent(X21, X26, X25) | ~ event(X21, X26) | ~ human_person(X21, X25) | ~ coffee(X21, X24)) & in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) | ~ actual_world(X21)) & ? [X14] : (! [X15, X16] : (sP0(X14, X15) | ~ in(X14, X15, X16) | ~ restaurant(X14, X16) | ~ customer(X14, X15)) & actual_world(X14))) | ~ sP2), inference(usedef, [], [e9])).
fof(e9, plain, (sP2 <=> (! [X21] : (? [X22, X23] : (! [X24, X25, X26, X27] : (~ see(X21, X27) | ~ nonreflexive(X21, X27) | ~ past(X21, X27) | ~ patient(X21, X27, X25) | ~ agent(X21, X27, X22) | ~ event(X21, X27) | ~ drink(X21, X26) | ~ nonreflexive(X21, X26) | ~ past(X21, X26) | ~ patient(X21, X26, X24) | ~ agent(X21, X26, X25) | ~ event(X21, X26) | ~ human_person(X21, X25) | ~ coffee(X21, X24)) & in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) | ~ actual_world(X21)) & ? [X14] : (! [X15, X16] : (sP0(X14, X15) | ~ in(X14, X15, X16) | ~ restaurant(X14, X16) | ~ customer(X14, X15)) & actual_world(X14)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f6, plain, ((! [X7] : (? [X8, X9] : (! [X10, X11, X12, X13] : (~ see(X7, X13) | ~ nonreflexive(X7, X13) | ~ past(X7, X13) | ~ patient(X7, X13, X10) | ~ agent(X7, X13, X8) | ~ event(X7, X13) | ~ drink(X7, X12) | ~ nonreflexive(X7, X12) | ~ past(X7, X12) | ~ patient(X7, X12, X11) | ~ agent(X7, X12, X10) | ~ event(X7, X12) | ~ coffee(X7, X11) | ~ human_person(X7, X10)) & in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) | ~ actual_world(X7)) & ? [X0] : (! [X1, X2] : (? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3)) | ~ in(X0, X1, X2) | ~ restaurant(X0, X2) | ~ customer(X0, X1)) & actual_world(X0))) | (! [X21] : (? [X22, X23] : (! [X24, X25, X26, X27] : (~ see(X21, X27) | ~ nonreflexive(X21, X27) | ~ past(X21, X27) | ~ patient(X21, X27, X25) | ~ agent(X21, X27, X22) | ~ event(X21, X27) | ~ drink(X21, X26) | ~ nonreflexive(X21, X26) | ~ past(X21, X26) | ~ patient(X21, X26, X24) | ~ agent(X21, X26, X25) | ~ event(X21, X26) | ~ human_person(X21, X25) | ~ coffee(X21, X24)) & in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) | ~ actual_world(X21)) & ? [X14] : (! [X15, X16] : (? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17)) | ~ in(X14, X15, X16) | ~ restaurant(X14, X16) | ~ customer(X14, X15)) & actual_world(X14)))), inference(flattening, [], [f5])).
fof(f5, plain, ((! [X7] : (? [X8, X9] : (! [X10, X11, X12, X13] : (~ see(X7, X13) | ~ nonreflexive(X7, X13) | ~ past(X7, X13) | ~ patient(X7, X13, X10) | ~ agent(X7, X13, X8) | ~ event(X7, X13) | ~ drink(X7, X12) | ~ nonreflexive(X7, X12) | ~ past(X7, X12) | ~ patient(X7, X12, X11) | ~ agent(X7, X12, X10) | ~ event(X7, X12) | ~ coffee(X7, X11) | ~ human_person(X7, X10)) & (in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8))) | ~ actual_world(X7)) & ? [X0] : (! [X1, X2] : (? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3)) | (~ in(X0, X1, X2) | ~ restaurant(X0, X2) | ~ customer(X0, X1))) & actual_world(X0))) | (! [X21] : (? [X22, X23] : (! [X24, X25, X26, X27] : (~ see(X21, X27) | ~ nonreflexive(X21, X27) | ~ past(X21, X27) | ~ patient(X21, X27, X25) | ~ agent(X21, X27, X22) | ~ event(X21, X27) | ~ drink(X21, X26) | ~ nonreflexive(X21, X26) | ~ past(X21, X26) | ~ patient(X21, X26, X24) | ~ agent(X21, X26, X25) | ~ event(X21, X26) | ~ human_person(X21, X25) | ~ coffee(X21, X24)) & (in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22))) | ~ actual_world(X21)) & ? [X14] : (! [X15, X16] : (? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17)) | (~ in(X14, X15, X16) | ~ restaurant(X14, X16) | ~ customer(X14, X15))) & actual_world(X14)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3))) & actual_world(X0)) => ? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X10) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X11) & agent(X7, X12, X10) & event(X7, X12) & coffee(X7, X11) & human_person(X7, X10))) & actual_world(X7))) & (? [X14] : (! [X15, X16] : ((in(X14, X15, X16) & restaurant(X14, X16) & customer(X14, X15)) => ? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17))) & actual_world(X14)) => ? [X21] : (! [X22, X23] : ((in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) => ? [X24, X25, X26, X27] : (see(X21, X27) & nonreflexive(X21, X27) & past(X21, X27) & patient(X21, X27, X25) & agent(X21, X27, X22) & event(X21, X27) & drink(X21, X26) & nonreflexive(X21, X26) & past(X21, X26) & patient(X21, X26, X24) & agent(X21, X26, X25) & event(X21, X26) & human_person(X21, X25) & coffee(X21, X24))) & actual_world(X21)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X3) & agent(X0, X5, X4) & event(X0, X5) & human_person(X0, X4) & coffee(X0, X3))) & actual_world(X0)) => ? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X10) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X11) & agent(X7, X12, X10) & event(X7, X12) & coffee(X7, X11) & human_person(X7, X10))) & actual_world(X7))) & (? [X14] : (! [X15, X16] : ((in(X14, X15, X16) & restaurant(X14, X16) & customer(X14, X15)) => ? [X17, X18, X19, X20] : (see(X14, X20) & nonreflexive(X14, X20) & past(X14, X20) & patient(X14, X20, X17) & agent(X14, X20, X15) & event(X14, X20) & drink(X14, X19) & nonreflexive(X14, X19) & past(X14, X19) & patient(X14, X19, X18) & agent(X14, X19, X17) & event(X14, X19) & coffee(X14, X18) & human_person(X14, X17))) & actual_world(X14)) => ? [X21] : (! [X22, X23] : ((in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) => ? [X24, X25, X26, X27] : (see(X21, X27) & nonreflexive(X21, X27) & past(X21, X27) & patient(X21, X27, X25) & agent(X21, X27, X22) & event(X21, X27) & drink(X21, X26) & nonreflexive(X21, X26) & past(X21, X26) & patient(X21, X26, X24) & agent(X21, X26, X25) & event(X21, X26) & human_person(X21, X25) & coffee(X21, X24))) & actual_world(X21)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X11) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X10) & agent(X7, X12, X11) & event(X7, X12) & human_person(X7, X11) & coffee(X7, X10))) & actual_world(X7)) => ? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X3) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X4) & agent(X0, X5, X3) & event(X0, X5) & coffee(X0, X4) & human_person(X0, X3))) & actual_world(X0))) & (? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X3) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X4) & agent(X0, X5, X3) & event(X0, X5) & coffee(X0, X4) & human_person(X0, X3))) & actual_world(X0)) => ? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X11) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X10) & agent(X7, X12, X11) & event(X7, X12) & human_person(X7, X11) & coffee(X7, X10))) & actual_world(X7)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X11) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X10) & agent(X7, X12, X11) & event(X7, X12) & human_person(X7, X11) & coffee(X7, X10))) & actual_world(X7)) => ? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X3) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X4) & agent(X0, X5, X3) & event(X0, X5) & coffee(X0, X4) & human_person(X0, X3))) & actual_world(X0))) & (? [X0] : (! [X1, X2] : ((in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => ? [X3, X4, X5, X6] : (see(X0, X6) & nonreflexive(X0, X6) & past(X0, X6) & patient(X0, X6, X3) & agent(X0, X6, X1) & event(X0, X6) & drink(X0, X5) & nonreflexive(X0, X5) & past(X0, X5) & patient(X0, X5, X4) & agent(X0, X5, X3) & event(X0, X5) & coffee(X0, X4) & human_person(X0, X3))) & actual_world(X0)) => ? [X7] : (! [X8, X9] : ((in(X7, X8, X9) & restaurant(X7, X9) & customer(X7, X8)) => ? [X10, X11, X12, X13] : (see(X7, X13) & nonreflexive(X7, X13) & past(X7, X13) & patient(X7, X13, X11) & agent(X7, X13, X8) & event(X7, X13) & drink(X7, X12) & nonreflexive(X7, X12) & past(X7, X12) & patient(X7, X12, X10) & agent(X7, X12, X11) & event(X7, X12) & human_person(X7, X11) & coffee(X7, X10))) & actual_world(X7)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP094+1.p', co1)).
fof(f116, plain, (spl17_1 | spl17_12), inference(avatar_split_clause, [], [f63, f114, f69])).
fof(f63, plain, ! [X8, X9] : (sP1(sK16, X8) | ~ in(sK16, X8, X9) | ~ restaurant(sK16, X9) | ~ customer(sK16, X8) | sP2), inference(cnf_transformation, [], [f27])).
fof(f112, plain, (spl17_1 | spl17_11), inference(avatar_split_clause, [], [f64, f110, f69])).
fof(f64, plain, ! [X0] : (customer(X0, sK14(X0)) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f27])).
fof(f108, plain, (spl17_1 | spl17_10), inference(avatar_split_clause, [], [f65, f106, f69])).
fof(f65, plain, ! [X0] : (restaurant(X0, sK15(X0)) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f27])).
fof(f104, plain, (spl17_1 | spl17_9), inference(avatar_split_clause, [], [f66, f102, f69])).
fof(f66, plain, ! [X0] : (in(X0, sK14(X0), sK15(X0)) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f27])).
fof(f100, plain, (spl17_1 | spl17_8), inference(avatar_split_clause, [], [f67, f98, f69])).
fof(f67, plain, ! [X6, X4, X0, X5, X3] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X3) | ~ agent(X0, X6, sK14(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X4) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ coffee(X0, X4) | ~ human_person(X0, X3) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f27])).
fof(f96, plain, (~ spl17_1 | spl17_7), inference(avatar_split_clause, [], [f28, f93, f69])).
fof(f28, plain, (actual_world(sK5) | ~ sP2), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ((! [X0] : ((! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, sK3(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X3) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ human_person(X0, X4) | ~ coffee(X0, X3)) & in(X0, sK3(X0), sK4(X0)) & restaurant(X0, sK4(X0)) & customer(X0, sK3(X0))) | ~ actual_world(X0)) & (! [X8, X9] : (sP0(sK5, X8) | ~ in(sK5, X8, X9) | ~ restaurant(sK5, X9) | ~ customer(sK5, X8)) & actual_world(sK5))) | ~ sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5])], [f12, f14, f13])).
fof(f13, plain, ! [X0] : (? [X1, X2] : (! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X3) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ human_person(X0, X4) | ~ coffee(X0, X3)) & in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) => (! [X6, X5, X4, X3] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, sK3(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X3) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ human_person(X0, X4) | ~ coffee(X0, X3)) & in(X0, sK3(X0), sK4(X0)) & restaurant(X0, sK4(X0)) & customer(X0, sK3(X0)))), introduced(choice_axiom, [])).
fof(f14, plain, (? [X7] : (! [X8, X9] : (sP0(X7, X8) | ~ in(X7, X8, X9) | ~ restaurant(X7, X9) | ~ customer(X7, X8)) & actual_world(X7)) => (! [X9, X8] : (sP0(sK5, X8) | ~ in(sK5, X8, X9) | ~ restaurant(sK5, X9) | ~ customer(sK5, X8)) & actual_world(sK5))), introduced(choice_axiom, [])).
fof(f12, plain, ((! [X0] : (? [X1, X2] : (! [X3, X4, X5, X6] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X3) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ human_person(X0, X4) | ~ coffee(X0, X3)) & in(X0, X1, X2) & restaurant(X0, X2) & customer(X0, X1)) | ~ actual_world(X0)) & ? [X7] : (! [X8, X9] : (sP0(X7, X8) | ~ in(X7, X8, X9) | ~ restaurant(X7, X9) | ~ customer(X7, X8)) & actual_world(X7))) | ~ sP2), inference(rectify, [], [f11])).
fof(f11, plain, ((! [X21] : (? [X22, X23] : (! [X24, X25, X26, X27] : (~ see(X21, X27) | ~ nonreflexive(X21, X27) | ~ past(X21, X27) | ~ patient(X21, X27, X25) | ~ agent(X21, X27, X22) | ~ event(X21, X27) | ~ drink(X21, X26) | ~ nonreflexive(X21, X26) | ~ past(X21, X26) | ~ patient(X21, X26, X24) | ~ agent(X21, X26, X25) | ~ event(X21, X26) | ~ human_person(X21, X25) | ~ coffee(X21, X24)) & in(X21, X22, X23) & restaurant(X21, X23) & customer(X21, X22)) | ~ actual_world(X21)) & ? [X14] : (! [X15, X16] : (sP0(X14, X15) | ~ in(X14, X15, X16) | ~ restaurant(X14, X16) | ~ customer(X14, X15)) & actual_world(X14))) | ~ sP2), inference(nnf_transformation, [], [f9])).
fof(f91, plain, (~ spl17_1 | spl17_6), inference(avatar_split_clause, [], [f29, f89, f69])).
fof(f29, plain, ! [X8, X9] : (sP0(sK5, X8) | ~ in(sK5, X8, X9) | ~ restaurant(sK5, X9) | ~ customer(sK5, X8) | ~ sP2), inference(cnf_transformation, [], [f15])).
fof(f87, plain, (~ spl17_1 | spl17_5), inference(avatar_split_clause, [], [f30, f85, f69])).
fof(f30, plain, ! [X0] : (customer(X0, sK3(X0)) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f15])).
fof(f83, plain, (~ spl17_1 | spl17_4), inference(avatar_split_clause, [], [f31, f81, f69])).
fof(f31, plain, ! [X0] : (restaurant(X0, sK4(X0)) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f15])).
fof(f79, plain, (~ spl17_1 | spl17_3), inference(avatar_split_clause, [], [f32, f77, f69])).
fof(f32, plain, ! [X0] : (in(X0, sK3(X0), sK4(X0)) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f15])).
fof(f75, plain, (~ spl17_1 | spl17_2), inference(avatar_split_clause, [], [f33, f73, f69])).
fof(f33, plain, ! [X6, X4, X0, X5, X3] : (~ see(X0, X6) | ~ nonreflexive(X0, X6) | ~ past(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, sK3(X0)) | ~ event(X0, X6) | ~ drink(X0, X5) | ~ nonreflexive(X0, X5) | ~ past(X0, X5) | ~ patient(X0, X5, X3) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ human_person(X0, X4) | ~ coffee(X0, X3) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f15])).