fof(f837, plain, $false, inference(avatar_sat_refutation, [], [f105, f109, f113, f117, f122, f127, f131, f135, f139, f143, f148, f153, f496, f836])).
fof(f836, plain, (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6 | ~ spl25_7), inference(avatar_contradiction_clause, [], [f835])).
fof(f835, plain, ($false | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6 | ~ spl25_7)), inference(subsumption_resolution, [], [f834, f126])).
fof(f126, plain, (actual_world(sK7) | ~ spl25_7), inference(avatar_component_clause, [], [f124])).
fof(f124, plain, (spl25_7 <=> actual_world(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl25_7])])).
fof(f834, plain, (~ actual_world(sK7) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5 | ~ spl25_6)), inference(resolution, [], [f833, f121])).
fof(f121, plain, (sP1(sK7) | ~ spl25_6), inference(avatar_component_clause, [], [f119])).
fof(f119, plain, (spl25_6 <=> sP1(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl25_6])])).
fof(f833, plain, (! [X0] : (~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f832, f76])).
fof(f76, plain, ! [X0] : (revenge(X0, sK18(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : ((of(X0, sK20(X0), sK18(X0)) & scream(X0, sK20(X0)) & nonreflexive(X0, sK20(X0)) & present(X0, sK20(X0)) & patient(X0, sK20(X0), sK19(X0)) & agent(X0, sK20(X0), sK15(X0)) & event(X0, sK20(X0)) & cry(X0, sK19(X0)) & revenge(X0, sK18(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK17(X0))) & group(X0, sK17(X0)) & six(X0, sK17(X0)) & sP0(sK16(X0), X0, sK15(X0), sK17(X0)) & cannon(X0, sK16(X0)) & of(X0, sK16(X0), sK15(X0)) & man(X0, sK15(X0)) & male(X0, sK15(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17, sK18, sK19, sK20])], [f26, f27])).
fof(f27, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP0(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) => (of(X0, sK20(X0), sK18(X0)) & scream(X0, sK20(X0)) & nonreflexive(X0, sK20(X0)) & present(X0, sK20(X0)) & patient(X0, sK20(X0), sK19(X0)) & agent(X0, sK20(X0), sK15(X0)) & event(X0, sK20(X0)) & cry(X0, sK19(X0)) & revenge(X0, sK18(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK17(X0))) & group(X0, sK17(X0)) & six(X0, sK17(X0)) & sP0(sK16(X0), X0, sK15(X0), sK17(X0)) & cannon(X0, sK16(X0)) & of(X0, sK16(X0), sK15(X0)) & man(X0, sK15(X0)) & male(X0, sK15(X0)))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP0(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21)) | ~ sP1(X20)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21)) | ~ sP1(X20)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X20] : (sP1(X20) <=> ? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & sP0(X22, X20, X21, X23) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f832, plain, (! [X0] : (~ actual_world(X0) | ~ sP1(X0) | ~ revenge(X0, sK18(X0))) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f831])).
fof(f831, plain, (! [X0] : (~ actual_world(X0) | ~ sP1(X0) | ~ revenge(X0, sK18(X0)) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f822, f84])).
fof(f84, plain, ! [X0] : (of(X0, sK20(X0), sK18(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f822, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ actual_world(X0) | ~ sP1(X0) | ~ revenge(X0, X1)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f821, f77])).
fof(f77, plain, ! [X0] : (cry(X0, sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f821, plain, (! [X0, X1] : (~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ sP1(X0) | ~ of(X0, sK20(X0), X1)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f820])).
fof(f820, plain, (! [X0, X1] : (~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ sP1(X0) | ~ of(X0, sK20(X0), X1) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f813, f80])).
fof(f80, plain, ! [X0] : (patient(X0, sK20(X0), sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f813, plain, (! [X2, X0, X1] : (~ patient(X0, sK20(X0), X1) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP1(X0) | ~ of(X0, sK20(X0), X2)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f812, f83])).
fof(f83, plain, ! [X0] : (scream(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f812, plain, (! [X2, X0, X1] : (~ sP1(X0) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f811, f82])).
fof(f82, plain, ! [X0] : (nonreflexive(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f811, plain, (! [X2, X0, X1] : (~ sP1(X0) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f810, f81])).
fof(f81, plain, ! [X0] : (present(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f810, plain, (! [X2, X0, X1] : (~ sP1(X0) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f809, f78])).
fof(f78, plain, ! [X0] : (event(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f809, plain, (! [X2, X0, X1] : (~ sP1(X0) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f806])).
fof(f806, plain, (! [X2, X0, X1] : (~ sP1(X0) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f776, f79])).
fof(f79, plain, ! [X0] : (agent(X0, sK20(X0), sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f776, plain, (! [X6, X4, X7, X5] : (~ agent(X4, X7, sK15(X4)) | ~ sP1(X4) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ actual_world(X4) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f775, f678])).
fof(f678, plain, (! [X0] : (shot(X0, sK5(X0, sK17(X0))) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f677])).
fof(f677, plain, (! [X0] : (~ actual_world(X0) | ~ sP1(X0) | shot(X0, sK5(X0, sK17(X0))) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f676, f75])).
fof(f75, plain, ! [X0, X7] : (~ member(X0, X7, sK17(X0)) | shot(X0, X7) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f676, plain, (! [X0] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f675, f76])).
fof(f675, plain, (! [X0] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, sK18(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f674])).
fof(f674, plain, (! [X0] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, sK18(X0)) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f661, f84])).
fof(f661, plain, (! [X0, X1] : (~ of(X0, sK20(X0), X1) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ revenge(X0, X1) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f660, f77])).
fof(f660, plain, (! [X0, X1] : (~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X1) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f659])).
fof(f659, plain, (! [X0, X1] : (~ cry(X0, sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X1) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f652, f80])).
fof(f652, plain, (! [X2, X0, X1] : (~ patient(X0, sK20(X0), X1) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f651, f83])).
fof(f651, plain, (! [X2, X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f650, f82])).
fof(f650, plain, (! [X2, X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f649, f81])).
fof(f649, plain, (! [X2, X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f648, f78])).
fof(f648, plain, (! [X2, X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f645])).
fof(f645, plain, (! [X2, X0, X1] : (member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, sK20(X0)) | ~ actual_world(X0) | ~ patient(X0, sK20(X0), X1) | ~ present(X0, sK20(X0)) | ~ nonreflexive(X0, sK20(X0)) | ~ scream(X0, sK20(X0)) | ~ of(X0, sK20(X0), X2) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f641, f79])).
fof(f641, plain, (! [X2, X0, X3, X1] : (~ agent(X0, X3, sK15(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ actual_world(X0) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f640, f70])).
fof(f70, plain, ! [X0] : (of(X0, sK16(X0), sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f640, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK15(X0)) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0) | ~ of(X0, sK16(X0), sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f639, f71])).
fof(f71, plain, ! [X0] : (cannon(X0, sK16(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f639, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK15(X0)) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0) | ~ cannon(X0, sK16(X0)) | ~ of(X0, sK16(X0), sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f638, f74])).
fof(f74, plain, ! [X0] : (group(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f638, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK15(X0)) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0) | ~ cannon(X0, sK16(X0)) | ~ of(X0, sK16(X0), sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f637, f73])).
fof(f73, plain, ! [X0] : (six(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f637, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK15(X0)) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0) | ~ cannon(X0, sK16(X0)) | ~ of(X0, sK16(X0), sK15(X0))) | (~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f636])).
fof(f636, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ cry(X0, X1) | ~ revenge(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK15(X0)) | ~ patient(X0, X3, X1) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X2) | ~ sP1(X0) | ~ cannon(X0, sK16(X0)) | ~ six(X0, sK17(X0)) | ~ group(X0, sK17(X0)) | member(X0, sK5(X0, sK17(X0)), sK17(X0)) | ~ actual_world(X0) | ~ of(X0, sK16(X0), sK15(X0)) | ~ sP1(X0)) | (~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f618, f549])).
fof(f549, plain, (! [X2, X0, X1] : (member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ cannon(X0, X1) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ actual_world(X0) | ~ of(X0, X1, sK15(X0)) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f548, f76])).
fof(f548, plain, (! [X2, X0, X1] : (~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ of(X0, X1, sK15(X0)) | ~ sP1(X0)) | ~ spl25_5), inference(duplicate_literal_removal, [], [f547])).
fof(f547, plain, (! [X2, X0, X1] : (~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ of(X0, X1, sK15(X0)) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl25_5), inference(resolution, [], [f536, f84])).
fof(f536, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK20(X0), X3) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, X1, sK15(X0)) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f535, f69])).
fof(f69, plain, ! [X0] : (man(X0, sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f535, plain, (! [X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0)) | ~ spl25_5), inference(subsumption_resolution, [], [f534, f68])).
fof(f68, plain, ! [X0] : (male(X0, sK15(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f534, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0)) | ~ spl25_5), inference(duplicate_literal_removal, [], [f533])).
fof(f533, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK20(X0), X3) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl25_5), inference(resolution, [], [f514, f79])).
fof(f514, plain, (! [X14, X12, X15, X13, X11] : (~ agent(X11, sK20(X11), X12) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ revenge(X11, X15) | ~ actual_world(X11) | ~ of(X11, sK20(X11), X15) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f513, f83])).
fof(f513, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f512, f82])).
fof(f512, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f511, f81])).
fof(f511, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ revenge(X11, X15) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f510, f78])).
fof(f510, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ revenge(X11, X15) | ~ event(X11, sK20(X11)) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(subsumption_resolution, [], [f502, f77])).
fof(f502, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ man(X11, X12) | ~ of(X11, X13, X12) | ~ cannon(X11, X13) | member(X11, sK6(X11, X12, X13, X14), X14) | ~ six(X11, X14) | ~ group(X11, X14) | member(X11, sK5(X11, X14), X14) | ~ cry(X11, sK19(X11)) | ~ revenge(X11, X15) | ~ event(X11, sK20(X11)) | ~ agent(X11, sK20(X11), X12) | ~ of(X11, sK20(X11), X15) | ~ present(X11, sK20(X11)) | ~ nonreflexive(X11, sK20(X11)) | ~ scream(X11, sK20(X11)) | ~ sP1(X11)) | ~ spl25_5), inference(resolution, [], [f116, f80])).
fof(f116, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ patient(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ of(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_5), inference(avatar_component_clause, [], [f115])).
fof(f115, plain, (spl25_5 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_5])])).
fof(f618, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ of(X0, X1, X2) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f617, f68])).
fof(f617, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f616, f70])).
fof(f616, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f615, f69])).
fof(f615, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(duplicate_literal_removal, [], [f614])).
fof(f614, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(resolution, [], [f591, f165])).
fof(f165, plain, ! [X0, X1] : (agent(X0, sK21(sK16(X0), X0, sK15(X0), X1), sK15(X0)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f86, f72])).
fof(f72, plain, ! [X0] : (sP0(sK16(X0), X0, sK15(X0), sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f86, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK21(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK21(X0, X1, X2, X4), X0) & fire(X1, sK21(X0, X1, X2, X4)) & nonreflexive(X1, sK21(X0, X1, X2, X4)) & present(X1, sK21(X0, X1, X2, X4)) & patient(X1, sK21(X0, X1, X2, X4), X4) & agent(X1, sK21(X0, X1, X2, X4), X2) & event(X1, sK21(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21])], [f30, f31])).
fof(f31, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK21(X0, X1, X2, X4), X0) & fire(X1, sK21(X0, X1, X2, X4)) & nonreflexive(X1, sK21(X0, X1, X2, X4)) & present(X1, sK21(X0, X1, X2, X4)) & patient(X1, sK21(X0, X1, X2, X4), X4) & agent(X1, sK21(X0, X1, X2, X4), X2) & event(X1, sK21(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(rectify, [], [f29])).
fof(f29, plain, ! [X22, X20, X21, X23] : (! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) | ~ sP0(X22, X20, X21, X23)), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ! [X22, X20, X21, X23] : (! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) | ~ sP0(X22, X20, X21, X23)), inference(usedef, [], [e6])).
fof(e6, plain, ! [X22, X20, X21, X23] : (sP0(X22, X20, X21, X23) <=> ! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f591, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(subsumption_resolution, [], [f590, f71])).
fof(f590, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(duplicate_literal_removal, [], [f589])).
fof(f589, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK5(X0, X2), X2) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_4), inference(resolution, [], [f546, f167])).
fof(f167, plain, ! [X0, X1] : (from_loc(X0, sK21(sK16(X0), X0, sK15(X0), X1), sK16(X0)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f91, f72])).
fof(f91, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK21(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f32])).
fof(f546, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ actual_world(X7) | ~ six(X7, X10) | ~ group(X7, X10) | member(X7, sK5(X7, X10), X10) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f545, f161])).
fof(f161, plain, ! [X0, X1] : (fire(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f90, f72])).
fof(f90, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f545, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | member(X7, sK5(X7, X10), X10) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f544, f160])).
fof(f160, plain, ! [X0, X1] : (nonreflexive(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f89, f72])).
fof(f89, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f544, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | member(X7, sK5(X7, X10), X10) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f543, f159])).
fof(f159, plain, ! [X0, X1] : (present(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f88, f72])).
fof(f88, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f543, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | member(X7, sK5(X7, X10), X10) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(subsumption_resolution, [], [f538, f158])).
fof(f158, plain, ! [X0, X1] : (event(X0, sK21(sK16(X0), X0, sK15(X0), X1)) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f85, f72])).
fof(f85, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK21(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f538, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ event(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | member(X7, sK5(X7, X10), X10) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_4), inference(resolution, [], [f112, f166])).
fof(f166, plain, ! [X0, X1] : (patient(X0, sK21(sK16(X0), X0, sK15(X0), X1), X1) | ~ member(X0, X1, sK17(X0)) | ~ sP1(X0)), inference(resolution, [], [f87, f72])).
fof(f87, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK21(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f32])).
fof(f112, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X5) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_4), inference(avatar_component_clause, [], [f111])).
fof(f111, plain, (spl25_4 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_4])])).
fof(f775, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP1(X4) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f774, f74])).
fof(f774, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP1(X4) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f773, f73])).
fof(f773, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP1(X4) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f772, f70])).
fof(f772, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK16(X4), sK15(X4)) | ~ actual_world(X4) | ~ sP1(X4) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f770, f71])).
fof(f770, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK16(X4)) | ~ of(X4, sK16(X4), sK15(X4)) | ~ actual_world(X4) | ~ sP1(X4) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f767])).
fof(f767, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK16(X4)) | ~ of(X4, sK16(X4), sK15(X4)) | ~ actual_world(X4) | ~ sP1(X4) | ~ actual_world(X4) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ shot(X4, sK5(X4, sK17(X4))) | ~ cry(X4, X5) | ~ revenge(X4, X6) | ~ event(X4, X7) | ~ agent(X4, X7, sK15(X4)) | ~ patient(X4, X7, X5) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7) | ~ of(X4, X7, X6) | ~ sP1(X4)) | (~ spl25_2 | ~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f763, f602])).
fof(f602, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ of(X0, X1, X2) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f601, f68])).
fof(f601, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f600, f70])).
fof(f600, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f599, f69])).
fof(f599, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(duplicate_literal_removal, [], [f598])).
fof(f598, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, sK16(X0), sK15(X0)) | ~ male(X0, sK15(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X2) | ~ event(X0, X1) | ~ agent(X0, X1, sK15(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK15(X0), sK16(X0), X3), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(resolution, [], [f580, f165])).
fof(f580, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK5(X0, X2)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(subsumption_resolution, [], [f579, f71])).
fof(f579, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK5(X0, X2)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(duplicate_literal_removal, [], [f578])).
fof(f578, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK16(X0), X1) | ~ cannon(X0, sK16(X0)) | ~ agent(X0, sK21(sK16(X0), X0, sK15(X0), sK6(X0, X1, sK16(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK5(X0, X2)) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X1, sK16(X0), X2), sK17(X0)) | ~ sP1(X0)) | ~ spl25_2), inference(resolution, [], [f532, f167])).
fof(f532, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ actual_world(X7) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK5(X7, X10)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f531, f161])).
fof(f531, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK5(X7, X10)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f530, f160])).
fof(f530, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK5(X7, X10)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f529, f159])).
fof(f529, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK5(X7, X10)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(subsumption_resolution, [], [f524, f158])).
fof(f524, plain, (! [X12, X10, X8, X7, X13, X11, X9] : (~ actual_world(X7) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ event(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ agent(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X8) | ~ of(X7, X11, X12) | ~ present(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ fire(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10))) | ~ from_loc(X7, sK21(sK16(X7), X7, sK15(X7), sK6(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK5(X7, X10)) | ~ cry(X7, X13) | ~ revenge(X7, X12) | ~ event(X7, X11) | ~ agent(X7, X11, X8) | ~ patient(X7, X11, X13) | ~ present(X7, X11) | ~ nonreflexive(X7, X11) | ~ scream(X7, X11) | ~ member(X7, sK6(X7, X8, X9, X10), sK17(X7)) | ~ sP1(X7)) | ~ spl25_2), inference(resolution, [], [f104, f166])).
fof(f104, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X5) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_2), inference(avatar_component_clause, [], [f103])).
fof(f103, plain, (spl25_2 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_2])])).
fof(f763, plain, (! [X0, X1] : (member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ cannon(X0, X1) | ~ of(X0, X1, sK15(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f762, f76])).
fof(f762, plain, (! [X0, X1] : (~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, X1, sK15(X0)) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f761])).
fof(f761, plain, (! [X0, X1] : (~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, X1, sK15(X0)) | ~ revenge(X0, sK18(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f736, f84])).
fof(f736, plain, (! [X2, X0, X1] : (~ of(X0, sK20(X0), X2) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, X1, sK15(X0)) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f735, f69])).
fof(f735, plain, (! [X2, X0, X1] : (~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f734, f68])).
fof(f734, plain, (! [X2, X0, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f733])).
fof(f733, plain, (! [X2, X0, X1] : (~ male(X0, sK15(X0)) | ~ man(X0, sK15(X0)) | ~ of(X0, X1, sK15(X0)) | ~ cannon(X0, X1) | member(X0, sK6(X0, sK15(X0), X1, sK17(X0)), sK17(X0)) | ~ of(X0, sK20(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f725, f79])).
fof(f725, plain, (! [X12, X10, X11, X9] : (~ agent(X9, sK20(X9), X10) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ revenge(X9, X12) | ~ actual_world(X9) | ~ sP1(X9)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f724, f83])).
fof(f724, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ revenge(X9, X12) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ scream(X9, sK20(X9))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f723, f82])).
fof(f723, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ revenge(X9, X12) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ nonreflexive(X9, sK20(X9)) | ~ scream(X9, sK20(X9))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f722, f81])).
fof(f722, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ revenge(X9, X12) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ present(X9, sK20(X9)) | ~ nonreflexive(X9, sK20(X9)) | ~ scream(X9, sK20(X9))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f721, f78])).
fof(f721, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ revenge(X9, X12) | ~ event(X9, sK20(X9)) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ present(X9, sK20(X9)) | ~ nonreflexive(X9, sK20(X9)) | ~ scream(X9, sK20(X9))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f712, f77])).
fof(f712, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ cry(X9, sK19(X9)) | ~ revenge(X9, X12) | ~ event(X9, sK20(X9)) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ present(X9, sK20(X9)) | ~ nonreflexive(X9, sK20(X9)) | ~ scream(X9, sK20(X9))) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f709])).
fof(f709, plain, (! [X12, X10, X11, X9] : (~ actual_world(X9) | ~ male(X9, X10) | ~ man(X9, X10) | ~ of(X9, X11, X10) | ~ cannon(X9, X11) | member(X9, sK6(X9, X10, X11, sK17(X9)), sK17(X9)) | ~ of(X9, sK20(X9), X12) | ~ cry(X9, sK19(X9)) | ~ revenge(X9, X12) | ~ event(X9, sK20(X9)) | ~ agent(X9, sK20(X9), X10) | ~ sP1(X9) | ~ present(X9, sK20(X9)) | ~ nonreflexive(X9, sK20(X9)) | ~ scream(X9, sK20(X9)) | ~ sP1(X9)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f684, f80])).
fof(f684, plain, (! [X6, X4, X8, X7, X5, X9] : (~ patient(X4, X7, X9) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ sP1(X4) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f683, f74])).
fof(f683, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(subsumption_resolution, [], [f681, f73])).
fof(f681, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(duplicate_literal_removal, [], [f680])).
fof(f680, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP1(X4) | ~ actual_world(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK6(X4, X5, X6, sK17(X4)), sK17(X4)) | ~ six(X4, sK17(X4)) | ~ group(X4, sK17(X4)) | ~ of(X4, X7, X8) | ~ cry(X4, X9) | ~ revenge(X4, X8) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_3 | ~ spl25_4 | ~ spl25_5)), inference(resolution, [], [f678, f108])).
fof(f108, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ shot(X0, sK5(X0, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ of(X0, X6, X5) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_3), inference(avatar_component_clause, [], [f107])).
fof(f107, plain, (spl25_3 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X4) | ~ revenge(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_3])])).
fof(f496, plain, (~ spl25_13 | ~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11 | ~ spl25_12), inference(avatar_split_clause, [], [f493, f145, f141, f137, f133, f129, f150])).
fof(f150, plain, (spl25_13 <=> actual_world(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl25_13])])).
fof(f129, plain, (spl25_8 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_8])])).
fof(f133, plain, (spl25_9 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_9])])).
fof(f137, plain, (spl25_10 <=> ! [X9, X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_10])])).
fof(f141, plain, (spl25_11 <=> ! [X1, X3, X5, X0, X2, X4, X6] : (~ of(X0, X6, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6))), introduced(avatar_definition, [new_symbols(naming, [spl25_11])])).
fof(f145, plain, (spl25_12 <=> sP3(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl25_12])])).
fof(f493, plain, (~ actual_world(sK24) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11 | ~ spl25_12)), inference(resolution, [], [f490, f147])).
fof(f147, plain, (sP3(sK24) | ~ spl25_12), inference(avatar_component_clause, [], [f145])).
fof(f490, plain, (! [X0] : (~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f489, f53])).
fof(f53, plain, ! [X0] : (revenge(X0, sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((of(X0, sK13(X0), sK12(X0)) & scream(X0, sK13(X0)) & nonreflexive(X0, sK13(X0)) & present(X0, sK13(X0)) & patient(X0, sK13(X0), sK11(X0)) & agent(X0, sK13(X0), sK8(X0)) & event(X0, sK13(X0)) & revenge(X0, sK12(X0)) & cry(X0, sK11(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK10(X0))) & group(X0, sK10(X0)) & six(X0, sK10(X0)) & sP2(sK9(X0), X0, sK8(X0), sK10(X0)) & cannon(X0, sK9(X0)) & of(X0, sK9(X0), sK8(X0)) & man(X0, sK8(X0)) & male(X0, sK8(X0))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10, sK11, sK12, sK13])], [f18, f19])).
fof(f19, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) => (of(X0, sK13(X0), sK12(X0)) & scream(X0, sK13(X0)) & nonreflexive(X0, sK13(X0)) & present(X0, sK13(X0)) & patient(X0, sK13(X0), sK11(X0)) & agent(X0, sK13(X0), sK8(X0)) & event(X0, sK13(X0)) & revenge(X0, sK12(X0)) & cry(X0, sK11(X0)) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, sK10(X0))) & group(X0, sK10(X0)) & six(X0, sK10(X0)) & sP2(sK9(X0), X0, sK8(X0), sK10(X0)) & cannon(X0, sK9(X0)) & of(X0, sK9(X0), sK8(X0)) & man(X0, sK8(X0)) & male(X0, sK8(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) | ~ sP3(X0)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) | ~ sP3(X0)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X0] : (sP3(X0) <=> ? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & sP2(X2, X0, X1, X3) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f489, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f488])).
fof(f488, plain, (! [X0] : (~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f484, f60])).
fof(f60, plain, ! [X0] : (of(X0, sK13(X0), sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f484, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f483, f52])).
fof(f52, plain, ! [X0] : (cry(X0, sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f483, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f482])).
fof(f482, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f467, f56])).
fof(f56, plain, ! [X0] : (patient(X0, sK13(X0), sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f467, plain, (! [X2, X0, X1] : (~ patient(X0, sK13(X0), X2) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f466, f59])).
fof(f59, plain, ! [X0] : (scream(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f466, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f465, f58])).
fof(f58, plain, ! [X0] : (nonreflexive(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f465, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f464, f57])).
fof(f57, plain, ! [X0] : (present(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f464, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f463, f54])).
fof(f54, plain, ! [X0] : (event(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f463, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f460])).
fof(f460, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ of(X0, sK13(X0), X1) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f419, f55])).
fof(f55, plain, ! [X0] : (agent(X0, sK13(X0), sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f419, plain, (! [X6, X4, X7, X5] : (~ agent(X4, X5, sK8(X4)) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ actual_world(X4) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f418, f328])).
fof(f328, plain, (! [X0] : (shot(X0, sK22(X0, sK10(X0))) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f327])).
fof(f327, plain, (! [X0] : (~ actual_world(X0) | ~ sP3(X0) | shot(X0, sK22(X0, sK10(X0))) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f323, f51])).
fof(f51, plain, ! [X0, X7] : (~ member(X0, X7, sK10(X0)) | shot(X0, X7) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f323, plain, (! [X0] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f322, f53])).
fof(f322, plain, (! [X0] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, sK12(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f321])).
fof(f321, plain, (! [X0] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, sK12(X0)) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f309, f60])).
fof(f309, plain, (! [X0, X1] : (~ of(X0, sK13(X0), X1) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f308, f52])).
fof(f308, plain, (! [X0, X1] : (~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f307])).
fof(f307, plain, (! [X0, X1] : (~ revenge(X0, X1) | ~ cry(X0, sK11(X0)) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f297, f56])).
fof(f297, plain, (! [X2, X0, X1] : (~ patient(X0, sK13(X0), X2) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f296, f59])).
fof(f296, plain, (! [X2, X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ scream(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f295, f58])).
fof(f295, plain, (! [X2, X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f294, f57])).
fof(f294, plain, (! [X2, X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f293, f54])).
fof(f293, plain, (! [X2, X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f290])).
fof(f290, plain, (! [X2, X0, X1] : (member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, sK13(X0)) | ~ actual_world(X0) | ~ patient(X0, sK13(X0), X2) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ of(X0, sK13(X0), X1) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f289, f55])).
fof(f289, plain, (! [X2, X0, X3, X1] : (~ agent(X0, X3, sK8(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ actual_world(X0) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f288, f46])).
fof(f46, plain, ! [X0] : (of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f288, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK8(X0)) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0) | ~ of(X0, sK9(X0), sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f287, f47])).
fof(f47, plain, ! [X0] : (cannon(X0, sK9(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f287, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK8(X0)) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0) | ~ cannon(X0, sK9(X0)) | ~ of(X0, sK9(X0), sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f286, f50])).
fof(f50, plain, ! [X0] : (group(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f286, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK8(X0)) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0) | ~ cannon(X0, sK9(X0)) | ~ of(X0, sK9(X0), sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f285, f49])).
fof(f49, plain, ! [X0] : (six(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f285, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK8(X0)) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0) | ~ cannon(X0, sK9(X0)) | ~ of(X0, sK9(X0), sK8(X0))) | (~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f284])).
fof(f284, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ revenge(X0, X1) | ~ cry(X0, X2) | ~ event(X0, X3) | ~ agent(X0, X3, sK8(X0)) | ~ patient(X0, X3, X2) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ of(X0, X3, X1) | ~ sP3(X0) | ~ cannon(X0, sK9(X0)) | ~ six(X0, sK10(X0)) | ~ group(X0, sK10(X0)) | member(X0, sK22(X0, sK10(X0)), sK10(X0)) | ~ actual_world(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0)) | (~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f272, f214])).
fof(f214, plain, (! [X2, X0, X1] : (member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ cannon(X0, X1) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ actual_world(X0) | ~ of(X0, X1, sK8(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f213, f53])).
fof(f213, plain, (! [X2, X0, X1] : (~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ of(X0, X1, sK8(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(duplicate_literal_removal, [], [f212])).
fof(f212, plain, (! [X2, X0, X1] : (~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ of(X0, X1, sK8(X0)) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f211, f60])).
fof(f211, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK13(X0), X3) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, X1, sK8(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f210, f45])).
fof(f45, plain, ! [X0] : (man(X0, sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f210, plain, (! [X2, X0, X3, X1] : (~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f209, f44])).
fof(f44, plain, ! [X0] : (male(X0, sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f209, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0)) | ~ spl25_11), inference(duplicate_literal_removal, [], [f208])).
fof(f208, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, X2), X2) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X3) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f174, f55])).
fof(f174, plain, (! [X4, X2, X0, X3, X1] : (~ agent(X0, sK13(X0), X1) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK13(X0), X4) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f173, f59])).
fof(f173, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f172, f58])).
fof(f172, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f171, f57])).
fof(f171, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f170, f54])).
fof(f170, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(subsumption_resolution, [], [f168, f52])).
fof(f168, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, sK11(X0)) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ of(X0, sK13(X0), X4) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | ~ spl25_11), inference(resolution, [], [f142, f56])).
fof(f142, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ patient(X0, X6, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ of(X0, X6, X4) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_11), inference(avatar_component_clause, [], [f141])).
fof(f272, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ of(X0, X1, X2) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f271, f44])).
fof(f271, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f270, f46])).
fof(f270, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, sK8(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f269, f45])).
fof(f269, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, sK8(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(duplicate_literal_removal, [], [f268])).
fof(f268, plain, (! [X4, X2, X0, X3, X1] : (~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ male(X0, sK8(X0)) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f248, f162])).
fof(f162, plain, ! [X0, X1] : (agent(X0, sK14(sK9(X0), X0, sK8(X0), X1), sK8(X0)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f62, f48])).
fof(f48, plain, ! [X0] : (sP2(sK9(X0), X0, sK8(X0), sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f62, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK14(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK14(X0, X1, X2, X4), X0) & fire(X1, sK14(X0, X1, X2, X4)) & nonreflexive(X1, sK14(X0, X1, X2, X4)) & present(X1, sK14(X0, X1, X2, X4)) & patient(X1, sK14(X0, X1, X2, X4), X4) & agent(X1, sK14(X0, X1, X2, X4), X2) & event(X1, sK14(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f22, f23])).
fof(f23, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK14(X0, X1, X2, X4), X0) & fire(X1, sK14(X0, X1, X2, X4)) & nonreflexive(X1, sK14(X0, X1, X2, X4)) & present(X1, sK14(X0, X1, X2, X4)) & patient(X1, sK14(X0, X1, X2, X4), X4) & agent(X1, sK14(X0, X1, X2, X4), X2) & event(X1, sK14(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(rectify, [], [f21])).
fof(f21, plain, ! [X2, X0, X1, X3] : (! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) | ~ sP2(X2, X0, X1, X3)), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ! [X2, X0, X1, X3] : (! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) | ~ sP2(X2, X0, X1, X3)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X2, X0, X1, X3] : (sP2(X2, X0, X1, X3) <=> ! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f248, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ male(X0, X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f247, f47])).
fof(f247, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(duplicate_literal_removal, [], [f246])).
fof(f246, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ actual_world(X0) | ~ six(X0, X2) | ~ group(X0, X2) | member(X0, sK22(X0, X2), X2) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ sP3(X0) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f193, f164])).
fof(f164, plain, ! [X0, X1] : (from_loc(X0, sK14(sK9(X0), X0, sK8(X0), X1), sK9(X0)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f67, f48])).
fof(f67, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK14(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f24])).
fof(f193, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X2) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X1) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X1, X2, X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f192, f157])).
fof(f157, plain, ! [X0, X1] : (fire(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f66, f48])).
fof(f66, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f192, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X1) | ~ of(X0, X4, X5) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X1, X2, X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f191, f156])).
fof(f156, plain, ! [X0, X1] : (nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f65, f48])).
fof(f65, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f191, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X1) | ~ of(X0, X4, X5) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X1, X2, X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f190, f155])).
fof(f155, plain, ! [X0, X1] : (present(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f64, f48])).
fof(f64, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f190, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X1) | ~ of(X0, X4, X5) | ~ present(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X1, X2, X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(subsumption_resolution, [], [f189, f154])).
fof(f154, plain, ! [X0, X1] : (event(X0, sK14(sK9(X0), X0, sK8(X0), X1)) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f61, f48])).
fof(f61, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK14(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f189, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X1) | ~ of(X0, X4, X5) | ~ present(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ nonreflexive(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ fire(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3))) | ~ from_loc(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, X2, X3)), X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK23(X0, X1, X2, X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_10), inference(resolution, [], [f138, f163])).
fof(f163, plain, ! [X0, X1] : (patient(X0, sK14(sK9(X0), X0, sK8(X0), X1), X1) | ~ member(X0, X1, sK10(X0)) | ~ sP3(X0)), inference(resolution, [], [f63, f48])).
fof(f63, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK14(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f24])).
fof(f138, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X4) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK22(X0, X3), X3) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_10), inference(avatar_component_clause, [], [f137])).
fof(f418, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f417, f50])).
fof(f417, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f416, f49])).
fof(f416, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f415, f46])).
fof(f415, plain, (! [X6, X4, X7, X5] : (~ of(X4, sK9(X4), sK8(X4)) | ~ actual_world(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f413, f47])).
fof(f413, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK9(X4)) | ~ of(X4, sK9(X4), sK8(X4)) | ~ actual_world(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f410])).
fof(f410, plain, (! [X6, X4, X7, X5] : (~ cannon(X4, sK9(X4)) | ~ of(X4, sK9(X4), sK8(X4)) | ~ actual_world(X4) | ~ sP3(X4) | ~ sP3(X4) | ~ of(X4, X5, X6) | ~ actual_world(X4) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ shot(X4, sK22(X4, sK10(X4))) | ~ revenge(X4, X6) | ~ cry(X4, X7) | ~ event(X4, X5) | ~ agent(X4, X5, sK8(X4)) | ~ patient(X4, X5, X7) | ~ present(X4, X5) | ~ nonreflexive(X4, X5) | ~ scream(X4, X5)) | (~ spl25_8 | ~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f404, f253])).
fof(f253, plain, (! [X4, X2, X0, X3, X1] : (~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0) | ~ of(X0, X1, X2) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1)) | ~ spl25_8), inference(subsumption_resolution, [], [f252, f46])).
fof(f252, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X1, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1)) | ~ spl25_8), inference(subsumption_resolution, [], [f251, f45])).
fof(f251, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X1, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1)) | ~ spl25_8), inference(subsumption_resolution, [], [f250, f44])).
fof(f250, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X1, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1)) | ~ spl25_8), inference(duplicate_literal_removal, [], [f249])).
fof(f249, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, sK9(X0), sK8(X0)) | ~ sP3(X0) | ~ of(X0, X1, X2) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X2) | ~ cry(X0, X4) | ~ event(X0, X1) | ~ agent(X0, X1, sK8(X0)) | ~ patient(X0, X1, X4) | ~ present(X0, X1) | ~ nonreflexive(X0, X1) | ~ scream(X0, X1) | ~ member(X0, sK23(X0, sK8(X0), sK9(X0), X3), sK10(X0)) | ~ sP3(X0)) | ~ spl25_8), inference(resolution, [], [f242, f162])).
fof(f242, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ sP3(X0) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK22(X0, X2)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3)) | ~ spl25_8), inference(subsumption_resolution, [], [f241, f47])).
fof(f241, plain, (! [X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK22(X0, X2)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3)) | ~ spl25_8), inference(duplicate_literal_removal, [], [f240])).
fof(f240, plain, (! [X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, sK9(X0), X1) | ~ cannon(X0, sK9(X0)) | ~ agent(X0, sK14(sK9(X0), X0, sK8(X0), sK23(X0, X1, sK9(X0), X2)), X1) | ~ of(X0, X3, X4) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ six(X0, X2) | ~ group(X0, X2) | ~ shot(X0, sK22(X0, X2)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X3) | ~ agent(X0, X3, X1) | ~ patient(X0, X3, X5) | ~ present(X0, X3) | ~ nonreflexive(X0, X3) | ~ scream(X0, X3) | ~ member(X0, sK23(X0, X1, sK9(X0), X2), sK10(X0)) | ~ sP3(X0)) | ~ spl25_8), inference(resolution, [], [f188, f164])).
fof(f188, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X7) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f187, f157])).
fof(f187, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X7) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f186, f156])).
fof(f186, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X7) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f185, f155])).
fof(f185, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ present(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X7) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(subsumption_resolution, [], [f181, f154])).
fof(f181, plain, (! [X6, X12, X10, X8, X7, X11, X9] : (~ member(X6, sK23(X6, X7, X8, X9), sK10(X6)) | ~ sP3(X6) | ~ actual_world(X6) | ~ male(X6, X7) | ~ man(X6, X7) | ~ of(X6, X8, X7) | ~ cannon(X6, X8) | ~ event(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ agent(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X7) | ~ of(X6, X10, X11) | ~ present(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ nonreflexive(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ fire(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9))) | ~ from_loc(X6, sK14(sK9(X6), X6, sK8(X6), sK23(X6, X7, X8, X9)), X8) | ~ six(X6, X9) | ~ group(X6, X9) | ~ shot(X6, sK22(X6, X9)) | ~ revenge(X6, X11) | ~ cry(X6, X12) | ~ event(X6, X10) | ~ agent(X6, X10, X7) | ~ patient(X6, X10, X12) | ~ present(X6, X10) | ~ nonreflexive(X6, X10) | ~ scream(X6, X10)) | ~ spl25_8), inference(resolution, [], [f163, f130])).
fof(f130, plain, (! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ event(X0, X9) | ~ agent(X0, X9, X1) | ~ of(X0, X6, X4) | ~ present(X0, X9) | ~ nonreflexive(X0, X9) | ~ fire(X0, X9) | ~ from_loc(X0, X9, X2) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK22(X0, X3)) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_8), inference(avatar_component_clause, [], [f129])).
fof(f404, plain, (! [X0, X1] : (member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ cannon(X0, X1) | ~ of(X0, X1, sK8(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f403, f53])).
fof(f403, plain, (! [X0, X1] : (~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, X1, sK8(X0)) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f402])).
fof(f402, plain, (! [X0, X1] : (~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, X1, sK8(X0)) | ~ revenge(X0, sK12(X0)) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f397, f60])).
fof(f397, plain, (! [X2, X0, X1] : (~ of(X0, sK13(X0), X2) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, X1, sK8(X0)) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f396, f45])).
fof(f396, plain, (! [X2, X0, X1] : (~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f395, f44])).
fof(f395, plain, (! [X2, X0, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f394])).
fof(f394, plain, (! [X2, X0, X1] : (~ male(X0, sK8(X0)) | ~ man(X0, sK8(X0)) | ~ of(X0, X1, sK8(X0)) | ~ cannon(X0, X1) | member(X0, sK23(X0, sK8(X0), X1, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X2) | ~ revenge(X0, X2) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f360, f55])).
fof(f360, plain, (! [X2, X0, X3, X1] : (~ agent(X0, sK13(X0), X1) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f359, f59])).
fof(f359, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f358, f58])).
fof(f358, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f357, f57])).
fof(f357, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f356, f54])).
fof(f356, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f355, f52])).
fof(f355, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, sK11(X0)) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0))) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f350])).
fof(f350, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, sK10(X0)), sK10(X0)) | ~ of(X0, sK13(X0), X3) | ~ revenge(X0, X3) | ~ cry(X0, sK11(X0)) | ~ event(X0, sK13(X0)) | ~ agent(X0, sK13(X0), X1) | ~ sP3(X0) | ~ present(X0, sK13(X0)) | ~ nonreflexive(X0, sK13(X0)) | ~ scream(X0, sK13(X0)) | ~ sP3(X0)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f336, f56])).
fof(f336, plain, (! [X6, X4, X8, X7, X5, X9] : (~ patient(X4, X7, X9) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ sP3(X4) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f335, f50])).
fof(f335, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(subsumption_resolution, [], [f333, f49])).
fof(f333, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(duplicate_literal_removal, [], [f332])).
fof(f332, plain, (! [X6, X4, X8, X7, X5, X9] : (~ sP3(X4) | ~ actual_world(X4) | ~ actual_world(X4) | ~ male(X4, X5) | ~ man(X4, X5) | ~ of(X4, X6, X5) | ~ cannon(X4, X6) | member(X4, sK23(X4, X5, X6, sK10(X4)), sK10(X4)) | ~ six(X4, sK10(X4)) | ~ group(X4, sK10(X4)) | ~ of(X4, X7, X8) | ~ revenge(X4, X8) | ~ cry(X4, X9) | ~ event(X4, X7) | ~ agent(X4, X7, X5) | ~ patient(X4, X7, X9) | ~ present(X4, X7) | ~ nonreflexive(X4, X7) | ~ scream(X4, X7)) | (~ spl25_9 | ~ spl25_10 | ~ spl25_11)), inference(resolution, [], [f328, f134])).
fof(f134, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ shot(X0, sK22(X0, X3)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | ~ of(X0, X6, X4) | ~ revenge(X0, X4) | ~ cry(X0, X5) | ~ event(X0, X6) | ~ agent(X0, X6, X1) | ~ patient(X0, X6, X5) | ~ present(X0, X6) | ~ nonreflexive(X0, X6) | ~ scream(X0, X6)) | ~ spl25_9), inference(avatar_component_clause, [], [f133])).
fof(f153, plain, (spl25_1 | spl25_13), inference(avatar_split_clause, [], [f92, f150, f99])).
fof(f99, plain, (spl25_1 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl25_1])])).
fof(f92, plain, (actual_world(sK24) | sP4), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | (~ shot(X0, sK22(X0, X3)) & member(X0, sK22(X0, X3), X3)) | ~ group(X0, X3) | ~ six(X0, X3) | (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK23(X0, X1, X2, X3), X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1)) | ~ actual_world(X0)) & (sP3(sK24) & actual_world(sK24))) | sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22, sK23, sK24])], [f33, f36, f35, f34])).
fof(f34, plain, ! [X3, X0] : (? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) => (~ shot(X0, sK22(X0, X3)) & member(X0, sK22(X0, X3), X3))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X3, X2, X1, X0] : (? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) => (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK23(X0, X1, X2, X3), X3))), introduced(choice_axiom, [])).
fof(f36, plain, (? [X10] : (sP3(X10) & actual_world(X10)) => (sP3(sK24) & actual_world(sK24))), introduced(choice_axiom, [])).
fof(f33, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1)) | ~ actual_world(X0)) & ? [X10] : (sP3(X10) & actual_world(X10))) | sP4), inference(rectify, [], [f11])).
fof(f11, plain, ((! [X10] : (! [X11, X12, X13, X14, X15, X16] : (~ of(X10, X16, X14) | ~ scream(X10, X16) | ~ nonreflexive(X10, X16) | ~ present(X10, X16) | ~ patient(X10, X16, X15) | ~ agent(X10, X16, X11) | ~ event(X10, X16) | ~ cry(X10, X15) | ~ revenge(X10, X14) | ? [X17] : (~ shot(X10, X17) & member(X10, X17, X13)) | ~ group(X10, X13) | ~ six(X10, X13) | ? [X18] : (! [X19] : (~ from_loc(X10, X19, X12) | ~ fire(X10, X19) | ~ nonreflexive(X10, X19) | ~ present(X10, X19) | ~ patient(X10, X19, X18) | ~ agent(X10, X19, X11) | ~ event(X10, X19)) & member(X10, X18, X13)) | ~ cannon(X10, X12) | ~ of(X10, X12, X11) | ~ man(X10, X11) | ~ male(X10, X11)) | ~ actual_world(X10)) & ? [X0] : (sP3(X0) & actual_world(X0))) | sP4), inference(definition_folding, [], [f5, e10, e9, e8, e7, e6])).
fof(f10, plain, ((! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X31) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X31) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f5, plain, ((! [X10] : (! [X11, X12, X13, X14, X15, X16] : (~ of(X10, X16, X14) | ~ scream(X10, X16) | ~ nonreflexive(X10, X16) | ~ present(X10, X16) | ~ patient(X10, X16, X15) | ~ agent(X10, X16, X11) | ~ event(X10, X16) | ~ cry(X10, X15) | ~ revenge(X10, X14) | ? [X17] : (~ shot(X10, X17) & member(X10, X17, X13)) | ~ group(X10, X13) | ~ six(X10, X13) | ? [X18] : (! [X19] : (~ from_loc(X10, X19, X12) | ~ fire(X10, X19) | ~ nonreflexive(X10, X19) | ~ present(X10, X19) | ~ patient(X10, X19, X18) | ~ agent(X10, X19, X11) | ~ event(X10, X19)) & member(X10, X18, X13)) | ~ cannon(X10, X12) | ~ of(X10, X12, X11) | ~ man(X10, X11) | ~ male(X10, X11)) | ~ actual_world(X10)) & ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (shot(X0, X7) | ~ member(X0, X7, X3)) & group(X0, X3) & six(X0, X3) & ! [X8] : (? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9)) | ~ member(X0, X8, X3)) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0))) | (! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X31) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31)) | ~ actual_world(X30)) & ? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (shot(X20, X27) | ~ member(X20, X27, X23)) & group(X20, X23) & six(X20, X23) & ! [X28] : (? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29)) | ~ member(X20, X28, X23)) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21)) & actual_world(X20)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (member(X0, X7, X3) => shot(X0, X7)) & group(X0, X3) & six(X0, X3) & ! [X8] : (member(X0, X8, X3) => ? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X14) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X15) & agent(X10, X16, X11) & event(X10, X16) & cry(X10, X15) & revenge(X10, X14) & ! [X17] : (member(X10, X17, X13) => shot(X10, X17)) & group(X10, X13) & six(X10, X13) & ! [X18] : (member(X10, X18, X13) => ? [X19] : (from_loc(X10, X19, X12) & fire(X10, X19) & nonreflexive(X10, X19) & present(X10, X19) & patient(X10, X19, X18) & agent(X10, X19, X11) & event(X10, X19))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10))) & (? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (member(X20, X27, X23) => shot(X20, X27)) & group(X20, X23) & six(X20, X23) & ! [X28] : (member(X20, X28, X23) => ? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29))) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21)) & actual_world(X20)) => ? [X30] : (? [X31, X32, X33, X34, X35, X36] : (of(X30, X36, X35) & scream(X30, X36) & nonreflexive(X30, X36) & present(X30, X36) & patient(X30, X36, X34) & agent(X30, X36, X31) & event(X30, X36) & revenge(X30, X35) & cry(X30, X34) & ! [X37] : (member(X30, X37, X33) => shot(X30, X37)) & group(X30, X33) & six(X30, X33) & ! [X38] : (member(X30, X38, X33) => ? [X39] : (from_loc(X30, X39, X32) & fire(X30, X39) & nonreflexive(X30, X39) & present(X30, X39) & patient(X30, X39, X38) & agent(X30, X39, X31) & event(X30, X39))) & cannon(X30, X32) & of(X30, X32, X31) & man(X30, X31) & male(X30, X31)) & actual_world(X30)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X5) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X4) & agent(X0, X6, X1) & event(X0, X6) & revenge(X0, X5) & cry(X0, X4) & ! [X7] : (member(X0, X7, X3) => shot(X0, X7)) & group(X0, X3) & six(X0, X3) & ! [X8] : (member(X0, X8, X3) => ? [X9] : (from_loc(X0, X9, X2) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X1) & event(X0, X9))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X14) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X15) & agent(X10, X16, X11) & event(X10, X16) & cry(X10, X15) & revenge(X10, X14) & ! [X17] : (member(X10, X17, X13) => shot(X10, X17)) & group(X10, X13) & six(X10, X13) & ! [X18] : (member(X10, X18, X13) => ? [X19] : (from_loc(X10, X19, X12) & fire(X10, X19) & nonreflexive(X10, X19) & present(X10, X19) & patient(X10, X19, X18) & agent(X10, X19, X11) & event(X10, X19))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10))) & (? [X20] : (? [X21, X22, X23, X24, X25, X26] : (of(X20, X26, X24) & scream(X20, X26) & nonreflexive(X20, X26) & present(X20, X26) & patient(X20, X26, X25) & agent(X20, X26, X21) & event(X20, X26) & cry(X20, X25) & revenge(X20, X24) & ! [X27] : (member(X20, X27, X23) => shot(X20, X27)) & group(X20, X23) & six(X20, X23) & ! [X28] : (member(X20, X28, X23) => ? [X29] : (from_loc(X20, X29, X22) & fire(X20, X29) & nonreflexive(X20, X29) & present(X20, X29) & patient(X20, X29, X28) & agent(X20, X29, X21) & event(X20, X29))) & cannon(X20, X22) & of(X20, X22, X21) & man(X20, X21) & male(X20, X21)) & actual_world(X20)) => ? [X30] : (? [X31, X32, X33, X34, X35, X36] : (of(X30, X36, X35) & scream(X30, X36) & nonreflexive(X30, X36) & present(X30, X36) & patient(X30, X36, X34) & agent(X30, X36, X31) & event(X30, X36) & revenge(X30, X35) & cry(X30, X34) & ! [X37] : (member(X30, X37, X33) => shot(X30, X37)) & group(X30, X33) & six(X30, X33) & ! [X38] : (member(X30, X38, X33) => ? [X39] : (from_loc(X30, X39, X32) & fire(X30, X39) & nonreflexive(X30, X39) & present(X30, X39) & patient(X30, X39, X38) & agent(X30, X39, X31) & event(X30, X39))) & cannon(X30, X32) & of(X30, X32, X31) & man(X30, X31) & male(X30, X31)) & actual_world(X30)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X11) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X11) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X11) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6] : (of(X0, X6, X4) & scream(X0, X6) & nonreflexive(X0, X6) & present(X0, X6) & patient(X0, X6, X5) & agent(X0, X6, X1) & event(X0, X6) & cry(X0, X5) & revenge(X0, X4) & ! [X9] : (member(X0, X9, X3) => shot(X0, X9)) & group(X0, X3) & six(X0, X3) & ! [X7] : (member(X0, X7, X3) => ? [X8] : (from_loc(X0, X8, X2) & fire(X0, X8) & nonreflexive(X0, X8) & present(X0, X8) & patient(X0, X8, X7) & agent(X0, X8, X1) & event(X0, X8))) & cannon(X0, X2) & of(X0, X2, X1) & man(X0, X1) & male(X0, X1)) & actual_world(X0)) => ? [X10] : (? [X11, X12, X13, X14, X15, X16] : (of(X10, X16, X15) & scream(X10, X16) & nonreflexive(X10, X16) & present(X10, X16) & patient(X10, X16, X14) & agent(X10, X16, X11) & event(X10, X16) & revenge(X10, X15) & cry(X10, X14) & ! [X19] : (member(X10, X19, X13) => shot(X10, X19)) & group(X10, X13) & six(X10, X13) & ! [X17] : (member(X10, X17, X13) => ? [X18] : (from_loc(X10, X18, X12) & fire(X10, X18) & nonreflexive(X10, X18) & present(X10, X18) & patient(X10, X18, X17) & agent(X10, X18, X11) & event(X10, X18))) & cannon(X10, X12) & of(X10, X12, X11) & man(X10, X11) & male(X10, X11)) & actual_world(X10)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP081+1.p', co1)).
fof(f148, plain, (spl25_1 | spl25_12), inference(avatar_split_clause, [], [f93, f145, f99])).
fof(f93, plain, (sP3(sK24) | sP4), inference(cnf_transformation, [], [f37])).
fof(f143, plain, (spl25_1 | spl25_11), inference(avatar_split_clause, [], [f94, f141, f99])).
fof(f94, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | member(X0, sK22(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f139, plain, (spl25_1 | spl25_10), inference(avatar_split_clause, [], [f95, f137, f99])).
fof(f95, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | member(X0, sK22(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f135, plain, (spl25_1 | spl25_9), inference(avatar_split_clause, [], [f96, f133, f99])).
fof(f96, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ shot(X0, sK22(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK23(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f131, plain, (spl25_1 | spl25_8), inference(avatar_split_clause, [], [f97, f129, f99])).
fof(f97, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X4) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X5) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ cry(X0, X5) | ~ revenge(X0, X4) | ~ shot(X0, sK22(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK23(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f127, plain, (~ spl25_1 | spl25_7), inference(avatar_split_clause, [], [f38, f124, f99])).
fof(f38, plain, (actual_world(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | (~ shot(X0, sK5(X0, X3)) & member(X0, sK5(X0, X3), X3)) | ~ group(X0, X3) | ~ six(X0, X3) | (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK6(X0, X1, X2, X3), X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1)) | ~ actual_world(X0)) & (sP1(sK7) & actual_world(sK7))) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6, sK7])], [f13, f16, f15, f14])).
fof(f14, plain, ! [X3, X0] : (? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) => (~ shot(X0, sK5(X0, X3)) & member(X0, sK5(X0, X3), X3))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X3, X2, X1, X0] : (? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) => (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, sK6(X0, X1, X2, X3), X3))), introduced(choice_axiom, [])).
fof(f16, plain, (? [X10] : (sP1(X10) & actual_world(X10)) => (sP1(sK7) & actual_world(sK7))), introduced(choice_axiom, [])).
fof(f13, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ? [X7] : (~ shot(X0, X7) & member(X0, X7, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ? [X8] : (! [X9] : (~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, X8) | ~ agent(X0, X9, X1) | ~ event(X0, X9)) & member(X0, X8, X3)) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1)) | ~ actual_world(X0)) & ? [X10] : (sP1(X10) & actual_world(X10))) | ~ sP4), inference(rectify, [], [f12])).
fof(f12, plain, ((! [X30] : (! [X31, X32, X33, X34, X35, X36] : (~ of(X30, X36, X35) | ~ scream(X30, X36) | ~ nonreflexive(X30, X36) | ~ present(X30, X36) | ~ patient(X30, X36, X34) | ~ agent(X30, X36, X31) | ~ event(X30, X36) | ~ revenge(X30, X35) | ~ cry(X30, X34) | ? [X37] : (~ shot(X30, X37) & member(X30, X37, X33)) | ~ group(X30, X33) | ~ six(X30, X33) | ? [X38] : (! [X39] : (~ from_loc(X30, X39, X32) | ~ fire(X30, X39) | ~ nonreflexive(X30, X39) | ~ present(X30, X39) | ~ patient(X30, X39, X38) | ~ agent(X30, X39, X31) | ~ event(X30, X39)) & member(X30, X38, X33)) | ~ cannon(X30, X32) | ~ of(X30, X32, X31) | ~ man(X30, X31) | ~ male(X30, X31)) | ~ actual_world(X30)) & ? [X20] : (sP1(X20) & actual_world(X20))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f122, plain, (~ spl25_1 | spl25_6), inference(avatar_split_clause, [], [f39, f119, f99])).
fof(f39, plain, (sP1(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f117, plain, (~ spl25_1 | spl25_5), inference(avatar_split_clause, [], [f40, f115, f99])).
fof(f40, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | member(X0, sK5(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f113, plain, (~ spl25_1 | spl25_4), inference(avatar_split_clause, [], [f41, f111, f99])).
fof(f41, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | member(X0, sK5(X0, X3), X3) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f109, plain, (~ spl25_1 | spl25_3), inference(avatar_split_clause, [], [f42, f107, f99])).
fof(f42, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ~ shot(X0, sK5(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f105, plain, (~ spl25_1 | spl25_2), inference(avatar_split_clause, [], [f43, f103, f99])).
fof(f43, plain, ! [X6, X4, X2, X0, X5, X3, X1, X9] : (~ of(X0, X6, X5) | ~ scream(X0, X6) | ~ nonreflexive(X0, X6) | ~ present(X0, X6) | ~ patient(X0, X6, X4) | ~ agent(X0, X6, X1) | ~ event(X0, X6) | ~ revenge(X0, X5) | ~ cry(X0, X4) | ~ shot(X0, sK5(X0, X3)) | ~ group(X0, X3) | ~ six(X0, X3) | ~ from_loc(X0, X9, X2) | ~ fire(X0, X9) | ~ nonreflexive(X0, X9) | ~ present(X0, X9) | ~ patient(X0, X9, sK6(X0, X1, X2, X3)) | ~ agent(X0, X9, X1) | ~ event(X0, X9) | ~ cannon(X0, X2) | ~ of(X0, X2, X1) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).