fof(f983, plain, $false, inference(avatar_sat_refutation, [], [f107, f111, f115, f119, f124, f129, f133, f137, f141, f145, f150, f155, f569, f982])).
fof(f982, plain, (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5 | ~ spl27_6 | ~ spl27_7), inference(avatar_contradiction_clause, [], [f981])).
fof(f981, plain, ($false | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5 | ~ spl27_6 | ~ spl27_7)), inference(subsumption_resolution, [], [f980, f128])).
fof(f128, plain, (actual_world(sK7) | ~ spl27_7), inference(avatar_component_clause, [], [f126])).
fof(f126, plain, (spl27_7 <=> actual_world(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl27_7])])).
fof(f980, plain, (~ actual_world(sK7) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5 | ~ spl27_6)), inference(resolution, [], [f977, f123])).
fof(f123, plain, (sP1(sK7) | ~ spl27_6), inference(avatar_component_clause, [], [f121])).
fof(f121, plain, (spl27_6 <=> sP1(sK7)), introduced(avatar_definition, [new_symbols(naming, [spl27_6])])).
fof(f977, plain, (! [X0] : (~ sP1(X0) | ~ actual_world(X0)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f976, f78])).
fof(f78, plain, ! [X0] : (revenge(X0, sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ! [X0] : ((of(X0, sK22(X0), sK20(X0)) & scream(X0, sK22(X0)) & nonreflexive(X0, sK22(X0)) & present(X0, sK22(X0)) & patient(X0, sK22(X0), sK21(X0)) & agent(X0, sK22(X0), sK16(X0)) & event(X0, sK22(X0)) & cry(X0, sK21(X0)) & revenge(X0, sK20(X0)) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, sK19(X0))) & group(X0, sK19(X0)) & six(X0, sK19(X0)) & sP0(sK18(X0), X0, sK17(X0), sK19(X0)) & cannon(X0, sK18(X0)) & of(X0, sK18(X0), sK17(X0)) & man(X0, sK17(X0)) & male(X0, sK17(X0)) & male(X0, sK16(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17, sK18, sK19, sK20, sK21, sK22])], [f26, f27])).
fof(f27, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP0(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) => (of(X0, sK22(X0), sK20(X0)) & scream(X0, sK22(X0)) & nonreflexive(X0, sK22(X0)) & present(X0, sK22(X0)) & patient(X0, sK22(X0), sK21(X0)) & agent(X0, sK22(X0), sK16(X0)) & event(X0, sK22(X0)) & cry(X0, sK21(X0)) & revenge(X0, sK20(X0)) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, sK19(X0))) & group(X0, sK19(X0)) & six(X0, sK19(X0)) & sP0(sK18(X0), X0, sK17(X0), sK19(X0)) & cannon(X0, sK18(X0)) & of(X0, sK18(X0), sK17(X0)) & man(X0, sK17(X0)) & male(X0, sK17(X0)) & male(X0, sK16(X0)))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP0(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) | ~ sP1(X0)), inference(rectify, [], [f25])).
fof(f25, plain, ! [X22] : (? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (shot(X22, X30) | ~ member(X22, X30, X26)) & group(X22, X26) & six(X22, X26) & sP0(X25, X22, X24, X26) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23)) | ~ sP1(X22)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X22] : (? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (shot(X22, X30) | ~ member(X22, X30, X26)) & group(X22, X26) & six(X22, X26) & sP0(X25, X22, X24, X26) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23)) | ~ sP1(X22)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X22] : (sP1(X22) <=> ? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (shot(X22, X30) | ~ member(X22, X30, X26)) & group(X22, X26) & six(X22, X26) & sP0(X25, X22, X24, X26) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f976, plain, (! [X0] : (~ revenge(X0, sK20(X0)) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f975])).
fof(f975, plain, (! [X0] : (~ revenge(X0, sK20(X0)) | ~ sP1(X0) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f966, f86])).
fof(f86, plain, ! [X0] : (of(X0, sK22(X0), sK20(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f966, plain, (! [X0, X1] : (~ of(X0, sK22(X0), X1) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f965, f69])).
fof(f69, plain, ! [X0] : (male(X0, sK16(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f965, plain, (! [X0, X1] : (~ of(X0, sK22(X0), X1) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0) | ~ male(X0, sK16(X0))) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f964])).
fof(f964, plain, (! [X0, X1] : (~ of(X0, sK22(X0), X1) | ~ revenge(X0, X1) | ~ sP1(X0) | ~ actual_world(X0) | ~ male(X0, sK16(X0)) | ~ sP1(X0)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f957, f81])).
fof(f81, plain, ! [X0] : (agent(X0, sK22(X0), sK16(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f957, plain, (! [X8, X7, X9] : (~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ revenge(X7, X8) | ~ sP1(X7) | ~ actual_world(X7) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f956, f85])).
fof(f85, plain, ! [X0] : (scream(X0, sK22(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f956, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f955, f84])).
fof(f84, plain, ! [X0] : (nonreflexive(X0, sK22(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f955, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f954, f83])).
fof(f83, plain, ! [X0] : (present(X0, sK22(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f954, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f953, f80])).
fof(f80, plain, ! [X0] : (event(X0, sK22(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f953, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f952, f79])).
fof(f79, plain, ! [X0] : (cry(X0, sK21(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f952, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ cry(X7, sK21(X7)) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f949])).
fof(f949, plain, (! [X8, X7, X9] : (~ sP1(X7) | ~ of(X7, sK22(X7), X8) | ~ cry(X7, sK21(X7)) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ actual_world(X7) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f916, f82])).
fof(f82, plain, ! [X0] : (patient(X0, sK22(X0), sK21(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f916, plain, (! [X6, X8, X7, X5, X9] : (~ patient(X5, X6, X8) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ actual_world(X5) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f915, f809])).
fof(f809, plain, (! [X0] : (shot(X0, sK5(X0, sK19(X0))) | ~ sP1(X0) | ~ actual_world(X0)) | (~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f808])).
fof(f808, plain, (! [X0] : (~ actual_world(X0) | ~ sP1(X0) | shot(X0, sK5(X0, sK19(X0))) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f805, f77])).
fof(f77, plain, ! [X0, X8] : (~ member(X0, X8, sK19(X0)) | shot(X0, X8) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f805, plain, (! [X0] : (member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f804, f78])).
fof(f804, plain, (! [X0] : (~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f803])).
fof(f803, plain, (! [X0] : (~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f802, f86])).
fof(f802, plain, (! [X0, X1] : (~ of(X0, sK22(X0), X1) | ~ revenge(X0, X1) | ~ actual_world(X0) | member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f801, f69])).
fof(f801, plain, (! [X0, X1] : (member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ of(X0, sK22(X0), X1) | ~ male(X0, sK16(X0)) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f800])).
fof(f800, plain, (! [X0, X1] : (member(X0, sK5(X0, sK19(X0)), sK19(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ of(X0, sK22(X0), X1) | ~ male(X0, sK16(X0)) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f777, f81])).
fof(f777, plain, (! [X8, X7, X9] : (~ agent(X7, sK22(X7), X9) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ revenge(X7, X8) | ~ actual_world(X7) | ~ of(X7, sK22(X7), X8) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f776, f85])).
fof(f776, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f775, f84])).
fof(f775, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f774, f83])).
fof(f774, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ revenge(X7, X8) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f773, f80])).
fof(f773, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f764, f79])).
fof(f764, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ cry(X7, sK21(X7)) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f761])).
fof(f761, plain, (! [X8, X7, X9] : (~ actual_world(X7) | member(X7, sK5(X7, sK19(X7)), sK19(X7)) | ~ cry(X7, sK21(X7)) | ~ revenge(X7, X8) | ~ event(X7, sK22(X7)) | ~ agent(X7, sK22(X7), X9) | ~ of(X7, sK22(X7), X8) | ~ present(X7, sK22(X7)) | ~ nonreflexive(X7, sK22(X7)) | ~ scream(X7, sK22(X7)) | ~ male(X7, X9) | ~ sP1(X7) | ~ sP1(X7)) | (~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f755, f82])).
fof(f755, plain, (! [X6, X8, X7, X5, X9] : (~ patient(X5, X6, X8) | ~ actual_world(X5) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ of(X5, X6, X7) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5)) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f754, f70])).
fof(f70, plain, ! [X0] : (male(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f754, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f753, f71])).
fof(f71, plain, ! [X0] : (man(X0, sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f753, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f752, f73])).
fof(f73, plain, ! [X0] : (cannon(X0, sK18(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f752, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f751, f72])).
fof(f72, plain, ! [X0] : (of(X0, sK18(X0), sK17(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f751, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f750, f76])).
fof(f76, plain, ! [X0] : (group(X0, sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f750, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ group(X5, sK19(X5)) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f742, f75])).
fof(f75, plain, ! [X0] : (six(X0, sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f742, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5))) | (~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f741])).
fof(f741, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5) | ~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | member(X5, sK5(X5, sK19(X5)), sK19(X5)) | ~ actual_world(X5) | ~ male(X5, sK17(X5)) | ~ sP1(X5)) | (~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f707, f621])).
fof(f621, plain, (! [X2, X0, X3, X1] : (member(X0, sK6(X0, X1, X2, X3), X3) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl27_5), inference(subsumption_resolution, [], [f620, f78])).
fof(f620, plain, (! [X2, X0, X3, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl27_5), inference(duplicate_literal_removal, [], [f619])).
fof(f619, plain, (! [X2, X0, X3, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl27_5), inference(resolution, [], [f608, f86])).
fof(f608, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK22(X0), X4) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl27_5), inference(subsumption_resolution, [], [f607, f69])).
fof(f607, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK16(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK22(X0), X4) | ~ sP1(X0)) | ~ spl27_5), inference(duplicate_literal_removal, [], [f606])).
fof(f606, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK16(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK22(X0), X4) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl27_5), inference(resolution, [], [f587, f81])).
fof(f587, plain, (! [X14, X17, X15, X13, X18, X16] : (~ agent(X13, sK22(X13), X14) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ revenge(X13, X18) | ~ actual_world(X13) | ~ of(X13, sK22(X13), X18) | ~ sP1(X13)) | ~ spl27_5), inference(subsumption_resolution, [], [f586, f85])).
fof(f586, plain, (! [X14, X17, X15, X13, X18, X16] : (~ actual_world(X13) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ revenge(X13, X18) | ~ agent(X13, sK22(X13), X14) | ~ of(X13, sK22(X13), X18) | ~ scream(X13, sK22(X13)) | ~ sP1(X13)) | ~ spl27_5), inference(subsumption_resolution, [], [f585, f84])).
fof(f585, plain, (! [X14, X17, X15, X13, X18, X16] : (~ actual_world(X13) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ revenge(X13, X18) | ~ agent(X13, sK22(X13), X14) | ~ of(X13, sK22(X13), X18) | ~ nonreflexive(X13, sK22(X13)) | ~ scream(X13, sK22(X13)) | ~ sP1(X13)) | ~ spl27_5), inference(subsumption_resolution, [], [f584, f83])).
fof(f584, plain, (! [X14, X17, X15, X13, X18, X16] : (~ actual_world(X13) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ revenge(X13, X18) | ~ agent(X13, sK22(X13), X14) | ~ of(X13, sK22(X13), X18) | ~ present(X13, sK22(X13)) | ~ nonreflexive(X13, sK22(X13)) | ~ scream(X13, sK22(X13)) | ~ sP1(X13)) | ~ spl27_5), inference(subsumption_resolution, [], [f583, f80])).
fof(f583, plain, (! [X14, X17, X15, X13, X18, X16] : (~ actual_world(X13) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ revenge(X13, X18) | ~ event(X13, sK22(X13)) | ~ agent(X13, sK22(X13), X14) | ~ of(X13, sK22(X13), X18) | ~ present(X13, sK22(X13)) | ~ nonreflexive(X13, sK22(X13)) | ~ scream(X13, sK22(X13)) | ~ sP1(X13)) | ~ spl27_5), inference(subsumption_resolution, [], [f575, f79])).
fof(f575, plain, (! [X14, X17, X15, X13, X18, X16] : (~ actual_world(X13) | ~ male(X13, X14) | ~ male(X13, X15) | ~ man(X13, X15) | ~ of(X13, X16, X15) | ~ cannon(X13, X16) | member(X13, sK6(X13, X15, X16, X17), X17) | ~ six(X13, X17) | ~ group(X13, X17) | member(X13, sK5(X13, X17), X17) | ~ cry(X13, sK21(X13)) | ~ revenge(X13, X18) | ~ event(X13, sK22(X13)) | ~ agent(X13, sK22(X13), X14) | ~ of(X13, sK22(X13), X18) | ~ present(X13, sK22(X13)) | ~ nonreflexive(X13, sK22(X13)) | ~ scream(X13, sK22(X13)) | ~ sP1(X13)) | ~ spl27_5), inference(resolution, [], [f118, f82])).
fof(f118, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ patient(X0, X7, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ of(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_5), inference(avatar_component_clause, [], [f117])).
fof(f117, plain, (spl27_5 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : (~ of(X0, X7, X6) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_5])])).
fof(f707, plain, (! [X4, X2, X0, X5, X3, X1] : (~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl27_4), inference(subsumption_resolution, [], [f706, f72])).
fof(f706, plain, (! [X4, X2, X0, X5, X3, X1] : (~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(subsumption_resolution, [], [f705, f71])).
fof(f705, plain, (! [X4, X2, X0, X5, X3, X1] : (~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(subsumption_resolution, [], [f704, f70])).
fof(f704, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK17(X0)) | ~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(duplicate_literal_removal, [], [f703])).
fof(f703, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK17(X0)) | ~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(resolution, [], [f661, f167])).
fof(f167, plain, ! [X0, X1] : (agent(X0, sK23(sK18(X0), X0, sK17(X0), X1), sK17(X0)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f88, f74])).
fof(f74, plain, ! [X0] : (sP0(sK18(X0), X0, sK17(X0), sK19(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f28])).
fof(f88, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK23(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK23(X0, X1, X2, X4), X0) & fire(X1, sK23(X0, X1, X2, X4)) & nonreflexive(X1, sK23(X0, X1, X2, X4)) & present(X1, sK23(X0, X1, X2, X4)) & patient(X1, sK23(X0, X1, X2, X4), X4) & agent(X1, sK23(X0, X1, X2, X4), X2) & event(X1, sK23(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23])], [f30, f31])).
fof(f31, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK23(X0, X1, X2, X4), X0) & fire(X1, sK23(X0, X1, X2, X4)) & nonreflexive(X1, sK23(X0, X1, X2, X4)) & present(X1, sK23(X0, X1, X2, X4)) & patient(X1, sK23(X0, X1, X2, X4), X4) & agent(X1, sK23(X0, X1, X2, X4), X2) & event(X1, sK23(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f30, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP0(X0, X1, X2, X3)), inference(rectify, [], [f29])).
fof(f29, plain, ! [X25, X22, X24, X26] : (! [X31] : (? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32)) | ~ member(X22, X31, X26)) | ~ sP0(X25, X22, X24, X26)), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ! [X25, X22, X24, X26] : (! [X31] : (? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32)) | ~ member(X22, X31, X26)) | ~ sP0(X25, X22, X24, X26)), inference(usedef, [], [e6])).
fof(e6, plain, ! [X25, X22, X24, X26] : (sP0(X25, X22, X24, X26) <=> ! [X31] : (? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32)) | ~ member(X22, X31, X26))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f661, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ male(X0, X1) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(subsumption_resolution, [], [f660, f73])).
fof(f660, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ cannon(X0, sK18(X0)) | ~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(duplicate_literal_removal, [], [f659])).
fof(f659, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ cannon(X0, sK18(X0)) | ~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK5(X0, X3), X3) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_4), inference(resolution, [], [f618, f169])).
fof(f169, plain, ! [X0, X1] : (from_loc(X0, sK23(sK18(X0), X0, sK17(X0), X1), sK18(X0)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f93, f74])).
fof(f93, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK23(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f32])).
fof(f618, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ actual_world(X8) | ~ six(X8, X12) | ~ group(X8, X12) | member(X8, sK5(X8, X12), X12) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_4), inference(subsumption_resolution, [], [f617, f163])).
fof(f163, plain, ! [X0, X1] : (fire(X0, sK23(sK18(X0), X0, sK17(X0), X1)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f92, f74])).
fof(f92, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK23(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f617, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | member(X8, sK5(X8, X12), X12) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_4), inference(subsumption_resolution, [], [f616, f162])).
fof(f162, plain, ! [X0, X1] : (nonreflexive(X0, sK23(sK18(X0), X0, sK17(X0), X1)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f91, f74])).
fof(f91, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK23(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f616, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | member(X8, sK5(X8, X12), X12) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_4), inference(subsumption_resolution, [], [f615, f161])).
fof(f161, plain, ! [X0, X1] : (present(X0, sK23(sK18(X0), X0, sK17(X0), X1)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f90, f74])).
fof(f90, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK23(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f615, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ present(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | member(X8, sK5(X8, X12), X12) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_4), inference(subsumption_resolution, [], [f610, f160])).
fof(f160, plain, ! [X0, X1] : (event(X0, sK23(sK18(X0), X0, sK17(X0), X1)) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f87, f74])).
fof(f87, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK23(X0, X1, X2, X4))), inference(cnf_transformation, [], [f32])).
fof(f610, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ event(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ present(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | member(X8, sK5(X8, X12), X12) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_4), inference(resolution, [], [f114, f168])).
fof(f168, plain, ! [X0, X1] : (patient(X0, sK23(sK18(X0), X0, sK17(X0), X1), X1) | ~ member(X0, X1, sK19(X0)) | ~ sP1(X0)), inference(resolution, [], [f89, f74])).
fof(f89, plain, ! [X4, X2, X0, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK23(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f32])).
fof(f114, plain, (! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ of(X0, X7, X6) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_4), inference(avatar_component_clause, [], [f113])).
fof(f113, plain, (spl27_4 <=> ! [X1, X3, X5, X7, X10, X0, X2, X4, X6] : (~ of(X0, X7, X6) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK5(X0, X4), X4) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_4])])).
fof(f915, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f914, f76])).
fof(f914, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f913, f75])).
fof(f913, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f912, f70])).
fof(f912, plain, (! [X6, X8, X7, X5, X9] : (~ male(X5, sK17(X5)) | ~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f911, f71])).
fof(f911, plain, (! [X6, X8, X7, X5, X9] : (~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5)) | ~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f910, f73])).
fof(f910, plain, (! [X6, X8, X7, X5, X9] : (~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5)) | ~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f908, f72])).
fof(f908, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5)) | ~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f905])).
fof(f905, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, sK18(X5), sK17(X5)) | ~ cannon(X5, sK18(X5)) | ~ man(X5, sK17(X5)) | ~ male(X5, sK17(X5)) | ~ actual_world(X5) | ~ sP1(X5) | ~ of(X5, X6, X7) | ~ actual_world(X5) | ~ six(X5, sK19(X5)) | ~ group(X5, sK19(X5)) | ~ shot(X5, sK5(X5, sK19(X5))) | ~ cry(X5, X8) | ~ revenge(X5, X7) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP1(X5)) | (~ spl27_2 | ~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f901, f681])).
fof(f681, plain, (! [X4, X2, X0, X5, X3, X1] : (~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP1(X0)) | ~ spl27_2), inference(subsumption_resolution, [], [f680, f72])).
fof(f680, plain, (! [X4, X2, X0, X5, X3, X1] : (~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(subsumption_resolution, [], [f679, f71])).
fof(f679, plain, (! [X4, X2, X0, X5, X3, X1] : (~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(subsumption_resolution, [], [f678, f70])).
fof(f678, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK17(X0)) | ~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(duplicate_literal_removal, [], [f677])).
fof(f677, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK17(X0)) | ~ man(X0, sK17(X0)) | ~ of(X0, sK18(X0), sK17(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X3) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, sK17(X0), sK18(X0), X4), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(resolution, [], [f648, f167])).
fof(f648, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ male(X0, X1) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(subsumption_resolution, [], [f647, f73])).
fof(f647, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ cannon(X0, sK18(X0)) | ~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(duplicate_literal_removal, [], [f646])).
fof(f646, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK18(X0), X2) | ~ cannon(X0, sK18(X0)) | ~ agent(X0, sK23(sK18(X0), X0, sK17(X0), sK6(X0, X2, sK18(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK5(X0, X3)) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0) | ~ member(X0, sK6(X0, X2, sK18(X0), X3), sK19(X0)) | ~ sP1(X0)) | ~ spl27_2), inference(resolution, [], [f605, f169])).
fof(f605, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ actual_world(X8) | ~ six(X8, X12) | ~ group(X8, X12) | ~ shot(X8, sK5(X8, X12)) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_2), inference(subsumption_resolution, [], [f604, f163])).
fof(f604, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | ~ shot(X8, sK5(X8, X12)) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_2), inference(subsumption_resolution, [], [f603, f162])).
fof(f603, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | ~ shot(X8, sK5(X8, X12)) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_2), inference(subsumption_resolution, [], [f602, f161])).
fof(f602, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ present(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | ~ shot(X8, sK5(X8, X12)) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_2), inference(subsumption_resolution, [], [f597, f160])).
fof(f597, plain, (! [X14, X12, X10, X8, X15, X13, X11, X9] : (~ actual_world(X8) | ~ male(X8, X9) | ~ male(X8, X10) | ~ man(X8, X10) | ~ of(X8, X11, X10) | ~ cannon(X8, X11) | ~ event(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ agent(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X10) | ~ of(X8, X13, X14) | ~ present(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ nonreflexive(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ fire(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12))) | ~ from_loc(X8, sK23(sK18(X8), X8, sK17(X8), sK6(X8, X10, X11, X12)), X11) | ~ six(X8, X12) | ~ group(X8, X12) | ~ shot(X8, sK5(X8, X12)) | ~ cry(X8, X15) | ~ revenge(X8, X14) | ~ event(X8, X13) | ~ agent(X8, X13, X9) | ~ patient(X8, X13, X15) | ~ present(X8, X13) | ~ nonreflexive(X8, X13) | ~ scream(X8, X13) | ~ member(X8, sK6(X8, X10, X11, X12), sK19(X8)) | ~ sP1(X8)) | ~ spl27_2), inference(resolution, [], [f106, f168])).
fof(f106, plain, (! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ of(X0, X7, X6) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_2), inference(avatar_component_clause, [], [f105])).
fof(f105, plain, (spl27_2 <=> ! [X1, X3, X5, X7, X10, X0, X2, X4, X6] : (~ of(X0, X7, X6) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_2])])).
fof(f901, plain, (! [X2, X0, X1] : (member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f900, f78])).
fof(f900, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f899])).
fof(f899, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK20(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f880, f86])).
fof(f880, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK22(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ male(X0, X1) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f879, f69])).
fof(f879, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK16(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ of(X0, sK22(X0), X3) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f878])).
fof(f878, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK16(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK6(X0, X1, X2, sK19(X0)), sK19(X0)) | ~ of(X0, sK22(X0), X3) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f871, f81])).
fof(f871, plain, (! [X14, X12, X15, X13, X11] : (~ agent(X11, sK22(X11), X12) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ revenge(X11, X15) | ~ actual_world(X11) | ~ sP1(X11)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f870, f85])).
fof(f870, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ revenge(X11, X15) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ scream(X11, sK22(X11))) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f869, f84])).
fof(f869, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ revenge(X11, X15) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ nonreflexive(X11, sK22(X11)) | ~ scream(X11, sK22(X11))) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f868, f83])).
fof(f868, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ revenge(X11, X15) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ present(X11, sK22(X11)) | ~ nonreflexive(X11, sK22(X11)) | ~ scream(X11, sK22(X11))) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f867, f80])).
fof(f867, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ revenge(X11, X15) | ~ event(X11, sK22(X11)) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ present(X11, sK22(X11)) | ~ nonreflexive(X11, sK22(X11)) | ~ scream(X11, sK22(X11))) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f858, f79])).
fof(f858, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ cry(X11, sK21(X11)) | ~ revenge(X11, X15) | ~ event(X11, sK22(X11)) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ present(X11, sK22(X11)) | ~ nonreflexive(X11, sK22(X11)) | ~ scream(X11, sK22(X11))) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f855])).
fof(f855, plain, (! [X14, X12, X15, X13, X11] : (~ actual_world(X11) | ~ male(X11, X12) | ~ male(X11, X13) | ~ man(X11, X13) | ~ of(X11, X14, X13) | ~ cannon(X11, X14) | member(X11, sK6(X11, X13, X14, sK19(X11)), sK19(X11)) | ~ of(X11, sK22(X11), X15) | ~ cry(X11, sK21(X11)) | ~ revenge(X11, X15) | ~ event(X11, sK22(X11)) | ~ agent(X11, sK22(X11), X12) | ~ sP1(X11) | ~ present(X11, sK22(X11)) | ~ nonreflexive(X11, sK22(X11)) | ~ scream(X11, sK22(X11)) | ~ sP1(X11)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f819, f82])).
fof(f819, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ patient(X10, X14, X16) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK6(X10, X12, X13, sK19(X10)), sK19(X10)) | ~ of(X10, X14, X15) | ~ cry(X10, X16) | ~ revenge(X10, X15) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ sP1(X10) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f818, f76])).
fof(f818, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP1(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK6(X10, X12, X13, sK19(X10)), sK19(X10)) | ~ group(X10, sK19(X10)) | ~ of(X10, X14, X15) | ~ cry(X10, X16) | ~ revenge(X10, X15) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(subsumption_resolution, [], [f815, f75])).
fof(f815, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP1(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK6(X10, X12, X13, sK19(X10)), sK19(X10)) | ~ six(X10, sK19(X10)) | ~ group(X10, sK19(X10)) | ~ of(X10, X14, X15) | ~ cry(X10, X16) | ~ revenge(X10, X15) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(duplicate_literal_removal, [], [f814])).
fof(f814, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP1(X10) | ~ actual_world(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK6(X10, X12, X13, sK19(X10)), sK19(X10)) | ~ six(X10, sK19(X10)) | ~ group(X10, sK19(X10)) | ~ of(X10, X14, X15) | ~ cry(X10, X16) | ~ revenge(X10, X15) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_3 | ~ spl27_4 | ~ spl27_5)), inference(resolution, [], [f809, f110])).
fof(f110, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ shot(X0, sK5(X0, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | ~ of(X0, X7, X6) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_3), inference(avatar_component_clause, [], [f109])).
fof(f109, plain, (spl27_3 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : (~ of(X0, X7, X6) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK5(X0, X4)) | ~ cry(X0, X5) | ~ revenge(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_3])])).
fof(f569, plain, (~ spl27_13 | ~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11 | ~ spl27_12), inference(avatar_split_clause, [], [f566, f147, f143, f139, f135, f131, f152])).
fof(f152, plain, (spl27_13 <=> actual_world(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl27_13])])).
fof(f131, plain, (spl27_8 <=> ! [X1, X3, X5, X7, X10, X0, X2, X4, X6] : (~ of(X0, X7, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_8])])).
fof(f135, plain, (spl27_9 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : (~ of(X0, X7, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_9])])).
fof(f139, plain, (spl27_10 <=> ! [X1, X3, X5, X7, X10, X0, X2, X4, X6] : (~ of(X0, X7, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_10])])).
fof(f143, plain, (spl27_11 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : (~ of(X0, X7, X5) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7))), introduced(avatar_definition, [new_symbols(naming, [spl27_11])])).
fof(f147, plain, (spl27_12 <=> sP3(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl27_12])])).
fof(f566, plain, (~ actual_world(sK26) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11 | ~ spl27_12)), inference(resolution, [], [f563, f149])).
fof(f149, plain, (sP3(sK26) | ~ spl27_12), inference(avatar_component_clause, [], [f147])).
fof(f563, plain, (! [X0] : (~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f562, f54])).
fof(f54, plain, ! [X0] : (revenge(X0, sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((of(X0, sK14(X0), sK13(X0)) & scream(X0, sK14(X0)) & nonreflexive(X0, sK14(X0)) & present(X0, sK14(X0)) & patient(X0, sK14(X0), sK12(X0)) & agent(X0, sK14(X0), sK8(X0)) & event(X0, sK14(X0)) & revenge(X0, sK13(X0)) & cry(X0, sK12(X0)) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, sK11(X0))) & group(X0, sK11(X0)) & six(X0, sK11(X0)) & sP2(sK10(X0), X0, sK9(X0), sK11(X0)) & cannon(X0, sK10(X0)) & of(X0, sK10(X0), sK9(X0)) & man(X0, sK9(X0)) & male(X0, sK9(X0)) & male(X0, sK8(X0))) | ~ sP3(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10, sK11, sK12, sK13, sK14])], [f18, f19])).
fof(f19, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP2(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) => (of(X0, sK14(X0), sK13(X0)) & scream(X0, sK14(X0)) & nonreflexive(X0, sK14(X0)) & present(X0, sK14(X0)) & patient(X0, sK14(X0), sK12(X0)) & agent(X0, sK14(X0), sK8(X0)) & event(X0, sK14(X0)) & revenge(X0, sK13(X0)) & cry(X0, sK12(X0)) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, sK11(X0))) & group(X0, sK11(X0)) & six(X0, sK11(X0)) & sP2(sK10(X0), X0, sK9(X0), sK11(X0)) & cannon(X0, sK10(X0)) & of(X0, sK10(X0), sK9(X0)) & man(X0, sK9(X0)) & male(X0, sK9(X0)) & male(X0, sK8(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP2(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) | ~ sP3(X0)), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP2(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) | ~ sP3(X0)), inference(usedef, [], [e9])).
fof(e9, plain, ! [X0] : (sP3(X0) <=> ? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & sP2(X3, X0, X2, X4) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f562, plain, (! [X0] : (~ revenge(X0, sK13(X0)) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f561])).
fof(f561, plain, (! [X0] : (~ revenge(X0, sK13(X0)) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f557, f61])).
fof(f61, plain, ! [X0] : (of(X0, sK14(X0), sK13(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f557, plain, (! [X0, X1] : (~ of(X0, sK14(X0), X1) | ~ revenge(X0, X1) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f556, f44])).
fof(f44, plain, ! [X0] : (male(X0, sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f556, plain, (! [X0, X1] : (~ male(X0, sK8(X0)) | ~ of(X0, sK14(X0), X1) | ~ revenge(X0, X1) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f555])).
fof(f555, plain, (! [X0, X1] : (~ male(X0, sK8(X0)) | ~ of(X0, sK14(X0), X1) | ~ revenge(X0, X1) | ~ sP3(X0) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f529, f56])).
fof(f56, plain, ! [X0] : (agent(X0, sK14(X0), sK8(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f529, plain, (! [X2, X0, X1] : (~ agent(X0, sK14(X0), X1) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f528, f60])).
fof(f60, plain, ! [X0] : (scream(X0, sK14(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f528, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ scream(X0, sK14(X0))) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f527, f59])).
fof(f59, plain, ! [X0] : (nonreflexive(X0, sK14(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f527, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f526, f58])).
fof(f58, plain, ! [X0] : (present(X0, sK14(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f526, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f525, f55])).
fof(f55, plain, ! [X0] : (event(X0, sK14(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f525, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f524, f53])).
fof(f53, plain, ! [X0] : (cry(X0, sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f524, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f519])).
fof(f519, plain, (! [X2, X0, X1] : (~ sP3(X0) | ~ male(X0, X1) | ~ of(X0, sK14(X0), X2) | ~ revenge(X0, X2) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ actual_world(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f487, f57])).
fof(f57, plain, ! [X0] : (patient(X0, sK14(X0), sK12(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f487, plain, (! [X6, X8, X7, X5, X9] : (~ patient(X5, X7, X9) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ actual_world(X5) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f486, f387])).
fof(f387, plain, (! [X0] : (shot(X0, sK24(X0, sK11(X0))) | ~ sP3(X0) | ~ actual_world(X0)) | (~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f386])).
fof(f386, plain, (! [X0] : (~ actual_world(X0) | ~ sP3(X0) | shot(X0, sK24(X0, sK11(X0))) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f377, f52])).
fof(f52, plain, ! [X0, X8] : (~ member(X0, X8, sK11(X0)) | shot(X0, X8) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f377, plain, (! [X0] : (member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f376, f54])).
fof(f376, plain, (! [X0] : (~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f375])).
fof(f375, plain, (! [X0] : (~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f352, f61])).
fof(f352, plain, (! [X0, X1] : (~ of(X0, sK14(X0), X1) | ~ revenge(X0, X1) | ~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f351, f44])).
fof(f351, plain, (! [X0, X1] : (member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X1) | ~ male(X0, sK8(X0)) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f350])).
fof(f350, plain, (! [X0, X1] : (member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X1) | ~ male(X0, sK8(X0)) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f328, f56])).
fof(f328, plain, (! [X2, X0, X1] : (~ agent(X0, sK14(X0), X2) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X1) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f327, f60])).
fof(f327, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f326, f59])).
fof(f326, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f325, f58])).
fof(f325, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f324, f55])).
fof(f324, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f323, f53])).
fof(f323, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f318])).
fof(f318, plain, (! [X2, X0, X1] : (~ actual_world(X0) | member(X0, sK24(X0, sK11(X0)), sK11(X0)) | ~ revenge(X0, X1) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X2) | ~ of(X0, sK14(X0), X1) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ male(X0, X2) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f317, f57])).
fof(f317, plain, (! [X6, X8, X7, X5, X9] : (~ patient(X5, X6, X8) | ~ actual_world(X5) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ of(X5, X6, X7) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5)) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f316, f45])).
fof(f45, plain, ! [X0] : (male(X0, sK9(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f316, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f315, f46])).
fof(f46, plain, ! [X0] : (man(X0, sK9(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f315, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f314, f48])).
fof(f48, plain, ! [X0] : (cannon(X0, sK10(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f314, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f313, f47])).
fof(f47, plain, ! [X0] : (of(X0, sK10(X0), sK9(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f313, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f312, f51])).
fof(f51, plain, ! [X0] : (group(X0, sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f312, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ group(X5, sK11(X5)) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f304, f50])).
fof(f50, plain, ! [X0] : (six(X0, sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f304, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5))) | (~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f303])).
fof(f303, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, X6, X7) | ~ actual_world(X5) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ revenge(X5, X7) | ~ cry(X5, X8) | ~ event(X5, X6) | ~ agent(X5, X6, X9) | ~ patient(X5, X6, X8) | ~ present(X5, X6) | ~ nonreflexive(X5, X6) | ~ scream(X5, X6) | ~ male(X5, X9) | ~ sP3(X5) | ~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | member(X5, sK24(X5, sK11(X5)), sK11(X5)) | ~ actual_world(X5) | ~ male(X5, sK9(X5)) | ~ sP3(X5)) | (~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f286, f215])).
fof(f215, plain, (! [X2, X0, X3, X1] : (member(X0, sK25(X0, X1, X2, X3), X3) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f214, f54])).
fof(f214, plain, (! [X2, X0, X3, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl27_11), inference(duplicate_literal_removal, [], [f213])).
fof(f213, plain, (! [X2, X0, X3, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl27_11), inference(resolution, [], [f212, f61])).
fof(f212, plain, (! [X4, X2, X0, X3, X1] : (~ of(X0, sK14(X0), X4) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f211, f44])).
fof(f211, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X4) | ~ sP3(X0)) | ~ spl27_11), inference(duplicate_literal_removal, [], [f210])).
fof(f210, plain, (! [X4, X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, X3), X3) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X4) | ~ sP3(X0) | ~ sP3(X0)) | ~ spl27_11), inference(resolution, [], [f176, f56])).
fof(f176, plain, (! [X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK14(X0), X1) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ actual_world(X0) | ~ of(X0, sK14(X0), X5) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f175, f60])).
fof(f175, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ agent(X0, sK14(X0), X1) | ~ of(X0, sK14(X0), X5) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f174, f59])).
fof(f174, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ agent(X0, sK14(X0), X1) | ~ of(X0, sK14(X0), X5) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f173, f58])).
fof(f173, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ agent(X0, sK14(X0), X1) | ~ of(X0, sK14(X0), X5) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f172, f55])).
fof(f172, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ of(X0, sK14(X0), X5) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | ~ spl27_11), inference(subsumption_resolution, [], [f170, f53])).
fof(f170, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ of(X0, sK14(X0), X5) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | ~ spl27_11), inference(resolution, [], [f144, f57])).
fof(f144, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ patient(X0, X7, X6) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ of(X0, X7, X5) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_11), inference(avatar_component_clause, [], [f143])).
fof(f286, plain, (! [X4, X2, X0, X5, X3, X1] : (~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ male(X0, X1) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f285, f47])).
fof(f285, plain, (! [X4, X2, X0, X5, X3, X1] : (~ of(X0, sK10(X0), sK9(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f284, f46])).
fof(f284, plain, (! [X4, X2, X0, X5, X3, X1] : (~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f283, f45])).
fof(f283, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK9(X0)) | ~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(duplicate_literal_removal, [], [f282])).
fof(f282, plain, (! [X4, X2, X0, X5, X3, X1] : (~ male(X0, sK9(X0)) | ~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ male(X0, X1) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(resolution, [], [f250, f164])).
fof(f164, plain, ! [X0, X1] : (agent(X0, sK15(sK10(X0), X0, sK9(X0), X1), sK9(X0)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f63, f49])).
fof(f49, plain, ! [X0] : (sP2(sK10(X0), X0, sK9(X0), sK11(X0)) | ~ sP3(X0)), inference(cnf_transformation, [], [f20])).
fof(f63, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | agent(X1, sK15(X0, X1, X2, X4), X2)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1, X2, X3] : (! [X4] : ((from_loc(X1, sK15(X0, X1, X2, X4), X0) & fire(X1, sK15(X0, X1, X2, X4)) & nonreflexive(X1, sK15(X0, X1, X2, X4)) & present(X1, sK15(X0, X1, X2, X4)) & patient(X1, sK15(X0, X1, X2, X4), X4) & agent(X1, sK15(X0, X1, X2, X4), X2) & event(X1, sK15(X0, X1, X2, X4))) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15])], [f22, f23])).
fof(f23, plain, ! [X4, X2, X1, X0] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) => (from_loc(X1, sK15(X0, X1, X2, X4), X0) & fire(X1, sK15(X0, X1, X2, X4)) & nonreflexive(X1, sK15(X0, X1, X2, X4)) & present(X1, sK15(X0, X1, X2, X4)) & patient(X1, sK15(X0, X1, X2, X4), X4) & agent(X1, sK15(X0, X1, X2, X4), X2) & event(X1, sK15(X0, X1, X2, X4)))), introduced(choice_axiom, [])).
fof(f22, plain, ! [X0, X1, X2, X3] : (! [X4] : (? [X5] : (from_loc(X1, X5, X0) & fire(X1, X5) & nonreflexive(X1, X5) & present(X1, X5) & patient(X1, X5, X4) & agent(X1, X5, X2) & event(X1, X5)) | ~ member(X1, X4, X3)) | ~ sP2(X0, X1, X2, X3)), inference(rectify, [], [f21])).
fof(f21, plain, ! [X3, X0, X2, X4] : (! [X9] : (? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10)) | ~ member(X0, X9, X4)) | ~ sP2(X3, X0, X2, X4)), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ! [X3, X0, X2, X4] : (! [X9] : (? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10)) | ~ member(X0, X9, X4)) | ~ sP2(X3, X0, X2, X4)), inference(usedef, [], [e8])).
fof(e8, plain, ! [X3, X0, X2, X4] : (sP2(X3, X0, X2, X4) <=> ! [X9] : (? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10)) | ~ member(X0, X9, X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f250, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ male(X0, X1) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f249, f48])).
fof(f249, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ cannon(X0, sK10(X0)) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(duplicate_literal_removal, [], [f248])).
fof(f248, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ cannon(X0, sK10(X0)) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ actual_world(X0) | ~ six(X0, X3) | ~ group(X0, X3) | member(X0, sK24(X0, X3), X3) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ sP3(X0) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(resolution, [], [f195, f166])).
fof(f166, plain, ! [X0, X1] : (from_loc(X0, sK15(sK10(X0), X0, sK9(X0), X1), sK10(X0)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f68, f49])).
fof(f68, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | from_loc(X1, sK15(X0, X1, X2, X4), X0)), inference(cnf_transformation, [], [f24])).
fof(f195, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ from_loc(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X3) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X2) | ~ of(X0, X5, X6) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X6) | ~ cry(X0, X7) | ~ event(X0, X5) | ~ agent(X0, X5, X1) | ~ patient(X0, X5, X7) | ~ present(X0, X5) | ~ nonreflexive(X0, X5) | ~ scream(X0, X5) | ~ member(X0, sK25(X0, X2, X3, X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f194, f159])).
fof(f159, plain, ! [X0, X1] : (fire(X0, sK15(sK10(X0), X0, sK9(X0), X1)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f67, f49])).
fof(f67, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | fire(X1, sK15(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f194, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X2) | ~ of(X0, X5, X6) | ~ fire(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ from_loc(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X6) | ~ cry(X0, X7) | ~ event(X0, X5) | ~ agent(X0, X5, X1) | ~ patient(X0, X5, X7) | ~ present(X0, X5) | ~ nonreflexive(X0, X5) | ~ scream(X0, X5) | ~ member(X0, sK25(X0, X2, X3, X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f193, f158])).
fof(f158, plain, ! [X0, X1] : (nonreflexive(X0, sK15(sK10(X0), X0, sK9(X0), X1)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f66, f49])).
fof(f66, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | nonreflexive(X1, sK15(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f193, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X2) | ~ of(X0, X5, X6) | ~ nonreflexive(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ fire(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ from_loc(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X6) | ~ cry(X0, X7) | ~ event(X0, X5) | ~ agent(X0, X5, X1) | ~ patient(X0, X5, X7) | ~ present(X0, X5) | ~ nonreflexive(X0, X5) | ~ scream(X0, X5) | ~ member(X0, sK25(X0, X2, X3, X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f192, f157])).
fof(f157, plain, ! [X0, X1] : (present(X0, sK15(sK10(X0), X0, sK9(X0), X1)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f65, f49])).
fof(f65, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | present(X1, sK15(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f192, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X2) | ~ of(X0, X5, X6) | ~ present(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ nonreflexive(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ fire(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ from_loc(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X6) | ~ cry(X0, X7) | ~ event(X0, X5) | ~ agent(X0, X5, X1) | ~ patient(X0, X5, X7) | ~ present(X0, X5) | ~ nonreflexive(X0, X5) | ~ scream(X0, X5) | ~ member(X0, sK25(X0, X2, X3, X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(subsumption_resolution, [], [f191, f156])).
fof(f156, plain, ! [X0, X1] : (event(X0, sK15(sK10(X0), X0, sK9(X0), X1)) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f62, f49])).
fof(f62, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | event(X1, sK15(X0, X1, X2, X4))), inference(cnf_transformation, [], [f24])).
fof(f191, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X2) | ~ of(X0, X5, X6) | ~ present(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ nonreflexive(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ fire(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4))) | ~ from_loc(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, X3, X4)), X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X6) | ~ cry(X0, X7) | ~ event(X0, X5) | ~ agent(X0, X5, X1) | ~ patient(X0, X5, X7) | ~ present(X0, X5) | ~ nonreflexive(X0, X5) | ~ scream(X0, X5) | ~ member(X0, sK25(X0, X2, X3, X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_10), inference(resolution, [], [f140, f165])).
fof(f165, plain, ! [X0, X1] : (patient(X0, sK15(sK10(X0), X0, sK9(X0), X1), X1) | ~ member(X0, X1, sK11(X0)) | ~ sP3(X0)), inference(resolution, [], [f64, f49])).
fof(f64, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3) | ~ member(X1, X4, X3) | patient(X1, sK15(X0, X1, X2, X4), X4)), inference(cnf_transformation, [], [f24])).
fof(f140, plain, (! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ of(X0, X7, X5) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | member(X0, sK24(X0, X4), X4) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_10), inference(avatar_component_clause, [], [f139])).
fof(f486, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f485, f51])).
fof(f485, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f484, f50])).
fof(f484, plain, (! [X6, X8, X7, X5, X9] : (~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f483, f45])).
fof(f483, plain, (! [X6, X8, X7, X5, X9] : (~ male(X5, sK9(X5)) | ~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f482, f46])).
fof(f482, plain, (! [X6, X8, X7, X5, X9] : (~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5)) | ~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f481, f48])).
fof(f481, plain, (! [X6, X8, X7, X5, X9] : (~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5)) | ~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f479, f47])).
fof(f479, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5)) | ~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ of(X5, X7, X8) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f476])).
fof(f476, plain, (! [X6, X8, X7, X5, X9] : (~ of(X5, sK10(X5), sK9(X5)) | ~ cannon(X5, sK10(X5)) | ~ man(X5, sK9(X5)) | ~ male(X5, sK9(X5)) | ~ actual_world(X5) | ~ sP3(X5) | ~ male(X5, X6) | ~ sP3(X5) | ~ of(X5, X7, X8) | ~ actual_world(X5) | ~ six(X5, sK11(X5)) | ~ group(X5, sK11(X5)) | ~ shot(X5, sK24(X5, sK11(X5))) | ~ revenge(X5, X8) | ~ cry(X5, X9) | ~ event(X5, X7) | ~ agent(X5, X7, X6) | ~ patient(X5, X7, X9) | ~ present(X5, X7) | ~ nonreflexive(X5, X7) | ~ scream(X5, X7)) | (~ spl27_8 | ~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f471, f260])).
fof(f260, plain, (! [X4, X2, X0, X5, X3, X1] : (~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ male(X0, X1) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ actual_world(X0) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl27_8), inference(subsumption_resolution, [], [f259, f47])).
fof(f259, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ of(X0, sK10(X0), sK9(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl27_8), inference(subsumption_resolution, [], [f258, f46])).
fof(f258, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl27_8), inference(subsumption_resolution, [], [f257, f45])).
fof(f257, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, sK9(X0)) | ~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2)) | ~ spl27_8), inference(duplicate_literal_removal, [], [f256])).
fof(f256, plain, (! [X4, X2, X0, X5, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, sK9(X0)) | ~ man(X0, sK9(X0)) | ~ of(X0, sK10(X0), sK9(X0)) | ~ sP3(X0) | ~ of(X0, X2, X3) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X3) | ~ cry(X0, X5) | ~ event(X0, X2) | ~ agent(X0, X2, X1) | ~ patient(X0, X2, X5) | ~ present(X0, X2) | ~ nonreflexive(X0, X2) | ~ scream(X0, X2) | ~ member(X0, sK25(X0, sK9(X0), sK10(X0), X4), sK11(X0)) | ~ sP3(X0)) | ~ spl27_8), inference(resolution, [], [f239, f164])).
fof(f239, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ sP3(X0) | ~ of(X0, X4, X5) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK24(X0, X3)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4)) | ~ spl27_8), inference(subsumption_resolution, [], [f238, f48])).
fof(f238, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ cannon(X0, sK10(X0)) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK24(X0, X3)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4)) | ~ spl27_8), inference(duplicate_literal_removal, [], [f237])).
fof(f237, plain, (! [X6, X4, X2, X0, X5, X3, X1] : (~ sP3(X0) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, sK10(X0), X2) | ~ cannon(X0, sK10(X0)) | ~ agent(X0, sK15(sK10(X0), X0, sK9(X0), sK25(X0, X2, sK10(X0), X3)), X2) | ~ of(X0, X4, X5) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ six(X0, X3) | ~ group(X0, X3) | ~ shot(X0, sK24(X0, X3)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X4) | ~ agent(X0, X4, X1) | ~ patient(X0, X4, X6) | ~ present(X0, X4) | ~ nonreflexive(X0, X4) | ~ scream(X0, X4) | ~ member(X0, sK25(X0, X2, sK10(X0), X3), sK11(X0)) | ~ sP3(X0)) | ~ spl27_8), inference(resolution, [], [f190, f166])).
fof(f190, plain, (! [X14, X12, X10, X8, X7, X13, X11, X9] : (~ from_loc(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X9) | ~ sP3(X7) | ~ actual_world(X7) | ~ male(X7, X11) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X8) | ~ of(X7, X12, X13) | ~ member(X7, sK25(X7, X8, X9, X10), sK11(X7)) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK24(X7, X10)) | ~ revenge(X7, X13) | ~ cry(X7, X14) | ~ event(X7, X12) | ~ agent(X7, X12, X11) | ~ patient(X7, X12, X14) | ~ present(X7, X12) | ~ nonreflexive(X7, X12) | ~ scream(X7, X12)) | ~ spl27_8), inference(subsumption_resolution, [], [f189, f159])).
fof(f189, plain, (! [X14, X12, X10, X8, X7, X13, X11, X9] : (~ member(X7, sK25(X7, X8, X9, X10), sK11(X7)) | ~ sP3(X7) | ~ actual_world(X7) | ~ male(X7, X11) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X8) | ~ of(X7, X12, X13) | ~ fire(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ from_loc(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK24(X7, X10)) | ~ revenge(X7, X13) | ~ cry(X7, X14) | ~ event(X7, X12) | ~ agent(X7, X12, X11) | ~ patient(X7, X12, X14) | ~ present(X7, X12) | ~ nonreflexive(X7, X12) | ~ scream(X7, X12)) | ~ spl27_8), inference(subsumption_resolution, [], [f188, f158])).
fof(f188, plain, (! [X14, X12, X10, X8, X7, X13, X11, X9] : (~ member(X7, sK25(X7, X8, X9, X10), sK11(X7)) | ~ sP3(X7) | ~ actual_world(X7) | ~ male(X7, X11) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X8) | ~ of(X7, X12, X13) | ~ nonreflexive(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ fire(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ from_loc(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK24(X7, X10)) | ~ revenge(X7, X13) | ~ cry(X7, X14) | ~ event(X7, X12) | ~ agent(X7, X12, X11) | ~ patient(X7, X12, X14) | ~ present(X7, X12) | ~ nonreflexive(X7, X12) | ~ scream(X7, X12)) | ~ spl27_8), inference(subsumption_resolution, [], [f187, f157])).
fof(f187, plain, (! [X14, X12, X10, X8, X7, X13, X11, X9] : (~ member(X7, sK25(X7, X8, X9, X10), sK11(X7)) | ~ sP3(X7) | ~ actual_world(X7) | ~ male(X7, X11) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ agent(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X8) | ~ of(X7, X12, X13) | ~ present(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ fire(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ from_loc(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK24(X7, X10)) | ~ revenge(X7, X13) | ~ cry(X7, X14) | ~ event(X7, X12) | ~ agent(X7, X12, X11) | ~ patient(X7, X12, X14) | ~ present(X7, X12) | ~ nonreflexive(X7, X12) | ~ scream(X7, X12)) | ~ spl27_8), inference(subsumption_resolution, [], [f183, f156])).
fof(f183, plain, (! [X14, X12, X10, X8, X7, X13, X11, X9] : (~ member(X7, sK25(X7, X8, X9, X10), sK11(X7)) | ~ sP3(X7) | ~ actual_world(X7) | ~ male(X7, X11) | ~ male(X7, X8) | ~ man(X7, X8) | ~ of(X7, X9, X8) | ~ cannon(X7, X9) | ~ event(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ agent(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X8) | ~ of(X7, X12, X13) | ~ present(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ nonreflexive(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ fire(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10))) | ~ from_loc(X7, sK15(sK10(X7), X7, sK9(X7), sK25(X7, X8, X9, X10)), X9) | ~ six(X7, X10) | ~ group(X7, X10) | ~ shot(X7, sK24(X7, X10)) | ~ revenge(X7, X13) | ~ cry(X7, X14) | ~ event(X7, X12) | ~ agent(X7, X12, X11) | ~ patient(X7, X12, X14) | ~ present(X7, X12) | ~ nonreflexive(X7, X12) | ~ scream(X7, X12)) | ~ spl27_8), inference(resolution, [], [f165, f132])).
fof(f132, plain, (! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | ~ event(X0, X10) | ~ agent(X0, X10, X2) | ~ of(X0, X7, X5) | ~ present(X0, X10) | ~ nonreflexive(X0, X10) | ~ fire(X0, X10) | ~ from_loc(X0, X10, X3) | ~ six(X0, X4) | ~ group(X0, X4) | ~ shot(X0, sK24(X0, X4)) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_8), inference(avatar_component_clause, [], [f131])).
fof(f471, plain, (! [X2, X0, X1] : (member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | ~ man(X0, X1) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f470, f54])).
fof(f470, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f469])).
fof(f469, plain, (! [X2, X0, X1] : (~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ male(X0, X1) | ~ revenge(X0, sK13(X0)) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f446, f61])).
fof(f446, plain, (! [X2, X0, X3, X1] : (~ of(X0, sK14(X0), X3) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ male(X0, X1) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f445, f44])).
fof(f445, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X3) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f444])).
fof(f444, plain, (! [X2, X0, X3, X1] : (~ male(X0, sK8(X0)) | ~ male(X0, X1) | ~ man(X0, X1) | ~ of(X0, X2, X1) | ~ cannon(X0, X2) | member(X0, sK25(X0, X1, X2, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X3) | ~ revenge(X0, X3) | ~ actual_world(X0) | ~ sP3(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f432, f56])).
fof(f432, plain, (! [X4, X2, X0, X3, X1] : (~ agent(X0, sK14(X0), X1) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ actual_world(X0) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f431, f60])).
fof(f431, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ scream(X0, sK14(X0))) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f430, f59])).
fof(f430, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f429, f58])).
fof(f429, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f428, f55])).
fof(f428, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f427, f53])).
fof(f427, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0))) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f422])).
fof(f422, plain, (! [X4, X2, X0, X3, X1] : (~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, sK11(X0)), sK11(X0)) | ~ of(X0, sK14(X0), X4) | ~ revenge(X0, X4) | ~ cry(X0, sK12(X0)) | ~ event(X0, sK14(X0)) | ~ agent(X0, sK14(X0), X1) | ~ sP3(X0) | ~ present(X0, sK14(X0)) | ~ nonreflexive(X0, sK14(X0)) | ~ scream(X0, sK14(X0)) | ~ sP3(X0)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f395, f57])).
fof(f395, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ patient(X10, X14, X16) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK25(X10, X12, X13, sK11(X10)), sK11(X10)) | ~ of(X10, X14, X15) | ~ revenge(X10, X15) | ~ cry(X10, X16) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ sP3(X10) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f394, f51])).
fof(f394, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP3(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK25(X10, X12, X13, sK11(X10)), sK11(X10)) | ~ group(X10, sK11(X10)) | ~ of(X10, X14, X15) | ~ revenge(X10, X15) | ~ cry(X10, X16) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(subsumption_resolution, [], [f391, f50])).
fof(f391, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP3(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK25(X10, X12, X13, sK11(X10)), sK11(X10)) | ~ six(X10, sK11(X10)) | ~ group(X10, sK11(X10)) | ~ of(X10, X14, X15) | ~ revenge(X10, X15) | ~ cry(X10, X16) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(duplicate_literal_removal, [], [f390])).
fof(f390, plain, (! [X14, X12, X10, X15, X13, X11, X16] : (~ sP3(X10) | ~ actual_world(X10) | ~ actual_world(X10) | ~ male(X10, X11) | ~ male(X10, X12) | ~ man(X10, X12) | ~ of(X10, X13, X12) | ~ cannon(X10, X13) | member(X10, sK25(X10, X12, X13, sK11(X10)), sK11(X10)) | ~ six(X10, sK11(X10)) | ~ group(X10, sK11(X10)) | ~ of(X10, X14, X15) | ~ revenge(X10, X15) | ~ cry(X10, X16) | ~ event(X10, X14) | ~ agent(X10, X14, X11) | ~ patient(X10, X14, X16) | ~ present(X10, X14) | ~ nonreflexive(X10, X14) | ~ scream(X10, X14)) | (~ spl27_9 | ~ spl27_10 | ~ spl27_11)), inference(resolution, [], [f387, f136])).
fof(f136, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ shot(X0, sK24(X0, X4)) | ~ actual_world(X0) | ~ male(X0, X1) | ~ male(X0, X2) | ~ man(X0, X2) | ~ of(X0, X3, X2) | ~ cannon(X0, X3) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ six(X0, X4) | ~ group(X0, X4) | ~ of(X0, X7, X5) | ~ revenge(X0, X5) | ~ cry(X0, X6) | ~ event(X0, X7) | ~ agent(X0, X7, X1) | ~ patient(X0, X7, X6) | ~ present(X0, X7) | ~ nonreflexive(X0, X7) | ~ scream(X0, X7)) | ~ spl27_9), inference(avatar_component_clause, [], [f135])).
fof(f155, plain, (spl27_1 | spl27_13), inference(avatar_split_clause, [], [f94, f152, f101])).
fof(f101, plain, (spl27_1 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl27_1])])).
fof(f94, plain, (actual_world(sK26) | sP4), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6, X7] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | (~ shot(X0, sK24(X0, X4)) & member(X0, sK24(X0, X4), X4)) | ~ group(X0, X4) | ~ six(X0, X4) | (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, sK25(X0, X2, X3, X4), X4)) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1)) | ~ actual_world(X0)) & (sP3(sK26) & actual_world(sK26))) | sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25, sK26])], [f33, f36, f35, f34])).
fof(f34, plain, ! [X4, X0] : (? [X8] : (~ shot(X0, X8) & member(X0, X8, X4)) => (~ shot(X0, sK24(X0, X4)) & member(X0, sK24(X0, X4), X4))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X4, X3, X2, X0] : (? [X9] : (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, X9) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, X9, X4)) => (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, sK25(X0, X2, X3, X4), X4))), introduced(choice_axiom, [])).
fof(f36, plain, (? [X11] : (sP3(X11) & actual_world(X11)) => (sP3(sK26) & actual_world(sK26))), introduced(choice_axiom, [])).
fof(f33, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6, X7] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ? [X8] : (~ shot(X0, X8) & member(X0, X8, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | ? [X9] : (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, X9) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, X9, X4)) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1)) | ~ actual_world(X0)) & ? [X11] : (sP3(X11) & actual_world(X11))) | sP4), inference(rectify, [], [f11])).
fof(f11, plain, ((! [X11] : (! [X12, X13, X14, X15, X16, X17, X18] : (~ of(X11, X18, X16) | ~ scream(X11, X18) | ~ nonreflexive(X11, X18) | ~ present(X11, X18) | ~ patient(X11, X18, X17) | ~ agent(X11, X18, X12) | ~ event(X11, X18) | ~ cry(X11, X17) | ~ revenge(X11, X16) | ? [X19] : (~ shot(X11, X19) & member(X11, X19, X15)) | ~ group(X11, X15) | ~ six(X11, X15) | ? [X20] : (! [X21] : (~ from_loc(X11, X21, X14) | ~ fire(X11, X21) | ~ nonreflexive(X11, X21) | ~ present(X11, X21) | ~ patient(X11, X21, X20) | ~ agent(X11, X21, X13) | ~ event(X11, X21)) & member(X11, X20, X15)) | ~ cannon(X11, X14) | ~ of(X11, X14, X13) | ~ man(X11, X13) | ~ male(X11, X13) | ~ male(X11, X12)) | ~ actual_world(X11)) & ? [X0] : (sP3(X0) & actual_world(X0))) | sP4), inference(definition_folding, [], [f5, e10, e9, e8, e7, e6])).
fof(f10, plain, ((! [X33] : (! [X34, X35, X36, X37, X38, X39, X40] : (~ of(X33, X40, X39) | ~ scream(X33, X40) | ~ nonreflexive(X33, X40) | ~ present(X33, X40) | ~ patient(X33, X40, X38) | ~ agent(X33, X40, X34) | ~ event(X33, X40) | ~ revenge(X33, X39) | ~ cry(X33, X38) | ? [X41] : (~ shot(X33, X41) & member(X33, X41, X37)) | ~ group(X33, X37) | ~ six(X33, X37) | ? [X42] : (! [X43] : (~ from_loc(X33, X43, X36) | ~ fire(X33, X43) | ~ nonreflexive(X33, X43) | ~ present(X33, X43) | ~ patient(X33, X43, X42) | ~ agent(X33, X43, X35) | ~ event(X33, X43)) & member(X33, X42, X37)) | ~ cannon(X33, X36) | ~ of(X33, X36, X35) | ~ man(X33, X35) | ~ male(X33, X35) | ~ male(X33, X34)) | ~ actual_world(X33)) & ? [X22] : (sP1(X22) & actual_world(X22))) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (! [X33] : (! [X34, X35, X36, X37, X38, X39, X40] : (~ of(X33, X40, X39) | ~ scream(X33, X40) | ~ nonreflexive(X33, X40) | ~ present(X33, X40) | ~ patient(X33, X40, X38) | ~ agent(X33, X40, X34) | ~ event(X33, X40) | ~ revenge(X33, X39) | ~ cry(X33, X38) | ? [X41] : (~ shot(X33, X41) & member(X33, X41, X37)) | ~ group(X33, X37) | ~ six(X33, X37) | ? [X42] : (! [X43] : (~ from_loc(X33, X43, X36) | ~ fire(X33, X43) | ~ nonreflexive(X33, X43) | ~ present(X33, X43) | ~ patient(X33, X43, X42) | ~ agent(X33, X43, X35) | ~ event(X33, X43)) & member(X33, X42, X37)) | ~ cannon(X33, X36) | ~ of(X33, X36, X35) | ~ man(X33, X35) | ~ male(X33, X35) | ~ male(X33, X34)) | ~ actual_world(X33)) & ? [X22] : (sP1(X22) & actual_world(X22)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f5, plain, ((! [X11] : (! [X12, X13, X14, X15, X16, X17, X18] : (~ of(X11, X18, X16) | ~ scream(X11, X18) | ~ nonreflexive(X11, X18) | ~ present(X11, X18) | ~ patient(X11, X18, X17) | ~ agent(X11, X18, X12) | ~ event(X11, X18) | ~ cry(X11, X17) | ~ revenge(X11, X16) | ? [X19] : (~ shot(X11, X19) & member(X11, X19, X15)) | ~ group(X11, X15) | ~ six(X11, X15) | ? [X20] : (! [X21] : (~ from_loc(X11, X21, X14) | ~ fire(X11, X21) | ~ nonreflexive(X11, X21) | ~ present(X11, X21) | ~ patient(X11, X21, X20) | ~ agent(X11, X21, X13) | ~ event(X11, X21)) & member(X11, X20, X15)) | ~ cannon(X11, X14) | ~ of(X11, X14, X13) | ~ man(X11, X13) | ~ male(X11, X13) | ~ male(X11, X12)) | ~ actual_world(X11)) & ? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (shot(X0, X8) | ~ member(X0, X8, X4)) & group(X0, X4) & six(X0, X4) & ! [X9] : (? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10)) | ~ member(X0, X9, X4)) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0))) | (! [X33] : (! [X34, X35, X36, X37, X38, X39, X40] : (~ of(X33, X40, X39) | ~ scream(X33, X40) | ~ nonreflexive(X33, X40) | ~ present(X33, X40) | ~ patient(X33, X40, X38) | ~ agent(X33, X40, X34) | ~ event(X33, X40) | ~ revenge(X33, X39) | ~ cry(X33, X38) | ? [X41] : (~ shot(X33, X41) & member(X33, X41, X37)) | ~ group(X33, X37) | ~ six(X33, X37) | ? [X42] : (! [X43] : (~ from_loc(X33, X43, X36) | ~ fire(X33, X43) | ~ nonreflexive(X33, X43) | ~ present(X33, X43) | ~ patient(X33, X43, X42) | ~ agent(X33, X43, X35) | ~ event(X33, X43)) & member(X33, X42, X37)) | ~ cannon(X33, X36) | ~ of(X33, X36, X35) | ~ man(X33, X35) | ~ male(X33, X35) | ~ male(X33, X34)) | ~ actual_world(X33)) & ? [X22] : (? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (shot(X22, X30) | ~ member(X22, X30, X26)) & group(X22, X26) & six(X22, X26) & ! [X31] : (? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32)) | ~ member(X22, X31, X26)) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23)) & actual_world(X22)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (member(X0, X8, X4) => shot(X0, X8)) & group(X0, X4) & six(X0, X4) & ! [X9] : (member(X0, X9, X4) => ? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0)) => ? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X16) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X17) & agent(X11, X18, X12) & event(X11, X18) & cry(X11, X17) & revenge(X11, X16) & ! [X19] : (member(X11, X19, X15) => shot(X11, X19)) & group(X11, X15) & six(X11, X15) & ! [X20] : (member(X11, X20, X15) => ? [X21] : (from_loc(X11, X21, X14) & fire(X11, X21) & nonreflexive(X11, X21) & present(X11, X21) & patient(X11, X21, X20) & agent(X11, X21, X13) & event(X11, X21))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11))) & (? [X22] : (? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (member(X22, X30, X26) => shot(X22, X30)) & group(X22, X26) & six(X22, X26) & ! [X31] : (member(X22, X31, X26) => ? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32))) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23)) & actual_world(X22)) => ? [X33] : (? [X34, X35, X36, X37, X38, X39, X40] : (of(X33, X40, X39) & scream(X33, X40) & nonreflexive(X33, X40) & present(X33, X40) & patient(X33, X40, X38) & agent(X33, X40, X34) & event(X33, X40) & revenge(X33, X39) & cry(X33, X38) & ! [X41] : (member(X33, X41, X37) => shot(X33, X41)) & group(X33, X37) & six(X33, X37) & ! [X42] : (member(X33, X42, X37) => ? [X43] : (from_loc(X33, X43, X36) & fire(X33, X43) & nonreflexive(X33, X43) & present(X33, X43) & patient(X33, X43, X42) & agent(X33, X43, X35) & event(X33, X43))) & cannon(X33, X36) & of(X33, X36, X35) & man(X33, X35) & male(X33, X35) & male(X33, X34)) & actual_world(X33)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X6) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X5) & agent(X0, X7, X1) & event(X0, X7) & revenge(X0, X6) & cry(X0, X5) & ! [X8] : (member(X0, X8, X4) => shot(X0, X8)) & group(X0, X4) & six(X0, X4) & ! [X9] : (member(X0, X9, X4) => ? [X10] : (from_loc(X0, X10, X3) & fire(X0, X10) & nonreflexive(X0, X10) & present(X0, X10) & patient(X0, X10, X9) & agent(X0, X10, X2) & event(X0, X10))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0)) => ? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X16) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X17) & agent(X11, X18, X12) & event(X11, X18) & cry(X11, X17) & revenge(X11, X16) & ! [X19] : (member(X11, X19, X15) => shot(X11, X19)) & group(X11, X15) & six(X11, X15) & ! [X20] : (member(X11, X20, X15) => ? [X21] : (from_loc(X11, X21, X14) & fire(X11, X21) & nonreflexive(X11, X21) & present(X11, X21) & patient(X11, X21, X20) & agent(X11, X21, X13) & event(X11, X21))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11))) & (? [X22] : (? [X23, X24, X25, X26, X27, X28, X29] : (of(X22, X29, X27) & scream(X22, X29) & nonreflexive(X22, X29) & present(X22, X29) & patient(X22, X29, X28) & agent(X22, X29, X23) & event(X22, X29) & cry(X22, X28) & revenge(X22, X27) & ! [X30] : (member(X22, X30, X26) => shot(X22, X30)) & group(X22, X26) & six(X22, X26) & ! [X31] : (member(X22, X31, X26) => ? [X32] : (from_loc(X22, X32, X25) & fire(X22, X32) & nonreflexive(X22, X32) & present(X22, X32) & patient(X22, X32, X31) & agent(X22, X32, X24) & event(X22, X32))) & cannon(X22, X25) & of(X22, X25, X24) & man(X22, X24) & male(X22, X24) & male(X22, X23)) & actual_world(X22)) => ? [X33] : (? [X34, X35, X36, X37, X38, X39, X40] : (of(X33, X40, X39) & scream(X33, X40) & nonreflexive(X33, X40) & present(X33, X40) & patient(X33, X40, X38) & agent(X33, X40, X34) & event(X33, X40) & revenge(X33, X39) & cry(X33, X38) & ! [X41] : (member(X33, X41, X37) => shot(X33, X41)) & group(X33, X37) & six(X33, X37) & ! [X42] : (member(X33, X42, X37) => ? [X43] : (from_loc(X33, X43, X36) & fire(X33, X43) & nonreflexive(X33, X43) & present(X33, X43) & patient(X33, X43, X42) & agent(X33, X43, X35) & event(X33, X43))) & cannon(X33, X36) & of(X33, X36, X35) & man(X33, X35) & male(X33, X35) & male(X33, X34)) & actual_world(X33)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X17) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X16) & agent(X11, X18, X12) & event(X11, X18) & revenge(X11, X17) & cry(X11, X16) & ! [X21] : (member(X11, X21, X15) => shot(X11, X21)) & group(X11, X15) & six(X11, X15) & ! [X19] : (member(X11, X19, X15) => ? [X20] : (from_loc(X11, X20, X14) & fire(X11, X20) & nonreflexive(X11, X20) & present(X11, X20) & patient(X11, X20, X19) & agent(X11, X20, X13) & event(X11, X20))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X10] : (member(X0, X10, X4) => shot(X0, X10)) & group(X0, X4) & six(X0, X4) & ! [X8] : (member(X0, X8, X4) => ? [X9] : (from_loc(X0, X9, X3) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X2) & event(X0, X9))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X10] : (member(X0, X10, X4) => shot(X0, X10)) & group(X0, X4) & six(X0, X4) & ! [X8] : (member(X0, X8, X4) => ? [X9] : (from_loc(X0, X9, X3) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X2) & event(X0, X9))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0)) => ? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X17) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X16) & agent(X11, X18, X12) & event(X11, X18) & revenge(X11, X17) & cry(X11, X16) & ! [X21] : (member(X11, X21, X15) => shot(X11, X21)) & group(X11, X15) & six(X11, X15) & ! [X19] : (member(X11, X19, X15) => ? [X20] : (from_loc(X11, X20, X14) & fire(X11, X20) & nonreflexive(X11, X20) & present(X11, X20) & patient(X11, X20, X19) & agent(X11, X20, X13) & event(X11, X20))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X17) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X16) & agent(X11, X18, X12) & event(X11, X18) & revenge(X11, X17) & cry(X11, X16) & ! [X21] : (member(X11, X21, X15) => shot(X11, X21)) & group(X11, X15) & six(X11, X15) & ! [X19] : (member(X11, X19, X15) => ? [X20] : (from_loc(X11, X20, X14) & fire(X11, X20) & nonreflexive(X11, X20) & present(X11, X20) & patient(X11, X20, X19) & agent(X11, X20, X13) & event(X11, X20))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11)) => ? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X10] : (member(X0, X10, X4) => shot(X0, X10)) & group(X0, X4) & six(X0, X4) & ! [X8] : (member(X0, X8, X4) => ? [X9] : (from_loc(X0, X9, X3) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X2) & event(X0, X9))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5, X6, X7] : (of(X0, X7, X5) & scream(X0, X7) & nonreflexive(X0, X7) & present(X0, X7) & patient(X0, X7, X6) & agent(X0, X7, X1) & event(X0, X7) & cry(X0, X6) & revenge(X0, X5) & ! [X10] : (member(X0, X10, X4) => shot(X0, X10)) & group(X0, X4) & six(X0, X4) & ! [X8] : (member(X0, X8, X4) => ? [X9] : (from_loc(X0, X9, X3) & fire(X0, X9) & nonreflexive(X0, X9) & present(X0, X9) & patient(X0, X9, X8) & agent(X0, X9, X2) & event(X0, X9))) & cannon(X0, X3) & of(X0, X3, X2) & man(X0, X2) & male(X0, X2) & male(X0, X1)) & actual_world(X0)) => ? [X11] : (? [X12, X13, X14, X15, X16, X17, X18] : (of(X11, X18, X17) & scream(X11, X18) & nonreflexive(X11, X18) & present(X11, X18) & patient(X11, X18, X16) & agent(X11, X18, X12) & event(X11, X18) & revenge(X11, X17) & cry(X11, X16) & ! [X21] : (member(X11, X21, X15) => shot(X11, X21)) & group(X11, X15) & six(X11, X15) & ! [X19] : (member(X11, X19, X15) => ? [X20] : (from_loc(X11, X20, X14) & fire(X11, X20) & nonreflexive(X11, X20) & present(X11, X20) & patient(X11, X20, X19) & agent(X11, X20, X13) & event(X11, X20))) & cannon(X11, X14) & of(X11, X14, X13) & man(X11, X13) & male(X11, X13) & male(X11, X12)) & actual_world(X11)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP079+1.p', co1)).
fof(f150, plain, (spl27_1 | spl27_12), inference(avatar_split_clause, [], [f95, f147, f101])).
fof(f95, plain, (sP3(sK26) | sP4), inference(cnf_transformation, [], [f37])).
fof(f145, plain, (spl27_1 | spl27_11), inference(avatar_split_clause, [], [f96, f143, f101])).
fof(f96, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | member(X0, sK24(X0, X4), X4) | ~ group(X0, X4) | ~ six(X0, X4) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f141, plain, (spl27_1 | spl27_10), inference(avatar_split_clause, [], [f97, f139, f101])).
fof(f97, plain, ! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | member(X0, sK24(X0, X4), X4) | ~ group(X0, X4) | ~ six(X0, X4) | ~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f137, plain, (spl27_1 | spl27_9), inference(avatar_split_clause, [], [f98, f135, f101])).
fof(f98, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ shot(X0, sK24(X0, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | member(X0, sK25(X0, X2, X3, X4), X4) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f133, plain, (spl27_1 | spl27_8), inference(avatar_split_clause, [], [f99, f131, f101])).
fof(f99, plain, ! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ of(X0, X7, X5) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X6) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ cry(X0, X6) | ~ revenge(X0, X5) | ~ shot(X0, sK24(X0, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | ~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK25(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | sP4), inference(cnf_transformation, [], [f37])).
fof(f129, plain, (~ spl27_1 | spl27_7), inference(avatar_split_clause, [], [f38, f126, f101])).
fof(f38, plain, (actual_world(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6, X7] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | (~ shot(X0, sK5(X0, X4)) & member(X0, sK5(X0, X4), X4)) | ~ group(X0, X4) | ~ six(X0, X4) | (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, sK6(X0, X2, X3, X4), X4)) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1)) | ~ actual_world(X0)) & (sP1(sK7) & actual_world(sK7))) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6, sK7])], [f13, f16, f15, f14])).
fof(f14, plain, ! [X4, X0] : (? [X8] : (~ shot(X0, X8) & member(X0, X8, X4)) => (~ shot(X0, sK5(X0, X4)) & member(X0, sK5(X0, X4), X4))), introduced(choice_axiom, [])).
fof(f15, plain, ! [X4, X3, X2, X0] : (? [X9] : (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, X9) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, X9, X4)) => (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, sK6(X0, X2, X3, X4), X4))), introduced(choice_axiom, [])).
fof(f16, plain, (? [X11] : (sP1(X11) & actual_world(X11)) => (sP1(sK7) & actual_world(sK7))), introduced(choice_axiom, [])).
fof(f13, plain, ((! [X0] : (! [X1, X2, X3, X4, X5, X6, X7] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | ? [X8] : (~ shot(X0, X8) & member(X0, X8, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | ? [X9] : (! [X10] : (~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, X9) | ~ agent(X0, X10, X2) | ~ event(X0, X10)) & member(X0, X9, X4)) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1)) | ~ actual_world(X0)) & ? [X11] : (sP1(X11) & actual_world(X11))) | ~ sP4), inference(rectify, [], [f12])).
fof(f12, plain, ((! [X33] : (! [X34, X35, X36, X37, X38, X39, X40] : (~ of(X33, X40, X39) | ~ scream(X33, X40) | ~ nonreflexive(X33, X40) | ~ present(X33, X40) | ~ patient(X33, X40, X38) | ~ agent(X33, X40, X34) | ~ event(X33, X40) | ~ revenge(X33, X39) | ~ cry(X33, X38) | ? [X41] : (~ shot(X33, X41) & member(X33, X41, X37)) | ~ group(X33, X37) | ~ six(X33, X37) | ? [X42] : (! [X43] : (~ from_loc(X33, X43, X36) | ~ fire(X33, X43) | ~ nonreflexive(X33, X43) | ~ present(X33, X43) | ~ patient(X33, X43, X42) | ~ agent(X33, X43, X35) | ~ event(X33, X43)) & member(X33, X42, X37)) | ~ cannon(X33, X36) | ~ of(X33, X36, X35) | ~ man(X33, X35) | ~ male(X33, X35) | ~ male(X33, X34)) | ~ actual_world(X33)) & ? [X22] : (sP1(X22) & actual_world(X22))) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f124, plain, (~ spl27_1 | spl27_6), inference(avatar_split_clause, [], [f39, f121, f101])).
fof(f39, plain, (sP1(sK7) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f119, plain, (~ spl27_1 | spl27_5), inference(avatar_split_clause, [], [f40, f117, f101])).
fof(f40, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | member(X0, sK5(X0, X4), X4) | ~ group(X0, X4) | ~ six(X0, X4) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f115, plain, (~ spl27_1 | spl27_4), inference(avatar_split_clause, [], [f41, f113, f101])).
fof(f41, plain, ! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | member(X0, sK5(X0, X4), X4) | ~ group(X0, X4) | ~ six(X0, X4) | ~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f111, plain, (~ spl27_1 | spl27_3), inference(avatar_split_clause, [], [f42, f109, f101])).
fof(f42, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | ~ shot(X0, sK5(X0, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | member(X0, sK6(X0, X2, X3, X4), X4) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).
fof(f107, plain, (~ spl27_1 | spl27_2), inference(avatar_split_clause, [], [f43, f105, f101])).
fof(f43, plain, ! [X6, X4, X2, X0, X10, X7, X5, X3, X1] : (~ of(X0, X7, X6) | ~ scream(X0, X7) | ~ nonreflexive(X0, X7) | ~ present(X0, X7) | ~ patient(X0, X7, X5) | ~ agent(X0, X7, X1) | ~ event(X0, X7) | ~ revenge(X0, X6) | ~ cry(X0, X5) | ~ shot(X0, sK5(X0, X4)) | ~ group(X0, X4) | ~ six(X0, X4) | ~ from_loc(X0, X10, X3) | ~ fire(X0, X10) | ~ nonreflexive(X0, X10) | ~ present(X0, X10) | ~ patient(X0, X10, sK6(X0, X2, X3, X4)) | ~ agent(X0, X10, X2) | ~ event(X0, X10) | ~ cannon(X0, X3) | ~ of(X0, X3, X2) | ~ man(X0, X2) | ~ male(X0, X2) | ~ male(X0, X1) | ~ actual_world(X0) | ~ sP4), inference(cnf_transformation, [], [f17])).