fof(f134, plain, $false, inference(avatar_sat_refutation, [], [f69, f74, f79, f80, f85, f90, f130, f133])).
fof(f133, plain, (~ spl13_2 | ~ spl13_3 | ~ spl13_4), inference(avatar_contradiction_clause, [], [f132])).
fof(f132, plain, ($false | (~ spl13_2 | ~ spl13_3 | ~ spl13_4)), inference(subsumption_resolution, [], [f131, f78])).
fof(f78, plain, (actual_world(sK3) | ~ spl13_4), inference(avatar_component_clause, [], [f76])).
fof(f76, plain, (spl13_4 <=> actual_world(sK3)), introduced(avatar_definition, [new_symbols(naming, [spl13_4])])).
fof(f131, plain, (~ actual_world(sK3) | (~ spl13_2 | ~ spl13_3)), inference(resolution, [], [f126, f73])).
fof(f73, plain, (sP0(sK3) | ~ spl13_3), inference(avatar_component_clause, [], [f71])).
fof(f71, plain, (spl13_3 <=> sP0(sK3)), introduced(avatar_definition, [new_symbols(naming, [spl13_3])])).
fof(f126, plain, (! [X0] : (~ sP0(X0) | ~ actual_world(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f125, f46])).
fof(f46, plain, ! [X0] : (placename(X0, sK8(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((in(X0, sK11(X0), sK9(X0)) & down(X0, sK11(X0), sK10(X0)) & barrel(X0, sK11(X0)) & present(X0, sK11(X0)) & agent(X0, sK11(X0), sK9(X0)) & event(X0, sK11(X0)) & lonely(X0, sK10(X0)) & street(X0, sK10(X0)) & old(X0, sK9(X0)) & dirty(X0, sK9(X0)) & white(X0, sK9(X0)) & chevy(X0, sK9(X0)) & placename(X0, sK8(X0)) & hollywood_placename(X0, sK8(X0)) & city(X0, sK9(X0)) & of(X0, sK8(X0), sK9(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10, sK11])], [f18, f19])).
fof(f19, plain, ! [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) => (in(X0, sK11(X0), sK9(X0)) & down(X0, sK11(X0), sK10(X0)) & barrel(X0, sK11(X0)) & present(X0, sK11(X0)) & agent(X0, sK11(X0), sK9(X0)) & event(X0, sK11(X0)) & lonely(X0, sK10(X0)) & street(X0, sK10(X0)) & old(X0, sK9(X0)) & dirty(X0, sK9(X0)) & white(X0, sK9(X0)) & chevy(X0, sK9(X0)) & placename(X0, sK8(X0)) & hollywood_placename(X0, sK8(X0)) & city(X0, sK9(X0)) & of(X0, sK8(X0), sK9(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) | ~ sP0(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X10] : (? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12)) | ~ sP0(X10)), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ! [X10] : (? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12)) | ~ sP0(X10)), inference(usedef, [], [e6])).
fof(e6, plain, ! [X10] : (sP0(X10) <=> ? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f125, plain, (! [X0] : (~ placename(X0, sK8(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f124, f45])).
fof(f45, plain, ! [X0] : (hollywood_placename(X0, sK8(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f124, plain, (! [X0] : (~ hollywood_placename(X0, sK8(X0)) | ~ placename(X0, sK8(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl13_2), inference(duplicate_literal_removal, [], [f123])).
fof(f123, plain, (! [X0] : (~ hollywood_placename(X0, sK8(X0)) | ~ placename(X0, sK8(X0)) | ~ actual_world(X0) | ~ sP0(X0) | ~ sP0(X0)) | ~ spl13_2), inference(resolution, [], [f118, f43])).
fof(f43, plain, ! [X0] : (of(X0, sK8(X0), sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f118, plain, (! [X0, X1] : (~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f117, f52])).
fof(f52, plain, ! [X0] : (lonely(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f117, plain, (! [X0, X1] : (~ lonely(X0, sK10(X0)) | ~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f116, f51])).
fof(f51, plain, ! [X0] : (street(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f116, plain, (! [X0, X1] : (~ street(X0, sK10(X0)) | ~ lonely(X0, sK10(X0)) | ~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl13_2), inference(duplicate_literal_removal, [], [f115])).
fof(f115, plain, (! [X0, X1] : (~ street(X0, sK10(X0)) | ~ lonely(X0, sK10(X0)) | ~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP0(X0) | ~ sP0(X0)) | ~ spl13_2), inference(resolution, [], [f110, f57])).
fof(f57, plain, ! [X0] : (down(X0, sK11(X0), sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f110, plain, (! [X4, X5, X3] : (~ down(X3, sK11(X3), X4) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ actual_world(X3) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f109, f56])).
fof(f56, plain, ! [X0] : (barrel(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f109, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f108, f55])).
fof(f55, plain, ! [X0] : (present(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f108, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f107, f54])).
fof(f54, plain, ! [X0] : (agent(X0, sK11(X0), sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f107, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f106, f53])).
fof(f53, plain, ! [X0] : (event(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f106, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f105, f50])).
fof(f50, plain, ! [X0] : (old(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f105, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ old(X3, sK9(X3)) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f104, f49])).
fof(f49, plain, ! [X0] : (dirty(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f104, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ dirty(X3, sK9(X3)) | ~ old(X3, sK9(X3)) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f103, f48])).
fof(f48, plain, ! [X0] : (white(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f103, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ white(X3, sK9(X3)) | ~ dirty(X3, sK9(X3)) | ~ old(X3, sK9(X3)) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f102, f47])).
fof(f47, plain, ! [X0] : (chevy(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f102, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ chevy(X3, sK9(X3)) | ~ white(X3, sK9(X3)) | ~ dirty(X3, sK9(X3)) | ~ old(X3, sK9(X3)) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(subsumption_resolution, [], [f92, f44])).
fof(f44, plain, ! [X0] : (city(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f92, plain, (! [X4, X5, X3] : (~ actual_world(X3) | ~ street(X3, X4) | ~ lonely(X3, X4) | ~ of(X3, X5, sK9(X3)) | ~ city(X3, sK9(X3)) | ~ hollywood_placename(X3, X5) | ~ placename(X3, X5) | ~ chevy(X3, sK9(X3)) | ~ white(X3, sK9(X3)) | ~ dirty(X3, sK9(X3)) | ~ old(X3, sK9(X3)) | ~ event(X3, sK11(X3)) | ~ agent(X3, sK11(X3), sK9(X3)) | ~ present(X3, sK11(X3)) | ~ barrel(X3, sK11(X3)) | ~ down(X3, sK11(X3), X4) | ~ sP0(X3)) | ~ spl13_2), inference(resolution, [], [f68, f58])).
fof(f58, plain, ! [X0] : (in(X0, sK11(X0), sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f68, plain, (! [X4, X2, X0, X3, X1] : (~ in(X0, X4, X3) | ~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, X3) | ~ city(X0, X3) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ event(X0, X4) | ~ agent(X0, X4, X3) | ~ present(X0, X4) | ~ barrel(X0, X4) | ~ down(X0, X4, X1)) | ~ spl13_2), inference(avatar_component_clause, [], [f67])).
fof(f67, plain, (spl13_2 <=> ! [X1, X3, X0, X2, X4] : (~ in(X0, X4, X3) | ~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, X3) | ~ city(X0, X3) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ event(X0, X4) | ~ agent(X0, X4, X3) | ~ present(X0, X4) | ~ barrel(X0, X4) | ~ down(X0, X4, X1))), introduced(avatar_definition, [new_symbols(naming, [spl13_2])])).
fof(f130, plain, (~ spl13_6 | ~ spl13_2 | ~ spl13_5), inference(avatar_split_clause, [], [f127, f82, f67, f87])).
fof(f87, plain, (spl13_6 <=> actual_world(sK12)), introduced(avatar_definition, [new_symbols(naming, [spl13_6])])).
fof(f82, plain, (spl13_5 <=> sP1(sK12)), introduced(avatar_definition, [new_symbols(naming, [spl13_5])])).
fof(f127, plain, (~ actual_world(sK12) | (~ spl13_2 | ~ spl13_5)), inference(resolution, [], [f122, f84])).
fof(f84, plain, (sP1(sK12) | ~ spl13_5), inference(avatar_component_clause, [], [f82])).
fof(f122, plain, (! [X0] : (~ sP1(X0) | ~ actual_world(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f121, f32])).
fof(f32, plain, ! [X0] : (placename(X0, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((in(X0, sK7(X0), sK6(X0)) & down(X0, sK7(X0), sK4(X0)) & barrel(X0, sK7(X0)) & present(X0, sK7(X0)) & agent(X0, sK7(X0), sK6(X0)) & event(X0, sK7(X0)) & old(X0, sK6(X0)) & dirty(X0, sK6(X0)) & white(X0, sK6(X0)) & chevy(X0, sK6(X0)) & placename(X0, sK5(X0)) & hollywood_placename(X0, sK5(X0)) & city(X0, sK6(X0)) & of(X0, sK5(X0), sK6(X0)) & lonely(X0, sK4(X0)) & street(X0, sK4(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6, sK7])], [f14, f15])).
fof(f15, plain, ! [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) => (in(X0, sK7(X0), sK6(X0)) & down(X0, sK7(X0), sK4(X0)) & barrel(X0, sK7(X0)) & present(X0, sK7(X0)) & agent(X0, sK7(X0), sK6(X0)) & event(X0, sK7(X0)) & old(X0, sK6(X0)) & dirty(X0, sK6(X0)) & white(X0, sK6(X0)) & chevy(X0, sK6(X0)) & placename(X0, sK5(X0)) & hollywood_placename(X0, sK5(X0)) & city(X0, sK6(X0)) & of(X0, sK5(X0), sK6(X0)) & lonely(X0, sK4(X0)) & street(X0, sK4(X0)))), introduced(choice_axiom, [])).
fof(f14, plain, ! [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) | ~ sP1(X0)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) | ~ sP1(X0)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X0] : (sP1(X0) <=> ? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f121, plain, (! [X0] : (~ placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f120, f31])).
fof(f31, plain, ! [X0] : (hollywood_placename(X0, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f120, plain, (! [X0] : (~ hollywood_placename(X0, sK5(X0)) | ~ placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(duplicate_literal_removal, [], [f119])).
fof(f119, plain, (! [X0] : (~ hollywood_placename(X0, sK5(X0)) | ~ placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl13_2), inference(resolution, [], [f114, f29])).
fof(f29, plain, ! [X0] : (of(X0, sK5(X0), sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f114, plain, (! [X0, X1] : (~ of(X0, X1, sK6(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f113, f28])).
fof(f28, plain, ! [X0] : (lonely(X0, sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f113, plain, (! [X0, X1] : (~ lonely(X0, sK4(X0)) | ~ of(X0, X1, sK6(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f112, f27])).
fof(f27, plain, ! [X0] : (street(X0, sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f112, plain, (! [X0, X1] : (~ street(X0, sK4(X0)) | ~ lonely(X0, sK4(X0)) | ~ of(X0, X1, sK6(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(duplicate_literal_removal, [], [f111])).
fof(f111, plain, (! [X0, X1] : (~ street(X0, sK4(X0)) | ~ lonely(X0, sK4(X0)) | ~ of(X0, X1, sK6(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl13_2), inference(resolution, [], [f101, f41])).
fof(f41, plain, ! [X0] : (down(X0, sK7(X0), sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f101, plain, (! [X2, X0, X1] : (~ down(X0, sK7(X0), X1) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f100, f40])).
fof(f40, plain, ! [X0] : (barrel(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f100, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f99, f39])).
fof(f39, plain, ! [X0] : (present(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f99, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f98, f38])).
fof(f38, plain, ! [X0] : (agent(X0, sK7(X0), sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f98, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f97, f37])).
fof(f37, plain, ! [X0] : (event(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f97, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f96, f36])).
fof(f36, plain, ! [X0] : (old(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f96, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ old(X0, sK6(X0)) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f95, f35])).
fof(f35, plain, ! [X0] : (dirty(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f95, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ dirty(X0, sK6(X0)) | ~ old(X0, sK6(X0)) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f94, f34])).
fof(f34, plain, ! [X0] : (white(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f94, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ white(X0, sK6(X0)) | ~ dirty(X0, sK6(X0)) | ~ old(X0, sK6(X0)) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f93, f33])).
fof(f33, plain, ! [X0] : (chevy(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f93, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ chevy(X0, sK6(X0)) | ~ white(X0, sK6(X0)) | ~ dirty(X0, sK6(X0)) | ~ old(X0, sK6(X0)) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(subsumption_resolution, [], [f91, f30])).
fof(f30, plain, ! [X0] : (city(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f91, plain, (! [X2, X0, X1] : (~ actual_world(X0) | ~ street(X0, X1) | ~ lonely(X0, X1) | ~ of(X0, X2, sK6(X0)) | ~ city(X0, sK6(X0)) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ chevy(X0, sK6(X0)) | ~ white(X0, sK6(X0)) | ~ dirty(X0, sK6(X0)) | ~ old(X0, sK6(X0)) | ~ event(X0, sK7(X0)) | ~ agent(X0, sK7(X0), sK6(X0)) | ~ present(X0, sK7(X0)) | ~ barrel(X0, sK7(X0)) | ~ down(X0, sK7(X0), X1) | ~ sP1(X0)) | ~ spl13_2), inference(resolution, [], [f68, f42])).
fof(f42, plain, ! [X0] : (in(X0, sK7(X0), sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f90, plain, (spl13_1 | spl13_6), inference(avatar_split_clause, [], [f59, f87, f63])).
fof(f63, plain, (spl13_1 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl13_1])])).
fof(f59, plain, (actual_world(sK12) | sP2), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((! [X0] : (! [X1, X2, X3, X4] : (~ in(X0, X4, X2) | ~ down(X0, X4, X3) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X2) | ~ event(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ old(X0, X2) | ~ dirty(X0, X2) | ~ white(X0, X2) | ~ chevy(X0, X2) | ~ placename(X0, X1) | ~ hollywood_placename(X0, X1) | ~ city(X0, X2) | ~ of(X0, X1, X2)) | ~ actual_world(X0)) & (sP1(sK12) & actual_world(sK12))) | sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12])], [f21, f22])).
fof(f22, plain, (? [X5] : (sP1(X5) & actual_world(X5)) => (sP1(sK12) & actual_world(sK12))), introduced(choice_axiom, [])).
fof(f21, plain, ((! [X0] : (! [X1, X2, X3, X4] : (~ in(X0, X4, X2) | ~ down(X0, X4, X3) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X2) | ~ event(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ old(X0, X2) | ~ dirty(X0, X2) | ~ white(X0, X2) | ~ chevy(X0, X2) | ~ placename(X0, X1) | ~ hollywood_placename(X0, X1) | ~ city(X0, X2) | ~ of(X0, X1, X2)) | ~ actual_world(X0)) & ? [X5] : (sP1(X5) & actual_world(X5))) | sP2), inference(rectify, [], [f9])).
fof(f9, plain, ((! [X5] : (! [X6, X7, X8, X9] : (~ in(X5, X9, X7) | ~ down(X5, X9, X8) | ~ barrel(X5, X9) | ~ present(X5, X9) | ~ agent(X5, X9, X7) | ~ event(X5, X9) | ~ lonely(X5, X8) | ~ street(X5, X8) | ~ old(X5, X7) | ~ dirty(X5, X7) | ~ white(X5, X7) | ~ chevy(X5, X7) | ~ placename(X5, X6) | ~ hollywood_placename(X5, X6) | ~ city(X5, X7) | ~ of(X5, X6, X7)) | ~ actual_world(X5)) & ? [X0] : (sP1(X0) & actual_world(X0))) | sP2), inference(definition_folding, [], [f5, e8, e7, e6])).
fof(f8, plain, ((! [X15] : (! [X16, X17, X18, X19] : (~ in(X15, X19, X18) | ~ down(X15, X19, X16) | ~ barrel(X15, X19) | ~ present(X15, X19) | ~ agent(X15, X19, X18) | ~ event(X15, X19) | ~ old(X15, X18) | ~ dirty(X15, X18) | ~ white(X15, X18) | ~ chevy(X15, X18) | ~ placename(X15, X17) | ~ hollywood_placename(X15, X17) | ~ city(X15, X18) | ~ of(X15, X17, X18) | ~ lonely(X15, X16) | ~ street(X15, X16)) | ~ actual_world(X15)) & ? [X10] : (sP0(X10) & actual_world(X10))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (! [X15] : (! [X16, X17, X18, X19] : (~ in(X15, X19, X18) | ~ down(X15, X19, X16) | ~ barrel(X15, X19) | ~ present(X15, X19) | ~ agent(X15, X19, X18) | ~ event(X15, X19) | ~ old(X15, X18) | ~ dirty(X15, X18) | ~ white(X15, X18) | ~ chevy(X15, X18) | ~ placename(X15, X17) | ~ hollywood_placename(X15, X17) | ~ city(X15, X18) | ~ of(X15, X17, X18) | ~ lonely(X15, X16) | ~ street(X15, X16)) | ~ actual_world(X15)) & ? [X10] : (sP0(X10) & actual_world(X10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f5, plain, ((! [X5] : (! [X6, X7, X8, X9] : (~ in(X5, X9, X7) | ~ down(X5, X9, X8) | ~ barrel(X5, X9) | ~ present(X5, X9) | ~ agent(X5, X9, X7) | ~ event(X5, X9) | ~ lonely(X5, X8) | ~ street(X5, X8) | ~ old(X5, X7) | ~ dirty(X5, X7) | ~ white(X5, X7) | ~ chevy(X5, X7) | ~ placename(X5, X6) | ~ hollywood_placename(X5, X6) | ~ city(X5, X7) | ~ of(X5, X6, X7)) | ~ actual_world(X5)) & ? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) & actual_world(X0))) | (! [X15] : (! [X16, X17, X18, X19] : (~ in(X15, X19, X18) | ~ down(X15, X19, X16) | ~ barrel(X15, X19) | ~ present(X15, X19) | ~ agent(X15, X19, X18) | ~ event(X15, X19) | ~ old(X15, X18) | ~ dirty(X15, X18) | ~ white(X15, X18) | ~ chevy(X15, X18) | ~ placename(X15, X17) | ~ hollywood_placename(X15, X17) | ~ city(X15, X18) | ~ of(X15, X17, X18) | ~ lonely(X15, X16) | ~ street(X15, X16)) | ~ actual_world(X15)) & ? [X10] : (? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12)) & actual_world(X10)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) & actual_world(X0)) => ? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X7) & down(X5, X9, X8) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X7) & event(X5, X9) & lonely(X5, X8) & street(X5, X8) & old(X5, X7) & dirty(X5, X7) & white(X5, X7) & chevy(X5, X7) & placename(X5, X6) & hollywood_placename(X5, X6) & city(X5, X7) & of(X5, X6, X7)) & actual_world(X5))) & (? [X10] : (? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12)) & actual_world(X10)) => ? [X15] : (? [X16, X17, X18, X19] : (in(X15, X19, X18) & down(X15, X19, X16) & barrel(X15, X19) & present(X15, X19) & agent(X15, X19, X18) & event(X15, X19) & old(X15, X18) & dirty(X15, X18) & white(X15, X18) & chevy(X15, X18) & placename(X15, X17) & hollywood_placename(X15, X17) & city(X15, X18) & of(X15, X17, X18) & lonely(X15, X16) & street(X15, X16)) & actual_world(X15)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X3) & down(X0, X4, X1) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X3) & event(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X3) & of(X0, X2, X3) & lonely(X0, X1) & street(X0, X1)) & actual_world(X0)) => ? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X7) & down(X5, X9, X8) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X7) & event(X5, X9) & lonely(X5, X8) & street(X5, X8) & old(X5, X7) & dirty(X5, X7) & white(X5, X7) & chevy(X5, X7) & placename(X5, X6) & hollywood_placename(X5, X6) & city(X5, X7) & of(X5, X6, X7)) & actual_world(X5))) & (? [X10] : (? [X11, X12, X13, X14] : (in(X10, X14, X12) & down(X10, X14, X13) & barrel(X10, X14) & present(X10, X14) & agent(X10, X14, X12) & event(X10, X14) & lonely(X10, X13) & street(X10, X13) & old(X10, X12) & dirty(X10, X12) & white(X10, X12) & chevy(X10, X12) & placename(X10, X11) & hollywood_placename(X10, X11) & city(X10, X12) & of(X10, X11, X12)) & actual_world(X10)) => ? [X15] : (? [X16, X17, X18, X19] : (in(X15, X19, X18) & down(X15, X19, X16) & barrel(X15, X19) & present(X15, X19) & agent(X15, X19, X18) & event(X15, X19) & old(X15, X18) & dirty(X15, X18) & white(X15, X18) & chevy(X15, X18) & placename(X15, X17) & hollywood_placename(X15, X17) & city(X15, X18) & of(X15, X17, X18) & lonely(X15, X16) & street(X15, X16)) & actual_world(X15)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X8) & down(X5, X9, X6) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X8) & event(X5, X9) & old(X5, X8) & dirty(X5, X8) & white(X5, X8) & chevy(X5, X8) & placename(X5, X7) & hollywood_placename(X5, X7) & city(X5, X8) & of(X5, X7, X8) & lonely(X5, X6) & street(X5, X6)) & actual_world(X5)) => ? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) & actual_world(X0)) => ? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X8) & down(X5, X9, X6) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X8) & event(X5, X9) & old(X5, X8) & dirty(X5, X8) & white(X5, X8) & chevy(X5, X8) & placename(X5, X7) & hollywood_placename(X5, X7) & city(X5, X8) & of(X5, X7, X8) & lonely(X5, X6) & street(X5, X6)) & actual_world(X5)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X8) & down(X5, X9, X6) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X8) & event(X5, X9) & old(X5, X8) & dirty(X5, X8) & white(X5, X8) & chevy(X5, X8) & placename(X5, X7) & hollywood_placename(X5, X7) & city(X5, X8) & of(X5, X7, X8) & lonely(X5, X6) & street(X5, X6)) & actual_world(X5)) => ? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4] : (in(X0, X4, X2) & down(X0, X4, X3) & barrel(X0, X4) & present(X0, X4) & agent(X0, X4, X2) & event(X0, X4) & lonely(X0, X3) & street(X0, X3) & old(X0, X2) & dirty(X0, X2) & white(X0, X2) & chevy(X0, X2) & placename(X0, X1) & hollywood_placename(X0, X1) & city(X0, X2) & of(X0, X1, X2)) & actual_world(X0)) => ? [X5] : (? [X6, X7, X8, X9] : (in(X5, X9, X8) & down(X5, X9, X6) & barrel(X5, X9) & present(X5, X9) & agent(X5, X9, X8) & event(X5, X9) & old(X5, X8) & dirty(X5, X8) & white(X5, X8) & chevy(X5, X8) & placename(X5, X7) & hollywood_placename(X5, X7) & city(X5, X8) & of(X5, X7, X8) & lonely(X5, X6) & street(X5, X6)) & actual_world(X5)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP122+1.p', co1)).
fof(f85, plain, (spl13_1 | spl13_5), inference(avatar_split_clause, [], [f60, f82, f63])).
fof(f60, plain, (sP1(sK12) | sP2), inference(cnf_transformation, [], [f23])).
fof(f80, plain, (spl13_1 | spl13_2), inference(avatar_split_clause, [], [f61, f67, f63])).
fof(f61, plain, ! [X4, X2, X0, X3, X1] : (~ in(X0, X4, X2) | ~ down(X0, X4, X3) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X2) | ~ event(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ old(X0, X2) | ~ dirty(X0, X2) | ~ white(X0, X2) | ~ chevy(X0, X2) | ~ placename(X0, X1) | ~ hollywood_placename(X0, X1) | ~ city(X0, X2) | ~ of(X0, X1, X2) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f23])).
fof(f79, plain, (~ spl13_1 | spl13_4), inference(avatar_split_clause, [], [f24, f76, f63])).
fof(f24, plain, (actual_world(sK3) | ~ sP2), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((! [X0] : (! [X1, X2, X3, X4] : (~ in(X0, X4, X3) | ~ down(X0, X4, X1) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X3) | ~ event(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X3) | ~ of(X0, X2, X3) | ~ lonely(X0, X1) | ~ street(X0, X1)) | ~ actual_world(X0)) & (sP0(sK3) & actual_world(sK3))) | ~ sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f11, f12])).
fof(f12, plain, (? [X5] : (sP0(X5) & actual_world(X5)) => (sP0(sK3) & actual_world(sK3))), introduced(choice_axiom, [])).
fof(f11, plain, ((! [X0] : (! [X1, X2, X3, X4] : (~ in(X0, X4, X3) | ~ down(X0, X4, X1) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X3) | ~ event(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X3) | ~ of(X0, X2, X3) | ~ lonely(X0, X1) | ~ street(X0, X1)) | ~ actual_world(X0)) & ? [X5] : (sP0(X5) & actual_world(X5))) | ~ sP2), inference(rectify, [], [f10])).
fof(f10, plain, ((! [X15] : (! [X16, X17, X18, X19] : (~ in(X15, X19, X18) | ~ down(X15, X19, X16) | ~ barrel(X15, X19) | ~ present(X15, X19) | ~ agent(X15, X19, X18) | ~ event(X15, X19) | ~ old(X15, X18) | ~ dirty(X15, X18) | ~ white(X15, X18) | ~ chevy(X15, X18) | ~ placename(X15, X17) | ~ hollywood_placename(X15, X17) | ~ city(X15, X18) | ~ of(X15, X17, X18) | ~ lonely(X15, X16) | ~ street(X15, X16)) | ~ actual_world(X15)) & ? [X10] : (sP0(X10) & actual_world(X10))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f74, plain, (~ spl13_1 | spl13_3), inference(avatar_split_clause, [], [f25, f71, f63])).
fof(f25, plain, (sP0(sK3) | ~ sP2), inference(cnf_transformation, [], [f13])).
fof(f69, plain, (~ spl13_1 | spl13_2), inference(avatar_split_clause, [], [f26, f67, f63])).
fof(f26, plain, ! [X4, X2, X0, X3, X1] : (~ in(X0, X4, X3) | ~ down(X0, X4, X1) | ~ barrel(X0, X4) | ~ present(X0, X4) | ~ agent(X0, X4, X3) | ~ event(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X3) | ~ of(X0, X2, X3) | ~ lonely(X0, X1) | ~ street(X0, X1) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f13])).