fof(f136, plain, $false, inference(avatar_sat_refutation, [], [f69, f74, f79, f80, f85, f90, f132, f135])).
fof(f135, plain, (~ spl15_2 | ~ spl15_3 | ~ spl15_4), inference(avatar_contradiction_clause, [], [f134])).
fof(f134, plain, ($false | (~ spl15_2 | ~ spl15_3 | ~ spl15_4)), inference(subsumption_resolution, [], [f133, f78])).
fof(f78, plain, (actual_world(sK3) | ~ spl15_4), inference(avatar_component_clause, [], [f76])).
fof(f76, plain, (spl15_4 <=> actual_world(sK3)), introduced(avatar_definition, [new_symbols(naming, [spl15_4])])).
fof(f133, plain, (~ actual_world(sK3) | (~ spl15_2 | ~ spl15_3)), inference(resolution, [], [f128, f73])).
fof(f73, plain, (sP0(sK3) | ~ spl15_3), inference(avatar_component_clause, [], [f71])).
fof(f71, plain, (spl15_3 <=> sP0(sK3)), introduced(avatar_definition, [new_symbols(naming, [spl15_3])])).
fof(f128, plain, (! [X0] : (~ sP0(X0) | ~ actual_world(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f127, f45])).
fof(f45, plain, ! [X0] : (hollywood_placename(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : ((in(X0, sK13(X0), sK9(X0)) & down(X0, sK13(X0), sK12(X0)) & barrel(X0, sK13(X0)) & present(X0, sK13(X0)) & agent(X0, sK13(X0), sK11(X0)) & event(X0, sK13(X0)) & lonely(X0, sK12(X0)) & street(X0, sK12(X0)) & old(X0, sK11(X0)) & dirty(X0, sK11(X0)) & white(X0, sK11(X0)) & chevy(X0, sK11(X0)) & placename(X0, sK10(X0)) & hollywood_placename(X0, sK10(X0)) & city(X0, sK9(X0)) & of(X0, sK10(X0), sK9(X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11, sK12, sK13])], [f18, f19])).
fof(f19, plain, ! [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) => (in(X0, sK13(X0), sK9(X0)) & down(X0, sK13(X0), sK12(X0)) & barrel(X0, sK13(X0)) & present(X0, sK13(X0)) & agent(X0, sK13(X0), sK11(X0)) & event(X0, sK13(X0)) & lonely(X0, sK12(X0)) & street(X0, sK12(X0)) & old(X0, sK11(X0)) & dirty(X0, sK11(X0)) & white(X0, sK11(X0)) & chevy(X0, sK11(X0)) & placename(X0, sK10(X0)) & hollywood_placename(X0, sK10(X0)) & city(X0, sK9(X0)) & of(X0, sK10(X0), sK9(X0)))), introduced(choice_axiom, [])).
fof(f18, plain, ! [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) | ~ sP0(X0)), inference(rectify, [], [f17])).
fof(f17, plain, ! [X12] : (? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13)) | ~ sP0(X12)), inference(nnf_transformation, [], [f6])).
fof(f6, plain, ! [X12] : (? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13)) | ~ sP0(X12)), inference(usedef, [], [e6])).
fof(e6, plain, ! [X12] : (sP0(X12) <=> ? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f127, plain, (! [X0] : (~ hollywood_placename(X0, sK10(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f126, f46])).
fof(f46, plain, ! [X0] : (placename(X0, sK10(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f126, plain, (! [X0] : (~ placename(X0, sK10(X0)) | ~ hollywood_placename(X0, sK10(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f125])).
fof(f125, plain, (! [X0] : (~ placename(X0, sK10(X0)) | ~ hollywood_placename(X0, sK10(X0)) | ~ actual_world(X0) | ~ sP0(X0) | ~ sP0(X0)) | ~ spl15_2), inference(resolution, [], [f120, f43])).
fof(f43, plain, ! [X0] : (of(X0, sK10(X0), sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f120, plain, (! [X0, X1] : (~ of(X0, X1, sK9(X0)) | ~ placename(X0, X1) | ~ hollywood_placename(X0, X1) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f119, f50])).
fof(f50, plain, ! [X0] : (old(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f119, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ old(X0, sK11(X0)) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f118, f49])).
fof(f49, plain, ! [X0] : (dirty(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f118, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ dirty(X0, sK11(X0)) | ~ old(X0, sK11(X0)) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f117, f48])).
fof(f48, plain, ! [X0] : (white(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f117, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ white(X0, sK11(X0)) | ~ dirty(X0, sK11(X0)) | ~ old(X0, sK11(X0)) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f116, f47])).
fof(f47, plain, ! [X0] : (chevy(X0, sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f116, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, sK11(X0)) | ~ white(X0, sK11(X0)) | ~ dirty(X0, sK11(X0)) | ~ old(X0, sK11(X0)) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f115])).
fof(f115, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, sK11(X0)) | ~ white(X0, sK11(X0)) | ~ dirty(X0, sK11(X0)) | ~ old(X0, sK11(X0)) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0) | ~ sP0(X0)) | ~ spl15_2), inference(resolution, [], [f108, f54])).
fof(f54, plain, ! [X0] : (agent(X0, sK13(X0), sK11(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f108, plain, (! [X2, X0, X1] : (~ agent(X0, sK13(X0), X2) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ of(X0, X1, sK9(X0)) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f107, f52])).
fof(f52, plain, ! [X0] : (lonely(X0, sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f107, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ lonely(X0, sK12(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK13(X0), X2) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f106, f51])).
fof(f51, plain, ! [X0] : (street(X0, sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f106, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, sK12(X0)) | ~ lonely(X0, sK12(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK13(X0), X2) | ~ actual_world(X0) | ~ sP0(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f105])).
fof(f105, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK9(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, sK12(X0)) | ~ lonely(X0, sK12(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK13(X0), X2) | ~ actual_world(X0) | ~ sP0(X0) | ~ sP0(X0)) | ~ spl15_2), inference(resolution, [], [f100, f57])).
fof(f57, plain, ! [X0] : (down(X0, sK13(X0), sK12(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f100, plain, (! [X6, X4, X7, X5] : (~ down(X4, sK13(X4), X6) | ~ of(X4, X5, sK9(X4)) | ~ hollywood_placename(X4, X5) | ~ placename(X4, X5) | ~ street(X4, X6) | ~ lonely(X4, X6) | ~ chevy(X4, X7) | ~ white(X4, X7) | ~ dirty(X4, X7) | ~ old(X4, X7) | ~ agent(X4, sK13(X4), X7) | ~ actual_world(X4) | ~ sP0(X4)) | ~ spl15_2), inference(subsumption_resolution, [], [f99, f56])).
fof(f56, plain, ! [X0] : (barrel(X0, sK13(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f99, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ of(X4, X5, sK9(X4)) | ~ hollywood_placename(X4, X5) | ~ placename(X4, X5) | ~ street(X4, X6) | ~ lonely(X4, X6) | ~ chevy(X4, X7) | ~ white(X4, X7) | ~ dirty(X4, X7) | ~ old(X4, X7) | ~ agent(X4, sK13(X4), X7) | ~ barrel(X4, sK13(X4)) | ~ down(X4, sK13(X4), X6) | ~ sP0(X4)) | ~ spl15_2), inference(subsumption_resolution, [], [f98, f55])).
fof(f55, plain, ! [X0] : (present(X0, sK13(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f98, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ of(X4, X5, sK9(X4)) | ~ hollywood_placename(X4, X5) | ~ placename(X4, X5) | ~ street(X4, X6) | ~ lonely(X4, X6) | ~ chevy(X4, X7) | ~ white(X4, X7) | ~ dirty(X4, X7) | ~ old(X4, X7) | ~ agent(X4, sK13(X4), X7) | ~ present(X4, sK13(X4)) | ~ barrel(X4, sK13(X4)) | ~ down(X4, sK13(X4), X6) | ~ sP0(X4)) | ~ spl15_2), inference(subsumption_resolution, [], [f97, f53])).
fof(f53, plain, ! [X0] : (event(X0, sK13(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f97, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ of(X4, X5, sK9(X4)) | ~ hollywood_placename(X4, X5) | ~ placename(X4, X5) | ~ street(X4, X6) | ~ lonely(X4, X6) | ~ chevy(X4, X7) | ~ white(X4, X7) | ~ dirty(X4, X7) | ~ old(X4, X7) | ~ event(X4, sK13(X4)) | ~ agent(X4, sK13(X4), X7) | ~ present(X4, sK13(X4)) | ~ barrel(X4, sK13(X4)) | ~ down(X4, sK13(X4), X6) | ~ sP0(X4)) | ~ spl15_2), inference(subsumption_resolution, [], [f92, f44])).
fof(f44, plain, ! [X0] : (city(X0, sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f92, plain, (! [X6, X4, X7, X5] : (~ actual_world(X4) | ~ of(X4, X5, sK9(X4)) | ~ city(X4, sK9(X4)) | ~ hollywood_placename(X4, X5) | ~ placename(X4, X5) | ~ street(X4, X6) | ~ lonely(X4, X6) | ~ chevy(X4, X7) | ~ white(X4, X7) | ~ dirty(X4, X7) | ~ old(X4, X7) | ~ event(X4, sK13(X4)) | ~ agent(X4, sK13(X4), X7) | ~ present(X4, sK13(X4)) | ~ barrel(X4, sK13(X4)) | ~ down(X4, sK13(X4), X6) | ~ sP0(X4)) | ~ spl15_2), inference(resolution, [], [f68, f58])).
fof(f58, plain, ! [X0] : (in(X0, sK13(X0), sK9(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f20])).
fof(f68, plain, (! [X4, X2, X0, X5, X3, X1] : (~ in(X0, X5, X1) | ~ actual_world(X0) | ~ of(X0, X2, X1) | ~ city(X0, X1) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ street(X0, X3) | ~ lonely(X0, X3) | ~ chevy(X0, X4) | ~ white(X0, X4) | ~ dirty(X0, X4) | ~ old(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X4) | ~ present(X0, X5) | ~ barrel(X0, X5) | ~ down(X0, X5, X3)) | ~ spl15_2), inference(avatar_component_clause, [], [f67])).
fof(f67, plain, (spl15_2 <=> ! [X1, X3, X5, X0, X2, X4] : (~ in(X0, X5, X1) | ~ actual_world(X0) | ~ of(X0, X2, X1) | ~ city(X0, X1) | ~ hollywood_placename(X0, X2) | ~ placename(X0, X2) | ~ street(X0, X3) | ~ lonely(X0, X3) | ~ chevy(X0, X4) | ~ white(X0, X4) | ~ dirty(X0, X4) | ~ old(X0, X4) | ~ event(X0, X5) | ~ agent(X0, X5, X4) | ~ present(X0, X5) | ~ barrel(X0, X5) | ~ down(X0, X5, X3))), introduced(avatar_definition, [new_symbols(naming, [spl15_2])])).
fof(f132, plain, (~ spl15_6 | ~ spl15_2 | ~ spl15_5), inference(avatar_split_clause, [], [f129, f82, f67, f87])).
fof(f87, plain, (spl15_6 <=> actual_world(sK14)), introduced(avatar_definition, [new_symbols(naming, [spl15_6])])).
fof(f82, plain, (spl15_5 <=> sP1(sK14)), introduced(avatar_definition, [new_symbols(naming, [spl15_5])])).
fof(f129, plain, (~ actual_world(sK14) | (~ spl15_2 | ~ spl15_5)), inference(resolution, [], [f124, f84])).
fof(f84, plain, (sP1(sK14) | ~ spl15_5), inference(avatar_component_clause, [], [f82])).
fof(f124, plain, (! [X0] : (~ sP1(X0) | ~ actual_world(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f123, f29])).
fof(f29, plain, ! [X0] : (hollywood_placename(X0, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : ((in(X0, sK8(X0), sK4(X0)) & down(X0, sK8(X0), sK6(X0)) & barrel(X0, sK8(X0)) & present(X0, sK8(X0)) & agent(X0, sK8(X0), sK7(X0)) & event(X0, sK8(X0)) & old(X0, sK7(X0)) & dirty(X0, sK7(X0)) & white(X0, sK7(X0)) & chevy(X0, sK7(X0)) & lonely(X0, sK6(X0)) & street(X0, sK6(X0)) & placename(X0, sK5(X0)) & hollywood_placename(X0, sK5(X0)) & city(X0, sK4(X0)) & of(X0, sK5(X0), sK4(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6, sK7, sK8])], [f14, f15])).
fof(f15, plain, ! [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) => (in(X0, sK8(X0), sK4(X0)) & down(X0, sK8(X0), sK6(X0)) & barrel(X0, sK8(X0)) & present(X0, sK8(X0)) & agent(X0, sK8(X0), sK7(X0)) & event(X0, sK8(X0)) & old(X0, sK7(X0)) & dirty(X0, sK7(X0)) & white(X0, sK7(X0)) & chevy(X0, sK7(X0)) & lonely(X0, sK6(X0)) & street(X0, sK6(X0)) & placename(X0, sK5(X0)) & hollywood_placename(X0, sK5(X0)) & city(X0, sK4(X0)) & of(X0, sK5(X0), sK4(X0)))), introduced(choice_axiom, [])).
fof(f14, plain, ! [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) | ~ sP1(X0)), inference(nnf_transformation, [], [f7])).
fof(f7, plain, ! [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) | ~ sP1(X0)), inference(usedef, [], [e7])).
fof(e7, plain, ! [X0] : (sP1(X0) <=> ? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f123, plain, (! [X0] : (~ hollywood_placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f122, f30])).
fof(f30, plain, ! [X0] : (placename(X0, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f122, plain, (! [X0] : (~ placename(X0, sK5(X0)) | ~ hollywood_placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f121])).
fof(f121, plain, (! [X0] : (~ placename(X0, sK5(X0)) | ~ hollywood_placename(X0, sK5(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl15_2), inference(resolution, [], [f114, f27])).
fof(f27, plain, ! [X0] : (of(X0, sK5(X0), sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f114, plain, (! [X0, X1] : (~ of(X0, X1, sK4(X0)) | ~ placename(X0, X1) | ~ hollywood_placename(X0, X1) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f113, f36])).
fof(f36, plain, ! [X0] : (old(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f113, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ old(X0, sK7(X0)) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f112, f35])).
fof(f35, plain, ! [X0] : (dirty(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f112, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ dirty(X0, sK7(X0)) | ~ old(X0, sK7(X0)) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f111, f34])).
fof(f34, plain, ! [X0] : (white(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f111, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ white(X0, sK7(X0)) | ~ dirty(X0, sK7(X0)) | ~ old(X0, sK7(X0)) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f110, f33])).
fof(f33, plain, ! [X0] : (chevy(X0, sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f110, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, sK7(X0)) | ~ white(X0, sK7(X0)) | ~ dirty(X0, sK7(X0)) | ~ old(X0, sK7(X0)) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f109])).
fof(f109, plain, (! [X0, X1] : (~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, sK7(X0)) | ~ white(X0, sK7(X0)) | ~ dirty(X0, sK7(X0)) | ~ old(X0, sK7(X0)) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl15_2), inference(resolution, [], [f104, f38])).
fof(f38, plain, ! [X0] : (agent(X0, sK8(X0), sK7(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f104, plain, (! [X2, X0, X1] : (~ agent(X0, sK8(X0), X2) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ of(X0, X1, sK4(X0)) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f103, f32])).
fof(f32, plain, ! [X0] : (lonely(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f103, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ lonely(X0, sK6(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK8(X0), X2) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f102, f31])).
fof(f31, plain, ! [X0] : (street(X0, sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f102, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, sK6(X0)) | ~ lonely(X0, sK6(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK8(X0), X2) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(duplicate_literal_removal, [], [f101])).
fof(f101, plain, (! [X2, X0, X1] : (~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, sK6(X0)) | ~ lonely(X0, sK6(X0)) | ~ chevy(X0, X2) | ~ white(X0, X2) | ~ dirty(X0, X2) | ~ old(X0, X2) | ~ agent(X0, sK8(X0), X2) | ~ actual_world(X0) | ~ sP1(X0) | ~ sP1(X0)) | ~ spl15_2), inference(resolution, [], [f96, f41])).
fof(f41, plain, ! [X0] : (down(X0, sK8(X0), sK6(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f96, plain, (! [X2, X0, X3, X1] : (~ down(X0, sK8(X0), X2) | ~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, X2) | ~ lonely(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ agent(X0, sK8(X0), X3) | ~ actual_world(X0) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f95, f40])).
fof(f40, plain, ! [X0] : (barrel(X0, sK8(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f95, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, X2) | ~ lonely(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ agent(X0, sK8(X0), X3) | ~ barrel(X0, sK8(X0)) | ~ down(X0, sK8(X0), X2) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f94, f39])).
fof(f39, plain, ! [X0] : (present(X0, sK8(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f94, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, X2) | ~ lonely(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ agent(X0, sK8(X0), X3) | ~ present(X0, sK8(X0)) | ~ barrel(X0, sK8(X0)) | ~ down(X0, sK8(X0), X2) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f93, f37])).
fof(f37, plain, ! [X0] : (event(X0, sK8(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f93, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ of(X0, X1, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, X2) | ~ lonely(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ event(X0, sK8(X0)) | ~ agent(X0, sK8(X0), X3) | ~ present(X0, sK8(X0)) | ~ barrel(X0, sK8(X0)) | ~ down(X0, sK8(X0), X2) | ~ sP1(X0)) | ~ spl15_2), inference(subsumption_resolution, [], [f91, f28])).
fof(f28, plain, ! [X0] : (city(X0, sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f91, plain, (! [X2, X0, X3, X1] : (~ actual_world(X0) | ~ of(X0, X1, sK4(X0)) | ~ city(X0, sK4(X0)) | ~ hollywood_placename(X0, X1) | ~ placename(X0, X1) | ~ street(X0, X2) | ~ lonely(X0, X2) | ~ chevy(X0, X3) | ~ white(X0, X3) | ~ dirty(X0, X3) | ~ old(X0, X3) | ~ event(X0, sK8(X0)) | ~ agent(X0, sK8(X0), X3) | ~ present(X0, sK8(X0)) | ~ barrel(X0, sK8(X0)) | ~ down(X0, sK8(X0), X2) | ~ sP1(X0)) | ~ spl15_2), inference(resolution, [], [f68, f42])).
fof(f42, plain, ! [X0] : (in(X0, sK8(X0), sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f16])).
fof(f90, plain, (spl15_1 | spl15_6), inference(avatar_split_clause, [], [f59, f87, f63])).
fof(f63, plain, (spl15_1 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl15_1])])).
fof(f59, plain, (actual_world(sK14) | sP2), inference(cnf_transformation, [], [f23])).
fof(f23, plain, ((! [X0] : (! [X1, X2, X3, X4, X5] : (~ in(X0, X5, X1) | ~ down(X0, X5, X4) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ lonely(X0, X4) | ~ street(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1)) | ~ actual_world(X0)) & (sP1(sK14) & actual_world(sK14))) | sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f21, f22])).
fof(f22, plain, (? [X6] : (sP1(X6) & actual_world(X6)) => (sP1(sK14) & actual_world(sK14))), introduced(choice_axiom, [])).
fof(f21, plain, ((! [X0] : (! [X1, X2, X3, X4, X5] : (~ in(X0, X5, X1) | ~ down(X0, X5, X4) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ lonely(X0, X4) | ~ street(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1)) | ~ actual_world(X0)) & ? [X6] : (sP1(X6) & actual_world(X6))) | sP2), inference(rectify, [], [f9])).
fof(f9, plain, ((! [X6] : (! [X7, X8, X9, X10, X11] : (~ in(X6, X11, X7) | ~ down(X6, X11, X10) | ~ barrel(X6, X11) | ~ present(X6, X11) | ~ agent(X6, X11, X9) | ~ event(X6, X11) | ~ lonely(X6, X10) | ~ street(X6, X10) | ~ old(X6, X9) | ~ dirty(X6, X9) | ~ white(X6, X9) | ~ chevy(X6, X9) | ~ placename(X6, X8) | ~ hollywood_placename(X6, X8) | ~ city(X6, X7) | ~ of(X6, X8, X7)) | ~ actual_world(X6)) & ? [X0] : (sP1(X0) & actual_world(X0))) | sP2), inference(definition_folding, [], [f5, e8, e7, e6])).
fof(f8, plain, ((! [X18] : (! [X19, X20, X21, X22, X23] : (~ in(X18, X23, X19) | ~ down(X18, X23, X21) | ~ barrel(X18, X23) | ~ present(X18, X23) | ~ agent(X18, X23, X22) | ~ event(X18, X23) | ~ old(X18, X22) | ~ dirty(X18, X22) | ~ white(X18, X22) | ~ chevy(X18, X22) | ~ lonely(X18, X21) | ~ street(X18, X21) | ~ placename(X18, X20) | ~ hollywood_placename(X18, X20) | ~ city(X18, X19) | ~ of(X18, X20, X19)) | ~ actual_world(X18)) & ? [X12] : (sP0(X12) & actual_world(X12))) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (! [X18] : (! [X19, X20, X21, X22, X23] : (~ in(X18, X23, X19) | ~ down(X18, X23, X21) | ~ barrel(X18, X23) | ~ present(X18, X23) | ~ agent(X18, X23, X22) | ~ event(X18, X23) | ~ old(X18, X22) | ~ dirty(X18, X22) | ~ white(X18, X22) | ~ chevy(X18, X22) | ~ lonely(X18, X21) | ~ street(X18, X21) | ~ placename(X18, X20) | ~ hollywood_placename(X18, X20) | ~ city(X18, X19) | ~ of(X18, X20, X19)) | ~ actual_world(X18)) & ? [X12] : (sP0(X12) & actual_world(X12)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f5, plain, ((! [X6] : (! [X7, X8, X9, X10, X11] : (~ in(X6, X11, X7) | ~ down(X6, X11, X10) | ~ barrel(X6, X11) | ~ present(X6, X11) | ~ agent(X6, X11, X9) | ~ event(X6, X11) | ~ lonely(X6, X10) | ~ street(X6, X10) | ~ old(X6, X9) | ~ dirty(X6, X9) | ~ white(X6, X9) | ~ chevy(X6, X9) | ~ placename(X6, X8) | ~ hollywood_placename(X6, X8) | ~ city(X6, X7) | ~ of(X6, X8, X7)) | ~ actual_world(X6)) & ? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0))) | (! [X18] : (! [X19, X20, X21, X22, X23] : (~ in(X18, X23, X19) | ~ down(X18, X23, X21) | ~ barrel(X18, X23) | ~ present(X18, X23) | ~ agent(X18, X23, X22) | ~ event(X18, X23) | ~ old(X18, X22) | ~ dirty(X18, X22) | ~ white(X18, X22) | ~ chevy(X18, X22) | ~ lonely(X18, X21) | ~ street(X18, X21) | ~ placename(X18, X20) | ~ hollywood_placename(X18, X20) | ~ city(X18, X19) | ~ of(X18, X20, X19)) | ~ actual_world(X18)) & ? [X12] : (? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13)) & actual_world(X12)))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ~ ((? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0)) => ? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X10) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X9) & event(X6, X11) & lonely(X6, X10) & street(X6, X10) & old(X6, X9) & dirty(X6, X9) & white(X6, X9) & chevy(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6))) & (? [X12] : (? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13)) & actual_world(X12)) => ? [X18] : (? [X19, X20, X21, X22, X23] : (in(X18, X23, X19) & down(X18, X23, X21) & barrel(X18, X23) & present(X18, X23) & agent(X18, X23, X22) & event(X18, X23) & old(X18, X22) & dirty(X18, X22) & white(X18, X22) & chevy(X18, X22) & lonely(X18, X21) & street(X18, X21) & placename(X18, X20) & hollywood_placename(X18, X20) & city(X18, X19) & of(X18, X20, X19)) & actual_world(X18)))), inference(flattening, [], [f3])).
fof(f3, plain, ~ ~ ~ ((? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X3) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X4) & event(X0, X5) & old(X0, X4) & dirty(X0, X4) & white(X0, X4) & chevy(X0, X4) & lonely(X0, X3) & street(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0)) => ? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X10) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X9) & event(X6, X11) & lonely(X6, X10) & street(X6, X10) & old(X6, X9) & dirty(X6, X9) & white(X6, X9) & chevy(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6))) & (? [X12] : (? [X13, X14, X15, X16, X17] : (in(X12, X17, X13) & down(X12, X17, X16) & barrel(X12, X17) & present(X12, X17) & agent(X12, X17, X15) & event(X12, X17) & lonely(X12, X16) & street(X12, X16) & old(X12, X15) & dirty(X12, X15) & white(X12, X15) & chevy(X12, X15) & placename(X12, X14) & hollywood_placename(X12, X14) & city(X12, X13) & of(X12, X14, X13)) & actual_world(X12)) => ? [X18] : (? [X19, X20, X21, X22, X23] : (in(X18, X23, X19) & down(X18, X23, X21) & barrel(X18, X23) & present(X18, X23) & agent(X18, X23, X22) & event(X18, X23) & old(X18, X22) & dirty(X18, X22) & white(X18, X22) & chevy(X18, X22) & lonely(X18, X21) & street(X18, X21) & placename(X18, X20) & hollywood_placename(X18, X20) & city(X18, X19) & of(X18, X20, X19)) & actual_world(X18)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ~ ~ ((? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X9) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X10) & event(X6, X11) & old(X6, X10) & dirty(X6, X10) & white(X6, X10) & chevy(X6, X10) & lonely(X6, X9) & street(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6)) => ? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0)) => ? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X9) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X10) & event(X6, X11) & old(X6, X10) & dirty(X6, X10) & white(X6, X10) & chevy(X6, X10) & lonely(X6, X9) & street(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ~ ~ ((? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X9) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X10) & event(X6, X11) & old(X6, X10) & dirty(X6, X10) & white(X6, X10) & chevy(X6, X10) & lonely(X6, X9) & street(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6)) => ? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0))) & (? [X0] : (? [X1, X2, X3, X4, X5] : (in(X0, X5, X1) & down(X0, X5, X4) & barrel(X0, X5) & present(X0, X5) & agent(X0, X5, X3) & event(X0, X5) & lonely(X0, X4) & street(X0, X4) & old(X0, X3) & dirty(X0, X3) & white(X0, X3) & chevy(X0, X3) & placename(X0, X2) & hollywood_placename(X0, X2) & city(X0, X1) & of(X0, X2, X1)) & actual_world(X0)) => ? [X6] : (? [X7, X8, X9, X10, X11] : (in(X6, X11, X7) & down(X6, X11, X9) & barrel(X6, X11) & present(X6, X11) & agent(X6, X11, X10) & event(X6, X11) & old(X6, X10) & dirty(X6, X10) & white(X6, X10) & chevy(X6, X10) & lonely(X6, X9) & street(X6, X9) & placename(X6, X8) & hollywood_placename(X6, X8) & city(X6, X7) & of(X6, X8, X7)) & actual_world(X6)))), file('/home/ubuntu/library/tptp/Problems/NLP/NLP117+1.p', co1)).
fof(f85, plain, (spl15_1 | spl15_5), inference(avatar_split_clause, [], [f60, f82, f63])).
fof(f60, plain, (sP1(sK14) | sP2), inference(cnf_transformation, [], [f23])).
fof(f80, plain, (spl15_1 | spl15_2), inference(avatar_split_clause, [], [f61, f67, f63])).
fof(f61, plain, ! [X4, X2, X0, X5, X3, X1] : (~ in(X0, X5, X1) | ~ down(X0, X5, X4) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X3) | ~ event(X0, X5) | ~ lonely(X0, X4) | ~ street(X0, X4) | ~ old(X0, X3) | ~ dirty(X0, X3) | ~ white(X0, X3) | ~ chevy(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1) | ~ actual_world(X0) | sP2), inference(cnf_transformation, [], [f23])).
fof(f79, plain, (~ spl15_1 | spl15_4), inference(avatar_split_clause, [], [f24, f76, f63])).
fof(f24, plain, (actual_world(sK3) | ~ sP2), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ((! [X0] : (! [X1, X2, X3, X4, X5] : (~ in(X0, X5, X1) | ~ down(X0, X5, X3) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ old(X0, X4) | ~ dirty(X0, X4) | ~ white(X0, X4) | ~ chevy(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1)) | ~ actual_world(X0)) & (sP0(sK3) & actual_world(sK3))) | ~ sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3])], [f11, f12])).
fof(f12, plain, (? [X6] : (sP0(X6) & actual_world(X6)) => (sP0(sK3) & actual_world(sK3))), introduced(choice_axiom, [])).
fof(f11, plain, ((! [X0] : (! [X1, X2, X3, X4, X5] : (~ in(X0, X5, X1) | ~ down(X0, X5, X3) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ old(X0, X4) | ~ dirty(X0, X4) | ~ white(X0, X4) | ~ chevy(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1)) | ~ actual_world(X0)) & ? [X6] : (sP0(X6) & actual_world(X6))) | ~ sP2), inference(rectify, [], [f10])).
fof(f10, plain, ((! [X18] : (! [X19, X20, X21, X22, X23] : (~ in(X18, X23, X19) | ~ down(X18, X23, X21) | ~ barrel(X18, X23) | ~ present(X18, X23) | ~ agent(X18, X23, X22) | ~ event(X18, X23) | ~ old(X18, X22) | ~ dirty(X18, X22) | ~ white(X18, X22) | ~ chevy(X18, X22) | ~ lonely(X18, X21) | ~ street(X18, X21) | ~ placename(X18, X20) | ~ hollywood_placename(X18, X20) | ~ city(X18, X19) | ~ of(X18, X20, X19)) | ~ actual_world(X18)) & ? [X12] : (sP0(X12) & actual_world(X12))) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f74, plain, (~ spl15_1 | spl15_3), inference(avatar_split_clause, [], [f25, f71, f63])).
fof(f25, plain, (sP0(sK3) | ~ sP2), inference(cnf_transformation, [], [f13])).
fof(f69, plain, (~ spl15_1 | spl15_2), inference(avatar_split_clause, [], [f26, f67, f63])).
fof(f26, plain, ! [X4, X2, X0, X5, X3, X1] : (~ in(X0, X5, X1) | ~ down(X0, X5, X3) | ~ barrel(X0, X5) | ~ present(X0, X5) | ~ agent(X0, X5, X4) | ~ event(X0, X5) | ~ old(X0, X4) | ~ dirty(X0, X4) | ~ white(X0, X4) | ~ chevy(X0, X4) | ~ lonely(X0, X3) | ~ street(X0, X3) | ~ placename(X0, X2) | ~ hollywood_placename(X0, X2) | ~ city(X0, X1) | ~ of(X0, X2, X1) | ~ actual_world(X0) | ~ sP2), inference(cnf_transformation, [], [f13])).