fof(f1167, plain, $false, inference(avatar_sat_refutation, [], [f1125, f1165])).
fof(f1165, plain, ~ spl41_16, inference(avatar_contradiction_clause, [], [f1164])).
fof(f1164, plain, ($false | ~ spl41_16), inference(subsumption_resolution, [], [f1162, f502])).
fof(f502, plain, aElementOf0(sz10, cS2076), inference(forward_demodulation, [], [f479, f391])).
fof(f391, plain, (stldt0(sbsmnsldt0(xS)) = cS2076), inference(cnf_transformation, [], [f213])).
fof(f213, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & (((aElementOf0(X2, sK31(X2)) & aElementOf0(sK31(X2), xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK31])], [f211, f212])).
fof(f212, plain, ! [X2] : (? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) => (aElementOf0(X2, sK31(X2)) & aElementOf0(sK31(X2), xS))), introduced(choice_axiom, [])).
fof(f211, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f210])).
fof(f210, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f209])).
fof(f209, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & (((smndt0(sz10) = X0) | (sz10 = X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1))) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | (! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2))) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(nnf_transformation, [], [f55])).
fof(f55, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))) & ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, sbsmnsldt0(xS)) <=> (? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f43])).
fof(f43, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM450+6.p', m__2079)).
fof(f479, plain, aElementOf0(sz10, stldt0(sbsmnsldt0(xS))), inference(equality_resolution, [], [f389])).
fof(f389, plain, ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ (sz10 = X0)), inference(cnf_transformation, [], [f213])).
fof(f1162, plain, (~ aElementOf0(sz10, cS2076) | ~ spl41_16), inference(trivial_inequality_removal, [], [f1160])).
fof(f1160, plain, (~ (sz00 = sz00) | ~ aElementOf0(sz10, cS2076) | ~ spl41_16), inference(superposition, [], [f514, f640])).
fof(f640, plain, ((sz00 = sK34(sz10)) | ~ spl41_16), inference(avatar_component_clause, [], [f638])).
fof(f638, plain, (spl41_16 <=> (sz00 = sK34(sz10))), introduced(avatar_definition, [new_symbols(naming, [spl41_16])])).
fof(f514, plain, ! [X0] : (~ (sz00 = sK34(X0)) | ~ aElementOf0(X0, cS2076)), inference(forward_demodulation, [], [f434, f391])).
fof(f434, plain, ! [X0] : (~ (sz00 = sK34(X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f229])).
fof(f229, plain, (! [X0] : ((aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)))) & sP9(sK34(X0), X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0))) & ~ (sz00 = sK34(X0)) & aInteger0(sK34(X0))) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X3] : ((aElementOf0(X3, stldt0(sbsmnsldt0(xS))) | aElementOf0(X3, sbsmnsldt0(xS)) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, sbsmnsldt0(xS)) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(sbsmnsldt0(xS))))) & ! [X4] : ((aElementOf0(X4, sbsmnsldt0(xS)) | ! [X5] : (~ aElementOf0(X4, X5) | ~ aElementOf0(X5, xS)) | ~ aInteger0(X4)) & (((aElementOf0(X4, sK35(X4)) & aElementOf0(sK35(X4), xS)) & aInteger0(X4)) | ~ aElementOf0(X4, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X7] : ((aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7)), stldt0(sbsmnsldt0(xS))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7)))) & sP8(sK36(X7), X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7))) & ~ (sz00 = sK36(X7)) & aInteger0(sK36(X7))) | ~ aElementOf0(X7, stldt0(sbsmnsldt0(xS)))) & ! [X10] : ((aElementOf0(X10, stldt0(sbsmnsldt0(xS))) | aElementOf0(X10, sbsmnsldt0(xS)) | ~ aInteger0(X10)) & ((~ aElementOf0(X10, sbsmnsldt0(xS)) & aInteger0(X10)) | ~ aElementOf0(X10, stldt0(sbsmnsldt0(xS))))) & ! [X11] : ((aElementOf0(X11, sbsmnsldt0(xS)) | ! [X12] : (~ aElementOf0(X11, X12) | ~ aElementOf0(X12, xS)) | ~ aInteger0(X11)) & (((aElementOf0(X11, sK37(X11)) & aElementOf0(sK37(X11), xS)) & aInteger0(X11)) | ~ aElementOf0(X11, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34, sK35, sK36, sK37])], [f224, f228, f227, f226, f225])).
fof(f225, plain, ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP9(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)))) & sP9(sK34(X0), X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0))) & ~ (sz00 = sK34(X0)) & aInteger0(sK34(X0)))), introduced(choice_axiom, [])).
fof(f226, plain, ! [X4] : (? [X6] : (aElementOf0(X4, X6) & aElementOf0(X6, xS)) => (aElementOf0(X4, sK35(X4)) & aElementOf0(sK35(X4), xS))), introduced(choice_axiom, [])).
fof(f227, plain, ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(sbsmnsldt0(xS))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & sP8(X8, X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7)), stldt0(sbsmnsldt0(xS))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7)))) & sP8(sK36(X7), X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, sK36(X7))) & ~ (sz00 = sK36(X7)) & aInteger0(sK36(X7)))), introduced(choice_axiom, [])).
fof(f228, plain, ! [X11] : (? [X13] : (aElementOf0(X11, X13) & aElementOf0(X13, xS)) => (aElementOf0(X11, sK37(X11)) & aElementOf0(sK37(X11), xS))), introduced(choice_axiom, [])).
fof(f224, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP9(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X3] : ((aElementOf0(X3, stldt0(sbsmnsldt0(xS))) | aElementOf0(X3, sbsmnsldt0(xS)) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, sbsmnsldt0(xS)) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(sbsmnsldt0(xS))))) & ! [X4] : ((aElementOf0(X4, sbsmnsldt0(xS)) | ! [X5] : (~ aElementOf0(X4, X5) | ~ aElementOf0(X5, xS)) | ~ aInteger0(X4)) & ((? [X6] : (aElementOf0(X4, X6) & aElementOf0(X6, xS)) & aInteger0(X4)) | ~ aElementOf0(X4, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(sbsmnsldt0(xS))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & sP8(X8, X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(sbsmnsldt0(xS)))) & ! [X10] : ((aElementOf0(X10, stldt0(sbsmnsldt0(xS))) | aElementOf0(X10, sbsmnsldt0(xS)) | ~ aInteger0(X10)) & ((~ aElementOf0(X10, sbsmnsldt0(xS)) & aInteger0(X10)) | ~ aElementOf0(X10, stldt0(sbsmnsldt0(xS))))) & ! [X11] : ((aElementOf0(X11, sbsmnsldt0(xS)) | ! [X12] : (~ aElementOf0(X11, X12) | ~ aElementOf0(X12, xS)) | ~ aInteger0(X11)) & ((? [X13] : (aElementOf0(X11, X13) & aElementOf0(X13, xS)) & aInteger0(X11)) | ~ aElementOf0(X11, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f223])).
fof(f223, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP9(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X6] : ((aElementOf0(X6, stldt0(sbsmnsldt0(xS))) | aElementOf0(X6, sbsmnsldt0(xS)) | ~ aInteger0(X6)) & ((~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6)) | ~ aElementOf0(X6, stldt0(sbsmnsldt0(xS))))) & ! [X7] : ((aElementOf0(X7, sbsmnsldt0(xS)) | ! [X8] : (~ aElementOf0(X7, X8) | ~ aElementOf0(X8, xS)) | ~ aInteger0(X7)) & ((? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7)) | ~ aElementOf0(X7, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & sP8(X10, X9) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS)))) & ! [X15] : ((aElementOf0(X15, stldt0(sbsmnsldt0(xS))) | aElementOf0(X15, sbsmnsldt0(xS)) | ~ aInteger0(X15)) & ((~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15)) | ~ aElementOf0(X15, stldt0(sbsmnsldt0(xS))))) & ! [X16] : ((aElementOf0(X16, sbsmnsldt0(xS)) | ! [X17] : (~ aElementOf0(X16, X17) | ~ aElementOf0(X17, xS)) | ~ aInteger0(X16)) & ((? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16)) | ~ aElementOf0(X16, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f222])).
fof(f222, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP9(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X6] : ((aElementOf0(X6, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X6, sbsmnsldt0(xS)) | ~ aInteger0(X6))) & ((~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6)) | ~ aElementOf0(X6, stldt0(sbsmnsldt0(xS))))) & ! [X7] : ((aElementOf0(X7, sbsmnsldt0(xS)) | (! [X8] : (~ aElementOf0(X7, X8) | ~ aElementOf0(X8, xS)) | ~ aInteger0(X7))) & ((? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7)) | ~ aElementOf0(X7, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & sP8(X10, X9) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS)))) & ! [X15] : ((aElementOf0(X15, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X15, sbsmnsldt0(xS)) | ~ aInteger0(X15))) & ((~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15)) | ~ aElementOf0(X15, stldt0(sbsmnsldt0(xS))))) & ! [X16] : ((aElementOf0(X16, sbsmnsldt0(xS)) | (! [X17] : (~ aElementOf0(X16, X17) | ~ aElementOf0(X17, xS)) | ~ aInteger0(X16))) & ((? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16)) | ~ aElementOf0(X16, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(nnf_transformation, [], [f134])).
fof(f134, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP9(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & sP8(X10, X9) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS)))) & ! [X15] : (aElementOf0(X15, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15))) & ! [X16] : (aElementOf0(X16, sbsmnsldt0(xS)) <=> (? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16))) & aSet0(sbsmnsldt0(xS))), inference(definition_folding, [], [f117, e133, e132])).
fof(f132, plain, ! [X10, X9] : (! [X12] : ((aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)) | (~ sdteqdtlpzmzozddtrp0(X12, X9, X10) & ~ aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ! [X13] : (~ (sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X13)) | ~ aInteger0(X13))) | ~ aInteger0(X12)) & ((sdteqdtlpzmzozddtrp0(X12, X9, X10) & aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ? [X14] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X14)) & aInteger0(X14)) & aInteger0(X12)) | ~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)))) | ~ sP8(X10, X9)), inference(usedef, [], [e132])).
fof(e132, plain, ! [X10, X9] : (sP8(X10, X9) <=> ! [X12] : ((aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)) | (~ sdteqdtlpzmzozddtrp0(X12, X9, X10) & ~ aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ! [X13] : (~ (sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X13)) | ~ aInteger0(X13))) | ~ aInteger0(X12)) & ((sdteqdtlpzmzozddtrp0(X12, X9, X10) & aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ? [X14] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X14)) & aInteger0(X14)) & aInteger0(X12)) | ~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f133, plain, ! [X1, X0] : (! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtasdt0(X1, X4) = sdtpldt0(X3, smndt0(X0))) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) | ~ sP9(X1, X0)), inference(usedef, [], [e133])).
fof(e133, plain, ! [X1, X0] : (sP9(X1, X0) <=> ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtasdt0(X1, X4) = sdtpldt0(X3, smndt0(X0))) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f117, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtasdt0(X1, X4) = sdtpldt0(X3, smndt0(X0))) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & ! [X12] : ((aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)) | (~ sdteqdtlpzmzozddtrp0(X12, X9, X10) & ~ aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ! [X13] : (~ (sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X13)) | ~ aInteger0(X13))) | ~ aInteger0(X12)) & ((sdteqdtlpzmzozddtrp0(X12, X9, X10) & aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ? [X14] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X14)) & aInteger0(X14)) & aInteger0(X12)) | ~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS)))) & ! [X15] : (aElementOf0(X15, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15))) & ! [X16] : (aElementOf0(X16, sbsmnsldt0(xS)) <=> (? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f116])).
fof(f116, plain, (! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | ((~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtasdt0(X1, X4) = sdtpldt0(X3, smndt0(X0))) | ~ aInteger0(X4))) | ~ aInteger0(X3))) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & ! [X12] : ((aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)) | ((~ sdteqdtlpzmzozddtrp0(X12, X9, X10) & ~ aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ! [X13] : (~ (sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X13)) | ~ aInteger0(X13))) | ~ aInteger0(X12))) & ((sdteqdtlpzmzozddtrp0(X12, X9, X10) & aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ? [X14] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X14)) & aInteger0(X14)) & aInteger0(X12)) | ~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS)))) & ! [X15] : (aElementOf0(X15, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15))) & ! [X16] : (aElementOf0(X16, sbsmnsldt0(xS)) <=> (? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16))) & aSet0(sbsmnsldt0(xS))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, (! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(sbsmnsldt0(xS)))) & ! [X3] : ((((sdteqdtlpzmzozddtrp0(X3, X0, X1) | aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) | ? [X4] : ((sdtasdt0(X1, X4) = sdtpldt0(X3, smndt0(X0))) & aInteger0(X4))) & aInteger0(X3)) => aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) => ? [X10] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X9, X10), stldt0(sbsmnsldt0(xS))) & ! [X11] : (aElementOf0(X11, szAzrzSzezqlpdtcmdtrp0(X9, X10)) => aElementOf0(X11, stldt0(sbsmnsldt0(xS)))) & ! [X12] : ((((sdteqdtlpzmzozddtrp0(X12, X9, X10) | aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) | ? [X13] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X13)) & aInteger0(X13))) & aInteger0(X12)) => aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10))) & (aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(X9, X10)) => (sdteqdtlpzmzozddtrp0(X12, X9, X10) & aDivisorOf0(X10, sdtpldt0(X12, smndt0(X9))) & ? [X14] : ((sdtpldt0(X12, smndt0(X9)) = sdtasdt0(X10, X14)) & aInteger0(X14)) & aInteger0(X12)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X9, X10)) & ~ (sz00 = X10) & aInteger0(X10))) & ! [X15] : (aElementOf0(X15, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X15, sbsmnsldt0(xS)) & aInteger0(X15))) & ! [X16] : (aElementOf0(X16, sbsmnsldt0(xS)) <=> (? [X17] : (aElementOf0(X16, X17) & aElementOf0(X17, xS)) & aInteger0(X16))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f45])).
fof(f45, plain, (! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(sbsmnsldt0(xS)))) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS)) & isClosed0(sbsmnsldt0(xS)) & isOpen0(stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(sbsmnsldt0(xS)))) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM450+6.p', m__2144)).
fof(f1125, plain, spl41_16, inference(avatar_split_clause, [], [f1124, f638])).
fof(f1124, plain, (sz00 = sK34(sz10)), inference(subsumption_resolution, [], [f1123, f502])).
fof(f1123, plain, ((sz00 = sK34(sz10)) | ~ aElementOf0(sz10, cS2076)), inference(subsumption_resolution, [], [f1121, f601])).
fof(f601, plain, aInteger0(sK34(sz10)), inference(resolution, [], [f515, f502])).
fof(f515, plain, ! [X0] : (~ aElementOf0(X0, cS2076) | aInteger0(sK34(X0))), inference(forward_demodulation, [], [f433, f391])).
fof(f433, plain, ! [X0] : (aInteger0(sK34(X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f229])).
fof(f1121, plain, ((sz00 = sK34(sz10)) | ~ aInteger0(sK34(sz10)) | ~ aElementOf0(sz10, cS2076)), inference(resolution, [], [f546, f509])).
fof(f509, plain, ! [X0] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)), cS2076) | ~ aElementOf0(X0, cS2076)), inference(forward_demodulation, [], [f508, f391])).
fof(f508, plain, ! [X0] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)), cS2076) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))), inference(forward_demodulation, [], [f438, f391])).
fof(f438, plain, ! [X0] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK34(X0)), stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f229])).
fof(f546, plain, ! [X0] : (~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), cS2076) | (sz00 = X0) | ~ aInteger0(X0)), inference(forward_demodulation, [], [f461, f391])).
fof(f461, plain, ! [X0] : (~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) | (sz00 = X0) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ! [X0] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) & (~ aElementOf0(sK40(X0), stldt0(sbsmnsldt0(xS))) & aElementOf0(sK40(X0), szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & sP12 & sP11 & aSet0(sbsmnsldt0(xS)) & sP10(X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) | (sz00 = X0) | ~ aInteger0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40])], [f241, f242])).
fof(f242, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))) & aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) => (~ aElementOf0(sK40(X0), stldt0(sbsmnsldt0(xS))) & aElementOf0(sK40(X0), szAzrzSzezqlpdtcmdtrp0(sz10, X0)))), introduced(choice_axiom, [])).
fof(f241, plain, ! [X0] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) & ? [X1] : (~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))) & aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & sP12 & sP11 & aSet0(sbsmnsldt0(xS)) & sP10(X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) | (sz00 = X0) | ~ aInteger0(X0)), inference(rectify, [], [f138])).
fof(f138, plain, ! [X0] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) & ? [X7] : (~ aElementOf0(X7, stldt0(sbsmnsldt0(xS))) & aElementOf0(X7, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & sP12 & sP11 & aSet0(sbsmnsldt0(xS)) & sP10(X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) | (sz00 = X0) | ~ aInteger0(X0)), inference(definition_folding, [], [f119, e137, e136, e135])).
fof(f135, plain, ! [X0] : (! [X1] : ((aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) | (~ sdteqdtlpzmzozddtrp0(X1, sz10, X0) & ~ aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ! [X2] : (~ (sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) | ~ aInteger0(X2))) | ~ aInteger0(X1)) & ((sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X3] : ((sdtpldt0(X1, smndt0(sz10)) = sdtasdt0(X0, X3)) & aInteger0(X3)) & aInteger0(X1)) | ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)))) | ~ sP10(X0)), inference(usedef, [], [e135])).
fof(e135, plain, ! [X0] : (sP10(X0) <=> ! [X1] : ((aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) | (~ sdteqdtlpzmzozddtrp0(X1, sz10, X0) & ~ aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ! [X2] : (~ (sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) | ~ aInteger0(X2))) | ~ aInteger0(X1)) & ((sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X3] : ((sdtpldt0(X1, smndt0(sz10)) = sdtasdt0(X0, X3)) & aInteger0(X3)) & aInteger0(X1)) | ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f136, plain, (! [X4] : (aElementOf0(X4, sbsmnsldt0(xS)) <=> (? [X5] : (aElementOf0(X4, X5) & aElementOf0(X5, xS)) & aInteger0(X4))) | ~ sP11), inference(usedef, [], [e136])).
fof(e136, plain, (sP11 <=> ! [X4] : (aElementOf0(X4, sbsmnsldt0(xS)) <=> (? [X5] : (aElementOf0(X4, X5) & aElementOf0(X5, xS)) & aInteger0(X4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f137, plain, (! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) | ~ sP12), inference(usedef, [], [e137])).
fof(e137, plain, (sP12 <=> ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f119, plain, ! [X0] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) & ? [X7] : (~ aElementOf0(X7, stldt0(sbsmnsldt0(xS))) & aElementOf0(X7, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) & ! [X4] : (aElementOf0(X4, sbsmnsldt0(xS)) <=> (? [X5] : (aElementOf0(X4, X5) & aElementOf0(X5, xS)) & aInteger0(X4))) & aSet0(sbsmnsldt0(xS)) & ! [X1] : ((aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) | (~ sdteqdtlpzmzozddtrp0(X1, sz10, X0) & ~ aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ! [X2] : (~ (sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) | ~ aInteger0(X2))) | ~ aInteger0(X1)) & ((sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X3] : ((sdtpldt0(X1, smndt0(sz10)) = sdtasdt0(X0, X3)) & aInteger0(X3)) & aInteger0(X1)) | ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) | (sz00 = X0) | ~ aInteger0(X0)), inference(flattening, [], [f118])).
fof(f118, plain, ! [X0] : (((((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) & ? [X7] : (~ aElementOf0(X7, stldt0(sbsmnsldt0(xS))) & aElementOf0(X7, szAzrzSzezqlpdtcmdtrp0(sz10, X0)))) & ! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6)))) & (! [X4] : (aElementOf0(X4, sbsmnsldt0(xS)) <=> (? [X5] : (aElementOf0(X4, X5) & aElementOf0(X5, xS)) & aInteger0(X4))) & aSet0(sbsmnsldt0(xS)))) & (! [X1] : ((aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) | ((~ sdteqdtlpzmzozddtrp0(X1, sz10, X0) & ~ aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ! [X2] : (~ (sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) | ~ aInteger0(X2))) | ~ aInteger0(X1))) & ((sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X3] : ((sdtpldt0(X1, smndt0(sz10)) = sdtasdt0(X0, X3)) & aInteger0(X3)) & aInteger0(X1)) | ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0)))) | (sz00 = X0) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ~ ? [X0] : (((! [X1] : ((((sdteqdtlpzmzozddtrp0(X1, sz10, X0) | aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) | ? [X2] : ((sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) & aInteger0(X2))) & aInteger0(X1)) => aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & (aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => (sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X3] : ((sdtpldt0(X1, smndt0(sz10)) = sdtasdt0(X0, X3)) & aInteger0(X3)) & aInteger0(X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) => ((! [X4] : (aElementOf0(X4, sbsmnsldt0(xS)) <=> (? [X5] : (aElementOf0(X4, X5) & aElementOf0(X5, xS)) & aInteger0(X4))) & aSet0(sbsmnsldt0(xS))) => (! [X6] : (aElementOf0(X6, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X6, sbsmnsldt0(xS)) & aInteger0(X6))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) | ! [X7] : (aElementOf0(X7, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => aElementOf0(X7, stldt0(sbsmnsldt0(xS)))))))) & ~ (sz00 = X0) & aInteger0(X0)), inference(rectify, [], [f47])).
fof(f47, plain, ~ ? [X0] : (((! [X1] : ((((sdteqdtlpzmzozddtrp0(X1, sz10, X0) | aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) | ? [X2] : ((sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) & aInteger0(X2))) & aInteger0(X1)) => aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & (aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => (sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X2] : ((sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) & aInteger0(X2)) & aInteger0(X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) => ((! [X1] : (aElementOf0(X1, sbsmnsldt0(xS)) <=> (? [X2] : (aElementOf0(X1, X2) & aElementOf0(X2, xS)) & aInteger0(X1))) & aSet0(sbsmnsldt0(xS))) => (! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) | ! [X1] : (aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => aElementOf0(X1, stldt0(sbsmnsldt0(xS)))))))) & ~ (sz00 = X0) & aInteger0(X0)), inference(negated_conjecture, [], [f46])).
fof(f46, plain, ~ ? [X0] : (((! [X1] : ((((sdteqdtlpzmzozddtrp0(X1, sz10, X0) | aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) | ? [X2] : ((sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) & aInteger0(X2))) & aInteger0(X1)) => aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0))) & (aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => (sdteqdtlpzmzozddtrp0(X1, sz10, X0) & aDivisorOf0(X0, sdtpldt0(X1, smndt0(sz10))) & ? [X2] : ((sdtasdt0(X0, X2) = sdtpldt0(X1, smndt0(sz10))) & aInteger0(X2)) & aInteger0(X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, X0))) => ((! [X1] : (aElementOf0(X1, sbsmnsldt0(xS)) <=> (? [X2] : (aElementOf0(X1, X2) & aElementOf0(X2, xS)) & aInteger0(X1))) & aSet0(sbsmnsldt0(xS))) => (! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, X0), stldt0(sbsmnsldt0(xS))) | ! [X1] : (aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(sz10, X0)) => aElementOf0(X1, stldt0(sbsmnsldt0(xS)))))))) & ~ (sz00 = X0) & aInteger0(X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM450+6.p', m__)).