fof(f369558, plain, $false, inference(avatar_sat_refutation, [], [f555, f574, f583, f2585, f325634, f369331, f369487])).
fof(f369487, plain, (~ spl38_1 | spl38_9 | ~ spl38_1628), inference(avatar_contradiction_clause, [], [f369486])).
fof(f369486, plain, ($false | (~ spl38_1 | spl38_9 | ~ spl38_1628)), inference(subsumption_resolution, [], [f369353, f2126])).
fof(f2126, plain, aDivisorOf0(xp, xp), inference(subsumption_resolution, [], [f2125, f434])).
fof(f434, plain, aInteger0(xp), inference(cnf_transformation, [], [f236])).
fof(f236, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & (((aElementOf0(X2, sK36(X2)) & aElementOf0(sK36(X2), xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | (~ sdteqdtlpzmzozddtrp0(X5, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X5, smndt0(sz10))) & ! [X6] : (~ (sdtasdt0(xp, X6) = sdtpldt0(X5, smndt0(sz10))) | ~ aInteger0(X6))) | ~ aInteger0(X5)) & ((sdteqdtlpzmzozddtrp0(X5, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X5, smndt0(sz10))) & ((sdtpldt0(X5, smndt0(sz10)) = sdtasdt0(xp, sK37(X5))) & aInteger0(sK37(X5))) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f233, f235, f234])).
fof(f234, plain, ! [X2] : (? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) => (aElementOf0(X2, sK36(X2)) & aElementOf0(sK36(X2), xS))), introduced(choice_axiom, [])).
fof(f235, plain, ! [X5] : (? [X7] : ((sdtpldt0(X5, smndt0(sz10)) = sdtasdt0(xp, X7)) & aInteger0(X7)) => ((sdtpldt0(X5, smndt0(sz10)) = sdtasdt0(xp, sK37(X5))) & aInteger0(sK37(X5)))), introduced(choice_axiom, [])).
fof(f233, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | (~ sdteqdtlpzmzozddtrp0(X5, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X5, smndt0(sz10))) & ! [X6] : (~ (sdtasdt0(xp, X6) = sdtpldt0(X5, smndt0(sz10))) | ~ aInteger0(X6))) | ~ aInteger0(X5)) & ((sdteqdtlpzmzozddtrp0(X5, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X5, smndt0(sz10))) & ? [X7] : ((sdtpldt0(X5, smndt0(sz10)) = sdtasdt0(xp, X7)) & aInteger0(X7)) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(rectify, [], [f232])).
fof(f232, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X4] : ((aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | (~ sdteqdtlpzmzozddtrp0(X4, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ! [X5] : (~ (sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X5)) | ~ aInteger0(X5))) | ~ aInteger0(X4)) & ((sdteqdtlpzmzozddtrp0(X4, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ? [X6] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X6)) & aInteger0(X6)) & aInteger0(X4)) | ~ aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(flattening, [], [f231])).
fof(f231, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1))) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | (! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2))) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X4] : ((aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | (~ sdteqdtlpzmzozddtrp0(X4, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ! [X5] : (~ (sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X5)) | ~ aInteger0(X5))) | ~ aInteger0(X4)) & ((sdteqdtlpzmzozddtrp0(X4, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ? [X6] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X6)) & aInteger0(X6)) & aInteger0(X4)) | ~ aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(nnf_transformation, [], [f121])).
fof(f121, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) & ! [X2] : (aElementOf0(X2, sbsmnsldt0(xS)) <=> (? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2))) & aSet0(sbsmnsldt0(xS)) & ! [X4] : ((aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | (~ sdteqdtlpzmzozddtrp0(X4, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ! [X5] : (~ (sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X5)) | ~ aInteger0(X5))) | ~ aInteger0(X4)) & ((sdteqdtlpzmzozddtrp0(X4, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ? [X6] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X6)) & aInteger0(X6)) & aInteger0(X4)) | ~ aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(flattening, [], [f120])).
fof(f120, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) & ! [X2] : (aElementOf0(X2, sbsmnsldt0(xS)) <=> (? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2))) & aSet0(sbsmnsldt0(xS)) & ! [X4] : ((aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | ((~ sdteqdtlpzmzozddtrp0(X4, sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ! [X5] : (~ (sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X5)) | ~ aInteger0(X5))) | ~ aInteger0(X4))) & ((sdteqdtlpzmzozddtrp0(X4, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ? [X6] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X6)) & aInteger0(X6)) & aInteger0(X4)) | ~ aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(ennf_transformation, [], [f58])).
fof(f58, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) => aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) & ! [X2] : (aElementOf0(X2, sbsmnsldt0(xS)) <=> (? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2))) & aSet0(sbsmnsldt0(xS)) & ! [X4] : ((((sdteqdtlpzmzozddtrp0(X4, sz10, xp) | aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) | ? [X5] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X5)) & aInteger0(X5))) & aInteger0(X4)) => aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & (aElementOf0(X4, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) => (sdteqdtlpzmzozddtrp0(X4, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X4, smndt0(sz10))) & ? [X6] : ((sdtpldt0(X4, smndt0(sz10)) = sdtasdt0(xp, X6)) & aInteger0(X6)) & aInteger0(X4)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), inference(rectify, [], [f46])).
fof(f46, plain, (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sz10, xp), stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) => aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS)) & ! [X0] : ((((sdteqdtlpzmzozddtrp0(X0, sz10, xp) | aDivisorOf0(xp, sdtpldt0(X0, smndt0(sz10))) | ? [X1] : ((sdtasdt0(xp, X1) = sdtpldt0(X0, smndt0(sz10))) & aInteger0(X1))) & aInteger0(X0)) => aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp))) & (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(sz10, xp)) => (sdteqdtlpzmzozddtrp0(X0, sz10, xp) & aDivisorOf0(xp, sdtpldt0(X0, smndt0(sz10))) & ? [X1] : ((sdtasdt0(xp, X1) = sdtpldt0(X0, smndt0(sz10))) & aInteger0(X1)) & aInteger0(X0)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ (sz00 = xp) & aInteger0(xp)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', m__2171)).
fof(f2125, plain, (aDivisorOf0(xp, xp) | ~ aInteger0(xp)), inference(subsumption_resolution, [], [f2124, f435])).
fof(f435, plain, ~ (sz00 = xp), inference(cnf_transformation, [], [f236])).
fof(f2124, plain, (aDivisorOf0(xp, xp) | (sz00 = xp) | ~ aInteger0(xp)), inference(subsumption_resolution, [], [f2023, f240])).
fof(f240, plain, aInteger0(sz10), inference(cnf_transformation, [], [f3])).
fof(f3, plain, aInteger0(sz10), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mIntOne)).
fof(f2023, plain, (aDivisorOf0(xp, xp) | ~ aInteger0(sz10) | (sz00 = xp) | ~ aInteger0(xp)), inference(superposition, [], [f481, f616])).
fof(f616, plain, (xp = sdtasdt0(xp, sz10)), inference(resolution, [], [f252, f434])).
fof(f252, plain, ! [X0] : (~ aInteger0(X0) | (sdtasdt0(X0, sz10) = X0)), inference(cnf_transformation, [], [f75])).
fof(f75, plain, ! [X0] : (((sdtasdt0(sz10, X0) = X0) & (sdtasdt0(X0, sz10) = X0)) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (aInteger0(X0) => ((sdtasdt0(sz10, X0) = X0) & (sdtasdt0(X0, sz10) = X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mMulOne)).
fof(f481, plain, ! [X2, X1] : (aDivisorOf0(X1, sdtasdt0(X1, X2)) | ~ aInteger0(X2) | (sz00 = X1) | ~ aInteger0(X1)), inference(subsumption_resolution, [], [f463, f243])).
fof(f243, plain, ! [X0, X1] : (aInteger0(sdtasdt0(X0, X1)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ! [X0, X1] : (aInteger0(sdtasdt0(X0, X1)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f63])).
fof(f63, plain, ! [X0, X1] : (aInteger0(sdtasdt0(X0, X1)) | (~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => aInteger0(sdtasdt0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mIntMult)).
fof(f463, plain, ! [X2, X1] : (aDivisorOf0(X1, sdtasdt0(X1, X2)) | ~ aInteger0(X2) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(sdtasdt0(X1, X2))), inference(equality_resolution, [], [f265])).
fof(f265, plain, ! [X2, X0, X1] : (aDivisorOf0(X1, X0) | ~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0] : (! [X1] : ((aDivisorOf0(X1, X0) | ! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1)) & ((((sdtasdt0(X1, sK11(X0, X1)) = X0) & aInteger0(sK11(X0, X1))) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aDivisorOf0(X1, X0))) | ~ aInteger0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11])], [f142, f143])).
fof(f143, plain, ! [X1, X0] : (? [X3] : ((sdtasdt0(X1, X3) = X0) & aInteger0(X3)) => ((sdtasdt0(X1, sK11(X0, X1)) = X0) & aInteger0(sK11(X0, X1)))), introduced(choice_axiom, [])).
fof(f142, plain, ! [X0] : (! [X1] : ((aDivisorOf0(X1, X0) | ! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1)) & ((? [X3] : ((sdtasdt0(X1, X3) = X0) & aInteger0(X3)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aDivisorOf0(X1, X0))) | ~ aInteger0(X0)), inference(rectify, [], [f141])).
fof(f141, plain, ! [X0] : (! [X1] : ((aDivisorOf0(X1, X0) | ! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1)) & ((? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aDivisorOf0(X1, X0))) | ~ aInteger0(X0)), inference(flattening, [], [f140])).
fof(f140, plain, ! [X0] : (! [X1] : ((aDivisorOf0(X1, X0) | (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))) & ((? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aDivisorOf0(X1, X0))) | ~ aInteger0(X0)), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : (! [X1] : (aDivisorOf0(X1, X0) <=> (? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1))) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0] : (aInteger0(X0) => ! [X1] : (aDivisorOf0(X1, X0) <=> (? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mDivisor)).
fof(f369353, plain, (~ aDivisorOf0(xp, xp) | (~ spl38_1 | spl38_9 | ~ spl38_1628)), inference(backward_demodulation, [], [f325850, f369352])).
fof(f369352, plain, ((xp = sdtpldt0(smndt0(sz10), sdtpldt0(sz10, xp))) | (~ spl38_1 | ~ spl38_1628)), inference(backward_demodulation, [], [f325820, f369351])).
fof(f369351, plain, ((xp = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ spl38_1), inference(forward_demodulation, [], [f369350, f596])).
fof(f596, plain, (xp = sdtpldt0(xp, sz00)), inference(resolution, [], [f246, f434])).
fof(f246, plain, ! [X0] : (~ aInteger0(X0) | (sdtpldt0(X0, sz00) = X0)), inference(cnf_transformation, [], [f69])).
fof(f69, plain, ! [X0] : (((sdtpldt0(sz00, X0) = X0) & (sdtpldt0(X0, sz00) = X0)) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0] : (aInteger0(X0) => ((sdtpldt0(sz00, X0) = X0) & (sdtpldt0(X0, sz00) = X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mAddZero)).
fof(f369350, plain, ((sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10)) = sdtpldt0(xp, sz00)) | ~ spl38_1), inference(forward_demodulation, [], [f369241, f265818])).
fof(f265818, plain, (sdtpldt0(sz10, xp) = sdtpldt0(xp, sz10)), inference(resolution, [], [f1270, f240])).
fof(f1270, plain, ! [X13] : (~ aInteger0(X13) | (sdtpldt0(X13, xp) = sdtpldt0(xp, X13))), inference(resolution, [], [f245, f434])).
fof(f245, plain, ! [X0, X1] : (~ aInteger0(X1) | (sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f67])).
fof(f67, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | (~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => (sdtpldt0(X0, X1) = sdtpldt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mAddComm)).
fof(f369241, plain, ((sdtpldt0(xp, sz00) = sdtpldt0(sdtpldt0(xp, sz10), smndt0(sz10))) | ~ spl38_1), inference(resolution, [], [f24141, f434])).
fof(f24141, plain, (! [X1] : (~ aInteger0(X1) | (sdtpldt0(X1, sz00) = sdtpldt0(sdtpldt0(X1, sz10), smndt0(sz10)))) | ~ spl38_1), inference(forward_demodulation, [], [f24036, f692])).
fof(f692, plain, (sz00 = sdtpldt0(sz10, smndt0(sz10))), inference(resolution, [], [f248, f240])).
fof(f248, plain, ! [X0] : (~ aInteger0(X0) | (sz00 = sdtpldt0(X0, smndt0(X0)))), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (((sz00 = sdtpldt0(smndt0(X0), X0)) & (sz00 = sdtpldt0(X0, smndt0(X0)))) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : (aInteger0(X0) => ((sz00 = sdtpldt0(smndt0(X0), X0)) & (sz00 = sdtpldt0(X0, smndt0(X0))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mAddNeg)).
fof(f24036, plain, (! [X1] : ((sdtpldt0(X1, sdtpldt0(sz10, smndt0(sz10))) = sdtpldt0(sdtpldt0(X1, sz10), smndt0(sz10))) | ~ aInteger0(X1)) | ~ spl38_1), inference(resolution, [], [f3308, f240])).
fof(f3308, plain, (! [X8, X7] : (~ aInteger0(X8) | (sdtpldt0(X7, sdtpldt0(X8, smndt0(sz10))) = sdtpldt0(sdtpldt0(X7, X8), smndt0(sz10))) | ~ aInteger0(X7)) | ~ spl38_1), inference(resolution, [], [f244, f484])).
fof(f484, plain, (aInteger0(smndt0(sz10)) | ~ spl38_1), inference(avatar_component_clause, [], [f483])).
fof(f483, plain, (spl38_1 <=> aInteger0(smndt0(sz10))), introduced(avatar_definition, [new_symbols(naming, [spl38_1])])).
fof(f244, plain, ! [X2, X0, X1] : (~ aInteger0(X2) | (sdtpldt0(X0, sdtpldt0(X1, X2)) = sdtpldt0(sdtpldt0(X0, X1), X2)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ! [X0, X1, X2] : ((sdtpldt0(X0, sdtpldt0(X1, X2)) = sdtpldt0(sdtpldt0(X0, X1), X2)) | ~ aInteger0(X2) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f65])).
fof(f65, plain, ! [X0, X1, X2] : ((sdtpldt0(X0, sdtpldt0(X1, X2)) = sdtpldt0(sdtpldt0(X0, X1), X2)) | (~ aInteger0(X2) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0, X1, X2] : ((aInteger0(X2) & aInteger0(X1) & aInteger0(X0)) => (sdtpldt0(X0, sdtpldt0(X1, X2)) = sdtpldt0(sdtpldt0(X0, X1), X2))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mAddAsso)).
fof(f325820, plain, ((sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10)) = sdtpldt0(smndt0(sz10), sdtpldt0(sz10, xp))) | (~ spl38_1 | ~ spl38_1628)), inference(resolution, [], [f325603, f59210])).
fof(f59210, plain, (! [X2] : (~ aInteger0(X2) | (sdtpldt0(X2, smndt0(sz10)) = sdtpldt0(smndt0(sz10), X2))) | ~ spl38_1), inference(resolution, [], [f484, f245])).
fof(f325603, plain, (aInteger0(sdtpldt0(sz10, xp)) | ~ spl38_1628), inference(avatar_component_clause, [], [f325602])).
fof(f325602, plain, (spl38_1628 <=> aInteger0(sdtpldt0(sz10, xp))), introduced(avatar_definition, [new_symbols(naming, [spl38_1628])])).
fof(f325850, plain, (~ aDivisorOf0(xp, sdtpldt0(smndt0(sz10), sdtpldt0(sz10, xp))) | (~ spl38_1 | spl38_9 | ~ spl38_1628)), inference(backward_demodulation, [], [f554, f325820])).
fof(f554, plain, (~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | spl38_9), inference(avatar_component_clause, [], [f552])).
fof(f552, plain, (spl38_9 <=> aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10)))), introduced(avatar_definition, [new_symbols(naming, [spl38_9])])).
fof(f369331, plain, (~ spl38_1 | spl38_13 | ~ spl38_133), inference(avatar_contradiction_clause, [], [f369330])).
fof(f369330, plain, ($false | (~ spl38_1 | spl38_13 | ~ spl38_133)), inference(subsumption_resolution, [], [f369328, f2129])).
fof(f2129, plain, (aDivisorOf0(xp, smndt0(xp)) | ~ spl38_1), inference(subsumption_resolution, [], [f2128, f434])).
fof(f2128, plain, (aDivisorOf0(xp, smndt0(xp)) | ~ aInteger0(xp) | ~ spl38_1), inference(subsumption_resolution, [], [f2127, f435])).
fof(f2127, plain, (aDivisorOf0(xp, smndt0(xp)) | (sz00 = xp) | ~ aInteger0(xp) | ~ spl38_1), inference(subsumption_resolution, [], [f2024, f484])).
fof(f2024, plain, (aDivisorOf0(xp, smndt0(xp)) | ~ aInteger0(smndt0(sz10)) | (sz00 = xp) | ~ aInteger0(xp)), inference(superposition, [], [f481, f875])).
fof(f875, plain, (smndt0(xp) = sdtasdt0(xp, smndt0(sz10))), inference(resolution, [], [f259, f434])).
fof(f259, plain, ! [X0] : (~ aInteger0(X0) | (smndt0(X0) = sdtasdt0(X0, smndt0(sz10)))), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : (((smndt0(X0) = sdtasdt0(X0, smndt0(sz10))) & (smndt0(X0) = sdtasdt0(smndt0(sz10), X0))) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (aInteger0(X0) => ((smndt0(X0) = sdtasdt0(X0, smndt0(sz10))) & (smndt0(X0) = sdtasdt0(smndt0(sz10), X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mMulMinOne)).
fof(f369328, plain, (~ aDivisorOf0(xp, smndt0(xp)) | (~ spl38_1 | spl38_13 | ~ spl38_133)), inference(backward_demodulation, [], [f573, f369327])).
fof(f369327, plain, ((smndt0(xp) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | (~ spl38_1 | ~ spl38_133)), inference(forward_demodulation, [], [f369326, f2587])).
fof(f2587, plain, ((smndt0(xp) = sdtpldt0(smndt0(xp), sz00)) | ~ spl38_133), inference(resolution, [], [f2571, f246])).
fof(f2571, plain, (aInteger0(smndt0(xp)) | ~ spl38_133), inference(avatar_component_clause, [], [f2570])).
fof(f2570, plain, (spl38_133 <=> aInteger0(smndt0(xp))), introduced(avatar_definition, [new_symbols(naming, [spl38_133])])).
fof(f369326, plain, ((sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10)) = sdtpldt0(smndt0(xp), sz00)) | (~ spl38_1 | ~ spl38_133)), inference(forward_demodulation, [], [f369222, f5585])).
fof(f5585, plain, ((sdtpldt0(sz10, smndt0(xp)) = sdtpldt0(smndt0(xp), sz10)) | ~ spl38_133), inference(resolution, [], [f1265, f2571])).
fof(f1265, plain, ! [X3] : (~ aInteger0(X3) | (sdtpldt0(X3, sz10) = sdtpldt0(sz10, X3))), inference(resolution, [], [f245, f240])).
fof(f369222, plain, ((sdtpldt0(smndt0(xp), sz00) = sdtpldt0(sdtpldt0(smndt0(xp), sz10), smndt0(sz10))) | (~ spl38_1 | ~ spl38_133)), inference(resolution, [], [f24141, f2571])).
fof(f573, plain, (~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | spl38_13), inference(avatar_component_clause, [], [f571])).
fof(f571, plain, (spl38_13 <=> aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10)))), introduced(avatar_definition, [new_symbols(naming, [spl38_13])])).
fof(f325634, plain, spl38_1628, inference(avatar_contradiction_clause, [], [f325633])).
fof(f325633, plain, ($false | spl38_1628), inference(subsumption_resolution, [], [f325632, f240])).
fof(f325632, plain, (~ aInteger0(sz10) | spl38_1628), inference(subsumption_resolution, [], [f325631, f434])).
fof(f325631, plain, (~ aInteger0(xp) | ~ aInteger0(sz10) | spl38_1628), inference(resolution, [], [f325604, f242])).
fof(f242, plain, ! [X0, X1] : (aInteger0(sdtpldt0(X0, X1)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ! [X0, X1] : (aInteger0(sdtpldt0(X0, X1)) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f61])).
fof(f61, plain, ! [X0, X1] : (aInteger0(sdtpldt0(X0, X1)) | (~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => aInteger0(sdtpldt0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mIntPlus)).
fof(f325604, plain, (~ aInteger0(sdtpldt0(sz10, xp)) | spl38_1628), inference(avatar_component_clause, [], [f325602])).
fof(f2585, plain, spl38_133, inference(avatar_contradiction_clause, [], [f2584])).
fof(f2584, plain, ($false | spl38_133), inference(subsumption_resolution, [], [f2583, f434])).
fof(f2583, plain, (~ aInteger0(xp) | spl38_133), inference(resolution, [], [f2572, f241])).
fof(f241, plain, ! [X0] : (aInteger0(smndt0(X0)) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ! [X0] : (aInteger0(smndt0(X0)) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : (aInteger0(X0) => aInteger0(smndt0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', mIntNeg)).
fof(f2572, plain, (~ aInteger0(smndt0(xp)) | spl38_133), inference(avatar_component_clause, [], [f2570])).
fof(f583, plain, spl38_1, inference(avatar_split_clause, [], [f580, f483])).
fof(f580, plain, aInteger0(smndt0(sz10)), inference(resolution, [], [f507, f502])).
fof(f502, plain, aElementOf0(smndt0(sz10), cS2076), inference(forward_demodulation, [], [f479, f386])).
fof(f386, plain, (stldt0(sbsmnsldt0(xS)) = cS2076), inference(cnf_transformation, [], [f214])).
fof(f214, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & (((aElementOf0(X2, sK29(X2)) & aElementOf0(sK29(X2), xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29])], [f212, f213])).
fof(f213, plain, ! [X2] : (? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) => (aElementOf0(X2, sK29(X2)) & aElementOf0(sK29(X2), xS))), introduced(choice_axiom, [])).
fof(f212, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f211])).
fof(f211, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & ((smndt0(sz10) = X0) | (sz10 = X0) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f210])).
fof(f210, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : ((aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) & (((smndt0(sz10) = X0) | (sz10 = X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1))) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | (! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2))) & ((? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(nnf_transformation, [], [f56])).
fof(f56, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))) & ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : (aElementOf0(X2, sbsmnsldt0(xS)) <=> (? [X3] : (aElementOf0(X2, X3) & aElementOf0(X3, xS)) & aInteger0(X2))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f43])).
fof(f43, plain, ((stldt0(sbsmnsldt0(xS)) = cS2076) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))) & ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', m__2079)).
fof(f479, plain, aElementOf0(smndt0(sz10), stldt0(sbsmnsldt0(xS))), inference(equality_resolution, [], [f385])).
fof(f385, plain, ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) | ~ (smndt0(sz10) = X0)), inference(cnf_transformation, [], [f214])).
fof(f507, plain, ! [X1] : (~ aElementOf0(X1, cS2076) | aInteger0(X1)), inference(forward_demodulation, [], [f380, f386])).
fof(f380, plain, ! [X1] : (aInteger0(X1) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f214])).
fof(f574, plain, (spl38_6 | ~ spl38_13), inference(avatar_split_clause, [], [f460, f571, f538])).
fof(f538, plain, (spl38_6 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl38_6])])).
fof(f460, plain, (~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | sP10), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ((~ aElementOf0(sdtpldt0(sz10, smndt0(xp)), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, smndt0(xp)), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) & ! [X0] : (~ (sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | ~ aInteger0(X0))) | sP10), inference(definition_folding, [], [f122, e138])).
fof(f138, plain, ((~ aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & ! [X1] : (~ (sdtasdt0(xp, X1) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ aInteger0(X1))) | ~ sP10), inference(usedef, [], [e138])).
fof(e138, plain, (sP10 <=> (~ aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & ! [X1] : (~ (sdtasdt0(xp, X1) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ aInteger0(X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f122, plain, ((~ aElementOf0(sdtpldt0(sz10, smndt0(xp)), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, smndt0(xp)), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) & ! [X0] : (~ (sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | ~ aInteger0(X0))) | (~ aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & ! [X1] : (~ (sdtasdt0(xp, X1) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ aInteger0(X1)))), inference(ennf_transformation, [], [f59])).
fof(f59, plain, ~ ((aElementOf0(sdtpldt0(sz10, smndt0(xp)), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, smndt0(xp)), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | ? [X0] : ((sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) & aInteger0(X0))) & (aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ? [X1] : ((sdtasdt0(xp, X1) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & aInteger0(X1)))), inference(rectify, [], [f48])).
fof(f48, plain, ~ ((aElementOf0(sdtpldt0(sz10, smndt0(xp)), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, smndt0(xp)), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | ? [X0] : ((sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) & aInteger0(X0))) & (aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ? [X0] : ((sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & aInteger0(X0)))), inference(negated_conjecture, [], [f47])).
fof(f47, plain, ~ ((aElementOf0(sdtpldt0(sz10, smndt0(xp)), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, smndt0(xp)), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) | ? [X0] : ((sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, smndt0(xp)), smndt0(sz10))) & aInteger0(X0))) & (aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) | sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) | aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ? [X0] : ((sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & aInteger0(X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM452+6.p', m__)).
fof(f555, plain, (~ spl38_6 | ~ spl38_9), inference(avatar_split_clause, [], [f456, f552, f538])).
fof(f456, plain, (~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ sP10), inference(cnf_transformation, [], [f238])).
fof(f238, plain, ((~ aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & ! [X0] : (~ (sdtasdt0(xp, X0) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ aInteger0(X0))) | ~ sP10), inference(rectify, [], [f237])).
fof(f237, plain, ((~ aElementOf0(sdtpldt0(sz10, xp), szAzrzSzezqlpdtcmdtrp0(sz10, xp)) & ~ sdteqdtlpzmzozddtrp0(sdtpldt0(sz10, xp), sz10, xp) & ~ aDivisorOf0(xp, sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) & ! [X1] : (~ (sdtasdt0(xp, X1) = sdtpldt0(sdtpldt0(sz10, xp), smndt0(sz10))) | ~ aInteger0(X1))) | ~ sP10), inference(nnf_transformation, [], [f138])).