fof(f5385, plain, $false, inference(avatar_sat_refutation, [], [f435, f440, f444, f475, f479, f483, f492, f550, f559, f561, f588, f599, f607, f611, f665, f1032, f1098, f1145, f5298, f5302, f5384])).
fof(f5384, plain, (~ spl37_14 | ~ spl37_158), inference(avatar_contradiction_clause, [], [f5383])).
fof(f5383, plain, ($false | (~ spl37_14 | ~ spl37_158)), inference(subsumption_resolution, [], [f5382, f322])).
fof(f322, plain, aInteger0(xq), inference(cnf_transformation, [], [f41])).
fof(f41, plain, (~ (sz00 = xq) & aInteger0(xq) & aInteger0(xa)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM444+6.p', m__1962)).
fof(f5382, plain, (~ aInteger0(xq) | (~ spl37_14 | ~ spl37_158)), inference(subsumption_resolution, [], [f5379, f323])).
fof(f323, plain, ~ (sz00 = xq), inference(cnf_transformation, [], [f41])).
fof(f5379, plain, ((sz00 = xq) | ~ aInteger0(xq) | (~ spl37_14 | ~ spl37_158)), inference(resolution, [], [f5297, f2890])).
fof(f2890, plain, (! [X0] : (~ sP9(sK31(X0)) | (sz00 = X0) | ~ aInteger0(X0)) | ~ spl37_14), inference(resolution, [], [f474, f365])).
fof(f365, plain, ! [X0] : (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP9(X0)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0] : ((aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP7 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP9(X0)), inference(rectify, [], [f203])).
fof(f203, plain, ! [X1] : ((aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP7 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP9(X1)), inference(nnf_transformation, [], [f118])).
fof(f118, plain, ! [X1] : ((aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP7 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP9(X1)), inference(usedef, [], [e118])).
fof(e118, plain, ! [X1] : (sP9(X1) <=> (aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP7 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f474, plain, (! [X1] : (~ aElementOf0(sK31(X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aInteger0(X1) | (sz00 = X1)) | ~ spl37_14), inference(avatar_component_clause, [], [f473])).
fof(f473, plain, (spl37_14 <=> ! [X1] : (~ aElementOf0(sK31(X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aInteger0(X1) | (sz00 = X1))), introduced(avatar_definition, [new_symbols(naming, [spl37_14])])).
fof(f5297, plain, (sP9(sK31(xq)) | ~ spl37_158), inference(avatar_component_clause, [], [f5295])).
fof(f5295, plain, (spl37_158 <=> sP9(sK31(xq))), introduced(avatar_definition, [new_symbols(naming, [spl37_158])])).
fof(f5302, plain, (~ spl37_15 | ~ spl37_16 | spl37_157), inference(avatar_contradiction_clause, [], [f5301])).
fof(f5301, plain, ($false | (~ spl37_15 | ~ spl37_16 | spl37_157)), inference(subsumption_resolution, [], [f5300, f322])).
fof(f5300, plain, (~ aInteger0(xq) | (~ spl37_15 | ~ spl37_16 | spl37_157)), inference(subsumption_resolution, [], [f5299, f323])).
fof(f5299, plain, ((sz00 = xq) | ~ aInteger0(xq) | (~ spl37_15 | ~ spl37_16 | spl37_157)), inference(resolution, [], [f5293, f985])).
fof(f985, plain, (! [X0] : (aInteger0(sK31(X0)) | (sz00 = X0) | ~ aInteger0(X0)) | (~ spl37_15 | ~ spl37_16)), inference(subsumption_resolution, [], [f984, f482])).
fof(f482, plain, (! [X1] : (sP11(X1, sK30) | ~ aInteger0(X1) | (sz00 = X1)) | ~ spl37_16), inference(avatar_component_clause, [], [f481])).
fof(f481, plain, (spl37_16 <=> ! [X1] : (sP11(X1, sK30) | ~ aInteger0(X1) | (sz00 = X1))), introduced(avatar_definition, [new_symbols(naming, [spl37_16])])).
fof(f984, plain, (! [X0] : (~ aInteger0(X0) | (sz00 = X0) | aInteger0(sK31(X0)) | ~ sP11(X0, sK30)) | ~ spl37_15), inference(resolution, [], [f478, f346])).
fof(f346, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)) | aInteger0(X2) | ~ sP11(X0, X1)), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ! [X0, X1] : (! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)) | (~ sdteqdtlpzmzozddtrp0(X2, X1, X0) & ~ aDivisorOf0(X0, sdtpldt0(X2, smndt0(X1))) & ! [X3] : (~ (sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, X3)) | ~ aInteger0(X3))) | ~ aInteger0(X2)) & ((sdteqdtlpzmzozddtrp0(X2, X1, X0) & aDivisorOf0(X0, sdtpldt0(X2, smndt0(X1))) & ((sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, sK33(X0, X1, X2))) & aInteger0(sK33(X0, X1, X2))) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)))) | ~ sP11(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK33])], [f196, f197])).
fof(f197, plain, ! [X2, X1, X0] : (? [X4] : ((sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, X4)) & aInteger0(X4)) => ((sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, sK33(X0, X1, X2))) & aInteger0(sK33(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f196, plain, ! [X0, X1] : (! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)) | (~ sdteqdtlpzmzozddtrp0(X2, X1, X0) & ~ aDivisorOf0(X0, sdtpldt0(X2, smndt0(X1))) & ! [X3] : (~ (sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, X3)) | ~ aInteger0(X3))) | ~ aInteger0(X2)) & ((sdteqdtlpzmzozddtrp0(X2, X1, X0) & aDivisorOf0(X0, sdtpldt0(X2, smndt0(X1))) & ? [X4] : ((sdtpldt0(X2, smndt0(X1)) = sdtasdt0(X0, X4)) & aInteger0(X4)) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)))) | ~ sP11(X0, X1)), inference(rectify, [], [f195])).
fof(f195, plain, ! [X14, X13] : (! [X15] : ((aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) | (~ sdteqdtlpzmzozddtrp0(X15, X13, X14) & ~ aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ! [X16] : (~ (sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) | ~ aInteger0(X16))) | ~ aInteger0(X15)) & ((sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)) | ~ aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)))) | ~ sP11(X14, X13)), inference(nnf_transformation, [], [f120])).
fof(f120, plain, ! [X14, X13] : (! [X15] : ((aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) | (~ sdteqdtlpzmzozddtrp0(X15, X13, X14) & ~ aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ! [X16] : (~ (sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) | ~ aInteger0(X16))) | ~ aInteger0(X15)) & ((sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)) | ~ aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)))) | ~ sP11(X14, X13)), inference(usedef, [], [e120])).
fof(e120, plain, ! [X14, X13] : (sP11(X14, X13) <=> ! [X15] : ((aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) | (~ sdteqdtlpzmzozddtrp0(X15, X13, X14) & ~ aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ! [X16] : (~ (sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) | ~ aInteger0(X16))) | ~ aInteger0(X15)) & ((sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)) | ~ aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f478, plain, (! [X1] : (aElementOf0(sK31(X1), szAzrzSzezqlpdtcmdtrp0(sK30, X1)) | ~ aInteger0(X1) | (sz00 = X1)) | ~ spl37_15), inference(avatar_component_clause, [], [f477])).
fof(f477, plain, (spl37_15 <=> ! [X1] : (aElementOf0(sK31(X1), szAzrzSzezqlpdtcmdtrp0(sK30, X1)) | ~ aInteger0(X1) | (sz00 = X1))), introduced(avatar_definition, [new_symbols(naming, [spl37_15])])).
fof(f5293, plain, (~ aInteger0(sK31(xq)) | spl37_157), inference(avatar_component_clause, [], [f5291])).
fof(f5291, plain, (spl37_157 <=> aInteger0(sK31(xq))), introduced(avatar_definition, [new_symbols(naming, [spl37_157])])).
fof(f5298, plain, (~ spl37_157 | spl37_158 | ~ spl37_15 | ~ spl37_16 | ~ spl37_18 | ~ spl37_41 | ~ spl37_42), inference(avatar_split_clause, [], [f5289, f609, f605, f489, f481, f477, f5295, f5291])).
fof(f489, plain, (spl37_18 <=> aElementOf0(sK30, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(avatar_definition, [new_symbols(naming, [spl37_18])])).
fof(f605, plain, (spl37_41 <=> ! [X0] : (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(avatar_definition, [new_symbols(naming, [spl37_41])])).
fof(f609, plain, (spl37_42 <=> ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(avatar_definition, [new_symbols(naming, [spl37_42])])).
fof(f5289, plain, (sP9(sK31(xq)) | ~ aInteger0(sK31(xq)) | (~ spl37_15 | ~ spl37_16 | ~ spl37_18 | ~ spl37_41 | ~ spl37_42)), inference(subsumption_resolution, [], [f5288, f667])).
fof(f667, plain, (aInteger0(sK30) | (~ spl37_18 | ~ spl37_42)), inference(resolution, [], [f610, f491])).
fof(f491, plain, (aElementOf0(sK30, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ spl37_18), inference(avatar_component_clause, [], [f489])).
fof(f610, plain, (! [X0] : (~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | aInteger0(X0)) | ~ spl37_42), inference(avatar_component_clause, [], [f609])).
fof(f5288, plain, (sP9(sK31(xq)) | ~ aInteger0(sK31(xq)) | ~ aInteger0(sK30) | (~ spl37_15 | ~ spl37_16 | ~ spl37_18 | ~ spl37_41)), inference(subsumption_resolution, [], [f5287, f943])).
fof(f943, plain, (~ sP8(sK30) | (~ spl37_18 | ~ spl37_41)), inference(resolution, [], [f936, f368])).
fof(f368, plain, ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ sP8(X0)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0] : ((~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP6 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP8(X0)), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ! [X0] : ((~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP6 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP8(X0)), inference(usedef, [], [e117])).
fof(e117, plain, ! [X0] : (sP8(X0) <=> (~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & sP6 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f936, plain, (~ aElementOf0(sK30, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ spl37_18 | ~ spl37_41)), inference(resolution, [], [f606, f491])).
fof(f606, plain, (! [X0] : (~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ spl37_41), inference(avatar_component_clause, [], [f605])).
fof(f5287, plain, (sP9(sK31(xq)) | sP8(sK30) | ~ aInteger0(sK31(xq)) | ~ aInteger0(sK30) | (~ spl37_15 | ~ spl37_16)), inference(subsumption_resolution, [], [f5286, f323])).
fof(f5286, plain, ((sz00 = xq) | sP9(sK31(xq)) | sP8(sK30) | ~ aInteger0(sK31(xq)) | ~ aInteger0(sK30) | (~ spl37_15 | ~ spl37_16)), inference(subsumption_resolution, [], [f5273, f322])).
fof(f5273, plain, (~ aInteger0(xq) | (sz00 = xq) | sP9(sK31(xq)) | sP8(sK30) | ~ aInteger0(sK31(xq)) | ~ aInteger0(sK30) | (~ spl37_15 | ~ spl37_16)), inference(resolution, [], [f1135, f388])).
fof(f388, plain, ! [X2, X1] : (~ sdteqdtlpzmzozddtrp0(X2, X1, xq) | sP9(X2) | sP8(X1) | ~ aInteger0(X2) | ~ aInteger0(X1)), inference(cnf_transformation, [], [f216])).
fof(f216, plain, (((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP13 & ! [X0] : ((aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aInteger0(X0)) & ((~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X0)) | ~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP12) | sP14) & ! [X1, X2] : (sP9(X2) | (~ sdteqdtlpzmzozddtrp0(X2, X1, xq) & ~ aDivisorOf0(xq, sdtpldt0(X2, smndt0(X1))) & ! [X3] : (~ (sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(X1))) | ~ aInteger0(X3))) | sP8(X1) | ~ aInteger0(X2) | ~ aInteger0(X1))), inference(rectify, [], [f215])).
fof(f215, plain, (((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP13 & ! [X12] : ((aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aInteger0(X12)) & ((~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12)) | ~ aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP12) | sP14) & ! [X0, X1] : (sP9(X1) | (~ sdteqdtlpzmzozddtrp0(X1, X0, xq) & ~ aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) & ! [X2] : (~ (sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) | ~ aInteger0(X2))) | sP8(X0) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(flattening, [], [f214])).
fof(f214, plain, (((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP13 & ! [X12] : ((aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | (aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aInteger0(X12))) & ((~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12)) | ~ aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP12) | sP14) & ! [X0, X1] : (sP9(X1) | (~ sdteqdtlpzmzozddtrp0(X1, X0, xq) & ~ aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) & ! [X2] : (~ (sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) | ~ aInteger0(X2))) | sP8(X0) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(nnf_transformation, [], [f124])).
fof(f124, plain, (((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP13 & ! [X12] : (aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & sP12) | sP14) & ! [X0, X1] : (sP9(X1) | (~ sdteqdtlpzmzozddtrp0(X1, X0, xq) & ~ aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) & ! [X2] : (~ (sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) | ~ aInteger0(X2))) | sP8(X0) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(definition_folding, [], [f105, e123, e122, e121, e120, e119, e118, e117, e116, e115])).
fof(f115, plain, (! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X3, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP6), inference(usedef, [], [e115])).
fof(e115, plain, (sP6 <=> ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X3, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f116, plain, (! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X6, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP7), inference(usedef, [], [e116])).
fof(e116, plain, (sP7 <=> ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X6, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f119, plain, (! [X19] : ((aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X19, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ! [X20] : (~ (sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X20)) | ~ aInteger0(X20))) | ~ aInteger0(X19)) & ((sdteqdtlpzmzozddtrp0(X19, xa, xq) & aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ? [X21] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X21)) & aInteger0(X21)) & aInteger0(X19)) | ~ aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP10), inference(usedef, [], [e119])).
fof(e119, plain, (sP10 <=> ! [X19] : ((aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X19, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ! [X20] : (~ (sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X20)) | ~ aInteger0(X20))) | ~ aInteger0(X19)) & ((sdteqdtlpzmzozddtrp0(X19, xa, xq) & aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ? [X21] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X21)) & aInteger0(X21)) & aInteger0(X19)) | ~ aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f121, plain, (! [X9] : ((aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X9, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ! [X10] : (~ (sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X10)) | ~ aInteger0(X10))) | ~ aInteger0(X9)) & ((sdteqdtlpzmzozddtrp0(X9, xa, xq) & aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ? [X11] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X11)) & aInteger0(X11)) & aInteger0(X9)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP12), inference(usedef, [], [e121])).
fof(e121, plain, (sP12 <=> ! [X9] : ((aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X9, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ! [X10] : (~ (sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X10)) | ~ aInteger0(X10))) | ~ aInteger0(X9)) & ((sdteqdtlpzmzozddtrp0(X9, xa, xq) & aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ? [X11] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X11)) & aInteger0(X11)) & aInteger0(X9)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f122, plain, (? [X13] : (! [X14] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X18] : (~ aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14))) & sP11(X14, X13) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14))) | (sz00 = X14) | ~ aInteger0(X14)) & aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP13), inference(usedef, [], [e122])).
fof(e122, plain, (sP13 <=> ? [X13] : (! [X14] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X18] : (~ aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14))) & sP11(X14, X13) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14))) | (sz00 = X14) | ~ aInteger0(X14)) & aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f123, plain, ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X23] : (~ aElementOf0(X23, cS1395) & aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X22] : (aElementOf0(X22, cS1395) <=> aInteger0(X22)) & aSet0(cS1395) & sP10 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP14), inference(usedef, [], [e123])).
fof(e123, plain, (sP14 <=> (~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X23] : (~ aElementOf0(X23, cS1395) & aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X22] : (aElementOf0(X22, cS1395) <=> aInteger0(X22)) & aSet0(cS1395) & sP10 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f105, plain, (((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X13] : (! [X14] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X18] : (~ aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14))) & ! [X15] : ((aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) | (~ sdteqdtlpzmzozddtrp0(X15, X13, X14) & ~ aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ! [X16] : (~ (sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) | ~ aInteger0(X16))) | ~ aInteger0(X15)) & ((sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)) | ~ aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14))) | (sz00 = X14) | ~ aInteger0(X14)) & aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & ! [X12] : (aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X9] : ((aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X9, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ! [X10] : (~ (sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X10)) | ~ aInteger0(X10))) | ~ aInteger0(X9)) & ((sdteqdtlpzmzozddtrp0(X9, xa, xq) & aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ? [X11] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X11)) & aInteger0(X11)) & aInteger0(X9)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq))))) | (~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X23] : (~ aElementOf0(X23, cS1395) & aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X22] : (aElementOf0(X22, cS1395) <=> aInteger0(X22)) & aSet0(cS1395) & ! [X19] : ((aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X19, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ! [X20] : (~ (sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X20)) | ~ aInteger0(X20))) | ~ aInteger0(X19)) & ((sdteqdtlpzmzozddtrp0(X19, xa, xq) & aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ? [X21] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X21)) & aInteger0(X21)) & aInteger0(X19)) | ~ aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & ! [X0, X1] : ((aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X6, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | (~ sdteqdtlpzmzozddtrp0(X1, X0, xq) & ~ aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) & ! [X2] : (~ (sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) | ~ aInteger0(X2))) | (~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X3, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(flattening, [], [f104])).
fof(f104, plain, ((((~ isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ((~ isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X13] : (! [X14] : (((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X18] : (~ aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14)))) & (! [X15] : ((aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) | ((~ sdteqdtlpzmzozddtrp0(X15, X13, X14) & ~ aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ! [X16] : (~ (sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) | ~ aInteger0(X16))) | ~ aInteger0(X15))) & ((sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)) | ~ aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14)))) | (sz00 = X14) | ~ aInteger0(X14)) & aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) & (! [X12] : (aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))))) & ! [X9] : ((aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((~ sdteqdtlpzmzozddtrp0(X9, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ! [X10] : (~ (sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X10)) | ~ aInteger0(X10))) | ~ aInteger0(X9))) & ((sdteqdtlpzmzozddtrp0(X9, xa, xq) & aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ? [X11] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X11)) & aInteger0(X11)) & aInteger0(X9)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq))))) | (((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X23] : (~ aElementOf0(X23, cS1395) & aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & (! [X22] : (aElementOf0(X22, cS1395) <=> aInteger0(X22)) & aSet0(cS1395))) & (! [X19] : ((aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((~ sdteqdtlpzmzozddtrp0(X19, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ! [X20] : (~ (sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X20)) | ~ aInteger0(X20))) | ~ aInteger0(X19))) & ((sdteqdtlpzmzozddtrp0(X19, xa, xq) & aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ? [X21] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X21)) & aInteger0(X21)) & aInteger0(X19)) | ~ aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) & ! [X0, X1] : (((aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((~ sdteqdtlpzmzozddtrp0(X6, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6))) & ((sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ((~ sdteqdtlpzmzozddtrp0(X1, X0, xq) & ~ aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) & ! [X2] : (~ (sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) | ~ aInteger0(X2))) | ((~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((~ sdteqdtlpzmzozddtrp0(X3, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3))) & ((sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))))) | (~ aInteger0(X1) | ~ aInteger0(X0)))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ~ (! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => (((sdteqdtlpzmzozddtrp0(X1, X0, xq) | aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) | ? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) & aInteger0(X2))) & ((! [X3] : ((((sdteqdtlpzmzozddtrp0(X3, xa, xq) | aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) | ? [X4] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) & aInteger0(X4))) & aInteger0(X3)) => aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => (aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X6] : ((((sdteqdtlpzmzozddtrp0(X6, xa, xq) | aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) | ? [X7] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) & aInteger0(X7))) & aInteger0(X6)) => aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => ((! [X9] : ((((sdteqdtlpzmzozddtrp0(X9, xa, xq) | aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) | ? [X10] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X10)) & aInteger0(X10))) & aInteger0(X9)) => aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X9, xa, xq) & aDivisorOf0(xq, sdtpldt0(X9, smndt0(xa))) & ? [X11] : ((sdtpldt0(X9, smndt0(xa)) = sdtasdt0(xq, X11)) & aInteger0(X11)) & aInteger0(X9)))) => (isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((! [X12] : (aElementOf0(X12, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X12, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X12))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) => (isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X13] : (aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ? [X14] : (((! [X15] : ((((sdteqdtlpzmzozddtrp0(X15, X13, X14) | aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) | ? [X16] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X16)) & aInteger0(X16))) & aInteger0(X15)) => aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14))) & (aElementOf0(X15, szAzrzSzezqlpdtcmdtrp0(X13, X14)) => (sdteqdtlpzmzozddtrp0(X15, X13, X14) & aDivisorOf0(X14, sdtpldt0(X15, smndt0(X13))) & ? [X17] : ((sdtpldt0(X15, smndt0(X13)) = sdtasdt0(X14, X17)) & aInteger0(X17)) & aInteger0(X15)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X18] : (aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14)) => aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))))) & ~ (sz00 = X14) & aInteger0(X14))))))) & ((! [X19] : ((((sdteqdtlpzmzozddtrp0(X19, xa, xq) | aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) | ? [X20] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X20)) & aInteger0(X20))) & aInteger0(X19)) => aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X19, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X19, xa, xq) & aDivisorOf0(xq, sdtpldt0(X19, smndt0(xa))) & ? [X21] : ((sdtpldt0(X19, smndt0(xa)) = sdtasdt0(xq, X21)) & aInteger0(X21)) & aInteger0(X19)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ((! [X22] : (aElementOf0(X22, cS1395) <=> aInteger0(X22)) & aSet0(cS1395)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) | ! [X23] : (aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => aElementOf0(X23, cS1395))))))), inference(rectify, [], [f43])).
fof(f43, plain, ~ (! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => (((sdteqdtlpzmzozddtrp0(X1, X0, xq) | aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) | ? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) & aInteger0(X2))) & ((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, xa, xq) | aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) | ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X2, xa, xq) & aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) & ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => (aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, xa, xq) | aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) | ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X2, xa, xq) & aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) & ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => ((! [X0] : ((((sdteqdtlpzmzozddtrp0(X0, xa, xq) | aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) | ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1))) & aInteger0(X0)) => aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1)) & aInteger0(X0)))) => (isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((! [X0] : (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X0))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) => (isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X0] : (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))))) & ~ (sz00 = X1) & aInteger0(X1))))))) & ((! [X0] : ((((sdteqdtlpzmzozddtrp0(X0, xa, xq) | aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) | ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1))) & aInteger0(X0)) => aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1)) & aInteger0(X0)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) | ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => aElementOf0(X0, cS1395))))))), inference(negated_conjecture, [], [f42])).
fof(f42, plain, ~ (! [X0, X1] : ((aInteger0(X1) & aInteger0(X0)) => (((sdteqdtlpzmzozddtrp0(X1, X0, xq) | aDivisorOf0(xq, sdtpldt0(X1, smndt0(X0))) | ? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X1, smndt0(X0))) & aInteger0(X2))) & ((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, xa, xq) | aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) | ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X2, xa, xq) & aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) & ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => (aElementOf0(X1, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ~ aElementOf0(X1, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, xa, xq) | aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) | ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X2, xa, xq) & aDivisorOf0(xq, sdtpldt0(X2, smndt0(xa))) & ? [X3] : ((sdtasdt0(xq, X3) = sdtpldt0(X2, smndt0(xa))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))) => ((! [X0] : ((((sdteqdtlpzmzozddtrp0(X0, xa, xq) | aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) | ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1))) & aInteger0(X0)) => aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1)) & aInteger0(X0)))) => (isClosed0(szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ((! [X0] : (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) <=> (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) & aInteger0(X0))) & aSet0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) => (isOpen0(stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X0] : (aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))))) & ~ (sz00 = X1) & aInteger0(X1))))))) & ((! [X0] : ((((sdteqdtlpzmzozddtrp0(X0, xa, xq) | aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) | ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1))) & aInteger0(X0)) => aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => (sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X1] : ((sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X1)) & aInteger0(X0)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) | ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) => aElementOf0(X0, cS1395))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM444+6.p', m__)).
fof(f1135, plain, (! [X11] : (sdteqdtlpzmzozddtrp0(sK31(X11), sK30, X11) | ~ aInteger0(X11) | (sz00 = X11)) | (~ spl37_15 | ~ spl37_16)), inference(subsumption_resolution, [], [f1134, f482])).
fof(f1134, plain, (! [X11] : (sdteqdtlpzmzozddtrp0(sK31(X11), sK30, X11) | ~ sP11(X11, sK30) | ~ aInteger0(X11) | (sz00 = X11)) | ~ spl37_15), inference(resolution, [], [f350, f478])).
fof(f350, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X1, X0)) | sdteqdtlpzmzozddtrp0(X2, X1, X0) | ~ sP11(X0, X1)), inference(cnf_transformation, [], [f198])).
fof(f1145, plain, (~ spl37_32 | ~ spl37_63), inference(avatar_contradiction_clause, [], [f1144])).
fof(f1144, plain, ($false | (~ spl37_32 | ~ spl37_63)), inference(subsumption_resolution, [], [f1141, f1048])).
fof(f1048, plain, (sP9(xa) | ~ spl37_32), inference(resolution, [], [f1035, f321])).
fof(f321, plain, aInteger0(xa), inference(cnf_transformation, [], [f41])).
fof(f1035, plain, (! [X0] : (~ aInteger0(X0) | sP9(X0)) | ~ spl37_32), inference(subsumption_resolution, [], [f1034, f322])).
fof(f1034, plain, (! [X0] : (sP9(X0) | ~ aInteger0(X0) | ~ aInteger0(xq)) | ~ spl37_32), inference(subsumption_resolution, [], [f1033, f323])).
fof(f1033, plain, (! [X0] : (sP9(X0) | ~ aInteger0(X0) | (sz00 = xq) | ~ aInteger0(xq)) | ~ spl37_32), inference(subsumption_resolution, [], [f1010, f554])).
fof(f554, plain, (! [X0] : ~ sP8(X0) | ~ spl37_32), inference(avatar_component_clause, [], [f553])).
fof(f553, plain, (spl37_32 <=> ! [X0] : ~ sP8(X0)), introduced(avatar_definition, [new_symbols(naming, [spl37_32])])).
fof(f1010, plain, ! [X0] : (sP9(X0) | sP8(X0) | ~ aInteger0(X0) | (sz00 = xq) | ~ aInteger0(xq)), inference(duplicate_literal_removal, [], [f1009])).
fof(f1009, plain, ! [X0] : (sP9(X0) | sP8(X0) | ~ aInteger0(X0) | ~ aInteger0(X0) | (sz00 = xq) | ~ aInteger0(xq) | ~ aInteger0(X0)), inference(resolution, [], [f388, f246])).
fof(f246, plain, ! [X0, X1] : (sdteqdtlpzmzozddtrp0(X0, X0, X1) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, ! [X0, X1] : (sdteqdtlpzmzozddtrp0(X0, X0, X1) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f77])).
fof(f77, plain, ! [X0, X1] : (sdteqdtlpzmzozddtrp0(X0, X0, X1) | ((sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X1] : ((~ (sz00 = X1) & aInteger0(X1) & aInteger0(X0)) => sdteqdtlpzmzozddtrp0(X0, X0, X1)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM444+6.p', mEquModRef)).
fof(f1141, plain, (~ sP9(xa) | ~ spl37_63), inference(resolution, [], [f999, f364])).
fof(f364, plain, ! [X0] : (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ sP9(X0)), inference(cnf_transformation, [], [f204])).
fof(f999, plain, (aElementOf0(xa, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ spl37_63), inference(avatar_component_clause, [], [f998])).
fof(f998, plain, (spl37_63 <=> aElementOf0(xa, szAzrzSzezqlpdtcmdtrp0(xa, xq))), introduced(avatar_definition, [new_symbols(naming, [spl37_63])])).
fof(f1098, plain, (~ spl37_20 | spl37_63), inference(avatar_contradiction_clause, [], [f1097])).
fof(f1097, plain, ($false | (~ spl37_20 | spl37_63)), inference(subsumption_resolution, [], [f1096, f321])).
fof(f1096, plain, (~ aInteger0(xa) | (~ spl37_20 | spl37_63)), inference(subsumption_resolution, [], [f1095, f322])).
fof(f1095, plain, (~ aInteger0(xq) | ~ aInteger0(xa) | (~ spl37_20 | spl37_63)), inference(subsumption_resolution, [], [f1094, f323])).
fof(f1094, plain, ((sz00 = xq) | ~ aInteger0(xq) | ~ aInteger0(xa) | (~ spl37_20 | spl37_63)), inference(resolution, [], [f1008, f246])).
fof(f1008, plain, (~ sdteqdtlpzmzozddtrp0(xa, xa, xq) | (~ spl37_20 | spl37_63)), inference(subsumption_resolution, [], [f1006, f321])).
fof(f1006, plain, (~ aInteger0(xa) | ~ sdteqdtlpzmzozddtrp0(xa, xa, xq) | (~ spl37_20 | spl37_63)), inference(resolution, [], [f1000, f499])).
fof(f499, plain, (! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aInteger0(X0) | ~ sdteqdtlpzmzozddtrp0(X0, xa, xq)) | ~ spl37_20), inference(avatar_component_clause, [], [f498])).
fof(f498, plain, (spl37_20 <=> ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aInteger0(X0) | ~ sdteqdtlpzmzozddtrp0(X0, xa, xq))), introduced(avatar_definition, [new_symbols(naming, [spl37_20])])).
fof(f1000, plain, (~ aElementOf0(xa, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | spl37_63), inference(avatar_component_clause, [], [f998])).
fof(f1032, plain, (~ spl37_30 | ~ spl37_32), inference(avatar_contradiction_clause, [], [f1015])).
fof(f1015, plain, ($false | (~ spl37_30 | ~ spl37_32)), inference(resolution, [], [f1014, f217])).
fof(f217, plain, aInteger0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aInteger0(sz00), file('/home/ubuntu/library/tptp/Problems/NUM/NUM444+6.p', mIntZero)).
fof(f1014, plain, (! [X0] : ~ aInteger0(X0) | (~ spl37_30 | ~ spl37_32)), inference(subsumption_resolution, [], [f1013, f322])).
fof(f1013, plain, (! [X0] : (~ aInteger0(X0) | ~ aInteger0(xq)) | (~ spl37_30 | ~ spl37_32)), inference(subsumption_resolution, [], [f1012, f323])).
fof(f1012, plain, (! [X0] : (~ aInteger0(X0) | (sz00 = xq) | ~ aInteger0(xq)) | (~ spl37_30 | ~ spl37_32)), inference(subsumption_resolution, [], [f1011, f554])).
fof(f1011, plain, (! [X0] : (sP8(X0) | ~ aInteger0(X0) | (sz00 = xq) | ~ aInteger0(xq)) | ~ spl37_30), inference(subsumption_resolution, [], [f1010, f545])).
fof(f545, plain, (! [X0] : ~ sP9(X0) | ~ spl37_30), inference(avatar_component_clause, [], [f544])).
fof(f544, plain, (spl37_30 <=> ! [X0] : ~ sP9(X0)), introduced(avatar_definition, [new_symbols(naming, [spl37_30])])).
fof(f665, plain, (spl37_5 | ~ spl37_6 | ~ spl37_7 | ~ spl37_27), inference(avatar_contradiction_clause, [], [f664])).
fof(f664, plain, ($false | (spl37_5 | ~ spl37_6 | ~ spl37_7 | ~ spl37_27)), inference(subsumption_resolution, [], [f662, f621])).
fof(f621, plain, (~ aInteger0(sK29) | (spl37_5 | ~ spl37_7)), inference(resolution, [], [f443, f434])).
fof(f434, plain, (~ aElementOf0(sK29, cS1395) | spl37_5), inference(avatar_component_clause, [], [f432])).
fof(f432, plain, (spl37_5 <=> aElementOf0(sK29, cS1395)), introduced(avatar_definition, [new_symbols(naming, [spl37_5])])).
fof(f443, plain, (! [X1] : (aElementOf0(X1, cS1395) | ~ aInteger0(X1)) | ~ spl37_7), inference(avatar_component_clause, [], [f442])).
fof(f442, plain, (spl37_7 <=> ! [X1] : (aElementOf0(X1, cS1395) | ~ aInteger0(X1))), introduced(avatar_definition, [new_symbols(naming, [spl37_7])])).
fof(f662, plain, (aInteger0(sK29) | (~ spl37_6 | ~ spl37_27)), inference(resolution, [], [f527, f439])).
fof(f439, plain, (aElementOf0(sK29, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ spl37_6), inference(avatar_component_clause, [], [f437])).
fof(f437, plain, (spl37_6 <=> aElementOf0(sK29, szAzrzSzezqlpdtcmdtrp0(xa, xq))), introduced(avatar_definition, [new_symbols(naming, [spl37_6])])).
fof(f527, plain, (! [X0] : (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | aInteger0(X0)) | ~ spl37_27), inference(avatar_component_clause, [], [f526])).
fof(f526, plain, (spl37_27 <=> ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(avatar_definition, [new_symbols(naming, [spl37_27])])).
fof(f611, plain, (spl37_3 | spl37_42), inference(avatar_split_clause, [], [f391, f609, f423])).
fof(f423, plain, (spl37_3 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl37_3])])).
fof(f391, plain, ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | sP14), inference(cnf_transformation, [], [f216])).
fof(f607, plain, (spl37_3 | spl37_41), inference(avatar_split_clause, [], [f392, f605, f423])).
fof(f392, plain, ! [X0] : (~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | sP14), inference(cnf_transformation, [], [f216])).
fof(f599, plain, (spl37_3 | spl37_12), inference(avatar_split_clause, [], [f394, f465, f423])).
fof(f465, plain, (spl37_12 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl37_12])])).
fof(f394, plain, (sP13 | sP14), inference(cnf_transformation, [], [f216])).
fof(f588, plain, (~ spl37_33 | spl37_27), inference(avatar_split_clause, [], [f378, f526, f556])).
fof(f556, plain, (spl37_33 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl37_33])])).
fof(f378, plain, ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ sP6), inference(cnf_transformation, [], [f213])).
fof(f213, plain, (! [X0] : ((aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X0, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ! [X1] : (~ (sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) | ~ aInteger0(X1))) | ~ aInteger0(X0)) & ((sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ((sdtpldt0(X0, smndt0(xa)) = sdtasdt0(xq, sK36(X0))) & aInteger0(sK36(X0))) & aInteger0(X0)) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36])], [f211, f212])).
fof(f212, plain, ! [X0] : (? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X2)) => ((sdtpldt0(X0, smndt0(xa)) = sdtasdt0(xq, sK36(X0))) & aInteger0(sK36(X0)))), introduced(choice_axiom, [])).
fof(f211, plain, (! [X0] : ((aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X0, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ! [X1] : (~ (sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) | ~ aInteger0(X1))) | ~ aInteger0(X0)) & ((sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X2)) & aInteger0(X0)) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP6), inference(rectify, [], [f210])).
fof(f210, plain, (! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X3, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, xa, xq) & aDivisorOf0(xq, sdtpldt0(X3, smndt0(xa))) & ? [X5] : ((sdtpldt0(X3, smndt0(xa)) = sdtasdt0(xq, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP6), inference(nnf_transformation, [], [f115])).
fof(f561, plain, (~ spl37_31 | spl37_20), inference(avatar_split_clause, [], [f377, f498, f547])).
fof(f547, plain, (spl37_31 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl37_31])])).
fof(f377, plain, ! [X0] : (aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ sdteqdtlpzmzozddtrp0(X0, xa, xq) | ~ aInteger0(X0) | ~ sP7), inference(cnf_transformation, [], [f209])).
fof(f209, plain, (! [X0] : ((aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X0, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ! [X1] : (~ (sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) | ~ aInteger0(X1))) | ~ aInteger0(X0)) & ((sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ((sdtpldt0(X0, smndt0(xa)) = sdtasdt0(xq, sK35(X0))) & aInteger0(sK35(X0))) & aInteger0(X0)) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP7), inference(skolemisation, [status(esa), new_symbols(skolem, [sK35])], [f207, f208])).
fof(f208, plain, ! [X0] : (? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X2)) => ((sdtpldt0(X0, smndt0(xa)) = sdtasdt0(xq, sK35(X0))) & aInteger0(sK35(X0)))), introduced(choice_axiom, [])).
fof(f207, plain, (! [X0] : ((aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X0, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ! [X1] : (~ (sdtasdt0(xq, X1) = sdtpldt0(X0, smndt0(xa))) | ~ aInteger0(X1))) | ~ aInteger0(X0)) & ((sdteqdtlpzmzozddtrp0(X0, xa, xq) & aDivisorOf0(xq, sdtpldt0(X0, smndt0(xa))) & ? [X2] : ((sdtasdt0(xq, X2) = sdtpldt0(X0, smndt0(xa))) & aInteger0(X2)) & aInteger0(X0)) | ~ aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP7), inference(rectify, [], [f206])).
fof(f206, plain, (! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | (~ sdteqdtlpzmzozddtrp0(X6, xa, xq) & ~ aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, xa, xq) & aDivisorOf0(xq, sdtpldt0(X6, smndt0(xa))) & ? [X8] : ((sdtpldt0(X6, smndt0(xa)) = sdtasdt0(xq, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP7), inference(nnf_transformation, [], [f116])).
fof(f559, plain, (spl37_32 | spl37_33), inference(avatar_split_clause, [], [f367, f556, f553])).
fof(f367, plain, ! [X0] : (sP6 | ~ sP8(X0)), inference(cnf_transformation, [], [f205])).
fof(f550, plain, (spl37_30 | spl37_31), inference(avatar_split_clause, [], [f363, f547, f544])).
fof(f363, plain, ! [X0] : (sP7 | ~ sP9(X0)), inference(cnf_transformation, [], [f204])).
fof(f492, plain, (~ spl37_12 | spl37_18), inference(avatar_split_clause, [], [f332, f489, f465])).
fof(f332, plain, (aElementOf0(sK30, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP13), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sK30, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & (~ aElementOf0(sK31(X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(sK31(X1), szAzrzSzezqlpdtcmdtrp0(sK30, X1))) & sP11(X1, sK30) & aSet0(szAzrzSzezqlpdtcmdtrp0(sK30, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(sK30, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP13), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31])], [f187, f189, f188])).
fof(f188, plain, (? [X0] : (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X2] : (~ aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP11(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) => (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sK30, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X2] : (~ aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sK30, X1))) & sP11(X1, sK30) & aSet0(szAzrzSzezqlpdtcmdtrp0(sK30, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(sK30, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))))), introduced(choice_axiom, [])).
fof(f189, plain, ! [X1] : (? [X2] : (~ aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sK30, X1))) => (~ aElementOf0(sK31(X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(sK31(X1), szAzrzSzezqlpdtcmdtrp0(sK30, X1)))), introduced(choice_axiom, [])).
fof(f187, plain, (? [X0] : (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X2] : (~ aElementOf0(X2, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP11(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(X0, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP13), inference(rectify, [], [f186])).
fof(f186, plain, (? [X13] : (! [X14] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X13, X14), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ? [X18] : (~ aElementOf0(X18, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) & aElementOf0(X18, szAzrzSzezqlpdtcmdtrp0(X13, X14))) & sP11(X14, X13) & aSet0(szAzrzSzezqlpdtcmdtrp0(X13, X14))) | (sz00 = X14) | ~ aInteger0(X14)) & aElementOf0(X13, stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq)))) | ~ sP13), inference(nnf_transformation, [], [f122])).
fof(f483, plain, (~ spl37_12 | spl37_16), inference(avatar_split_clause, [], [f334, f481, f465])).
fof(f334, plain, ! [X1] : (sP11(X1, sK30) | (sz00 = X1) | ~ aInteger0(X1) | ~ sP13), inference(cnf_transformation, [], [f190])).
fof(f479, plain, (~ spl37_12 | spl37_15), inference(avatar_split_clause, [], [f335, f477, f465])).
fof(f335, plain, ! [X1] : (aElementOf0(sK31(X1), szAzrzSzezqlpdtcmdtrp0(sK30, X1)) | (sz00 = X1) | ~ aInteger0(X1) | ~ sP13), inference(cnf_transformation, [], [f190])).
fof(f475, plain, (~ spl37_12 | spl37_14), inference(avatar_split_clause, [], [f336, f473, f465])).
fof(f336, plain, ! [X1] : (~ aElementOf0(sK31(X1), stldt0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | (sz00 = X1) | ~ aInteger0(X1) | ~ sP13), inference(cnf_transformation, [], [f190])).
fof(f444, plain, (~ spl37_3 | spl37_7), inference(avatar_split_clause, [], [f328, f442, f423])).
fof(f328, plain, ! [X1] : (aElementOf0(X1, cS1395) | ~ aInteger0(X1) | ~ sP14), inference(cnf_transformation, [], [f185])).
fof(f185, plain, ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & (~ aElementOf0(sK29, cS1395) & aElementOf0(sK29, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP10 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP14), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29])], [f183, f184])).
fof(f184, plain, (? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) => (~ aElementOf0(sK29, cS1395) & aElementOf0(sK29, szAzrzSzezqlpdtcmdtrp0(xa, xq)))), introduced(choice_axiom, [])).
fof(f183, plain, ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP10 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP14), inference(rectify, [], [f182])).
fof(f182, plain, ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(xa, xq), cS1395) & ? [X23] : (~ aElementOf0(X23, cS1395) & aElementOf0(X23, szAzrzSzezqlpdtcmdtrp0(xa, xq))) & ! [X22] : ((aElementOf0(X22, cS1395) | ~ aInteger0(X22)) & (aInteger0(X22) | ~ aElementOf0(X22, cS1395))) & aSet0(cS1395) & sP10 & aSet0(szAzrzSzezqlpdtcmdtrp0(xa, xq))) | ~ sP14), inference(nnf_transformation, [], [f123])).
fof(f440, plain, (~ spl37_3 | spl37_6), inference(avatar_split_clause, [], [f329, f437, f423])).
fof(f329, plain, (aElementOf0(sK29, szAzrzSzezqlpdtcmdtrp0(xa, xq)) | ~ sP14), inference(cnf_transformation, [], [f185])).
fof(f435, plain, (~ spl37_3 | ~ spl37_5), inference(avatar_split_clause, [], [f330, f432, f423])).
fof(f330, plain, (~ aElementOf0(sK29, cS1395) | ~ sP14), inference(cnf_transformation, [], [f185])).