fof(f6839, plain, $false, inference(avatar_sat_refutation, [], [f393, f414, f419, f420, f427, f5636, f6205, f6817, f6834])).
fof(f6834, plain, (spl31_169 | ~ spl31_7), inference(avatar_split_clause, [], [f6833, f411, f5642])).
fof(f5642, plain, (spl31_169 <=> aInteger0(sK28)), introduced(avatar_definition, [new_symbols(naming, [spl31_169])])).
fof(f411, plain, (spl31_7 <=> (smndt0(sz10) = sK28)), introduced(avatar_definition, [new_symbols(naming, [spl31_7])])).
fof(f6833, plain, (aInteger0(sK28) | ~ spl31_7), inference(subsumption_resolution, [], [f6630, f207])).
fof(f207, plain, aInteger0(sz10), inference(cnf_transformation, [], [f3])).
fof(f3, plain, aInteger0(sz10), file('/home/ubuntu/library/tptp/Problems/NUM/NUM448+5.p', mIntOne)).
fof(f6630, plain, (aInteger0(sK28) | ~ aInteger0(sz10) | ~ spl31_7), inference(superposition, [], [f208, f412])).
fof(f412, plain, ((smndt0(sz10) = sK28) | ~ spl31_7), inference(avatar_component_clause, [], [f411])).
fof(f208, plain, ! [X0] : (aInteger0(smndt0(X0)) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ! [X0] : (aInteger0(smndt0(X0)) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : (aInteger0(X0) => aInteger0(smndt0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM448+5.p', mIntNeg)).
fof(f6817, plain, (~ spl31_2 | spl31_6 | ~ spl31_7 | ~ spl31_169), inference(avatar_contradiction_clause, [], [f6816])).
fof(f6816, plain, ($false | (~ spl31_2 | spl31_6 | ~ spl31_7 | ~ spl31_169)), inference(subsumption_resolution, [], [f6813, f6285])).
fof(f6285, plain, (~ sP8(sK28) | (~ spl31_2 | ~ spl31_7)), inference(backward_demodulation, [], [f435, f412])).
fof(f435, plain, (~ sP8(smndt0(sz10)) | ~ spl31_2), inference(subsumption_resolution, [], [f434, f344])).
fof(f344, plain, ! [X0] : (isPrime0(sK26(X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ! [X0] : ((isPrime0(sK26(X0)) & aDivisorOf0(sK26(X0), X0) & ((sdtasdt0(sK26(X0), sK27(X0)) = X0) & aInteger0(sK27(X0))) & ~ (sz00 = sK26(X0)) & aInteger0(sK26(X0))) | ~ sP8(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27])], [f195, f197, f196])).
fof(f196, plain, ! [X0] : (? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)) => (isPrime0(sK26(X0)) & aDivisorOf0(sK26(X0), X0) & ? [X2] : ((sdtasdt0(sK26(X0), X2) = X0) & aInteger0(X2)) & ~ (sz00 = sK26(X0)) & aInteger0(sK26(X0)))), introduced(choice_axiom, [])).
fof(f197, plain, ! [X0] : (? [X2] : ((sdtasdt0(sK26(X0), X2) = X0) & aInteger0(X2)) => ((sdtasdt0(sK26(X0), sK27(X0)) = X0) & aInteger0(sK27(X0)))), introduced(choice_axiom, [])).
fof(f195, plain, ! [X0] : (? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ sP8(X0)), inference(rectify, [], [f194])).
fof(f194, plain, ! [X0] : (? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5)) | ~ sP8(X0)), inference(nnf_transformation, [], [f124])).
fof(f124, plain, ! [X0] : (? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5)) | ~ sP8(X0)), inference(usedef, [], [e124])).
fof(e124, plain, ! [X0] : (sP8(X0) <=> ? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f434, plain, (~ isPrime0(sK26(smndt0(sz10))) | ~ sP8(smndt0(sz10)) | ~ spl31_2), inference(resolution, [], [f392, f343])).
fof(f343, plain, ! [X0] : (aDivisorOf0(sK26(X0), X0) | ~ sP8(X0)), inference(cnf_transformation, [], [f198])).
fof(f392, plain, (! [X2] : (~ aDivisorOf0(X2, smndt0(sz10)) | ~ isPrime0(X2)) | ~ spl31_2), inference(avatar_component_clause, [], [f391])).
fof(f391, plain, (spl31_2 <=> ! [X2] : (~ isPrime0(X2) | ~ aDivisorOf0(X2, smndt0(sz10)))), introduced(avatar_definition, [new_symbols(naming, [spl31_2])])).
fof(f6813, plain, (sP8(sK28) | (spl31_6 | ~ spl31_169)), inference(resolution, [], [f6635, f424])).
fof(f424, plain, ! [X5] : (~ aElementOf0(X5, sbsmnsldt0(xS)) | sP8(X5)), inference(subsumption_resolution, [], [f346, f354])).
fof(f354, plain, ! [X2] : (~ aElementOf0(X2, sbsmnsldt0(xS)) | aInteger0(X2)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & (((~ (smndt0(sz10) = sK28) & ~ (sz10 = sK28)) | ~ aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))) & ((smndt0(sz10) = sK28) | (sz10 = sK28) | aElementOf0(sK28, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & (((aElementOf0(X2, sK29(X2)) & aElementOf0(sK29(X2), xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X5] : ((((aElementOf0(X5, sbsmnsldt0(xS)) & (aElementOf0(X5, sK30(X5)) & aElementOf0(sK30(X5), xS))) | ! [X7] : (~ isPrime0(X7) | (~ aDivisorOf0(X7, X5) & (! [X8] : (~ (sdtasdt0(X7, X8) = X5) | ~ aInteger0(X8)) | (sz00 = X7) | ~ aInteger0(X7))))) & (sP8(X5) | (~ aElementOf0(X5, sbsmnsldt0(xS)) & ! [X9] : (~ aElementOf0(X5, X9) | ~ aElementOf0(X9, xS))))) | ~ aInteger0(X5))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29, sK30])], [f201, f204, f203, f202])).
fof(f202, plain, (? [X0] : (((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ((smndt0(sz10) = X0) | (sz10 = X0) | aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) => (((~ (smndt0(sz10) = sK28) & ~ (sz10 = sK28)) | ~ aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))) & ((smndt0(sz10) = sK28) | (sz10 = sK28) | aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))))), introduced(choice_axiom, [])).
fof(f203, plain, ! [X2] : (? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) => (aElementOf0(X2, sK29(X2)) & aElementOf0(sK29(X2), xS))), introduced(choice_axiom, [])).
fof(f204, plain, ! [X5] : (? [X6] : (aElementOf0(X5, X6) & aElementOf0(X6, xS)) => (aElementOf0(X5, sK30(X5)) & aElementOf0(sK30(X5), xS))), introduced(choice_axiom, [])).
fof(f201, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X0] : (((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ~ aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ((smndt0(sz10) = X0) | (sz10 = X0) | aElementOf0(X0, stldt0(sbsmnsldt0(xS))))) & ! [X1] : ((aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sbsmnsldt0(xS)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X2] : ((aElementOf0(X2, sbsmnsldt0(xS)) | ! [X3] : (~ aElementOf0(X2, X3) | ~ aElementOf0(X3, xS)) | ~ aInteger0(X2)) & ((? [X4] : (aElementOf0(X2, X4) & aElementOf0(X4, xS)) & aInteger0(X2)) | ~ aElementOf0(X2, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X5] : ((((aElementOf0(X5, sbsmnsldt0(xS)) & ? [X6] : (aElementOf0(X5, X6) & aElementOf0(X6, xS))) | ! [X7] : (~ isPrime0(X7) | (~ aDivisorOf0(X7, X5) & (! [X8] : (~ (sdtasdt0(X7, X8) = X5) | ~ aInteger0(X8)) | (sz00 = X7) | ~ aInteger0(X7))))) & (sP8(X5) | (~ aElementOf0(X5, sbsmnsldt0(xS)) & ! [X9] : (~ aElementOf0(X5, X9) | ~ aElementOf0(X9, xS))))) | ~ aInteger0(X5))), inference(rectify, [], [f200])).
fof(f200, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X10] : (((~ (smndt0(sz10) = X10) & ~ (sz10 = X10)) | ~ aElementOf0(X10, stldt0(sbsmnsldt0(xS)))) & ((smndt0(sz10) = X10) | (sz10 = X10) | aElementOf0(X10, stldt0(sbsmnsldt0(xS))))) & ! [X9] : ((aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | aElementOf0(X9, sbsmnsldt0(xS)) | ~ aInteger0(X9)) & ((~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X7] : ((aElementOf0(X7, sbsmnsldt0(xS)) | ! [X8] : (~ aElementOf0(X7, X8) | ~ aElementOf0(X8, xS)) | ~ aInteger0(X7)) & ((? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7)) | ~ aElementOf0(X7, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X0] : ((((aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS))) | ! [X1] : (~ isPrime0(X1) | (~ aDivisorOf0(X1, X0) & (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))))) & (sP8(X0) | (~ aElementOf0(X0, sbsmnsldt0(xS)) & ! [X4] : (~ aElementOf0(X0, X4) | ~ aElementOf0(X4, xS))))) | ~ aInteger0(X0))), inference(flattening, [], [f199])).
fof(f199, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X10] : (((~ (smndt0(sz10) = X10) & ~ (sz10 = X10)) | ~ aElementOf0(X10, stldt0(sbsmnsldt0(xS)))) & (((smndt0(sz10) = X10) | (sz10 = X10)) | aElementOf0(X10, stldt0(sbsmnsldt0(xS))))) & ! [X9] : ((aElementOf0(X9, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X9, sbsmnsldt0(xS)) | ~ aInteger0(X9))) & ((~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9)) | ~ aElementOf0(X9, stldt0(sbsmnsldt0(xS))))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X7] : ((aElementOf0(X7, sbsmnsldt0(xS)) | (! [X8] : (~ aElementOf0(X7, X8) | ~ aElementOf0(X8, xS)) | ~ aInteger0(X7))) & ((? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7)) | ~ aElementOf0(X7, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS)) & ! [X0] : ((((aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS))) | ! [X1] : (~ isPrime0(X1) | (~ aDivisorOf0(X1, X0) & (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))))) & (sP8(X0) | (~ aElementOf0(X0, sbsmnsldt0(xS)) & ! [X4] : (~ aElementOf0(X0, X4) | ~ aElementOf0(X4, xS))))) | ~ aInteger0(X0))), inference(nnf_transformation, [], [f125])).
fof(f125, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X10] : ~ (aElementOf0(X10, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X10) | (sz10 = X10))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & ! [X0] : ((((aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS))) | ! [X1] : (~ isPrime0(X1) | (~ aDivisorOf0(X1, X0) & (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))))) & (sP8(X0) | (~ aElementOf0(X0, sbsmnsldt0(xS)) & ! [X4] : (~ aElementOf0(X0, X4) | ~ aElementOf0(X4, xS))))) | ~ aInteger0(X0))), inference(definition_folding, [], [f111, e124])).
fof(f111, plain, (~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X10] : ~ (aElementOf0(X10, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X10) | (sz10 = X10))) & ! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9))) & aSet0(stldt0(sbsmnsldt0(xS))) & ! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)) & ! [X0] : ((((aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS))) | ! [X1] : (~ isPrime0(X1) | (~ aDivisorOf0(X1, X0) & (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))))) & (? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5)) | (~ aElementOf0(X0, sbsmnsldt0(xS)) & ! [X4] : (~ aElementOf0(X0, X4) | ~ aElementOf0(X4, xS))))) | ~ aInteger0(X0))), inference(flattening, [], [f110])).
fof(f110, plain, ((((~ (stldt0(sbsmnsldt0(xS)) = cS2076) & ? [X10] : ~ (aElementOf0(X10, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X10) | (sz10 = X10)))) & (! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9))) & aSet0(stldt0(sbsmnsldt0(xS))))) & (! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS)))) & ! [X0] : ((((aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS))) | ! [X1] : (~ isPrime0(X1) | (~ aDivisorOf0(X1, X0) & (! [X2] : (~ (sdtasdt0(X1, X2) = X0) | ~ aInteger0(X2)) | (sz00 = X1) | ~ aInteger0(X1))))) & (? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5)) | (~ aElementOf0(X0, sbsmnsldt0(xS)) & ! [X4] : (~ aElementOf0(X0, X4) | ~ aElementOf0(X4, xS))))) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ~ (! [X0] : (aInteger0(X0) => ((? [X1] : (isPrime0(X1) & (aDivisorOf0(X1, X0) | (? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)))) => (aElementOf0(X0, sbsmnsldt0(xS)) & ? [X3] : (aElementOf0(X0, X3) & aElementOf0(X3, xS)))) & ((aElementOf0(X0, sbsmnsldt0(xS)) | ? [X4] : (aElementOf0(X0, X4) & aElementOf0(X4, xS))) => ? [X5] : (isPrime0(X5) & aDivisorOf0(X5, X0) & ? [X6] : ((sdtasdt0(X5, X6) = X0) & aInteger0(X6)) & ~ (sz00 = X5) & aInteger0(X5))))) => ((! [X7] : (aElementOf0(X7, sbsmnsldt0(xS)) <=> (? [X8] : (aElementOf0(X7, X8) & aElementOf0(X8, xS)) & aInteger0(X7))) & aSet0(sbsmnsldt0(xS))) => ((! [X9] : (aElementOf0(X9, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X9, sbsmnsldt0(xS)) & aInteger0(X9))) & aSet0(stldt0(sbsmnsldt0(xS)))) => ((stldt0(sbsmnsldt0(xS)) = cS2076) | ! [X10] : (aElementOf0(X10, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X10) | (sz10 = X10))))))), inference(rectify, [], [f44])).
fof(f44, plain, ~ (! [X0] : (aInteger0(X0) => ((? [X1] : (isPrime0(X1) & (aDivisorOf0(X1, X0) | (? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)))) => (aElementOf0(X0, sbsmnsldt0(xS)) & ? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)))) & ((aElementOf0(X0, sbsmnsldt0(xS)) | ? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS))) => ? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1))))) => ((! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))) => ((! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & aSet0(stldt0(sbsmnsldt0(xS)))) => ((stldt0(sbsmnsldt0(xS)) = cS2076) | ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))))))), inference(negated_conjecture, [], [f43])).
fof(f43, plain, ~ (! [X0] : (aInteger0(X0) => ((? [X1] : (isPrime0(X1) & (aDivisorOf0(X1, X0) | (? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1)))) => (aElementOf0(X0, sbsmnsldt0(xS)) & ? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)))) & ((aElementOf0(X0, sbsmnsldt0(xS)) | ? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS))) => ? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aInteger0(X2)) & ~ (sz00 = X1) & aInteger0(X1))))) => ((! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))) => ((! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) & aSet0(stldt0(sbsmnsldt0(xS)))) => ((stldt0(sbsmnsldt0(xS)) = cS2076) | ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> ((smndt0(sz10) = X0) | (sz10 = X0))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM448+5.p', m__)).
fof(f346, plain, ! [X5] : (sP8(X5) | ~ aElementOf0(X5, sbsmnsldt0(xS)) | ~ aInteger0(X5)), inference(cnf_transformation, [], [f205])).
fof(f6635, plain, (aElementOf0(sK28, sbsmnsldt0(xS)) | (spl31_6 | ~ spl31_169)), inference(subsumption_resolution, [], [f6633, f5643])).
fof(f5643, plain, (aInteger0(sK28) | ~ spl31_169), inference(avatar_component_clause, [], [f5642])).
fof(f6633, plain, (aElementOf0(sK28, sbsmnsldt0(xS)) | ~ aInteger0(sK28) | spl31_6), inference(resolution, [], [f409, f361])).
fof(f361, plain, ! [X1] : (aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)), inference(cnf_transformation, [], [f205])).
fof(f409, plain, (~ aElementOf0(sK28, stldt0(sbsmnsldt0(xS))) | spl31_6), inference(avatar_component_clause, [], [f407])).
fof(f407, plain, (spl31_6 <=> aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))), introduced(avatar_definition, [new_symbols(naming, [spl31_6])])).
fof(f6205, plain, (spl31_6 | ~ spl31_8), inference(avatar_contradiction_clause, [], [f6204])).
fof(f6204, plain, ($false | (spl31_6 | ~ spl31_8)), inference(subsumption_resolution, [], [f6201, f432])).
fof(f432, plain, ~ sP8(sz10), inference(subsumption_resolution, [], [f431, f344])).
fof(f431, plain, (~ sP8(sz10) | ~ isPrime0(sK26(sz10))), inference(resolution, [], [f343, f394])).
fof(f394, plain, ! [X2] : (~ aDivisorOf0(X2, sz10) | ~ isPrime0(X2)), inference(subsumption_resolution, [], [f369, f207])).
fof(f369, plain, ! [X2] : (~ isPrime0(X2) | ~ aDivisorOf0(X2, sz10) | ~ aInteger0(sz10)), inference(equality_resolution, [], [f240])).
fof(f240, plain, ! [X2, X0] : (~ (sz10 = X0) | ~ isPrime0(X2) | ~ aDivisorOf0(X2, X0) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ! [X0] : ((((isPrime0(sK10(X0)) & aDivisorOf0(sK10(X0), X0)) | (smndt0(sz10) = X0) | (sz10 = X0)) & ((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ! [X2] : (~ isPrime0(X2) | ~ aDivisorOf0(X2, X0)))) | ~ aInteger0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10])], [f134, f135])).
fof(f135, plain, ! [X0] : (? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) => (isPrime0(sK10(X0)) & aDivisorOf0(sK10(X0), X0))), introduced(choice_axiom, [])).
fof(f134, plain, ! [X0] : (((? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) | (smndt0(sz10) = X0) | (sz10 = X0)) & ((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ! [X2] : (~ isPrime0(X2) | ~ aDivisorOf0(X2, X0)))) | ~ aInteger0(X0)), inference(rectify, [], [f133])).
fof(f133, plain, ! [X0] : (((? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) | (smndt0(sz10) = X0) | (sz10 = X0)) & ((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ! [X1] : (~ isPrime0(X1) | ~ aDivisorOf0(X1, X0)))) | ~ aInteger0(X0)), inference(flattening, [], [f132])).
fof(f132, plain, ! [X0] : (((? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) | ((smndt0(sz10) = X0) | (sz10 = X0))) & ((~ (smndt0(sz10) = X0) & ~ (sz10 = X0)) | ! [X1] : (~ isPrime0(X1) | ~ aDivisorOf0(X1, X0)))) | ~ aInteger0(X0)), inference(nnf_transformation, [], [f87])).
fof(f87, plain, ! [X0] : ((? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) <=> (~ (smndt0(sz10) = X0) & ~ (sz10 = X0))) | ~ aInteger0(X0)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (aInteger0(X0) => (? [X1] : (isPrime0(X1) & aDivisorOf0(X1, X0)) <=> (~ (smndt0(sz10) = X0) & ~ (sz10 = X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM448+5.p', mPrimeDivisor)).
fof(f6201, plain, (sP8(sz10) | (spl31_6 | ~ spl31_8)), inference(resolution, [], [f6198, f424])).
fof(f6198, plain, (aElementOf0(sz10, sbsmnsldt0(xS)) | (spl31_6 | ~ spl31_8)), inference(subsumption_resolution, [], [f6196, f207])).
fof(f6196, plain, (aElementOf0(sz10, sbsmnsldt0(xS)) | ~ aInteger0(sz10) | (spl31_6 | ~ spl31_8)), inference(resolution, [], [f6045, f361])).
fof(f6045, plain, (~ aElementOf0(sz10, stldt0(sbsmnsldt0(xS))) | (spl31_6 | ~ spl31_8)), inference(forward_demodulation, [], [f409, f417])).
fof(f417, plain, ((sz10 = sK28) | ~ spl31_8), inference(avatar_component_clause, [], [f416])).
fof(f416, plain, (spl31_8 <=> (sz10 = sK28)), introduced(avatar_definition, [new_symbols(naming, [spl31_8])])).
fof(f5636, plain, (~ spl31_6 | spl31_7 | spl31_8), inference(avatar_contradiction_clause, [], [f5635])).
fof(f5635, plain, ($false | (~ spl31_6 | spl31_7 | spl31_8)), inference(subsumption_resolution, [], [f5634, f413])).
fof(f413, plain, (~ (smndt0(sz10) = sK28) | spl31_7), inference(avatar_component_clause, [], [f411])).
fof(f5634, plain, ((smndt0(sz10) = sK28) | (~ spl31_6 | spl31_8)), inference(subsumption_resolution, [], [f5633, f1722])).
fof(f1722, plain, (aInteger0(sK28) | ~ spl31_6), inference(resolution, [], [f408, f359])).
fof(f359, plain, ! [X1] : (~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | aInteger0(X1)), inference(cnf_transformation, [], [f205])).
fof(f408, plain, (aElementOf0(sK28, stldt0(sbsmnsldt0(xS))) | ~ spl31_6), inference(avatar_component_clause, [], [f407])).
fof(f5633, plain, (~ aInteger0(sK28) | (smndt0(sz10) = sK28) | (~ spl31_6 | spl31_8)), inference(subsumption_resolution, [], [f5618, f418])).
fof(f418, plain, (~ (sz10 = sK28) | spl31_8), inference(avatar_component_clause, [], [f416])).
fof(f5618, plain, ((sz10 = sK28) | ~ aInteger0(sK28) | (smndt0(sz10) = sK28) | ~ spl31_6), inference(resolution, [], [f1938, f1721])).
fof(f1721, plain, (~ aElementOf0(sK28, sbsmnsldt0(xS)) | ~ spl31_6), inference(resolution, [], [f408, f360])).
fof(f360, plain, ! [X1] : (~ aElementOf0(X1, stldt0(sbsmnsldt0(xS))) | ~ aElementOf0(X1, sbsmnsldt0(xS))), inference(cnf_transformation, [], [f205])).
fof(f1938, plain, ! [X1] : (aElementOf0(X1, sbsmnsldt0(xS)) | (sz10 = X1) | ~ aInteger0(X1) | (smndt0(sz10) = X1)), inference(subsumption_resolution, [], [f1936, f243])).
fof(f243, plain, ! [X0] : (isPrime0(sK10(X0)) | (smndt0(sz10) = X0) | (sz10 = X0) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f136])).
fof(f1936, plain, ! [X1] : ((smndt0(sz10) = X1) | (sz10 = X1) | ~ aInteger0(X1) | ~ isPrime0(sK10(X1)) | aElementOf0(X1, sbsmnsldt0(xS))), inference(duplicate_literal_removal, [], [f1927])).
fof(f1927, plain, ! [X1] : ((smndt0(sz10) = X1) | (sz10 = X1) | ~ aInteger0(X1) | ~ isPrime0(sK10(X1)) | aElementOf0(X1, sbsmnsldt0(xS)) | ~ aInteger0(X1)), inference(resolution, [], [f242, f352])).
fof(f352, plain, ! [X7, X5] : (~ aDivisorOf0(X7, X5) | ~ isPrime0(X7) | aElementOf0(X5, sbsmnsldt0(xS)) | ~ aInteger0(X5)), inference(cnf_transformation, [], [f205])).
fof(f242, plain, ! [X0] : (aDivisorOf0(sK10(X0), X0) | (smndt0(sz10) = X0) | (sz10 = X0) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f136])).
fof(f427, plain, spl31_1, inference(avatar_contradiction_clause, [], [f426])).
fof(f426, plain, ($false | spl31_1), inference(subsumption_resolution, [], [f425, f207])).
fof(f425, plain, (~ aInteger0(sz10) | spl31_1), inference(resolution, [], [f208, f389])).
fof(f389, plain, (~ aInteger0(smndt0(sz10)) | spl31_1), inference(avatar_component_clause, [], [f387])).
fof(f387, plain, (spl31_1 <=> aInteger0(smndt0(sz10))), introduced(avatar_definition, [new_symbols(naming, [spl31_1])])).
fof(f420, plain, (spl31_6 | spl31_8 | spl31_7), inference(avatar_split_clause, [], [f362, f411, f416, f407])).
fof(f362, plain, ((smndt0(sz10) = sK28) | (sz10 = sK28) | aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f205])).
fof(f419, plain, (~ spl31_6 | ~ spl31_8), inference(avatar_split_clause, [], [f363, f416, f407])).
fof(f363, plain, (~ (sz10 = sK28) | ~ aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f205])).
fof(f414, plain, (~ spl31_6 | ~ spl31_7), inference(avatar_split_clause, [], [f364, f411, f407])).
fof(f364, plain, (~ (smndt0(sz10) = sK28) | ~ aElementOf0(sK28, stldt0(sbsmnsldt0(xS)))), inference(cnf_transformation, [], [f205])).
fof(f393, plain, (~ spl31_1 | spl31_2), inference(avatar_split_clause, [], [f368, f391, f387])).
fof(f368, plain, ! [X2] : (~ isPrime0(X2) | ~ aDivisorOf0(X2, smndt0(sz10)) | ~ aInteger0(smndt0(sz10))), inference(equality_resolution, [], [f241])).
fof(f241, plain, ! [X2, X0] : (~ (smndt0(sz10) = X0) | ~ isPrime0(X2) | ~ aDivisorOf0(X2, X0) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f136])).