fof(f1251179, plain, $false, inference(avatar_sat_refutation, [], [f11903, f17541, f17606, f49385, f96911, f288785, f767948, f767955, f768235, f1146581, f1251178])).
fof(f1251178, plain, (spl14_5172 | ~ spl14_8665), inference(avatar_contradiction_clause, [], [f1251177])).
fof(f1251177, plain, ($false | (spl14_5172 | ~ spl14_8665)), inference(subsumption_resolution, [], [f1251176, f232])).
fof(f232, plain, aNaturalNumber0(xn), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (aNaturalNumber0(xp) & aNaturalNumber0(xm) & aNaturalNumber0(xn)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1837)).
fof(f1251176, plain, (~ aNaturalNumber0(xn) | (spl14_5172 | ~ spl14_8665)), inference(subsumption_resolution, [], [f1251175, f265])).
fof(f265, plain, ~ (xn = xr), inference(cnf_transformation, [], [f161])).
fof(f161, plain, (sdtlseqdt0(xr, xn) & ((xn = sdtpldt0(xr, sK12)) & aNaturalNumber0(sK12)) & ~ (xn = xr)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12])], [f44, f160])).
fof(f160, plain, (? [X0] : ((xn = sdtpldt0(xr, X0)) & aNaturalNumber0(X0)) => ((xn = sdtpldt0(xr, sK12)) & aNaturalNumber0(sK12))), introduced(choice_axiom, [])).
fof(f44, plain, (sdtlseqdt0(xr, xn) & ? [X0] : ((xn = sdtpldt0(xr, X0)) & aNaturalNumber0(X0)) & ~ (xn = xr)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1894)).
fof(f1251175, plain, ((xn = xr) | ~ aNaturalNumber0(xn) | (spl14_5172 | ~ spl14_8665)), inference(subsumption_resolution, [], [f1251174, f268])).
fof(f268, plain, sdtlseqdt0(xr, xn), inference(cnf_transformation, [], [f161])).
fof(f1251174, plain, (~ sdtlseqdt0(xr, xn) | (xn = xr) | ~ aNaturalNumber0(xn) | (spl14_5172 | ~ spl14_8665)), inference(resolution, [], [f1202647, f247436])).
fof(f247436, plain, (~ sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm)) | spl14_5172), inference(avatar_component_clause, [], [f247434])).
fof(f247434, plain, (spl14_5172 <=> sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm))), introduced(avatar_definition, [new_symbols(naming, [spl14_5172])])).
fof(f1202647, plain, (! [X7] : (sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(X7, xm)) | ~ sdtlseqdt0(xr, X7) | (xr = X7) | ~ aNaturalNumber0(X7)) | ~ spl14_8665), inference(subsumption_resolution, [], [f1202646, f765565])).
fof(f765565, plain, (aNaturalNumber0(xr) | ~ spl14_8665), inference(avatar_component_clause, [], [f765564])).
fof(f765564, plain, (spl14_8665 <=> aNaturalNumber0(xr)), introduced(avatar_definition, [new_symbols(naming, [spl14_8665])])).
fof(f1202646, plain, ! [X7] : (sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(X7, xm)) | ~ sdtlseqdt0(xr, X7) | (xr = X7) | ~ aNaturalNumber0(X7) | ~ aNaturalNumber0(xr)), inference(subsumption_resolution, [], [f1202607, f233])).
fof(f233, plain, aNaturalNumber0(xm), inference(cnf_transformation, [], [f39])).
fof(f1202607, plain, ! [X7] : (sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(X7, xm)) | ~ aNaturalNumber0(xm) | ~ sdtlseqdt0(xr, X7) | (xr = X7) | ~ aNaturalNumber0(X7) | ~ aNaturalNumber0(xr)), inference(superposition, [], [f202, f9691])).
fof(f9691, plain, (sdtpldt0(xr, xm) = sdtpldt0(xm, xr)), inference(resolution, [], [f441, f262])).
fof(f262, plain, aNaturalNumber0(xr), inference(cnf_transformation, [], [f43])).
fof(f43, plain, ((xr = sdtmndt0(xn, xp)) & (xn = sdtpldt0(xp, xr)) & aNaturalNumber0(xr)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1883)).
fof(f441, plain, ! [X11] : (~ aNaturalNumber0(X11) | (sdtpldt0(X11, xm) = sdtpldt0(xm, X11))), inference(resolution, [], [f169, f233])).
fof(f169, plain, ! [X0, X1] : (~ aNaturalNumber0(X1) | (sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f58])).
fof(f58, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f57])).
fof(f57, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtpldt0(X0, X1) = sdtpldt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mAddComm)).
fof(f202, plain, ! [X2, X0, X1] : (sdtlseqdt0(sdtpldt0(X0, X2), sdtpldt0(X1, X2)) | ~ aNaturalNumber0(X2) | ~ sdtlseqdt0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f90])).
fof(f90, plain, ! [X0, X1] : (! [X2] : ((sdtlseqdt0(sdtpldt0(X0, X2), sdtpldt0(X1, X2)) & ~ (sdtpldt0(X1, X2) = sdtpldt0(X0, X2)) & sdtlseqdt0(sdtpldt0(X2, X0), sdtpldt0(X2, X1)) & ~ (sdtpldt0(X2, X0) = sdtpldt0(X2, X1))) | ~ aNaturalNumber0(X2)) | ~ sdtlseqdt0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f89])).
fof(f89, plain, ! [X0, X1] : ((! [X2] : ((sdtlseqdt0(sdtpldt0(X0, X2), sdtpldt0(X1, X2)) & ~ (sdtpldt0(X1, X2) = sdtpldt0(X0, X2)) & sdtlseqdt0(sdtpldt0(X2, X0), sdtpldt0(X2, X1)) & ~ (sdtpldt0(X2, X0) = sdtpldt0(X2, X1))) | ~ aNaturalNumber0(X2)) | (~ sdtlseqdt0(X0, X1) | (X0 = X1))) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X0, X1) & ~ (X0 = X1)) => ! [X2] : (aNaturalNumber0(X2) => (sdtlseqdt0(sdtpldt0(X0, X2), sdtpldt0(X1, X2)) & ~ (sdtpldt0(X1, X2) = sdtpldt0(X0, X2)) & sdtlseqdt0(sdtpldt0(X2, X0), sdtpldt0(X2, X1)) & ~ (sdtpldt0(X2, X0) = sdtpldt0(X2, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mMonAdd)).
fof(f1146581, plain, (~ spl14_983 | ~ spl14_241 | spl14_1476), inference(avatar_split_clause, [], [f318732, f86609, f17533, f72741])).
fof(f72741, plain, (spl14_983 <=> iLess0(sdtpldt0(sdtpldt0(xm, xr), xp), sdtpldt0(sdtpldt0(xn, xm), xp))), introduced(avatar_definition, [new_symbols(naming, [spl14_983])])).
fof(f17533, plain, (spl14_241 <=> ! [X18, X19] : (~ (sdtasdt0(X18, X19) = sdtasdt0(xm, xr)) | ~ aNaturalNumber0(X18) | ~ aNaturalNumber0(X19) | doDivides0(xp, X19) | ~ iLess0(sdtpldt0(sdtpldt0(X18, X19), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(X18, xp))), introduced(avatar_definition, [new_symbols(naming, [spl14_241])])).
fof(f86609, plain, (spl14_1476 <=> sP1(xr, xp)), introduced(avatar_definition, [new_symbols(naming, [spl14_1476])])).
fof(f318732, plain, (~ iLess0(sdtpldt0(sdtpldt0(xm, xr), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ spl14_241 | spl14_1476)), inference(forward_demodulation, [], [f318731, f9691])).
fof(f318731, plain, (~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ spl14_241 | spl14_1476)), inference(subsumption_resolution, [], [f318730, f86610])).
fof(f86610, plain, (~ sP1(xr, xp) | spl14_1476), inference(avatar_component_clause, [], [f86609])).
fof(f318730, plain, (~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(xr, xp) | ~ spl14_241), inference(subsumption_resolution, [], [f318729, f275])).
fof(f275, plain, ~ doDivides0(xp, xm), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (~ doDivides0(xp, xm) & ! [X0] : (~ (xm = sdtasdt0(xp, X0)) | ~ aNaturalNumber0(X0)) & ~ doDivides0(xp, xr) & ! [X1] : (~ (xr = sdtasdt0(xp, X1)) | ~ aNaturalNumber0(X1))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ~ (doDivides0(xp, xm) | ? [X0] : ((xm = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) | doDivides0(xp, xr) | ? [X1] : ((xr = sdtasdt0(xp, X1)) & aNaturalNumber0(X1))), inference(rectify, [], [f47])).
fof(f47, plain, ~ (doDivides0(xp, xm) | ? [X0] : ((xm = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) | doDivides0(xp, xr) | ? [X0] : ((sdtasdt0(xp, X0) = xr) & aNaturalNumber0(X0))), inference(negated_conjecture, [], [f46])).
fof(f46, plain, ~ (doDivides0(xp, xm) | ? [X0] : ((xm = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) | doDivides0(xp, xr) | ? [X0] : ((sdtasdt0(xp, X0) = xr) & aNaturalNumber0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__)).
fof(f318729, plain, (doDivides0(xp, xm) | ~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(xr, xp) | ~ spl14_241), inference(subsumption_resolution, [], [f318728, f233])).
fof(f318728, plain, (~ aNaturalNumber0(xm) | doDivides0(xp, xm) | ~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(xr, xp) | ~ spl14_241), inference(subsumption_resolution, [], [f318673, f262])).
fof(f318673, plain, (~ aNaturalNumber0(xr) | ~ aNaturalNumber0(xm) | doDivides0(xp, xm) | ~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(xr, xp) | ~ spl14_241), inference(trivial_inequality_removal, [], [f318666])).
fof(f318666, plain, (~ (sdtasdt0(xm, xr) = sdtasdt0(xm, xr)) | ~ aNaturalNumber0(xr) | ~ aNaturalNumber0(xm) | doDivides0(xp, xm) | ~ iLess0(sdtpldt0(sdtpldt0(xr, xm), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(xr, xp) | ~ spl14_241), inference(superposition, [], [f17534, f13682])).
fof(f13682, plain, (sdtasdt0(xr, xm) = sdtasdt0(xm, xr)), inference(resolution, [], [f456, f262])).
fof(f456, plain, ! [X11] : (~ aNaturalNumber0(X11) | (sdtasdt0(X11, xm) = sdtasdt0(xm, X11))), inference(resolution, [], [f173, f233])).
fof(f173, plain, ! [X0, X1] : (~ aNaturalNumber0(X1) | (sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f62])).
fof(f62, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtasdt0(X0, X1) = sdtasdt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mMulComm)).
fof(f17534, plain, (! [X19, X18] : (~ (sdtasdt0(X18, X19) = sdtasdt0(xm, xr)) | ~ aNaturalNumber0(X18) | ~ aNaturalNumber0(X19) | doDivides0(xp, X19) | ~ iLess0(sdtpldt0(sdtpldt0(X18, X19), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | sP1(X18, xp)) | ~ spl14_241), inference(avatar_component_clause, [], [f17533])).
fof(f768235, plain, (spl14_5171 | ~ spl14_5172 | ~ spl14_186 | ~ spl14_636 | spl14_983 | ~ spl14_1484), inference(avatar_split_clause, [], [f768234, f96673, f72741, f49291, f11813, f247434, f247430])).
fof(f247430, plain, (spl14_5171 <=> (sdtpldt0(xn, xm) = sdtpldt0(xm, xr))), introduced(avatar_definition, [new_symbols(naming, [spl14_5171])])).
fof(f11813, plain, (spl14_186 <=> aNaturalNumber0(xp)), introduced(avatar_definition, [new_symbols(naming, [spl14_186])])).
fof(f49291, plain, (spl14_636 <=> aNaturalNumber0(sdtpldt0(xn, xm))), introduced(avatar_definition, [new_symbols(naming, [spl14_636])])).
fof(f96673, plain, (spl14_1484 <=> aNaturalNumber0(sdtpldt0(xm, xr))), introduced(avatar_definition, [new_symbols(naming, [spl14_1484])])).
fof(f768234, plain, (~ sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm)) | (sdtpldt0(xn, xm) = sdtpldt0(xm, xr)) | (~ spl14_186 | ~ spl14_636 | spl14_983 | ~ spl14_1484)), inference(subsumption_resolution, [], [f761911, f96674])).
fof(f96674, plain, (aNaturalNumber0(sdtpldt0(xm, xr)) | ~ spl14_1484), inference(avatar_component_clause, [], [f96673])).
fof(f761911, plain, (~ sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm)) | (sdtpldt0(xn, xm) = sdtpldt0(xm, xr)) | ~ aNaturalNumber0(sdtpldt0(xm, xr)) | (~ spl14_186 | ~ spl14_636 | spl14_983)), inference(subsumption_resolution, [], [f761910, f11814])).
fof(f11814, plain, (aNaturalNumber0(xp) | ~ spl14_186), inference(avatar_component_clause, [], [f11813])).
fof(f761910, plain, (~ sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm)) | (sdtpldt0(xn, xm) = sdtpldt0(xm, xr)) | ~ aNaturalNumber0(sdtpldt0(xm, xr)) | ~ aNaturalNumber0(xp) | (~ spl14_636 | spl14_983)), inference(subsumption_resolution, [], [f288766, f49292])).
fof(f49292, plain, (aNaturalNumber0(sdtpldt0(xn, xm)) | ~ spl14_636), inference(avatar_component_clause, [], [f49291])).
fof(f288766, plain, (~ sdtlseqdt0(sdtpldt0(xm, xr), sdtpldt0(xn, xm)) | (sdtpldt0(xn, xm) = sdtpldt0(xm, xr)) | ~ aNaturalNumber0(sdtpldt0(xn, xm)) | ~ aNaturalNumber0(sdtpldt0(xm, xr)) | ~ aNaturalNumber0(xp) | spl14_983), inference(resolution, [], [f72743, f2059])).
fof(f2059, plain, ! [X12, X10, X11] : (iLess0(sdtpldt0(X11, X10), sdtpldt0(X12, X10)) | ~ sdtlseqdt0(X11, X12) | (X11 = X12) | ~ aNaturalNumber0(X12) | ~ aNaturalNumber0(X11) | ~ aNaturalNumber0(X10)), inference(subsumption_resolution, [], [f2058, f167])).
fof(f167, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f54])).
fof(f54, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f53])).
fof(f53, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtpldt0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mSortsB)).
fof(f2058, plain, ! [X12, X10, X11] : (~ aNaturalNumber0(X10) | ~ sdtlseqdt0(X11, X12) | (X11 = X12) | ~ aNaturalNumber0(X12) | ~ aNaturalNumber0(X11) | iLess0(sdtpldt0(X11, X10), sdtpldt0(X12, X10)) | ~ aNaturalNumber0(sdtpldt0(X11, X10))), inference(subsumption_resolution, [], [f2057, f167])).
fof(f2057, plain, ! [X12, X10, X11] : (~ aNaturalNumber0(X10) | ~ sdtlseqdt0(X11, X12) | (X11 = X12) | ~ aNaturalNumber0(X12) | ~ aNaturalNumber0(X11) | iLess0(sdtpldt0(X11, X10), sdtpldt0(X12, X10)) | ~ aNaturalNumber0(sdtpldt0(X12, X10)) | ~ aNaturalNumber0(sdtpldt0(X11, X10))), inference(subsumption_resolution, [], [f2021, f182])).
fof(f182, plain, ! [X2, X0, X1] : (~ (sdtpldt0(X1, X0) = sdtpldt0(X2, X0)) | (X1 = X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ! [X0, X1, X2] : ((X1 = X2) | (~ (sdtpldt0(X1, X0) = sdtpldt0(X2, X0)) & ~ (sdtpldt0(X0, X1) = sdtpldt0(X0, X2))) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f70])).
fof(f70, plain, ! [X0, X1, X2] : (((X1 = X2) | (~ (sdtpldt0(X1, X0) = sdtpldt0(X2, X0)) & ~ (sdtpldt0(X0, X1) = sdtpldt0(X0, X2)))) | (~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0, X1, X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (((sdtpldt0(X1, X0) = sdtpldt0(X2, X0)) | (sdtpldt0(X0, X1) = sdtpldt0(X0, X2))) => (X1 = X2))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mAddCanc)).
fof(f2021, plain, ! [X12, X10, X11] : (~ aNaturalNumber0(X10) | ~ sdtlseqdt0(X11, X12) | (X11 = X12) | ~ aNaturalNumber0(X12) | ~ aNaturalNumber0(X11) | iLess0(sdtpldt0(X11, X10), sdtpldt0(X12, X10)) | (sdtpldt0(X11, X10) = sdtpldt0(X12, X10)) | ~ aNaturalNumber0(sdtpldt0(X12, X10)) | ~ aNaturalNumber0(sdtpldt0(X11, X10))), inference(resolution, [], [f202, f210])).
fof(f210, plain, ! [X0, X1] : (~ sdtlseqdt0(X0, X1) | iLess0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ! [X0, X1] : (iLess0(X0, X1) | ~ sdtlseqdt0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f97])).
fof(f97, plain, ! [X0, X1] : ((iLess0(X0, X1) | (~ sdtlseqdt0(X0, X1) | (X0 = X1))) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X0, X1) & ~ (X0 = X1)) => iLess0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mIH_03)).
fof(f72743, plain, (~ iLess0(sdtpldt0(sdtpldt0(xm, xr), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | spl14_983), inference(avatar_component_clause, [], [f72741])).
fof(f767955, plain, spl14_8665, inference(avatar_split_clause, [], [f262, f765564])).
fof(f767948, plain, ~ spl14_5171, inference(avatar_contradiction_clause, [], [f767947])).
fof(f767947, plain, ($false | ~ spl14_5171), inference(subsumption_resolution, [], [f767946, f265])).
fof(f767946, plain, ((xn = xr) | ~ spl14_5171), inference(forward_demodulation, [], [f767945, f14456])).
fof(f14456, plain, (xn = sdtmndt0(sdtpldt0(xn, xm), xm)), inference(forward_demodulation, [], [f14434, f9659])).
fof(f9659, plain, (sdtpldt0(xn, xm) = sdtpldt0(xm, xn)), inference(resolution, [], [f440, f233])).
fof(f440, plain, ! [X10] : (~ aNaturalNumber0(X10) | (sdtpldt0(X10, xn) = sdtpldt0(xn, X10))), inference(resolution, [], [f169, f232])).
fof(f14434, plain, (xn = sdtmndt0(sdtpldt0(xm, xn), xm)), inference(resolution, [], [f540, f233])).
fof(f540, plain, ! [X13] : (~ aNaturalNumber0(X13) | (xn = sdtmndt0(sdtpldt0(X13, xn), X13))), inference(resolution, [], [f291, f232])).
fof(f291, plain, ! [X2, X0] : (~ aNaturalNumber0(X2) | (sdtmndt0(sdtpldt0(X0, X2), X0) = X2) | ~ aNaturalNumber0(X0)), inference(subsumption_resolution, [], [f290, f167])).
fof(f290, plain, ! [X2, X0] : ((sdtmndt0(sdtpldt0(X0, X2), X0) = X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(sdtpldt0(X0, X2)) | ~ aNaturalNumber0(X0)), inference(subsumption_resolution, [], [f277, f289])).
fof(f289, plain, ! [X2, X0] : (sdtlseqdt0(X0, sdtpldt0(X0, X2)) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X0)), inference(subsumption_resolution, [], [f276, f167])).
fof(f276, plain, ! [X2, X0] : (sdtlseqdt0(X0, sdtpldt0(X0, X2)) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(sdtpldt0(X0, X2)) | ~ aNaturalNumber0(X0)), inference(equality_resolution, [], [f190])).
fof(f190, plain, ! [X2, X0, X1] : (sdtlseqdt0(X0, X1) | ~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ! [X0, X1] : (((sdtlseqdt0(X0, X1) | ! [X2] : (~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2))) & (((sdtpldt0(X0, sK2(X0, X1)) = X1) & aNaturalNumber0(sK2(X0, X1))) | ~ sdtlseqdt0(X0, X1))) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2])], [f126, f127])).
fof(f127, plain, ! [X1, X0] : (? [X3] : ((sdtpldt0(X0, X3) = X1) & aNaturalNumber0(X3)) => ((sdtpldt0(X0, sK2(X0, X1)) = X1) & aNaturalNumber0(sK2(X0, X1)))), introduced(choice_axiom, [])).
fof(f126, plain, ! [X0, X1] : (((sdtlseqdt0(X0, X1) | ! [X2] : (~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2))) & (? [X3] : ((sdtpldt0(X0, X3) = X1) & aNaturalNumber0(X3)) | ~ sdtlseqdt0(X0, X1))) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(rectify, [], [f125])).
fof(f125, plain, ! [X0, X1] : (((sdtlseqdt0(X0, X1) | ! [X2] : (~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2))) & (? [X2] : ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2)) | ~ sdtlseqdt0(X0, X1))) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(nnf_transformation, [], [f79])).
fof(f79, plain, ! [X0, X1] : ((sdtlseqdt0(X0, X1) <=> ? [X2] : ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2))) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f78])).
fof(f78, plain, ! [X0, X1] : ((sdtlseqdt0(X0, X1) <=> ? [X2] : ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2))) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtlseqdt0(X0, X1) <=> ? [X2] : ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mDefLE)).
fof(f277, plain, ! [X2, X0] : ((sdtmndt0(sdtpldt0(X0, X2), X0) = X2) | ~ aNaturalNumber0(X2) | ~ sdtlseqdt0(X0, sdtpldt0(X0, X2)) | ~ aNaturalNumber0(sdtpldt0(X0, X2)) | ~ aNaturalNumber0(X0)), inference(equality_resolution, [], [f193])).
fof(f193, plain, ! [X2, X0, X1] : ((sdtmndt0(X1, X0) = X2) | ~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2) | ~ sdtlseqdt0(X0, X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f130])).
fof(f130, plain, ! [X0, X1] : (! [X2] : (((sdtmndt0(X1, X0) = X2) | ~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2)) & (((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2)) | ~ (sdtmndt0(X1, X0) = X2))) | ~ sdtlseqdt0(X0, X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f129])).
fof(f129, plain, ! [X0, X1] : (! [X2] : (((sdtmndt0(X1, X0) = X2) | (~ (sdtpldt0(X0, X2) = X1) | ~ aNaturalNumber0(X2))) & (((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2)) | ~ (sdtmndt0(X1, X0) = X2))) | ~ sdtlseqdt0(X0, X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(nnf_transformation, [], [f81])).
fof(f81, plain, ! [X0, X1] : (! [X2] : ((sdtmndt0(X1, X0) = X2) <=> ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2))) | ~ sdtlseqdt0(X0, X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f80])).
fof(f80, plain, ! [X0, X1] : ((! [X2] : ((sdtmndt0(X1, X0) = X2) <=> ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2))) | ~ sdtlseqdt0(X0, X1)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtlseqdt0(X0, X1) => ! [X2] : ((sdtmndt0(X1, X0) = X2) <=> ((sdtpldt0(X0, X2) = X1) & aNaturalNumber0(X2))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', mDefDiff)).
fof(f767945, plain, ((xr = sdtmndt0(sdtpldt0(xn, xm), xm)) | ~ spl14_5171), inference(forward_demodulation, [], [f18694, f247432])).
fof(f247432, plain, ((sdtpldt0(xn, xm) = sdtpldt0(xm, xr)) | ~ spl14_5171), inference(avatar_component_clause, [], [f247430])).
fof(f18694, plain, (xr = sdtmndt0(sdtpldt0(xm, xr), xm)), inference(resolution, [], [f543, f233])).
fof(f543, plain, ! [X16] : (~ aNaturalNumber0(X16) | (xr = sdtmndt0(sdtpldt0(X16, xr), X16))), inference(resolution, [], [f291, f262])).
fof(f288785, plain, ~ spl14_1476, inference(avatar_contradiction_clause, [], [f288784])).
fof(f288784, plain, ($false | ~ spl14_1476), inference(subsumption_resolution, [], [f288769, f273])).
fof(f273, plain, ~ doDivides0(xp, xr), inference(cnf_transformation, [], [f121])).
fof(f288769, plain, (doDivides0(xp, xr) | ~ spl14_1476), inference(resolution, [], [f86611, f237])).
fof(f237, plain, ! [X0, X1] : (~ sP1(X0, X1) | doDivides0(X1, X0)), inference(cnf_transformation, [], [f147])).
fof(f147, plain, ! [X0, X1] : ((doDivides0(X1, X0) & ((sdtasdt0(X1, sK6(X0, X1)) = X0) & aNaturalNumber0(sK6(X0, X1)))) | ~ sP1(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f145, f146])).
fof(f146, plain, ! [X1, X0] : (? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) => ((sdtasdt0(X1, sK6(X0, X1)) = X0) & aNaturalNumber0(sK6(X0, X1)))), introduced(choice_axiom, [])).
fof(f145, plain, ! [X0, X1] : ((doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2))) | ~ sP1(X0, X1)), inference(rectify, [], [f144])).
fof(f144, plain, ! [X0, X2] : ((doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ sP1(X0, X2)), inference(nnf_transformation, [], [f123])).
fof(f123, plain, ! [X0, X2] : ((doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ sP1(X0, X2)), inference(usedef, [], [e123])).
fof(e123, plain, ! [X0, X2] : (sP1(X0, X2) <=> (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f86611, plain, (sP1(xr, xp) | ~ spl14_1476), inference(avatar_component_clause, [], [f86609])).
fof(f96911, plain, spl14_1484, inference(avatar_contradiction_clause, [], [f96910])).
fof(f96910, plain, ($false | spl14_1484), inference(subsumption_resolution, [], [f96909, f233])).
fof(f96909, plain, (~ aNaturalNumber0(xm) | spl14_1484), inference(subsumption_resolution, [], [f96908, f262])).
fof(f96908, plain, (~ aNaturalNumber0(xr) | ~ aNaturalNumber0(xm) | spl14_1484), inference(resolution, [], [f96675, f167])).
fof(f96675, plain, (~ aNaturalNumber0(sdtpldt0(xm, xr)) | spl14_1484), inference(avatar_component_clause, [], [f96673])).
fof(f49385, plain, spl14_636, inference(avatar_contradiction_clause, [], [f49384])).
fof(f49384, plain, ($false | spl14_636), inference(subsumption_resolution, [], [f49383, f232])).
fof(f49383, plain, (~ aNaturalNumber0(xn) | spl14_636), inference(subsumption_resolution, [], [f49382, f233])).
fof(f49382, plain, (~ aNaturalNumber0(xm) | ~ aNaturalNumber0(xn) | spl14_636), inference(resolution, [], [f49293, f167])).
fof(f49293, plain, (~ aNaturalNumber0(sdtpldt0(xn, xm)) | spl14_636), inference(avatar_component_clause, [], [f49291])).
fof(f17606, plain, ~ spl14_11, inference(avatar_contradiction_clause, [], [f17605])).
fof(f17605, plain, ($false | ~ spl14_11), inference(subsumption_resolution, [], [f17601, f255])).
fof(f255, plain, isPrime0(xp), inference(cnf_transformation, [], [f157])).
fof(f157, plain, (doDivides0(xp, sdtasdt0(xn, xm)) & ((sdtasdt0(xn, xm) = sdtasdt0(xp, sK10)) & aNaturalNumber0(sK10)) & isPrime0(xp) & ! [X1] : ((xp = X1) | (sz10 = X1) | (~ doDivides0(X1, xp) & ! [X2] : (~ (sdtasdt0(X1, X2) = xp) | ~ aNaturalNumber0(X2))) | ~ aNaturalNumber0(X1)) & ~ (sz10 = xp) & ~ (sz00 = xp)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK10])], [f120, f156])).
fof(f156, plain, (? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) => ((sdtasdt0(xn, xm) = sdtasdt0(xp, sK10)) & aNaturalNumber0(sK10))), introduced(choice_axiom, [])).
fof(f120, plain, (doDivides0(xp, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) & isPrime0(xp) & ! [X1] : ((xp = X1) | (sz10 = X1) | (~ doDivides0(X1, xp) & ! [X2] : (~ (sdtasdt0(X1, X2) = xp) | ~ aNaturalNumber0(X2))) | ~ aNaturalNumber0(X1)) & ~ (sz10 = xp) & ~ (sz00 = xp)), inference(flattening, [], [f119])).
fof(f119, plain, (doDivides0(xp, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) & isPrime0(xp) & ! [X1] : (((xp = X1) | (sz10 = X1)) | ((~ doDivides0(X1, xp) & ! [X2] : (~ (sdtasdt0(X1, X2) = xp) | ~ aNaturalNumber0(X2))) | ~ aNaturalNumber0(X1))) & ~ (sz10 = xp) & ~ (sz00 = xp)), inference(ennf_transformation, [], [f51])).
fof(f51, plain, (doDivides0(xp, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) & isPrime0(xp) & ! [X1] : (((doDivides0(X1, xp) | ? [X2] : ((sdtasdt0(X1, X2) = xp) & aNaturalNumber0(X2))) & aNaturalNumber0(X1)) => ((xp = X1) | (sz10 = X1))) & ~ (sz10 = xp) & ~ (sz00 = xp)), inference(rectify, [], [f41])).
fof(f41, plain, (doDivides0(xp, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xp, X0)) & aNaturalNumber0(X0)) & isPrime0(xp) & ! [X0] : (((doDivides0(X0, xp) | ? [X1] : ((sdtasdt0(X0, X1) = xp) & aNaturalNumber0(X1))) & aNaturalNumber0(X0)) => ((xp = X0) | (sz10 = X0))) & ~ (sz10 = xp) & ~ (sz00 = xp)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1860)).
fof(f17601, plain, (~ isPrime0(xp) | ~ spl14_11), inference(resolution, [], [f770, f244])).
fof(f244, plain, ! [X0] : (~ sP0(X0) | ~ isPrime0(X0)), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ! [X0] : ((~ isPrime0(X0) & ((~ (sK7(X0) = X0) & ~ (sz10 = sK7(X0)) & doDivides0(sK7(X0), X0) & ((sdtasdt0(sK7(X0), sK8(X0)) = X0) & aNaturalNumber0(sK8(X0))) & aNaturalNumber0(sK7(X0))) | (sz10 = X0) | (sz00 = X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f149, f151, f150])).
fof(f150, plain, ! [X0] : (? [X1] : (~ (X0 = X1) & ~ (sz10 = X1) & doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (~ (sK7(X0) = X0) & ~ (sz10 = sK7(X0)) & doDivides0(sK7(X0), X0) & ? [X2] : ((sdtasdt0(sK7(X0), X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(sK7(X0)))), introduced(choice_axiom, [])).
fof(f151, plain, ! [X0] : (? [X2] : ((sdtasdt0(sK7(X0), X2) = X0) & aNaturalNumber0(X2)) => ((sdtasdt0(sK7(X0), sK8(X0)) = X0) & aNaturalNumber0(sK8(X0)))), introduced(choice_axiom, [])).
fof(f149, plain, ! [X0] : ((~ isPrime0(X0) & (? [X1] : (~ (X0 = X1) & ~ (sz10 = X1) & doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) | (sz10 = X0) | (sz00 = X0))) | ~ sP0(X0)), inference(rectify, [], [f148])).
fof(f148, plain, ! [X2] : ((~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ sP0(X2)), inference(nnf_transformation, [], [f122])).
fof(f122, plain, ! [X2] : ((~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ sP0(X2)), inference(usedef, [], [e122])).
fof(e122, plain, ! [X2] : (sP0(X2) <=> (~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f770, plain, (sP0(xp) | ~ spl14_11), inference(avatar_component_clause, [], [f769])).
fof(f769, plain, (spl14_11 <=> sP0(xp)), introduced(avatar_definition, [new_symbols(naming, [spl14_11])])).
fof(f17541, plain, (spl14_11 | spl14_241 | ~ spl14_186), inference(avatar_split_clause, [], [f17540, f11813, f17533, f769])).
fof(f17540, plain, (! [X19, X18] : (~ (sdtasdt0(X18, X19) = sdtasdt0(xm, xr)) | sP1(X18, xp) | ~ iLess0(sdtpldt0(sdtpldt0(X18, X19), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | doDivides0(xp, X19) | sP0(xp) | ~ aNaturalNumber0(X19) | ~ aNaturalNumber0(X18)) | ~ spl14_186), inference(subsumption_resolution, [], [f13845, f11814])).
fof(f13845, plain, ! [X19, X18] : (~ (sdtasdt0(X18, X19) = sdtasdt0(xm, xr)) | sP1(X18, xp) | ~ iLess0(sdtpldt0(sdtpldt0(X18, X19), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | doDivides0(xp, X19) | sP0(xp) | ~ aNaturalNumber0(xp) | ~ aNaturalNumber0(X19) | ~ aNaturalNumber0(X18)), inference(subsumption_resolution, [], [f13811, f269])).
fof(f269, plain, aNaturalNumber0(sK13), inference(cnf_transformation, [], [f163])).
fof(f163, plain, (doDivides0(xp, sdtasdt0(xr, xm)) & ((sdtasdt0(xr, xm) = sdtasdt0(xp, sK13)) & aNaturalNumber0(sK13))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13])], [f45, f162])).
fof(f162, plain, (? [X0] : ((sdtasdt0(xp, X0) = sdtasdt0(xr, xm)) & aNaturalNumber0(X0)) => ((sdtasdt0(xr, xm) = sdtasdt0(xp, sK13)) & aNaturalNumber0(sK13))), introduced(choice_axiom, [])).
fof(f45, plain, (doDivides0(xp, sdtasdt0(xr, xm)) & ? [X0] : ((sdtasdt0(xp, X0) = sdtasdt0(xr, xm)) & aNaturalNumber0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1913)).
fof(f13811, plain, ! [X19, X18] : (~ (sdtasdt0(X18, X19) = sdtasdt0(xm, xr)) | sP1(X18, xp) | ~ iLess0(sdtpldt0(sdtpldt0(X18, X19), xp), sdtpldt0(sdtpldt0(xn, xm), xp)) | doDivides0(xp, X19) | ~ aNaturalNumber0(sK13) | sP0(xp) | ~ aNaturalNumber0(xp) | ~ aNaturalNumber0(X19) | ~ aNaturalNumber0(X18)), inference(superposition, [], [f249, f13699])).
fof(f13699, plain, (sdtasdt0(xp, sK13) = sdtasdt0(xm, xr)), inference(backward_demodulation, [], [f270, f13682])).
fof(f270, plain, (sdtasdt0(xr, xm) = sdtasdt0(xp, sK13)), inference(cnf_transformation, [], [f163])).
fof(f249, plain, ! [X4, X2, X0, X1] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | doDivides0(X2, X1) | ~ aNaturalNumber0(X4) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ((sdtasdt0(X2, sK9(X1, X2)) = X1) & aNaturalNumber0(sK9(X1, X2)))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X4] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | ~ aNaturalNumber0(X4))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9])], [f153, f154])).
fof(f154, plain, ! [X2, X1] : (? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3)) => ((sdtasdt0(X2, sK9(X1, X2)) = X1) & aNaturalNumber0(sK9(X1, X2)))), introduced(choice_axiom, [])).
fof(f153, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X4] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | ~ aNaturalNumber0(X4))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(rectify, [], [f124])).
fof(f124, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(definition_folding, [], [f118, e123, e122])).
fof(f118, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | (~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f117])).
fof(f117, plain, ! [X0, X1, X2] : (((((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7)))) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp))) | ((~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | (~ isPrime0(X2) & (? [X4] : ((~ (X2 = X4) & ~ (sz10 = X4)) & (doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4))) | (sz10 = X2) | (sz00 = X2))))) | (~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f50])).
fof(f50, plain, ! [X0, X1, X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (((doDivides0(X2, sdtasdt0(X0, X1)) | ? [X3] : ((sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) & aNaturalNumber0(X3))) & (isPrime0(X2) | (! [X4] : ((doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) => ((X2 = X4) | (sz10 = X4))) & ~ (sz10 = X2) & ~ (sz00 = X2)))) => (iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) => ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))))))), inference(rectify, [], [f40])).
fof(f40, plain, ! [X0, X1, X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (((doDivides0(X2, sdtasdt0(X0, X1)) | ? [X3] : ((sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) & aNaturalNumber0(X3))) & (isPrime0(X2) | (! [X3] : ((doDivides0(X3, X2) & ? [X4] : ((sdtasdt0(X3, X4) = X2) & aNaturalNumber0(X4)) & aNaturalNumber0(X3)) => ((X2 = X3) | (sz10 = X3))) & ~ (sz10 = X2) & ~ (sz00 = X2)))) => (iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) => ((doDivides0(X2, X1) & ? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3))) | (doDivides0(X2, X0) & ? [X3] : ((sdtasdt0(X2, X3) = X0) & aNaturalNumber0(X3))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM493+3.p', m__1799)).
fof(f11903, plain, spl14_186, inference(avatar_split_clause, [], [f234, f11813])).
fof(f234, plain, aNaturalNumber0(xp), inference(cnf_transformation, [], [f39])).