fof(f2997, plain, $false, inference(avatar_sat_refutation, [], [f1304, f1342, f1363, f2750, f2995])).
fof(f2995, plain, (~ spl32_61 | ~ spl32_70), inference(avatar_contradiction_clause, [], [f2994])).
fof(f2994, plain, ($false | (~ spl32_61 | ~ spl32_70)), inference(subsumption_resolution, [], [f2993, f340])).
fof(f340, plain, aSet0(xS), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((xS = cS2043) & ! [X0] : ((aElementOf0(X0, xS) | ! [X1] : ((~ (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & sP7(X1) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) | ~ isPrime0(X1) | (sz00 = X1) | ~ aInteger0(X1))) & (((szAzrzSzezqlpdtcmdtrp0(sz00, sK26(X0)) = X0) & sP6(sK26(X0)) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, sK26(X0))) & isPrime0(sK26(X0)) & ~ (sz00 = sK26(X0)) & aInteger0(sK26(X0))) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26])], [f197, f198])).
fof(f198, plain, ! [X0] : (? [X2] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X2) = X0) & sP6(X2) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X2)) & isPrime0(X2) & ~ (sz00 = X2) & aInteger0(X2)) => ((szAzrzSzezqlpdtcmdtrp0(sz00, sK26(X0)) = X0) & sP6(sK26(X0)) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, sK26(X0))) & isPrime0(sK26(X0)) & ~ (sz00 = sK26(X0)) & aInteger0(sK26(X0)))), introduced(choice_axiom, [])).
fof(f197, plain, ((xS = cS2043) & ! [X0] : ((aElementOf0(X0, xS) | ! [X1] : ((~ (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & sP7(X1) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) | ~ isPrime0(X1) | (sz00 = X1) | ~ aInteger0(X1))) & (? [X2] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X2) = X0) & sP6(X2) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X2)) & isPrime0(X2) & ~ (sz00 = X2) & aInteger0(X2)) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(rectify, [], [f127])).
fof(f127, plain, ((xS = cS2043) & ! [X0] : ((aElementOf0(X0, xS) | ! [X1] : ((~ (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & sP7(X1) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) | ~ isPrime0(X1) | (sz00 = X1) | ~ aInteger0(X1))) & (? [X5] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X5) = X0) & sP6(X5) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X5)) & isPrime0(X5) & ~ (sz00 = X5) & aInteger0(X5)) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(definition_folding, [], [f113, e126, e125])).
fof(f125, plain, ! [X5] : (! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)) | (~ sdteqdtlpzmzozddtrp0(X6, sz00, X5) & ~ aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, sz00, X5) & aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ? [X8] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)))) | ~ sP6(X5)), inference(usedef, [], [e125])).
fof(e125, plain, ! [X5] : (sP6(X5) <=> ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)) | (~ sdteqdtlpzmzozddtrp0(X6, sz00, X5) & ~ aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, sz00, X5) & aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ? [X8] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f126, plain, ! [X1] : (! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) | (~ sdteqdtlpzmzozddtrp0(X2, sz00, X1) & ~ aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ! [X3] : (~ (sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) | ~ aInteger0(X3))) | ~ aInteger0(X2)) & ((sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X4] : ((sdtpldt0(X2, smndt0(sz00)) = sdtasdt0(X1, X4)) & aInteger0(X4)) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)))) | ~ sP7(X1)), inference(usedef, [], [e126])).
fof(e126, plain, ! [X1] : (sP7(X1) <=> ! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) | (~ sdteqdtlpzmzozddtrp0(X2, sz00, X1) & ~ aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ! [X3] : (~ (sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) | ~ aInteger0(X3))) | ~ aInteger0(X2)) & ((sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X4] : ((sdtpldt0(X2, smndt0(sz00)) = sdtasdt0(X1, X4)) & aInteger0(X4)) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f113, plain, ((xS = cS2043) & ! [X0] : ((aElementOf0(X0, xS) | ! [X1] : ((~ (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & ! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) | (~ sdteqdtlpzmzozddtrp0(X2, sz00, X1) & ~ aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ! [X3] : (~ (sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) | ~ aInteger0(X3))) | ~ aInteger0(X2)) & ((sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X4] : ((sdtpldt0(X2, smndt0(sz00)) = sdtasdt0(X1, X4)) & aInteger0(X4)) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) | ~ isPrime0(X1) | (sz00 = X1) | ~ aInteger0(X1))) & (? [X5] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X5) = X0) & ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)) | (~ sdteqdtlpzmzozddtrp0(X6, sz00, X5) & ~ aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6)) & ((sdteqdtlpzmzozddtrp0(X6, sz00, X5) & aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ? [X8] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X5)) & isPrime0(X5) & ~ (sz00 = X5) & aInteger0(X5)) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(flattening, [], [f112])).
fof(f112, plain, ((xS = cS2043) & ! [X0] : ((aElementOf0(X0, xS) | ! [X1] : ((~ (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & (! [X2] : ((aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) | ((~ sdteqdtlpzmzozddtrp0(X2, sz00, X1) & ~ aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ! [X3] : (~ (sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) | ~ aInteger0(X3))) | ~ aInteger0(X2))) & ((sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X4] : ((sdtpldt0(X2, smndt0(sz00)) = sdtasdt0(X1, X4)) & aInteger0(X4)) & aInteger0(X2)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1)))) | ~ isPrime0(X1) | (sz00 = X1) | ~ aInteger0(X1))) & (? [X5] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X5) = X0) & ! [X6] : ((aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)) | ((~ sdteqdtlpzmzozddtrp0(X6, sz00, X5) & ~ aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ! [X7] : (~ (sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X7)) | ~ aInteger0(X7))) | ~ aInteger0(X6))) & ((sdteqdtlpzmzozddtrp0(X6, sz00, X5) & aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ? [X8] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X8)) & aInteger0(X8)) & aInteger0(X6)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X5)) & isPrime0(X5) & ~ (sz00 = X5) & aInteger0(X5)) | ~ aElementOf0(X0, xS))) & aSet0(xS)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ((xS = cS2043) & ! [X0] : ((? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, sz00, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) => (sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X4] : ((sdtpldt0(X2, smndt0(sz00)) = sdtasdt0(X1, X4)) & aInteger0(X4)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) => (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0)) & isPrime0(X1) & ~ (sz00 = X1) & aInteger0(X1)) => aElementOf0(X0, xS)) & (aElementOf0(X0, xS) => ? [X5] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X5) = X0) & ! [X6] : ((((sdteqdtlpzmzozddtrp0(X6, sz00, X5) | aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) | ? [X7] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X7)) & aInteger0(X7))) & aInteger0(X6)) => aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5))) & (aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(sz00, X5)) => (sdteqdtlpzmzozddtrp0(X6, sz00, X5) & aDivisorOf0(X5, sdtpldt0(X6, smndt0(sz00))) & ? [X8] : ((sdtpldt0(X6, smndt0(sz00)) = sdtasdt0(X5, X8)) & aInteger0(X8)) & aInteger0(X6)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X5)) & isPrime0(X5) & ~ (sz00 = X5) & aInteger0(X5)))) & aSet0(xS)), inference(rectify, [], [f42])).
fof(f42, plain, ((xS = cS2043) & ! [X0] : ((? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, sz00, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) => (sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1))) => (szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0)) & isPrime0(X1) & ~ (sz00 = X1) & aInteger0(X1)) => aElementOf0(X0, xS)) & (aElementOf0(X0, xS) => ? [X1] : ((szAzrzSzezqlpdtcmdtrp0(sz00, X1) = X0) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, sz00, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sz00, X1)) => (sdteqdtlpzmzozddtrp0(X2, sz00, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(sz00))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(sz00))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(sz00, X1)) & isPrime0(X1) & ~ (sz00 = X1) & aInteger0(X1)))) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', m__2046)).
fof(f2993, plain, (~ aSet0(xS) | (~ spl32_61 | ~ spl32_70)), inference(subsumption_resolution, [], [f2992, f364])).
fof(f364, plain, isFinite0(xS), inference(cnf_transformation, [], [f44])).
fof(f44, plain, isFinite0(xS), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', m__2117)).
fof(f2992, plain, (~ isFinite0(xS) | ~ aSet0(xS) | (~ spl32_61 | ~ spl32_70)), inference(subsumption_resolution, [], [f2991, f388])).
fof(f388, plain, ~ isClosed0(sbsmnsldt0(xS)), inference(cnf_transformation, [], [f215])).
fof(f215, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sK29, X1), stldt0(sbsmnsldt0(xS))) & (~ aElementOf0(sK30(X1), stldt0(sbsmnsldt0(xS))) & aElementOf0(sK30(X1), szAzrzSzezqlpdtcmdtrp0(sK29, X1))) & sP8(X1, sK29) & aSet0(szAzrzSzezqlpdtcmdtrp0(sK29, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(sK29, stldt0(sbsmnsldt0(xS)))) & ! [X3] : ((aElementOf0(X3, stldt0(sbsmnsldt0(xS))) | aElementOf0(X3, sbsmnsldt0(xS)) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, sbsmnsldt0(xS)) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(sbsmnsldt0(xS))))) & ! [X4] : ((aElementOf0(X4, sbsmnsldt0(xS)) | ! [X5] : (~ aElementOf0(X4, X5) | ~ aElementOf0(X5, xS)) | ~ aInteger0(X4)) & (((aElementOf0(X4, sK31(X4)) & aElementOf0(sK31(X4), xS)) & aInteger0(X4)) | ~ aElementOf0(X4, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30, sK31])], [f211, f214, f213, f212])).
fof(f212, plain, (? [X0] : (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ? [X2] : (~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP8(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) => (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(sK29, X1), stldt0(sbsmnsldt0(xS))) & ? [X2] : (~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sK29, X1))) & sP8(X1, sK29) & aSet0(szAzrzSzezqlpdtcmdtrp0(sK29, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(sK29, stldt0(sbsmnsldt0(xS))))), introduced(choice_axiom, [])).
fof(f213, plain, ! [X1] : (? [X2] : (~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(sK29, X1))) => (~ aElementOf0(sK30(X1), stldt0(sbsmnsldt0(xS))) & aElementOf0(sK30(X1), szAzrzSzezqlpdtcmdtrp0(sK29, X1)))), introduced(choice_axiom, [])).
fof(f214, plain, ! [X4] : (? [X6] : (aElementOf0(X4, X6) & aElementOf0(X6, xS)) => (aElementOf0(X4, sK31(X4)) & aElementOf0(sK31(X4), xS))), introduced(choice_axiom, [])).
fof(f211, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X0] : (! [X1] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) & ? [X2] : (~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))) & aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP8(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) | (sz00 = X1) | ~ aInteger0(X1)) & aElementOf0(X0, stldt0(sbsmnsldt0(xS)))) & ! [X3] : ((aElementOf0(X3, stldt0(sbsmnsldt0(xS))) | aElementOf0(X3, sbsmnsldt0(xS)) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, sbsmnsldt0(xS)) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(sbsmnsldt0(xS))))) & ! [X4] : ((aElementOf0(X4, sbsmnsldt0(xS)) | ! [X5] : (~ aElementOf0(X4, X5) | ~ aElementOf0(X5, xS)) | ~ aInteger0(X4)) & ((? [X6] : (aElementOf0(X4, X6) & aElementOf0(X6, xS)) & aInteger0(X4)) | ~ aElementOf0(X4, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(rectify, [], [f210])).
fof(f210, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X3] : (! [X4] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) & ? [X8] : (~ aElementOf0(X8, stldt0(sbsmnsldt0(xS))) & aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4))) & sP8(X4, X3) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4))) | (sz00 = X4) | ~ aInteger0(X4)) & aElementOf0(X3, stldt0(sbsmnsldt0(xS)))) & ! [X2] : ((aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | aElementOf0(X2, sbsmnsldt0(xS)) | ~ aInteger0(X2)) & ((~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2)) | ~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))))) & ! [X0] : ((aElementOf0(X0, sbsmnsldt0(xS)) | ! [X1] : (~ aElementOf0(X0, X1) | ~ aElementOf0(X1, xS)) | ~ aInteger0(X0)) & ((? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0)) | ~ aElementOf0(X0, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f209])).
fof(f209, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X3] : (! [X4] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) & ? [X8] : (~ aElementOf0(X8, stldt0(sbsmnsldt0(xS))) & aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4))) & sP8(X4, X3) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4))) | (sz00 = X4) | ~ aInteger0(X4)) & aElementOf0(X3, stldt0(sbsmnsldt0(xS)))) & ! [X2] : ((aElementOf0(X2, stldt0(sbsmnsldt0(xS))) | (aElementOf0(X2, sbsmnsldt0(xS)) | ~ aInteger0(X2))) & ((~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2)) | ~ aElementOf0(X2, stldt0(sbsmnsldt0(xS))))) & ! [X0] : ((aElementOf0(X0, sbsmnsldt0(xS)) | (! [X1] : (~ aElementOf0(X0, X1) | ~ aElementOf0(X1, xS)) | ~ aInteger0(X0))) & ((? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0)) | ~ aElementOf0(X0, sbsmnsldt0(xS)))) & aSet0(sbsmnsldt0(xS))), inference(nnf_transformation, [], [f129])).
fof(f129, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X3] : (! [X4] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) & ? [X8] : (~ aElementOf0(X8, stldt0(sbsmnsldt0(xS))) & aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4))) & sP8(X4, X3) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4))) | (sz00 = X4) | ~ aInteger0(X4)) & aElementOf0(X3, stldt0(sbsmnsldt0(xS)))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))), inference(definition_folding, [], [f115, e128])).
fof(f128, plain, ! [X4, X3] : (! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)) | (~ sdteqdtlpzmzozddtrp0(X5, X3, X4) & ~ aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ! [X6] : (~ (sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X6)) | ~ aInteger0(X6))) | ~ aInteger0(X5)) & ((sdteqdtlpzmzozddtrp0(X5, X3, X4) & aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ? [X7] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X7)) & aInteger0(X7)) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)))) | ~ sP8(X4, X3)), inference(usedef, [], [e128])).
fof(e128, plain, ! [X4, X3] : (sP8(X4, X3) <=> ! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)) | (~ sdteqdtlpzmzozddtrp0(X5, X3, X4) & ~ aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ! [X6] : (~ (sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X6)) | ~ aInteger0(X6))) | ~ aInteger0(X5)) & ((sdteqdtlpzmzozddtrp0(X5, X3, X4) & aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ? [X7] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X7)) & aInteger0(X7)) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f115, plain, (~ isClosed0(sbsmnsldt0(xS)) & ~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X3] : (! [X4] : ((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) & ? [X8] : (~ aElementOf0(X8, stldt0(sbsmnsldt0(xS))) & aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4))) & ! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)) | (~ sdteqdtlpzmzozddtrp0(X5, X3, X4) & ~ aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ! [X6] : (~ (sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X6)) | ~ aInteger0(X6))) | ~ aInteger0(X5)) & ((sdteqdtlpzmzozddtrp0(X5, X3, X4) & aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ? [X7] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X7)) & aInteger0(X7)) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4))) | (sz00 = X4) | ~ aInteger0(X4)) & aElementOf0(X3, stldt0(sbsmnsldt0(xS)))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2))) & ! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))), inference(flattening, [], [f114])).
fof(f114, plain, ((~ isClosed0(sbsmnsldt0(xS)) & ((~ isOpen0(stldt0(sbsmnsldt0(xS))) & ? [X3] : (! [X4] : (((~ aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) & ? [X8] : (~ aElementOf0(X8, stldt0(sbsmnsldt0(xS))) & aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4)))) & (! [X5] : ((aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)) | ((~ sdteqdtlpzmzozddtrp0(X5, X3, X4) & ~ aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ! [X6] : (~ (sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X6)) | ~ aInteger0(X6))) | ~ aInteger0(X5))) & ((sdteqdtlpzmzozddtrp0(X5, X3, X4) & aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ? [X7] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X7)) & aInteger0(X7)) & aInteger0(X5)) | ~ aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4)))) | (sz00 = X4) | ~ aInteger0(X4)) & aElementOf0(X3, stldt0(sbsmnsldt0(xS))))) & ! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2))))) & (! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS)))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ~ ((! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))) => (isClosed0(sbsmnsldt0(xS)) | (! [X2] : (aElementOf0(X2, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X2, sbsmnsldt0(xS)) & aInteger0(X2))) => (isOpen0(stldt0(sbsmnsldt0(xS))) | ! [X3] : (aElementOf0(X3, stldt0(sbsmnsldt0(xS))) => ? [X4] : (((! [X5] : ((((sdteqdtlpzmzozddtrp0(X5, X3, X4) | aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) | ? [X6] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X6)) & aInteger0(X6))) & aInteger0(X5)) => aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4))) & (aElementOf0(X5, szAzrzSzezqlpdtcmdtrp0(X3, X4)) => (sdteqdtlpzmzozddtrp0(X5, X3, X4) & aDivisorOf0(X4, sdtpldt0(X5, smndt0(X3))) & ? [X7] : ((sdtpldt0(X5, smndt0(X3)) = sdtasdt0(X4, X7)) & aInteger0(X7)) & aInteger0(X5)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X3, X4))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X3, X4), stldt0(sbsmnsldt0(xS))) | ! [X8] : (aElementOf0(X8, szAzrzSzezqlpdtcmdtrp0(X3, X4)) => aElementOf0(X8, stldt0(sbsmnsldt0(xS)))))) & ~ (sz00 = X4) & aInteger0(X4))))))), inference(rectify, [], [f46])).
fof(f46, plain, ~ ((! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))) => (isClosed0(sbsmnsldt0(xS)) | (! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) => (isOpen0(stldt0(sbsmnsldt0(xS))) | ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) => ? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) | ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(sbsmnsldt0(xS)))))) & ~ (sz00 = X1) & aInteger0(X1))))))), inference(negated_conjecture, [], [f45])).
fof(f45, plain, ~ ((! [X0] : (aElementOf0(X0, sbsmnsldt0(xS)) <=> (? [X1] : (aElementOf0(X0, X1) & aElementOf0(X1, xS)) & aInteger0(X0))) & aSet0(sbsmnsldt0(xS))) => (isClosed0(sbsmnsldt0(xS)) | (! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) <=> (~ aElementOf0(X0, sbsmnsldt0(xS)) & aInteger0(X0))) => (isOpen0(stldt0(sbsmnsldt0(xS))) | ! [X0] : (aElementOf0(X0, stldt0(sbsmnsldt0(xS))) => ? [X1] : (((! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1))) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(sbsmnsldt0(xS))) | ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(sbsmnsldt0(xS)))))) & ~ (sz00 = X1) & aInteger0(X1))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', m__)).
fof(f2991, plain, (isClosed0(sbsmnsldt0(xS)) | ~ isFinite0(xS) | ~ aSet0(xS) | (~ spl32_61 | ~ spl32_70)), inference(subsumption_resolution, [], [f2984, f1303])).
fof(f1303, plain, (isClosed0(sK23(xS)) | ~ spl32_61), inference(avatar_component_clause, [], [f1301])).
fof(f1301, plain, (spl32_61 <=> isClosed0(sK23(xS))), introduced(avatar_definition, [new_symbols(naming, [spl32_61])])).
fof(f2984, plain, (~ isClosed0(sK23(xS)) | isClosed0(sbsmnsldt0(xS)) | ~ isFinite0(xS) | ~ aSet0(xS) | ~ spl32_70), inference(resolution, [], [f1488, f321])).
fof(f321, plain, ! [X0] : (~ aSubsetOf0(sK23(X0), cS1395) | ~ isClosed0(sK23(X0)) | isClosed0(sbsmnsldt0(X0)) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f188])).
fof(f188, plain, ! [X0] : (isClosed0(sbsmnsldt0(X0)) | ((~ isClosed0(sK23(X0)) | ~ aSubsetOf0(sK23(X0), cS1395)) & aElementOf0(sK23(X0), X0)) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23])], [f109, f187])).
fof(f187, plain, ! [X0] : (? [X1] : ((~ isClosed0(X1) | ~ aSubsetOf0(X1, cS1395)) & aElementOf0(X1, X0)) => ((~ isClosed0(sK23(X0)) | ~ aSubsetOf0(sK23(X0), cS1395)) & aElementOf0(sK23(X0), X0))), introduced(choice_axiom, [])).
fof(f109, plain, ! [X0] : (isClosed0(sbsmnsldt0(X0)) | ? [X1] : ((~ isClosed0(X1) | ~ aSubsetOf0(X1, cS1395)) & aElementOf0(X1, X0)) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(flattening, [], [f108])).
fof(f108, plain, ! [X0] : (isClosed0(sbsmnsldt0(X0)) | (? [X1] : ((~ isClosed0(X1) | ~ aSubsetOf0(X1, cS1395)) & aElementOf0(X1, X0)) | ~ isFinite0(X0) | ~ aSet0(X0))), inference(ennf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : ((! [X1] : (aElementOf0(X1, X0) => (isClosed0(X1) & aSubsetOf0(X1, cS1395))) & isFinite0(X0) & aSet0(X0)) => isClosed0(sbsmnsldt0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', mUnionSClosed)).
fof(f1488, plain, (aSubsetOf0(sK23(xS), cS1395) | ~ spl32_70), inference(avatar_component_clause, [], [f1487])).
fof(f1487, plain, (spl32_70 <=> aSubsetOf0(sK23(xS), cS1395)), introduced(avatar_definition, [new_symbols(naming, [spl32_70])])).
fof(f2750, plain, (spl32_70 | spl32_58 | ~ spl32_59), inference(avatar_split_clause, [], [f2749, f1287, f1265, f1487])).
fof(f1265, plain, (spl32_58 <=> (sz00 = sK26(sK23(xS)))), introduced(avatar_definition, [new_symbols(naming, [spl32_58])])).
fof(f1287, plain, (spl32_59 <=> aElementOf0(sK23(xS), xS)), introduced(avatar_definition, [new_symbols(naming, [spl32_59])])).
fof(f2749, plain, (aSubsetOf0(sK23(xS), cS1395) | (spl32_58 | ~ spl32_59)), inference(subsumption_resolution, [], [f2748, f216])).
fof(f216, plain, aInteger0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aInteger0(sz00), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', mIntZero)).
fof(f2748, plain, (aSubsetOf0(sK23(xS), cS1395) | ~ aInteger0(sz00) | (spl32_58 | ~ spl32_59)), inference(subsumption_resolution, [], [f2747, f1354])).
fof(f1354, plain, (aInteger0(sK26(sK23(xS))) | ~ spl32_59), inference(resolution, [], [f1288, f341])).
fof(f341, plain, ! [X0] : (~ aElementOf0(X0, xS) | aInteger0(sK26(X0))), inference(cnf_transformation, [], [f199])).
fof(f1288, plain, (aElementOf0(sK23(xS), xS) | ~ spl32_59), inference(avatar_component_clause, [], [f1287])).
fof(f2747, plain, (aSubsetOf0(sK23(xS), cS1395) | ~ aInteger0(sK26(sK23(xS))) | ~ aInteger0(sz00) | (spl32_58 | ~ spl32_59)), inference(subsumption_resolution, [], [f2720, f1266])).
fof(f1266, plain, (~ (sz00 = sK26(sK23(xS))) | spl32_58), inference(avatar_component_clause, [], [f1265])).
fof(f2720, plain, (aSubsetOf0(sK23(xS), cS1395) | (sz00 = sK26(sK23(xS))) | ~ aInteger0(sK26(sK23(xS))) | ~ aInteger0(sz00) | ~ spl32_59), inference(superposition, [], [f322, f1351])).
fof(f1351, plain, ((sK23(xS) = szAzrzSzezqlpdtcmdtrp0(sz00, sK26(sK23(xS)))) | ~ spl32_59), inference(resolution, [], [f1288, f346])).
fof(f346, plain, ! [X0] : (~ aElementOf0(X0, xS) | (szAzrzSzezqlpdtcmdtrp0(sz00, sK26(X0)) = X0)), inference(cnf_transformation, [], [f199])).
fof(f322, plain, ! [X0, X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), cS1395) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ! [X0, X1] : ((isClosed0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), cS1395)) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(flattening, [], [f110])).
fof(f110, plain, ! [X0, X1] : ((isClosed0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), cS1395)) | ((sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0))), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0, X1] : ((~ (sz00 = X1) & aInteger0(X1) & aInteger0(X0)) => (isClosed0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), cS1395))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM449+6.p', mArSeqClosed)).
fof(f1363, plain, (~ spl32_58 | ~ spl32_59), inference(avatar_contradiction_clause, [], [f1362])).
fof(f1362, plain, ($false | (~ spl32_58 | ~ spl32_59)), inference(subsumption_resolution, [], [f1361, f1288])).
fof(f1361, plain, (~ aElementOf0(sK23(xS), xS) | ~ spl32_58), inference(trivial_inequality_removal, [], [f1359])).
fof(f1359, plain, (~ (sz00 = sz00) | ~ aElementOf0(sK23(xS), xS) | ~ spl32_58), inference(superposition, [], [f342, f1267])).
fof(f1267, plain, ((sz00 = sK26(sK23(xS))) | ~ spl32_58), inference(avatar_component_clause, [], [f1265])).
fof(f342, plain, ! [X0] : (~ (sz00 = sK26(X0)) | ~ aElementOf0(X0, xS)), inference(cnf_transformation, [], [f199])).
fof(f1342, plain, spl32_59, inference(avatar_contradiction_clause, [], [f1341])).
fof(f1341, plain, ($false | spl32_59), inference(subsumption_resolution, [], [f1340, f340])).
fof(f1340, plain, (~ aSet0(xS) | spl32_59), inference(subsumption_resolution, [], [f1339, f364])).
fof(f1339, plain, (~ isFinite0(xS) | ~ aSet0(xS) | spl32_59), inference(subsumption_resolution, [], [f1338, f388])).
fof(f1338, plain, (isClosed0(sbsmnsldt0(xS)) | ~ isFinite0(xS) | ~ aSet0(xS) | spl32_59), inference(resolution, [], [f1289, f320])).
fof(f320, plain, ! [X0] : (aElementOf0(sK23(X0), X0) | isClosed0(sbsmnsldt0(X0)) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f188])).
fof(f1289, plain, (~ aElementOf0(sK23(xS), xS) | spl32_59), inference(avatar_component_clause, [], [f1287])).
fof(f1304, plain, (spl32_58 | spl32_61), inference(avatar_split_clause, [], [f1299, f1301, f1265])).
fof(f1299, plain, (isClosed0(sK23(xS)) | (sz00 = sK26(sK23(xS)))), inference(subsumption_resolution, [], [f1298, f216])).
fof(f1298, plain, (isClosed0(sK23(xS)) | (sz00 = sK26(sK23(xS))) | ~ aInteger0(sz00)), inference(subsumption_resolution, [], [f1284, f1089])).
fof(f1089, plain, aInteger0(sK26(sK23(xS))), inference(subsumption_resolution, [], [f1088, f340])).
fof(f1088, plain, (~ aSet0(xS) | aInteger0(sK26(sK23(xS)))), inference(subsumption_resolution, [], [f1087, f364])).
fof(f1087, plain, (~ isFinite0(xS) | ~ aSet0(xS) | aInteger0(sK26(sK23(xS)))), inference(subsumption_resolution, [], [f1000, f388])).
fof(f1000, plain, (isClosed0(sbsmnsldt0(xS)) | ~ isFinite0(xS) | ~ aSet0(xS) | aInteger0(sK26(sK23(xS)))), inference(resolution, [], [f320, f341])).
fof(f1284, plain, (isClosed0(sK23(xS)) | (sz00 = sK26(sK23(xS))) | ~ aInteger0(sK26(sK23(xS))) | ~ aInteger0(sz00)), inference(superposition, [], [f323, f1256])).
fof(f1256, plain, (sK23(xS) = szAzrzSzezqlpdtcmdtrp0(sz00, sK26(sK23(xS)))), inference(subsumption_resolution, [], [f1255, f340])).
fof(f1255, plain, ((sK23(xS) = szAzrzSzezqlpdtcmdtrp0(sz00, sK26(sK23(xS)))) | ~ aSet0(xS)), inference(subsumption_resolution, [], [f1254, f364])).
fof(f1254, plain, ((sK23(xS) = szAzrzSzezqlpdtcmdtrp0(sz00, sK26(sK23(xS)))) | ~ isFinite0(xS) | ~ aSet0(xS)), inference(subsumption_resolution, [], [f1253, f388])).
fof(f1253, plain, ((sK23(xS) = szAzrzSzezqlpdtcmdtrp0(sz00, sK26(sK23(xS)))) | isClosed0(sbsmnsldt0(xS)) | ~ isFinite0(xS) | ~ aSet0(xS)), inference(resolution, [], [f346, f320])).
fof(f323, plain, ! [X0, X1] : (isClosed0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (sz00 = X1) | ~ aInteger0(X1) | ~ aInteger0(X0)), inference(cnf_transformation, [], [f111])).