fof(f2299, plain, $false, inference(avatar_sat_refutation, [], [f467, f472, f491, f496, f509, f513, f517, f535, f536, f537, f542, f550, f554, f558, f568, f579, f585, f598, f602, f829, f1246, f1250, f2298])).
fof(f2298, plain, (~ spl38_14 | spl38_18 | ~ spl38_19 | ~ spl38_20 | ~ spl38_21 | ~ spl38_25), inference(avatar_contradiction_clause, [], [f2297])).
fof(f2297, plain, ($false | (~ spl38_14 | spl38_18 | ~ spl38_19 | ~ spl38_20 | ~ spl38_21 | ~ spl38_25)), inference(subsumption_resolution, [], [f2296, f1253])).
fof(f1253, plain, (~ aElementOf0(sK37, xA) | ~ spl38_19), inference(resolution, [], [f529, f360])).
fof(f360, plain, ! [X7] : (~ aElementOf0(X7, stldt0(xA)) | ~ aElementOf0(X7, xA)), inference(cnf_transformation, [], [f196])).
fof(f196, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : ((aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0)), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0)))) & sP7(sK33(X0), X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0))) & ~ (sz00 = sK33(X0)) & aInteger0(sK33(X0))) | ~ aElementOf0(X0, stldt0(xB))) & ! [X3] : ((aElementOf0(X3, stldt0(xB)) | aElementOf0(X3, xB) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, xB) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(xB)))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X4] : ((aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4)), stldt0(xA)) & ! [X6] : (aElementOf0(X6, stldt0(xA)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4)))) & sP6(sK34(X4), X4) & aSet0(szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4))) & ~ (sz00 = sK34(X4)) & aInteger0(sK34(X4))) | ~ aElementOf0(X4, stldt0(xA))) & ! [X7] : ((aElementOf0(X7, stldt0(xA)) | aElementOf0(X7, xA) | ~ aInteger0(X7)) & ((~ aElementOf0(X7, xA) & aInteger0(X7)) | ~ aElementOf0(X7, stldt0(xA)))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X8] : (aElementOf0(X8, cS1395) | ~ aElementOf0(X8, xB)) & aSet0(xB) & ! [X9] : ((aElementOf0(X9, cS1395) | ~ aInteger0(X9)) & (aInteger0(X9) | ~ aElementOf0(X9, cS1395))) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X10] : (aElementOf0(X10, cS1395) | ~ aElementOf0(X10, xA)) & aSet0(xA) & ! [X11] : ((aElementOf0(X11, cS1395) | ~ aInteger0(X11)) & (aInteger0(X11) | ~ aElementOf0(X11, cS1395))) & aSet0(cS1395)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK33, sK34])], [f193, f195, f194])).
fof(f194, plain, ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP7(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0)), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0)))) & sP7(sK33(X0), X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, sK33(X0))) & ~ (sz00 = sK33(X0)) & aInteger0(sK33(X0)))), introduced(choice_axiom, [])).
fof(f195, plain, ! [X4] : (? [X5] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X4, X5), stldt0(xA)) & ! [X6] : (aElementOf0(X6, stldt0(xA)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(X4, X5))) & sP6(X5, X4) & aSet0(szAzrzSzezqlpdtcmdtrp0(X4, X5)) & ~ (sz00 = X5) & aInteger0(X5)) => (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4)), stldt0(xA)) & ! [X6] : (aElementOf0(X6, stldt0(xA)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4)))) & sP6(sK34(X4), X4) & aSet0(szAzrzSzezqlpdtcmdtrp0(X4, sK34(X4))) & ~ (sz00 = sK34(X4)) & aInteger0(sK34(X4)))), introduced(choice_axiom, [])).
fof(f193, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP7(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X3] : ((aElementOf0(X3, stldt0(xB)) | aElementOf0(X3, xB) | ~ aInteger0(X3)) & ((~ aElementOf0(X3, xB) & aInteger0(X3)) | ~ aElementOf0(X3, stldt0(xB)))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X4] : (? [X5] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X4, X5), stldt0(xA)) & ! [X6] : (aElementOf0(X6, stldt0(xA)) | ~ aElementOf0(X6, szAzrzSzezqlpdtcmdtrp0(X4, X5))) & sP6(X5, X4) & aSet0(szAzrzSzezqlpdtcmdtrp0(X4, X5)) & ~ (sz00 = X5) & aInteger0(X5)) | ~ aElementOf0(X4, stldt0(xA))) & ! [X7] : ((aElementOf0(X7, stldt0(xA)) | aElementOf0(X7, xA) | ~ aInteger0(X7)) & ((~ aElementOf0(X7, xA) & aInteger0(X7)) | ~ aElementOf0(X7, stldt0(xA)))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X8] : (aElementOf0(X8, cS1395) | ~ aElementOf0(X8, xB)) & aSet0(xB) & ! [X9] : ((aElementOf0(X9, cS1395) | ~ aInteger0(X9)) & (aInteger0(X9) | ~ aElementOf0(X9, cS1395))) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X10] : (aElementOf0(X10, cS1395) | ~ aElementOf0(X10, xA)) & aSet0(xA) & ! [X11] : ((aElementOf0(X11, cS1395) | ~ aInteger0(X11)) & (aInteger0(X11) | ~ aElementOf0(X11, cS1395))) & aSet0(cS1395)), inference(rectify, [], [f192])).
fof(f192, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP7(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X6] : ((aElementOf0(X6, stldt0(xB)) | aElementOf0(X6, xB) | ~ aInteger0(X6)) & ((~ aElementOf0(X6, xB) & aInteger0(X6)) | ~ aElementOf0(X6, stldt0(xB)))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, stldt0(xA)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & sP6(X8, X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(xA))) & ! [X13] : ((aElementOf0(X13, stldt0(xA)) | aElementOf0(X13, xA) | ~ aInteger0(X13)) & ((~ aElementOf0(X13, xA) & aInteger0(X13)) | ~ aElementOf0(X13, stldt0(xA)))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, cS1395) | ~ aElementOf0(X14, xB)) & aSet0(xB) & ! [X15] : ((aElementOf0(X15, cS1395) | ~ aInteger0(X15)) & (aInteger0(X15) | ~ aElementOf0(X15, cS1395))) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, cS1395) | ~ aElementOf0(X16, xA)) & aSet0(xA) & ! [X17] : ((aElementOf0(X17, cS1395) | ~ aInteger0(X17)) & (aInteger0(X17) | ~ aElementOf0(X17, cS1395))) & aSet0(cS1395)), inference(flattening, [], [f191])).
fof(f191, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP7(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X6] : ((aElementOf0(X6, stldt0(xB)) | (aElementOf0(X6, xB) | ~ aInteger0(X6))) & ((~ aElementOf0(X6, xB) & aInteger0(X6)) | ~ aElementOf0(X6, stldt0(xB)))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, stldt0(xA)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & sP6(X8, X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(xA))) & ! [X13] : ((aElementOf0(X13, stldt0(xA)) | (aElementOf0(X13, xA) | ~ aInteger0(X13))) & ((~ aElementOf0(X13, xA) & aInteger0(X13)) | ~ aElementOf0(X13, stldt0(xA)))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, cS1395) | ~ aElementOf0(X14, xB)) & aSet0(xB) & ! [X15] : ((aElementOf0(X15, cS1395) | ~ aInteger0(X15)) & (aInteger0(X15) | ~ aElementOf0(X15, cS1395))) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, cS1395) | ~ aElementOf0(X16, xA)) & aSet0(xA) & ! [X17] : ((aElementOf0(X17, cS1395) | ~ aInteger0(X17)) & (aInteger0(X17) | ~ aElementOf0(X17, cS1395))) & aSet0(cS1395)), inference(nnf_transformation, [], [f115])).
fof(f115, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & sP7(X1, X0) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X6] : (aElementOf0(X6, stldt0(xB)) <=> (~ aElementOf0(X6, xB) & aInteger0(X6))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, stldt0(xA)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & sP6(X8, X7) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(xA))) & ! [X13] : (aElementOf0(X13, stldt0(xA)) <=> (~ aElementOf0(X13, xA) & aInteger0(X13))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, cS1395) | ~ aElementOf0(X14, xB)) & aSet0(xB) & ! [X15] : (aElementOf0(X15, cS1395) <=> aInteger0(X15)) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, cS1395) | ~ aElementOf0(X16, xA)) & aSet0(xA) & ! [X17] : (aElementOf0(X17, cS1395) <=> aInteger0(X17)) & aSet0(cS1395)), inference(definition_folding, [], [f101, e114, e113])).
fof(f113, plain, ! [X8, X7] : (! [X10] : ((aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)) | (~ sdteqdtlpzmzozddtrp0(X10, X7, X8) & ~ aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ! [X11] : (~ (sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X11)) | ~ aInteger0(X11))) | ~ aInteger0(X10)) & ((sdteqdtlpzmzozddtrp0(X10, X7, X8) & aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ? [X12] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X12)) & aInteger0(X12)) & aInteger0(X10)) | ~ aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)))) | ~ sP6(X8, X7)), inference(usedef, [], [e113])).
fof(e113, plain, ! [X8, X7] : (sP6(X8, X7) <=> ! [X10] : ((aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)) | (~ sdteqdtlpzmzozddtrp0(X10, X7, X8) & ~ aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ! [X11] : (~ (sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X11)) | ~ aInteger0(X11))) | ~ aInteger0(X10)) & ((sdteqdtlpzmzozddtrp0(X10, X7, X8) & aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ? [X12] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X12)) & aInteger0(X12)) & aInteger0(X10)) | ~ aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f114, plain, ! [X1, X0] : (! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) | ~ sP7(X1, X0)), inference(usedef, [], [e114])).
fof(e114, plain, ! [X1, X0] : (sP7(X1, X0) <=> ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f101, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | (~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3)) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X6] : (aElementOf0(X6, stldt0(xB)) <=> (~ aElementOf0(X6, xB) & aInteger0(X6))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, stldt0(xA)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & ! [X10] : ((aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)) | (~ sdteqdtlpzmzozddtrp0(X10, X7, X8) & ~ aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ! [X11] : (~ (sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X11)) | ~ aInteger0(X11))) | ~ aInteger0(X10)) & ((sdteqdtlpzmzozddtrp0(X10, X7, X8) & aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ? [X12] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X12)) & aInteger0(X12)) & aInteger0(X10)) | ~ aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(xA))) & ! [X13] : (aElementOf0(X13, stldt0(xA)) <=> (~ aElementOf0(X13, xA) & aInteger0(X13))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, cS1395) | ~ aElementOf0(X14, xB)) & aSet0(xB) & ! [X15] : (aElementOf0(X15, cS1395) <=> aInteger0(X15)) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, cS1395) | ~ aElementOf0(X16, xA)) & aSet0(xA) & ! [X17] : (aElementOf0(X17, cS1395) <=> aInteger0(X17)) & aSet0(cS1395)), inference(flattening, [], [f100])).
fof(f100, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, stldt0(xB)) | ~ aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & ! [X3] : ((aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) | ((~ sdteqdtlpzmzozddtrp0(X3, X0, X1) & ~ aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ! [X4] : (~ (sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X4)) | ~ aInteger0(X4))) | ~ aInteger0(X3))) & ((sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)) | ~ aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1)) | ~ aElementOf0(X0, stldt0(xB))) & ! [X6] : (aElementOf0(X6, stldt0(xB)) <=> (~ aElementOf0(X6, xB) & aInteger0(X6))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, stldt0(xA)) | ~ aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & ! [X10] : ((aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)) | ((~ sdteqdtlpzmzozddtrp0(X10, X7, X8) & ~ aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ! [X11] : (~ (sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X11)) | ~ aInteger0(X11))) | ~ aInteger0(X10))) & ((sdteqdtlpzmzozddtrp0(X10, X7, X8) & aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ? [X12] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X12)) & aInteger0(X12)) & aInteger0(X10)) | ~ aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8)) | ~ aElementOf0(X7, stldt0(xA))) & ! [X13] : (aElementOf0(X13, stldt0(xA)) <=> (~ aElementOf0(X13, xA) & aInteger0(X13))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, cS1395) | ~ aElementOf0(X14, xB)) & aSet0(xB) & ! [X15] : (aElementOf0(X15, cS1395) <=> aInteger0(X15)) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, cS1395) | ~ aElementOf0(X16, xA)) & aSet0(xA) & ! [X17] : (aElementOf0(X17, cS1395) <=> aInteger0(X17)) & aSet0(cS1395)), inference(ennf_transformation, [], [f48])).
fof(f48, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (aElementOf0(X0, stldt0(xB)) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(xB))) & ! [X3] : ((((sdteqdtlpzmzozddtrp0(X3, X0, X1) | aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) | ? [X4] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X4)) & aInteger0(X4))) & aInteger0(X3)) => aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X3, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X3, X0, X1) & aDivisorOf0(X1, sdtpldt0(X3, smndt0(X0))) & ? [X5] : ((sdtpldt0(X3, smndt0(X0)) = sdtasdt0(X1, X5)) & aInteger0(X5)) & aInteger0(X3)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X6] : (aElementOf0(X6, stldt0(xB)) <=> (~ aElementOf0(X6, xB) & aInteger0(X6))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X7] : (aElementOf0(X7, stldt0(xA)) => ? [X8] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X7, X8), stldt0(xA)) & ! [X9] : (aElementOf0(X9, szAzrzSzezqlpdtcmdtrp0(X7, X8)) => aElementOf0(X9, stldt0(xA))) & ! [X10] : ((((sdteqdtlpzmzozddtrp0(X10, X7, X8) | aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) | ? [X11] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X11)) & aInteger0(X11))) & aInteger0(X10)) => aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8))) & (aElementOf0(X10, szAzrzSzezqlpdtcmdtrp0(X7, X8)) => (sdteqdtlpzmzozddtrp0(X10, X7, X8) & aDivisorOf0(X8, sdtpldt0(X10, smndt0(X7))) & ? [X12] : ((sdtpldt0(X10, smndt0(X7)) = sdtasdt0(X8, X12)) & aInteger0(X12)) & aInteger0(X10)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X7, X8)) & ~ (sz00 = X8) & aInteger0(X8))) & ! [X13] : (aElementOf0(X13, stldt0(xA)) <=> (~ aElementOf0(X13, xA) & aInteger0(X13))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X14] : (aElementOf0(X14, xB) => aElementOf0(X14, cS1395)) & aSet0(xB) & ! [X15] : (aElementOf0(X15, cS1395) <=> aInteger0(X15)) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X16] : (aElementOf0(X16, xA) => aElementOf0(X16, cS1395)) & aSet0(xA) & ! [X17] : (aElementOf0(X17, cS1395) <=> aInteger0(X17)) & aSet0(cS1395)), inference(rectify, [], [f39])).
fof(f39, plain, (isClosed0(xB) & isOpen0(stldt0(xB)) & ! [X0] : (aElementOf0(X0, stldt0(xB)) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xB)) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(xB))) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X0] : (aElementOf0(X0, stldt0(xB)) <=> (~ aElementOf0(X0, xB) & aInteger0(X0))) & aSet0(stldt0(xB)) & isClosed0(xA) & isOpen0(stldt0(xA)) & ! [X0] : (aElementOf0(X0, stldt0(xA)) => ? [X1] : (aSubsetOf0(szAzrzSzezqlpdtcmdtrp0(X0, X1), stldt0(xA)) & ! [X2] : (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => aElementOf0(X2, stldt0(xA))) & ! [X2] : ((((sdteqdtlpzmzozddtrp0(X2, X0, X1) | aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) | ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3))) & aInteger0(X2)) => aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1))) & (aElementOf0(X2, szAzrzSzezqlpdtcmdtrp0(X0, X1)) => (sdteqdtlpzmzozddtrp0(X2, X0, X1) & aDivisorOf0(X1, sdtpldt0(X2, smndt0(X0))) & ? [X3] : ((sdtasdt0(X1, X3) = sdtpldt0(X2, smndt0(X0))) & aInteger0(X3)) & aInteger0(X2)))) & aSet0(szAzrzSzezqlpdtcmdtrp0(X0, X1)) & ~ (sz00 = X1) & aInteger0(X1))) & ! [X0] : (aElementOf0(X0, stldt0(xA)) <=> (~ aElementOf0(X0, xA) & aInteger0(X0))) & aSet0(stldt0(xA)) & aSubsetOf0(xB, cS1395) & ! [X0] : (aElementOf0(X0, xB) => aElementOf0(X0, cS1395)) & aSet0(xB) & ! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395) & aSubsetOf0(xA, cS1395) & ! [X0] : (aElementOf0(X0, xA) => aElementOf0(X0, cS1395)) & aSet0(xA) & ! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM440+6.p', m__1826)).
fof(f529, plain, (aElementOf0(sK37, stldt0(xA)) | ~ spl38_19), inference(avatar_component_clause, [], [f528])).
fof(f528, plain, (spl38_19 <=> aElementOf0(sK37, stldt0(xA))), introduced(avatar_definition, [new_symbols(naming, [spl38_19])])).
fof(f2296, plain, (aElementOf0(sK37, xA) | (~ spl38_14 | spl38_18 | ~ spl38_20 | ~ spl38_21 | ~ spl38_25)), inference(subsumption_resolution, [], [f2294, f1259])).
fof(f1259, plain, (~ aElementOf0(sK37, xB) | ~ spl38_20), inference(resolution, [], [f533, f372])).
fof(f372, plain, ! [X3] : (~ aElementOf0(X3, stldt0(xB)) | ~ aElementOf0(X3, xB)), inference(cnf_transformation, [], [f196])).
fof(f533, plain, (aElementOf0(sK37, stldt0(xB)) | ~ spl38_20), inference(avatar_component_clause, [], [f532])).
fof(f532, plain, (spl38_20 <=> aElementOf0(sK37, stldt0(xB))), introduced(avatar_definition, [new_symbols(naming, [spl38_20])])).
fof(f2294, plain, (aElementOf0(sK37, xB) | aElementOf0(sK37, xA) | (~ spl38_14 | spl38_18 | ~ spl38_21 | ~ spl38_25)), inference(resolution, [], [f2074, f557])).
fof(f557, plain, (! [X0] : (~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | aElementOf0(X0, xB) | aElementOf0(X0, xA)) | ~ spl38_25), inference(avatar_component_clause, [], [f556])).
fof(f556, plain, (spl38_25 <=> ! [X0] : (aElementOf0(X0, xB) | ~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | aElementOf0(X0, xA))), introduced(avatar_definition, [new_symbols(naming, [spl38_25])])).
fof(f2074, plain, (aElementOf0(sK37, sdtbsmnsldt0(xA, xB)) | (~ spl38_14 | spl38_18 | ~ spl38_21)), inference(subsumption_resolution, [], [f2072, f541])).
fof(f541, plain, (aInteger0(sK37) | ~ spl38_21), inference(avatar_component_clause, [], [f539])).
fof(f539, plain, (spl38_21 <=> aInteger0(sK37)), introduced(avatar_definition, [new_symbols(naming, [spl38_21])])).
fof(f2072, plain, (~ aInteger0(sK37) | aElementOf0(sK37, sdtbsmnsldt0(xA, xB)) | (~ spl38_14 | spl38_18)), inference(resolution, [], [f508, f526])).
fof(f526, plain, (~ aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | spl38_18), inference(avatar_component_clause, [], [f524])).
fof(f524, plain, (spl38_18 <=> aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB)))), introduced(avatar_definition, [new_symbols(naming, [spl38_18])])).
fof(f508, plain, (! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | ~ aInteger0(X0) | aElementOf0(X0, sdtbsmnsldt0(xA, xB))) | ~ spl38_14), inference(avatar_component_clause, [], [f507])).
fof(f507, plain, (spl38_14 <=> ! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | ~ aInteger0(X0) | aElementOf0(X0, sdtbsmnsldt0(xA, xB)))), introduced(avatar_definition, [new_symbols(naming, [spl38_14])])).
fof(f1250, plain, (~ spl38_15 | ~ spl38_18 | spl38_19 | ~ spl38_21 | ~ spl38_24), inference(avatar_contradiction_clause, [], [f1249])).
fof(f1249, plain, ($false | (~ spl38_15 | ~ spl38_18 | spl38_19 | ~ spl38_21 | ~ spl38_24)), inference(subsumption_resolution, [], [f1248, f541])).
fof(f1248, plain, (~ aInteger0(sK37) | (~ spl38_15 | ~ spl38_18 | spl38_19 | ~ spl38_21 | ~ spl38_24)), inference(subsumption_resolution, [], [f1247, f1243])).
fof(f1243, plain, (~ aElementOf0(sK37, xA) | (~ spl38_15 | ~ spl38_18 | ~ spl38_21 | ~ spl38_24)), inference(subsumption_resolution, [], [f1241, f541])).
fof(f1241, plain, (~ aInteger0(sK37) | ~ aElementOf0(sK37, xA) | (~ spl38_15 | ~ spl38_18 | ~ spl38_24)), inference(resolution, [], [f1234, f553])).
fof(f553, plain, (! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0) | ~ aElementOf0(X0, xA)) | ~ spl38_24), inference(avatar_component_clause, [], [f552])).
fof(f552, plain, (spl38_24 <=> ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0) | ~ aElementOf0(X0, xA))), introduced(avatar_definition, [new_symbols(naming, [spl38_24])])).
fof(f1234, plain, (~ aElementOf0(sK37, sdtbsmnsldt0(xA, xB)) | (~ spl38_15 | ~ spl38_18)), inference(resolution, [], [f512, f525])).
fof(f525, plain, (aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ spl38_18), inference(avatar_component_clause, [], [f524])).
fof(f512, plain, (! [X0] : (~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | ~ aElementOf0(X0, sdtbsmnsldt0(xA, xB))) | ~ spl38_15), inference(avatar_component_clause, [], [f511])).
fof(f511, plain, (spl38_15 <=> ! [X0] : (~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))))), introduced(avatar_definition, [new_symbols(naming, [spl38_15])])).
fof(f1247, plain, (aElementOf0(sK37, xA) | ~ aInteger0(sK37) | spl38_19), inference(resolution, [], [f530, f361])).
fof(f361, plain, ! [X7] : (aElementOf0(X7, stldt0(xA)) | aElementOf0(X7, xA) | ~ aInteger0(X7)), inference(cnf_transformation, [], [f196])).
fof(f530, plain, (~ aElementOf0(sK37, stldt0(xA)) | spl38_19), inference(avatar_component_clause, [], [f528])).
fof(f1246, plain, (~ spl38_15 | ~ spl38_18 | spl38_20 | ~ spl38_21 | ~ spl38_23), inference(avatar_contradiction_clause, [], [f1245])).
fof(f1245, plain, ($false | (~ spl38_15 | ~ spl38_18 | spl38_20 | ~ spl38_21 | ~ spl38_23)), inference(subsumption_resolution, [], [f1244, f857])).
fof(f857, plain, (aElementOf0(sK37, xB) | (spl38_20 | ~ spl38_21)), inference(subsumption_resolution, [], [f856, f541])).
fof(f856, plain, (aElementOf0(sK37, xB) | ~ aInteger0(sK37) | spl38_20), inference(resolution, [], [f373, f534])).
fof(f534, plain, (~ aElementOf0(sK37, stldt0(xB)) | spl38_20), inference(avatar_component_clause, [], [f532])).
fof(f373, plain, ! [X3] : (aElementOf0(X3, stldt0(xB)) | aElementOf0(X3, xB) | ~ aInteger0(X3)), inference(cnf_transformation, [], [f196])).
fof(f1244, plain, (~ aElementOf0(sK37, xB) | (~ spl38_15 | ~ spl38_18 | ~ spl38_21 | ~ spl38_23)), inference(subsumption_resolution, [], [f1242, f541])).
fof(f1242, plain, (~ aInteger0(sK37) | ~ aElementOf0(sK37, xB) | (~ spl38_15 | ~ spl38_18 | ~ spl38_23)), inference(resolution, [], [f1234, f549])).
fof(f549, plain, (! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0) | ~ aElementOf0(X0, xB)) | ~ spl38_23), inference(avatar_component_clause, [], [f548])).
fof(f548, plain, (spl38_23 <=> ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0) | ~ aElementOf0(X0, xB))), introduced(avatar_definition, [new_symbols(naming, [spl38_23])])).
fof(f829, plain, (spl38_21 | ~ spl38_16 | ~ spl38_18), inference(avatar_split_clause, [], [f818, f524, f515, f539])).
fof(f515, plain, (spl38_16 <=> ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))))), introduced(avatar_definition, [new_symbols(naming, [spl38_16])])).
fof(f818, plain, (aInteger0(sK37) | (~ spl38_16 | ~ spl38_18)), inference(resolution, [], [f516, f525])).
fof(f516, plain, (! [X0] : (~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | aInteger0(X0)) | ~ spl38_16), inference(avatar_component_clause, [], [f515])).
fof(f602, plain, (spl38_10 | ~ spl38_11), inference(avatar_contradiction_clause, [], [f601])).
fof(f601, plain, ($false | (spl38_10 | ~ spl38_11)), inference(subsumption_resolution, [], [f600, f599])).
fof(f599, plain, (~ aInteger0(sK36) | spl38_10), inference(resolution, [], [f490, f354])).
fof(f354, plain, ! [X9] : (aElementOf0(X9, cS1395) | ~ aInteger0(X9)), inference(cnf_transformation, [], [f196])).
fof(f490, plain, (~ aElementOf0(sK36, cS1395) | spl38_10), inference(avatar_component_clause, [], [f488])).
fof(f488, plain, (spl38_10 <=> aElementOf0(sK36, cS1395)), introduced(avatar_definition, [new_symbols(naming, [spl38_10])])).
fof(f600, plain, (aInteger0(sK36) | ~ spl38_11), inference(resolution, [], [f495, f359])).
fof(f359, plain, ! [X7] : (~ aElementOf0(X7, stldt0(xA)) | aInteger0(X7)), inference(cnf_transformation, [], [f196])).
fof(f495, plain, (aElementOf0(sK36, stldt0(xA)) | ~ spl38_11), inference(avatar_component_clause, [], [f493])).
fof(f493, plain, (spl38_11 <=> aElementOf0(sK36, stldt0(xA))), introduced(avatar_definition, [new_symbols(naming, [spl38_11])])).
fof(f598, plain, (spl38_5 | ~ spl38_6), inference(avatar_contradiction_clause, [], [f597])).
fof(f597, plain, ($false | (spl38_5 | ~ spl38_6)), inference(subsumption_resolution, [], [f596, f591])).
fof(f591, plain, (~ aInteger0(sK35) | spl38_5), inference(resolution, [], [f354, f466])).
fof(f466, plain, (~ aElementOf0(sK35, cS1395) | spl38_5), inference(avatar_component_clause, [], [f464])).
fof(f464, plain, (spl38_5 <=> aElementOf0(sK35, cS1395)), introduced(avatar_definition, [new_symbols(naming, [spl38_5])])).
fof(f596, plain, (aInteger0(sK35) | ~ spl38_6), inference(resolution, [], [f371, f471])).
fof(f471, plain, (aElementOf0(sK35, stldt0(xB)) | ~ spl38_6), inference(avatar_component_clause, [], [f469])).
fof(f469, plain, (spl38_6 <=> aElementOf0(sK35, stldt0(xB))), introduced(avatar_definition, [new_symbols(naming, [spl38_6])])).
fof(f371, plain, ! [X3] : (~ aElementOf0(X3, stldt0(xB)) | aInteger0(X3)), inference(cnf_transformation, [], [f196])).
fof(f585, plain, (spl38_8 | spl38_3 | spl38_22), inference(avatar_split_clause, [], [f422, f544, f455, f479])).
fof(f479, plain, (spl38_8 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl38_8])])).
fof(f455, plain, (spl38_3 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl38_3])])).
fof(f544, plain, (spl38_22 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl38_22])])).
fof(f422, plain, (sP10 | sP16 | sP15), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ((~ (stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) & sP11 & sP14 & sP13 & sP12 & aSet0(stldt0(sdtbsmnsldt0(xA, xB))) & sP10 & aSet0(sdtbsmnsldt0(xA, xB))) | sP16 | sP15), inference(definition_folding, [], [f103, e124, e123, e122, e121, e120, e119, e118, e117, e116])).
fof(f116, plain, (! [X8] : (aElementOf0(X8, stldt0(xA)) <=> (~ aElementOf0(X8, xA) & aInteger0(X8))) | ~ sP8), inference(usedef, [], [e116])).
fof(e116, plain, (sP8 <=> ! [X8] : (aElementOf0(X8, stldt0(xA)) <=> (~ aElementOf0(X8, xA) & aInteger0(X8)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f117, plain, (! [X5] : (aElementOf0(X5, stldt0(xB)) <=> (~ aElementOf0(X5, xB) & aInteger0(X5))) | ~ sP9), inference(usedef, [], [e117])).
fof(e117, plain, (sP9 <=> ! [X5] : (aElementOf0(X5, stldt0(xB)) <=> (~ aElementOf0(X5, xB) & aInteger0(X5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f118, plain, (! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) | ~ sP10), inference(usedef, [], [e118])).
fof(e118, plain, (sP10 <=> ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f119, plain, (? [X4] : ~ (aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4))) | ~ sP11), inference(usedef, [], [e119])).
fof(e119, plain, (sP11 <=> ? [X4] : ~ (aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f120, plain, (! [X1] : (aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1))) | ~ sP12), inference(usedef, [], [e120])).
fof(e120, plain, (sP12 <=> ! [X1] : (aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f121, plain, (! [X2] : (aElementOf0(X2, stldt0(xA)) <=> (~ aElementOf0(X2, xA) & aInteger0(X2))) | ~ sP13), inference(usedef, [], [e121])).
fof(e121, plain, (sP13 <=> ! [X2] : (aElementOf0(X2, stldt0(xA)) <=> (~ aElementOf0(X2, xA) & aInteger0(X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f122, plain, (! [X3] : (aElementOf0(X3, stldt0(xB)) <=> (~ aElementOf0(X3, xB) & aInteger0(X3))) | ~ sP14), inference(usedef, [], [e122])).
fof(e122, plain, (sP14 <=> ! [X3] : (aElementOf0(X3, stldt0(xB)) <=> (~ aElementOf0(X3, xB) & aInteger0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f123, plain, ((~ aSubsetOf0(stldt0(xA), cS1395) & ? [X10] : (~ aElementOf0(X10, cS1395) & aElementOf0(X10, stldt0(xA))) & ! [X9] : (aElementOf0(X9, cS1395) <=> aInteger0(X9)) & aSet0(cS1395) & sP8 & aSet0(stldt0(xA))) | ~ sP15), inference(usedef, [], [e123])).
fof(e123, plain, (sP15 <=> (~ aSubsetOf0(stldt0(xA), cS1395) & ? [X10] : (~ aElementOf0(X10, cS1395) & aElementOf0(X10, stldt0(xA))) & ! [X9] : (aElementOf0(X9, cS1395) <=> aInteger0(X9)) & aSet0(cS1395) & sP8 & aSet0(stldt0(xA)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f124, plain, ((~ aSubsetOf0(stldt0(xB), cS1395) & ? [X7] : (~ aElementOf0(X7, cS1395) & aElementOf0(X7, stldt0(xB))) & ! [X6] : (aElementOf0(X6, cS1395) <=> aInteger0(X6)) & aSet0(cS1395) & sP9 & aSet0(stldt0(xB))) | ~ sP16), inference(usedef, [], [e124])).
fof(e124, plain, (sP16 <=> (~ aSubsetOf0(stldt0(xB), cS1395) & ? [X7] : (~ aElementOf0(X7, cS1395) & aElementOf0(X7, stldt0(xB))) & ! [X6] : (aElementOf0(X6, cS1395) <=> aInteger0(X6)) & aSet0(cS1395) & sP9 & aSet0(stldt0(xB)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f103, plain, ((~ (stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) & ? [X4] : ~ (aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4))) & ! [X3] : (aElementOf0(X3, stldt0(xB)) <=> (~ aElementOf0(X3, xB) & aInteger0(X3))) & ! [X2] : (aElementOf0(X2, stldt0(xA)) <=> (~ aElementOf0(X2, xA) & aInteger0(X2))) & ! [X1] : (aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1))) & aSet0(stldt0(sdtbsmnsldt0(xA, xB))) & ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) & aSet0(sdtbsmnsldt0(xA, xB))) | (~ aSubsetOf0(stldt0(xB), cS1395) & ? [X7] : (~ aElementOf0(X7, cS1395) & aElementOf0(X7, stldt0(xB))) & ! [X6] : (aElementOf0(X6, cS1395) <=> aInteger0(X6)) & aSet0(cS1395) & ! [X5] : (aElementOf0(X5, stldt0(xB)) <=> (~ aElementOf0(X5, xB) & aInteger0(X5))) & aSet0(stldt0(xB))) | (~ aSubsetOf0(stldt0(xA), cS1395) & ? [X10] : (~ aElementOf0(X10, cS1395) & aElementOf0(X10, stldt0(xA))) & ! [X9] : (aElementOf0(X9, cS1395) <=> aInteger0(X9)) & aSet0(cS1395) & ! [X8] : (aElementOf0(X8, stldt0(xA)) <=> (~ aElementOf0(X8, xA) & aInteger0(X8))) & aSet0(stldt0(xA)))), inference(flattening, [], [f102])).
fof(f102, plain, ((((((~ (stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) & ? [X4] : ~ (aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4)))) & ! [X3] : (aElementOf0(X3, stldt0(xB)) <=> (~ aElementOf0(X3, xB) & aInteger0(X3)))) & ! [X2] : (aElementOf0(X2, stldt0(xA)) <=> (~ aElementOf0(X2, xA) & aInteger0(X2)))) & (! [X1] : (aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1))) & aSet0(stldt0(sdtbsmnsldt0(xA, xB))))) & (! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) & aSet0(sdtbsmnsldt0(xA, xB)))) | (((~ aSubsetOf0(stldt0(xB), cS1395) & ? [X7] : (~ aElementOf0(X7, cS1395) & aElementOf0(X7, stldt0(xB)))) & (! [X6] : (aElementOf0(X6, cS1395) <=> aInteger0(X6)) & aSet0(cS1395))) & (! [X5] : (aElementOf0(X5, stldt0(xB)) <=> (~ aElementOf0(X5, xB) & aInteger0(X5))) & aSet0(stldt0(xB)))) | (((~ aSubsetOf0(stldt0(xA), cS1395) & ? [X10] : (~ aElementOf0(X10, cS1395) & aElementOf0(X10, stldt0(xA)))) & (! [X9] : (aElementOf0(X9, cS1395) <=> aInteger0(X9)) & aSet0(cS1395))) & (! [X8] : (aElementOf0(X8, stldt0(xA)) <=> (~ aElementOf0(X8, xA) & aInteger0(X8))) & aSet0(stldt0(xA))))), inference(ennf_transformation, [], [f49])).
fof(f49, plain, ~ (((! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) & aSet0(sdtbsmnsldt0(xA, xB))) => ((! [X1] : (aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1))) & aSet0(stldt0(sdtbsmnsldt0(xA, xB)))) => (! [X2] : (aElementOf0(X2, stldt0(xA)) <=> (~ aElementOf0(X2, xA) & aInteger0(X2))) => (! [X3] : (aElementOf0(X3, stldt0(xB)) <=> (~ aElementOf0(X3, xB) & aInteger0(X3))) => ((stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) | ! [X4] : (aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4)))))))) & ((! [X5] : (aElementOf0(X5, stldt0(xB)) <=> (~ aElementOf0(X5, xB) & aInteger0(X5))) & aSet0(stldt0(xB))) => ((! [X6] : (aElementOf0(X6, cS1395) <=> aInteger0(X6)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xB), cS1395) | ! [X7] : (aElementOf0(X7, stldt0(xB)) => aElementOf0(X7, cS1395))))) & ((! [X8] : (aElementOf0(X8, stldt0(xA)) <=> (~ aElementOf0(X8, xA) & aInteger0(X8))) & aSet0(stldt0(xA))) => ((! [X9] : (aElementOf0(X9, cS1395) <=> aInteger0(X9)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xA), cS1395) | ! [X10] : (aElementOf0(X10, stldt0(xA)) => aElementOf0(X10, cS1395)))))), inference(rectify, [], [f41])).
fof(f41, plain, ~ (((! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) & aSet0(sdtbsmnsldt0(xA, xB))) => ((! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) & aInteger0(X0))) & aSet0(stldt0(sdtbsmnsldt0(xA, xB)))) => (! [X0] : (aElementOf0(X0, stldt0(xA)) <=> (~ aElementOf0(X0, xA) & aInteger0(X0))) => (! [X0] : (aElementOf0(X0, stldt0(xB)) <=> (~ aElementOf0(X0, xB) & aInteger0(X0))) => ((stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) | ! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X0, stldt0(xB)) & aElementOf0(X0, stldt0(xA)) & aInteger0(X0)))))))) & ((! [X0] : (aElementOf0(X0, stldt0(xB)) <=> (~ aElementOf0(X0, xB) & aInteger0(X0))) & aSet0(stldt0(xB))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xB), cS1395) | ! [X0] : (aElementOf0(X0, stldt0(xB)) => aElementOf0(X0, cS1395))))) & ((! [X0] : (aElementOf0(X0, stldt0(xA)) <=> (~ aElementOf0(X0, xA) & aInteger0(X0))) & aSet0(stldt0(xA))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xA), cS1395) | ! [X0] : (aElementOf0(X0, stldt0(xA)) => aElementOf0(X0, cS1395)))))), inference(negated_conjecture, [], [f40])).
fof(f40, plain, ~ (((! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) <=> ((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0))) & aSet0(sdtbsmnsldt0(xA, xB))) => ((! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) <=> (~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) & aInteger0(X0))) & aSet0(stldt0(sdtbsmnsldt0(xA, xB)))) => (! [X0] : (aElementOf0(X0, stldt0(xA)) <=> (~ aElementOf0(X0, xA) & aInteger0(X0))) => (! [X0] : (aElementOf0(X0, stldt0(xB)) <=> (~ aElementOf0(X0, xB) & aInteger0(X0))) => ((stldt0(sdtbsmnsldt0(xA, xB)) = sdtslmnbsdt0(stldt0(xA), stldt0(xB))) | ! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) <=> (aElementOf0(X0, stldt0(xB)) & aElementOf0(X0, stldt0(xA)) & aInteger0(X0)))))))) & ((! [X0] : (aElementOf0(X0, stldt0(xB)) <=> (~ aElementOf0(X0, xB) & aInteger0(X0))) & aSet0(stldt0(xB))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xB), cS1395) | ! [X0] : (aElementOf0(X0, stldt0(xB)) => aElementOf0(X0, cS1395))))) & ((! [X0] : (aElementOf0(X0, stldt0(xA)) <=> (~ aElementOf0(X0, xA) & aInteger0(X0))) & aSet0(stldt0(xA))) => ((! [X0] : (aElementOf0(X0, cS1395) <=> aInteger0(X0)) & aSet0(cS1395)) => (aSubsetOf0(stldt0(xA), cS1395) | ! [X0] : (aElementOf0(X0, stldt0(xA)) => aElementOf0(X0, cS1395)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM440+6.p', m__)).
fof(f579, plain, (spl38_8 | spl38_3 | spl38_13), inference(avatar_split_clause, [], [f424, f503, f455, f479])).
fof(f503, plain, (spl38_13 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl38_13])])).
fof(f424, plain, (sP12 | sP16 | sP15), inference(cnf_transformation, [], [f125])).
fof(f568, plain, (spl38_8 | spl38_3 | spl38_17), inference(avatar_split_clause, [], [f427, f520, f455, f479])).
fof(f520, plain, (spl38_17 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl38_17])])).
fof(f427, plain, (sP11 | sP16 | sP15), inference(cnf_transformation, [], [f125])).
fof(f558, plain, (~ spl38_22 | spl38_25), inference(avatar_split_clause, [], [f412, f556, f544])).
fof(f412, plain, ! [X0] : (aElementOf0(X0, xB) | aElementOf0(X0, xA) | ~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ sP10), inference(cnf_transformation, [], [f220])).
fof(f220, plain, (! [X0] : ((aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | (~ aElementOf0(X0, xB) & ~ aElementOf0(X0, xA)) | ~ aInteger0(X0)) & (((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0)) | ~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)))) | ~ sP10), inference(flattening, [], [f219])).
fof(f219, plain, (! [X0] : ((aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ((~ aElementOf0(X0, xB) & ~ aElementOf0(X0, xA)) | ~ aInteger0(X0))) & (((aElementOf0(X0, xB) | aElementOf0(X0, xA)) & aInteger0(X0)) | ~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)))) | ~ sP10), inference(nnf_transformation, [], [f118])).
fof(f554, plain, (~ spl38_22 | spl38_24), inference(avatar_split_clause, [], [f413, f552, f544])).
fof(f413, plain, ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aElementOf0(X0, xA) | ~ aInteger0(X0) | ~ sP10), inference(cnf_transformation, [], [f220])).
fof(f550, plain, (~ spl38_22 | spl38_23), inference(avatar_split_clause, [], [f414, f548, f544])).
fof(f414, plain, ! [X0] : (aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aElementOf0(X0, xB) | ~ aInteger0(X0) | ~ sP10), inference(cnf_transformation, [], [f220])).
fof(f542, plain, (~ spl38_17 | spl38_18 | spl38_21), inference(avatar_split_clause, [], [f407, f539, f524, f520])).
fof(f407, plain, (aInteger0(sK37) | aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP11), inference(cnf_transformation, [], [f218])).
fof(f218, plain, (((~ aElementOf0(sK37, stldt0(xB)) | ~ aElementOf0(sK37, stldt0(xA)) | ~ aInteger0(sK37) | ~ aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(sK37, stldt0(xB)) & aElementOf0(sK37, stldt0(xA)) & aInteger0(sK37)) | aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP11), inference(skolemisation, [status(esa), new_symbols(skolem, [sK37])], [f216, f217])).
fof(f217, plain, (? [X0] : ((~ aElementOf0(X0, stldt0(xB)) | ~ aElementOf0(X0, stldt0(xA)) | ~ aInteger0(X0) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(X0, stldt0(xB)) & aElementOf0(X0, stldt0(xA)) & aInteger0(X0)) | aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))))) => ((~ aElementOf0(sK37, stldt0(xB)) | ~ aElementOf0(sK37, stldt0(xA)) | ~ aInteger0(sK37) | ~ aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(sK37, stldt0(xB)) & aElementOf0(sK37, stldt0(xA)) & aInteger0(sK37)) | aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB)))))), introduced(choice_axiom, [])).
fof(f216, plain, (? [X0] : ((~ aElementOf0(X0, stldt0(xB)) | ~ aElementOf0(X0, stldt0(xA)) | ~ aInteger0(X0) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(X0, stldt0(xB)) & aElementOf0(X0, stldt0(xA)) & aInteger0(X0)) | aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP11), inference(rectify, [], [f215])).
fof(f215, plain, (? [X4] : ((~ aElementOf0(X4, stldt0(xB)) | ~ aElementOf0(X4, stldt0(xA)) | ~ aInteger0(X4) | ~ aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4)) | aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP11), inference(flattening, [], [f214])).
fof(f214, plain, (? [X4] : (((~ aElementOf0(X4, stldt0(xB)) | ~ aElementOf0(X4, stldt0(xA)) | ~ aInteger0(X4)) | ~ aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB)))) & ((aElementOf0(X4, stldt0(xB)) & aElementOf0(X4, stldt0(xA)) & aInteger0(X4)) | aElementOf0(X4, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP11), inference(nnf_transformation, [], [f119])).
fof(f537, plain, (~ spl38_17 | spl38_18 | spl38_19), inference(avatar_split_clause, [], [f408, f528, f524, f520])).
fof(f408, plain, (aElementOf0(sK37, stldt0(xA)) | aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP11), inference(cnf_transformation, [], [f218])).
fof(f536, plain, (~ spl38_17 | spl38_18 | spl38_20), inference(avatar_split_clause, [], [f409, f532, f524, f520])).
fof(f409, plain, (aElementOf0(sK37, stldt0(xB)) | aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP11), inference(cnf_transformation, [], [f218])).
fof(f535, plain, (~ spl38_17 | ~ spl38_18 | ~ spl38_19 | ~ spl38_20), inference(avatar_split_clause, [], [f518, f532, f528, f524, f520])).
fof(f518, plain, (~ aElementOf0(sK37, stldt0(xB)) | ~ aElementOf0(sK37, stldt0(xA)) | ~ aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP11), inference(subsumption_resolution, [], [f410, f359])).
fof(f410, plain, (~ aElementOf0(sK37, stldt0(xB)) | ~ aElementOf0(sK37, stldt0(xA)) | ~ aInteger0(sK37) | ~ aElementOf0(sK37, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP11), inference(cnf_transformation, [], [f218])).
fof(f517, plain, (~ spl38_13 | spl38_16), inference(avatar_split_clause, [], [f404, f515, f503])).
fof(f404, plain, ! [X0] : (aInteger0(X0) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP12), inference(cnf_transformation, [], [f213])).
fof(f213, plain, (! [X0] : ((aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0)) & ((~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) & aInteger0(X0)) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP12), inference(rectify, [], [f212])).
fof(f212, plain, (! [X1] : ((aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) | aElementOf0(X1, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X1)) & ((~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP12), inference(flattening, [], [f211])).
fof(f211, plain, (! [X1] : ((aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))) | (aElementOf0(X1, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X1))) & ((~ aElementOf0(X1, sdtbsmnsldt0(xA, xB)) & aInteger0(X1)) | ~ aElementOf0(X1, stldt0(sdtbsmnsldt0(xA, xB))))) | ~ sP12), inference(nnf_transformation, [], [f120])).
fof(f513, plain, (~ spl38_13 | spl38_15), inference(avatar_split_clause, [], [f405, f511, f503])).
fof(f405, plain, ! [X0] : (~ aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | ~ sP12), inference(cnf_transformation, [], [f213])).
fof(f509, plain, (~ spl38_13 | spl38_14), inference(avatar_split_clause, [], [f406, f507, f503])).
fof(f406, plain, ! [X0] : (aElementOf0(X0, stldt0(sdtbsmnsldt0(xA, xB))) | aElementOf0(X0, sdtbsmnsldt0(xA, xB)) | ~ aInteger0(X0) | ~ sP12), inference(cnf_transformation, [], [f213])).
fof(f496, plain, (~ spl38_8 | spl38_11), inference(avatar_split_clause, [], [f395, f493, f479])).
fof(f395, plain, (aElementOf0(sK36, stldt0(xA)) | ~ sP15), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ((~ aSubsetOf0(stldt0(xA), cS1395) & (~ aElementOf0(sK36, cS1395) & aElementOf0(sK36, stldt0(xA))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP8 & aSet0(stldt0(xA))) | ~ sP15), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36])], [f202, f203])).
fof(f203, plain, (? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, stldt0(xA))) => (~ aElementOf0(sK36, cS1395) & aElementOf0(sK36, stldt0(xA)))), introduced(choice_axiom, [])).
fof(f202, plain, ((~ aSubsetOf0(stldt0(xA), cS1395) & ? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, stldt0(xA))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP8 & aSet0(stldt0(xA))) | ~ sP15), inference(rectify, [], [f201])).
fof(f201, plain, ((~ aSubsetOf0(stldt0(xA), cS1395) & ? [X10] : (~ aElementOf0(X10, cS1395) & aElementOf0(X10, stldt0(xA))) & ! [X9] : ((aElementOf0(X9, cS1395) | ~ aInteger0(X9)) & (aInteger0(X9) | ~ aElementOf0(X9, cS1395))) & aSet0(cS1395) & sP8 & aSet0(stldt0(xA))) | ~ sP15), inference(nnf_transformation, [], [f123])).
fof(f491, plain, (~ spl38_8 | ~ spl38_10), inference(avatar_split_clause, [], [f396, f488, f479])).
fof(f396, plain, (~ aElementOf0(sK36, cS1395) | ~ sP15), inference(cnf_transformation, [], [f204])).
fof(f472, plain, (~ spl38_3 | spl38_6), inference(avatar_split_clause, [], [f387, f469, f455])).
fof(f387, plain, (aElementOf0(sK35, stldt0(xB)) | ~ sP16), inference(cnf_transformation, [], [f200])).
fof(f200, plain, ((~ aSubsetOf0(stldt0(xB), cS1395) & (~ aElementOf0(sK35, cS1395) & aElementOf0(sK35, stldt0(xB))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP9 & aSet0(stldt0(xB))) | ~ sP16), inference(skolemisation, [status(esa), new_symbols(skolem, [sK35])], [f198, f199])).
fof(f199, plain, (? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, stldt0(xB))) => (~ aElementOf0(sK35, cS1395) & aElementOf0(sK35, stldt0(xB)))), introduced(choice_axiom, [])).
fof(f198, plain, ((~ aSubsetOf0(stldt0(xB), cS1395) & ? [X0] : (~ aElementOf0(X0, cS1395) & aElementOf0(X0, stldt0(xB))) & ! [X1] : ((aElementOf0(X1, cS1395) | ~ aInteger0(X1)) & (aInteger0(X1) | ~ aElementOf0(X1, cS1395))) & aSet0(cS1395) & sP9 & aSet0(stldt0(xB))) | ~ sP16), inference(rectify, [], [f197])).
fof(f197, plain, ((~ aSubsetOf0(stldt0(xB), cS1395) & ? [X7] : (~ aElementOf0(X7, cS1395) & aElementOf0(X7, stldt0(xB))) & ! [X6] : ((aElementOf0(X6, cS1395) | ~ aInteger0(X6)) & (aInteger0(X6) | ~ aElementOf0(X6, cS1395))) & aSet0(cS1395) & sP9 & aSet0(stldt0(xB))) | ~ sP16), inference(nnf_transformation, [], [f124])).
fof(f467, plain, (~ spl38_3 | ~ spl38_5), inference(avatar_split_clause, [], [f388, f464, f455])).
fof(f388, plain, (~ aElementOf0(sK35, cS1395) | ~ sP16), inference(cnf_transformation, [], [f200])).