fof(f855290, plain, $false, inference(avatar_sat_refutation, [], [f618, f627, f632, f1676, f1744, f1749, f1775, f1792, f5309, f5311, f5317, f855289])).
fof(f855289, plain, (spl42_4 | ~ spl42_5 | ~ spl42_62 | ~ spl42_66 | ~ spl42_67 | ~ spl42_74), inference(avatar_contradiction_clause, [], [f855288])).
fof(f855288, plain, ($false | (spl42_4 | ~ spl42_5 | ~ spl42_62 | ~ spl42_66 | ~ spl42_67 | ~ spl42_74)), inference(subsumption_resolution, [], [f855285, f626])).
fof(f626, plain, (~ aElementOf0(sK41, szNzAzT0) | spl42_4), inference(avatar_component_clause, [], [f624])).
fof(f624, plain, (spl42_4 <=> aElementOf0(sK41, szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_4])])).
fof(f855285, plain, (aElementOf0(sK41, szNzAzT0) | (~ spl42_5 | ~ spl42_62 | ~ spl42_66 | ~ spl42_67 | ~ spl42_74)), inference(resolution, [], [f840805, f631])).
fof(f631, plain, (aElementOf0(sK41, sdtlpdtrp0(xN, sK40)) | ~ spl42_5), inference(avatar_component_clause, [], [f629])).
fof(f629, plain, (spl42_5 <=> aElementOf0(sK41, sdtlpdtrp0(xN, sK40))), introduced(avatar_definition, [new_symbols(naming, [spl42_5])])).
fof(f840805, plain, (! [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, sK40)) | aElementOf0(X0, szNzAzT0)) | (~ spl42_62 | ~ spl42_66 | ~ spl42_67 | ~ spl42_74)), inference(subsumption_resolution, [], [f840804, f1655])).
fof(f1655, plain, (aElementOf0(sK14(sK40), szNzAzT0) | ~ spl42_62), inference(avatar_component_clause, [], [f1654])).
fof(f1654, plain, (spl42_62 <=> aElementOf0(sK14(sK40), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_62])])).
fof(f840804, plain, (! [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, sK40)) | aElementOf0(X0, szNzAzT0) | ~ aElementOf0(sK14(sK40), szNzAzT0)) | (~ spl42_66 | ~ spl42_67 | ~ spl42_74)), inference(subsumption_resolution, [], [f840776, f1675])).
fof(f1675, plain, (iLess0(sK14(sK40), sK40) | ~ spl42_66), inference(avatar_component_clause, [], [f1673])).
fof(f1673, plain, (spl42_66 <=> iLess0(sK14(sK40), sK40)), introduced(avatar_definition, [new_symbols(naming, [spl42_66])])).
fof(f840776, plain, (! [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, sK40)) | aElementOf0(X0, szNzAzT0) | ~ iLess0(sK14(sK40), sK40) | ~ aElementOf0(sK14(sK40), szNzAzT0)) | (~ spl42_67 | ~ spl42_74)), inference(resolution, [], [f9914, f549])).
fof(f549, plain, ! [X2, X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, X2)) | aElementOf0(X3, szNzAzT0) | ~ iLess0(X2, sK40) | ~ aElementOf0(X2, szNzAzT0)), inference(cnf_transformation, [], [f328])).
fof(f328, plain, ((~ isCountable0(sdtlpdtrp0(xN, sK40)) | (~ aSubsetOf0(sdtlpdtrp0(xN, sK40), szNzAzT0) & ((~ aElementOf0(sK41, szNzAzT0) & aElementOf0(sK41, sdtlpdtrp0(xN, sK40))) | ~ aSet0(sdtlpdtrp0(xN, sK40))))) & ! [X2] : ((isCountable0(sdtlpdtrp0(xN, X2)) & aSubsetOf0(sdtlpdtrp0(xN, X2), szNzAzT0) & ! [X3] : (aElementOf0(X3, szNzAzT0) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X2))) & aSet0(sdtlpdtrp0(xN, X2))) | ~ iLess0(X2, sK40) | ~ aElementOf0(X2, szNzAzT0)) & aElementOf0(sK40, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40, sK41])], [f325, f327, f326])).
fof(f326, plain, (? [X0] : ((~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0))))) & ! [X2] : ((isCountable0(sdtlpdtrp0(xN, X2)) & aSubsetOf0(sdtlpdtrp0(xN, X2), szNzAzT0) & ! [X3] : (aElementOf0(X3, szNzAzT0) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X2))) & aSet0(sdtlpdtrp0(xN, X2))) | ~ iLess0(X2, X0) | ~ aElementOf0(X2, szNzAzT0)) & aElementOf0(X0, szNzAzT0)) => ((~ isCountable0(sdtlpdtrp0(xN, sK40)) | (~ aSubsetOf0(sdtlpdtrp0(xN, sK40), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, sK40))) | ~ aSet0(sdtlpdtrp0(xN, sK40))))) & ! [X2] : ((isCountable0(sdtlpdtrp0(xN, X2)) & aSubsetOf0(sdtlpdtrp0(xN, X2), szNzAzT0) & ! [X3] : (aElementOf0(X3, szNzAzT0) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X2))) & aSet0(sdtlpdtrp0(xN, X2))) | ~ iLess0(X2, sK40) | ~ aElementOf0(X2, szNzAzT0)) & aElementOf0(sK40, szNzAzT0))), introduced(choice_axiom, [])).
fof(f327, plain, (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, sK40))) => (~ aElementOf0(sK41, szNzAzT0) & aElementOf0(sK41, sdtlpdtrp0(xN, sK40)))), introduced(choice_axiom, [])).
fof(f325, plain, ? [X0] : ((~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0))))) & ! [X2] : ((isCountable0(sdtlpdtrp0(xN, X2)) & aSubsetOf0(sdtlpdtrp0(xN, X2), szNzAzT0) & ! [X3] : (aElementOf0(X3, szNzAzT0) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X2))) & aSet0(sdtlpdtrp0(xN, X2))) | ~ iLess0(X2, X0) | ~ aElementOf0(X2, szNzAzT0)) & aElementOf0(X0, szNzAzT0)), inference(rectify, [], [f202])).
fof(f202, plain, ? [X0] : ((~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X3] : (~ aElementOf0(X3, szNzAzT0) & aElementOf0(X3, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0))))) & ! [X1] : ((isCountable0(sdtlpdtrp0(xN, X1)) & aSubsetOf0(sdtlpdtrp0(xN, X1), szNzAzT0) & ! [X2] : (aElementOf0(X2, szNzAzT0) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X1))) & aSet0(sdtlpdtrp0(xN, X1))) | ~ iLess0(X1, X0) | ~ aElementOf0(X1, szNzAzT0)) & aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f201])).
fof(f201, plain, ? [X0] : (((~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X3] : (~ aElementOf0(X3, szNzAzT0) & aElementOf0(X3, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0))))) & ! [X1] : (((isCountable0(sdtlpdtrp0(xN, X1)) & aSubsetOf0(sdtlpdtrp0(xN, X1), szNzAzT0) & ! [X2] : (aElementOf0(X2, szNzAzT0) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X1))) & aSet0(sdtlpdtrp0(xN, X1))) | ~ iLess0(X1, X0)) | ~ aElementOf0(X1, szNzAzT0))) & aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f94])).
fof(f94, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => (! [X1] : (aElementOf0(X1, szNzAzT0) => (iLess0(X1, X0) => (isCountable0(sdtlpdtrp0(xN, X1)) & aSubsetOf0(sdtlpdtrp0(xN, X1), szNzAzT0) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => aElementOf0(X2, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X1))))) => (isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X3] : (aElementOf0(X3, sdtlpdtrp0(xN, X0)) => aElementOf0(X3, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))))), inference(rectify, [], [f83])).
fof(f83, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => (! [X1] : (aElementOf0(X1, szNzAzT0) => (iLess0(X1, X0) => (isCountable0(sdtlpdtrp0(xN, X1)) & aSubsetOf0(sdtlpdtrp0(xN, X1), szNzAzT0) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => aElementOf0(X2, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X1))))) => (isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))))), inference(negated_conjecture, [], [f82])).
fof(f82, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => (! [X1] : (aElementOf0(X1, szNzAzT0) => (iLess0(X1, X0) => (isCountable0(sdtlpdtrp0(xN, X1)) & aSubsetOf0(sdtlpdtrp0(xN, X1), szNzAzT0) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => aElementOf0(X2, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X1))))) => (isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM569+3.p', m__)).
fof(f9914, plain, (! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, sK14(sK40))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, sK40))) | (~ spl42_67 | ~ spl42_74)), inference(subsumption_resolution, [], [f9892, f1679])).
fof(f1679, plain, (sP9(sK14(sK40)) | ~ spl42_67), inference(avatar_component_clause, [], [f1678])).
fof(f1678, plain, (spl42_67 <=> sP9(sK14(sK40))), introduced(avatar_definition, [new_symbols(naming, [spl42_67])])).
fof(f9892, plain, (! [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, sK40)) | ~ sP9(sK14(sK40)) | aElementOf0(X2, sdtlpdtrp0(xN, sK14(sK40)))) | ~ spl42_74), inference(superposition, [], [f3462, f1748])).
fof(f1748, plain, ((sK40 = szszuzczcdt0(sK14(sK40))) | ~ spl42_74), inference(avatar_component_clause, [], [f1746])).
fof(f1746, plain, (spl42_74 <=> (sK40 = szszuzczcdt0(sK14(sK40)))), introduced(avatar_definition, [new_symbols(naming, [spl42_74])])).
fof(f3462, plain, ! [X0, X1] : (~ aElementOf0(X0, sdtlpdtrp0(xN, szszuzczcdt0(X1))) | ~ sP9(X1) | aElementOf0(X0, sdtlpdtrp0(xN, X1))), inference(subsumption_resolution, [], [f3454, f532])).
fof(f532, plain, ! [X0] : (~ sP9(X0) | sP8(X0)), inference(cnf_transformation, [], [f319])).
fof(f319, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(rectify, [], [f318])).
fof(f318, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(nnf_transformation, [], [f215])).
fof(f215, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e215])).
fof(e215, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f3454, plain, ! [X0, X1] : (~ aElementOf0(X0, sdtlpdtrp0(xN, szszuzczcdt0(X1))) | ~ sP9(X1) | aElementOf0(X0, sdtlpdtrp0(xN, X1)) | ~ sP8(X1)), inference(resolution, [], [f534, f538])).
fof(f538, plain, ! [X0, X1] : (~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f322])).
fof(f322, plain, ! [X0] : (! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ aElement0(X1)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(rectify, [], [f321])).
fof(f321, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(flattening, [], [f320])).
fof(f320, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(nnf_transformation, [], [f214])).
fof(f214, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e214])).
fof(e214, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f534, plain, ! [X0, X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ sP9(X0)), inference(cnf_transformation, [], [f319])).
fof(f5317, plain, (~ spl42_67 | spl42_3 | ~ spl42_74), inference(avatar_split_clause, [], [f3276, f1746, f620, f1678])).
fof(f620, plain, (spl42_3 <=> aSet0(sdtlpdtrp0(xN, sK40))), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f3276, plain, (aSet0(sdtlpdtrp0(xN, sK40)) | ~ sP9(sK14(sK40)) | ~ spl42_74), inference(superposition, [], [f533, f1748])).
fof(f533, plain, ! [X0] : (aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ sP9(X0)), inference(cnf_transformation, [], [f319])).
fof(f5311, plain, (~ spl42_67 | spl42_2 | ~ spl42_74), inference(avatar_split_clause, [], [f3278, f1746, f615, f1678])).
fof(f615, plain, (spl42_2 <=> isCountable0(sdtlpdtrp0(xN, sK40))), introduced(avatar_definition, [new_symbols(naming, [spl42_2])])).
fof(f3278, plain, (isCountable0(sdtlpdtrp0(xN, sK40)) | ~ sP9(sK14(sK40)) | ~ spl42_74), inference(superposition, [], [f536, f1748])).
fof(f536, plain, ! [X0] : (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ sP9(X0)), inference(cnf_transformation, [], [f319])).
fof(f5309, plain, (spl42_67 | ~ spl42_62 | ~ spl42_66), inference(avatar_split_clause, [], [f5306, f1673, f1654, f1678])).
fof(f5306, plain, (sP9(sK14(sK40)) | (~ spl42_62 | ~ spl42_66)), inference(subsumption_resolution, [], [f5305, f1655])).
fof(f5305, plain, (~ aElementOf0(sK14(sK40), szNzAzT0) | sP9(sK14(sK40)) | ~ spl42_66), inference(resolution, [], [f2356, f1675])).
fof(f2356, plain, ! [X0] : (~ iLess0(X0, sK40) | ~ aElementOf0(X0, szNzAzT0) | sP9(X0)), inference(subsumption_resolution, [], [f2355, f551])).
fof(f551, plain, ! [X2] : (isCountable0(sdtlpdtrp0(xN, X2)) | ~ iLess0(X2, sK40) | ~ aElementOf0(X2, szNzAzT0)), inference(cnf_transformation, [], [f328])).
fof(f2355, plain, ! [X0] : (~ iLess0(X0, sK40) | ~ aElementOf0(X0, szNzAzT0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | sP9(X0)), inference(duplicate_literal_removal, [], [f2348])).
fof(f2348, plain, ! [X0] : (~ iLess0(X0, sK40) | ~ aElementOf0(X0, szNzAzT0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | sP9(X0) | ~ aElementOf0(X0, szNzAzT0)), inference(resolution, [], [f550, f546])).
fof(f546, plain, ! [X0] : (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | sP9(X0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f324])).
fof(f324, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f216, f323])).
fof(f323, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f216, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f200, e215, e214])).
fof(f200, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f199])).
fof(f199, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f93])).
fof(f93, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM569+3.p', m__3623)).
fof(f550, plain, ! [X2] : (aSubsetOf0(sdtlpdtrp0(xN, X2), szNzAzT0) | ~ iLess0(X2, sK40) | ~ aElementOf0(X2, szNzAzT0)), inference(cnf_transformation, [], [f328])).
fof(f1792, plain, (spl42_1 | ~ spl42_29), inference(avatar_contradiction_clause, [], [f1791])).
fof(f1791, plain, ($false | (spl42_1 | ~ spl42_29)), inference(subsumption_resolution, [], [f1790, f477])).
fof(f477, plain, aSubsetOf0(xS, szNzAzT0), inference(cnf_transformation, [], [f194])).
fof(f194, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xS)) & aSet0(xS)), inference(ennf_transformation, [], [f75])).
fof(f75, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, szNzAzT0)) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM569+3.p', m__3435)).
fof(f1790, plain, (~ aSubsetOf0(xS, szNzAzT0) | (spl42_1 | ~ spl42_29)), inference(forward_demodulation, [], [f1789, f543])).
fof(f543, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f324])).
fof(f1789, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, sz00), szNzAzT0) | (spl42_1 | ~ spl42_29)), inference(forward_demodulation, [], [f613, f1117])).
fof(f1117, plain, ((sz00 = sK40) | ~ spl42_29), inference(avatar_component_clause, [], [f1116])).
fof(f1116, plain, (spl42_29 <=> (sz00 = sK40)), introduced(avatar_definition, [new_symbols(naming, [spl42_29])])).
fof(f613, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, sK40), szNzAzT0) | spl42_1), inference(avatar_component_clause, [], [f611])).
fof(f611, plain, (spl42_1 <=> aSubsetOf0(sdtlpdtrp0(xN, sK40), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f1775, plain, (spl42_2 | ~ spl42_29), inference(avatar_contradiction_clause, [], [f1774])).
fof(f1774, plain, ($false | (spl42_2 | ~ spl42_29)), inference(subsumption_resolution, [], [f1773, f478])).
fof(f478, plain, isCountable0(xS), inference(cnf_transformation, [], [f194])).
fof(f1773, plain, (~ isCountable0(xS) | (spl42_2 | ~ spl42_29)), inference(forward_demodulation, [], [f1755, f543])).
fof(f1755, plain, (~ isCountable0(sdtlpdtrp0(xN, sz00)) | (spl42_2 | ~ spl42_29)), inference(backward_demodulation, [], [f617, f1117])).
fof(f617, plain, (~ isCountable0(sdtlpdtrp0(xN, sK40)) | spl42_2), inference(avatar_component_clause, [], [f615])).
fof(f1749, plain, (spl42_74 | spl42_29), inference(avatar_split_clause, [], [f1474, f1116, f1746])).
fof(f1474, plain, ((sz00 = sK40) | (sK40 = szszuzczcdt0(sK14(sK40)))), inference(resolution, [], [f381, f547])).
fof(f547, plain, aElementOf0(sK40, szNzAzT0), inference(cnf_transformation, [], [f328])).
fof(f381, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | (sz00 = X0) | (szszuzczcdt0(sK14(X0)) = X0)), inference(cnf_transformation, [], [f240])).
fof(f240, plain, ! [X0] : (((szszuzczcdt0(sK14(X0)) = X0) & aElementOf0(sK14(X0), szNzAzT0)) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f128, f239])).
fof(f239, plain, ! [X0] : (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) => ((szszuzczcdt0(sK14(X0)) = X0) & aElementOf0(sK14(X0), szNzAzT0))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X0] : (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f127])).
fof(f127, plain, ! [X0] : ((? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM569+3.p', mNatExtra)).
fof(f1744, plain, (spl42_29 | spl42_62), inference(avatar_split_clause, [], [f1743, f1654, f1116])).
fof(f1743, plain, ((sz00 = sK40) | spl42_62), inference(subsumption_resolution, [], [f1739, f547])).
fof(f1739, plain, ((sz00 = sK40) | ~ aElementOf0(sK40, szNzAzT0) | spl42_62), inference(resolution, [], [f1656, f380])).
fof(f380, plain, ! [X0] : (aElementOf0(sK14(X0), szNzAzT0) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f240])).
fof(f1656, plain, (~ aElementOf0(sK14(sK40), szNzAzT0) | spl42_62), inference(avatar_component_clause, [], [f1654])).
fof(f1676, plain, (~ spl42_62 | spl42_66 | spl42_29), inference(avatar_split_clause, [], [f1645, f1116, f1673, f1654])).
fof(f1645, plain, (iLess0(sK14(sK40), sK40) | ~ aElementOf0(sK14(sK40), szNzAzT0) | spl42_29), inference(superposition, [], [f392, f1491])).
fof(f1491, plain, ((sK40 = szszuzczcdt0(sK14(sK40))) | spl42_29), inference(subsumption_resolution, [], [f1474, f1118])).
fof(f1118, plain, (~ (sz00 = sK40) | spl42_29), inference(avatar_component_clause, [], [f1116])).
fof(f392, plain, ! [X0] : (iLess0(X0, szszuzczcdt0(X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f142])).
fof(f142, plain, ! [X0] : (iLess0(X0, szszuzczcdt0(X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => iLess0(X0, szszuzczcdt0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM569+3.p', mIH)).
fof(f632, plain, (~ spl42_3 | spl42_5 | ~ spl42_2), inference(avatar_split_clause, [], [f552, f615, f629, f620])).
fof(f552, plain, (~ isCountable0(sdtlpdtrp0(xN, sK40)) | aElementOf0(sK41, sdtlpdtrp0(xN, sK40)) | ~ aSet0(sdtlpdtrp0(xN, sK40))), inference(cnf_transformation, [], [f328])).
fof(f627, plain, (~ spl42_3 | ~ spl42_4 | ~ spl42_2), inference(avatar_split_clause, [], [f553, f615, f624, f620])).
fof(f553, plain, (~ isCountable0(sdtlpdtrp0(xN, sK40)) | ~ aElementOf0(sK41, szNzAzT0) | ~ aSet0(sdtlpdtrp0(xN, sK40))), inference(cnf_transformation, [], [f328])).
fof(f618, plain, (~ spl42_1 | ~ spl42_2), inference(avatar_split_clause, [], [f554, f615, f611])).
fof(f554, plain, (~ isCountable0(sdtlpdtrp0(xN, sK40)) | ~ aSubsetOf0(sdtlpdtrp0(xN, sK40), szNzAzT0)), inference(cnf_transformation, [], [f328])).