fof(f496152, plain, $false, inference(avatar_sat_refutation, [], [f668, f673, f780, f1801, f14201, f104664, f256874, f496129])).
fof(f496129, plain, spl44_1133, inference(avatar_split_clause, [], [f496128, f34192])).
fof(f34192, plain, (spl44_1133 <=> ! [X89] : (aElementOf0(X89, xS) | ~ aElementOf0(X89, sK42))), introduced(avatar_definition, [new_symbols(naming, [spl44_1133])])).
fof(f496128, plain, ! [X0] : (~ aElementOf0(X0, sK42) | aElementOf0(X0, xS)), inference(subsumption_resolution, [], [f496094, f573])).
fof(f573, plain, aElementOf0(sK41, szNzAzT0), inference(cnf_transformation, [], [f343])).
fof(f343, plain, ((~ aElementOf0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | (~ aSubsetOf0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))), xS) & (~ aElementOf0(sK43, xS) & aElementOf0(sK43, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) & ~ aElementOf0(X3, sK42)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) | aElementOf0(X3, sK42)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aElementOf0(sK42, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41))), xk)) & (xk = sbrdtbr0(sK42)) & aSubsetOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ aElementOf0(X5, sK42)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, sK41)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, sK41)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aSet0(sK42)) & aElementOf0(sK41, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41, sK42, sK43])], [f339, f342, f341, f340])).
fof(f340, plain, (? [X0] : (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & ~ aElementOf0(X3, X1)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | aElementOf0(X3, X1)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X5, X1)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, X0)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) & aElementOf0(X0, szNzAzT0)) => (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) & ~ aElementOf0(X3, X1)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) | aElementOf0(X3, X1)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ aElementOf0(X5, X1)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, sK41)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, sK41)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aSet0(X1)) & aElementOf0(sK41, szNzAzT0))), introduced(choice_axiom, [])).
fof(f341, plain, (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) & ~ aElementOf0(X3, X1)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) | aElementOf0(X3, X1)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ aElementOf0(X5, X1)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, sK41)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, sK41)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aSet0(X1)) => (~ aElementOf0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | (~ aSubsetOf0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) & ~ aElementOf0(X3, sK42)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3) | aElementOf0(X3, sK42)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aElementOf0(sK42, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41))), xk)) & (xk = sbrdtbr0(sK42)) & aSubsetOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ aElementOf0(X5, sK42)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, sK41)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, sK41)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK41)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)) & aSet0(sK42))), introduced(choice_axiom, [])).
fof(f342, plain, (? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) => (~ aElementOf0(sK43, xS) & aElementOf0(sK43, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))))), introduced(choice_axiom, [])).
fof(f339, plain, ? [X0] : (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))))) & ! [X3] : ((aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & ~ aElementOf0(X3, X1)) | ~ aElement0(X3)) & ((((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | aElementOf0(X3, X1)) & aElement0(X3)) | ~ aElementOf0(X3, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X5, X1)) & ! [X6] : ((aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0)) | ~ aElement0(X6)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) & aElementOf0(X6, sdtlpdtrp0(xN, X0)) & aElement0(X6)) | ~ aElementOf0(X6, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X7) | ~ aElementOf0(X7, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) & aElementOf0(X0, szNzAzT0)), inference(rectify, [], [f338])).
fof(f338, plain, ? [X0] : (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X7] : (~ aElementOf0(X7, xS) & aElementOf0(X7, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))))) & ! [X6] : ((aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) & ~ aElementOf0(X6, X1)) | ~ aElement0(X6)) & ((((szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | aElementOf0(X6, X1)) & aElement0(X6)) | ~ aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X5] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X5) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, X1)) & ! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) & aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f337])).
fof(f337, plain, ? [X0] : (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X7] : (~ aElementOf0(X7, xS) & aElementOf0(X7, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))))) & ! [X6] : ((aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) & ~ aElementOf0(X6, X1)) | ~ aElement0(X6))) & ((((szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | aElementOf0(X6, X1)) & aElement0(X6)) | ~ aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X5] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X5) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, X1)) & ! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) & aElementOf0(X0, szNzAzT0)), inference(nnf_transformation, [], [f211])).
fof(f211, plain, ? [X0] : (? [X1] : (~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X7] : (~ aElementOf0(X7, xS) & aElementOf0(X7, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))))) & ! [X6] : (aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | aElementOf0(X6, X1)) & aElement0(X6))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X5] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X5) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, X1)) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) & aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f210])).
fof(f210, plain, ? [X0] : (? [X1] : ((((~ aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) & (~ (xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) | (~ aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) & ? [X7] : (~ aElementOf0(X7, xS) & aElementOf0(X7, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))))) & (! [X6] : (aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | aElementOf0(X6, X1)) & aElement0(X6))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))))) & (! [X5] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X5) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))) & (aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, X1)) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1))) & aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => ! [X1] : ((aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) => ((! [X5] : (aElementOf0(X5, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X5)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) => ((! [X6] : (aElementOf0(X6, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X6) | aElementOf0(X6, X1)) & aElement0(X6))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) => (aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) | ((xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & (aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) | ! [X7] : (aElementOf0(X7, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) => aElementOf0(X7, xS))))))))), inference(rectify, [], [f86])).
fof(f86, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => ! [X1] : ((aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, X0)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) => ((! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) => ((! [X2] : (aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) | aElementOf0(X2, X1)) & aElement0(X2))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) => (aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) | ((xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & (aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) | ! [X2] : (aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) => aElementOf0(X2, xS))))))))), inference(negated_conjecture, [], [f85])).
fof(f85, plain, ~ ! [X0] : (aElementOf0(X0, szNzAzT0) => ! [X1] : ((aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, X0)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aSet0(X1)) => ((! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) => ((! [X2] : (aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) | aElementOf0(X2, X1)) & aElement0(X2))) & aSet0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) => (aElementOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), slbdtsldtrb0(xS, xK)) | ((xK = sbrdtbr0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & (aSubsetOf0(sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))), xS) | ! [X2] : (aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) => aElementOf0(X2, xS))))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__)).
fof(f496094, plain, ! [X0] : (~ aElementOf0(X0, sK42) | ~ aElementOf0(sK41, szNzAzT0) | aElementOf0(X0, xS)), inference(resolution, [], [f7529, f11407])).
fof(f11407, plain, ! [X4, X5] : (~ aElementOf0(X5, sdtlpdtrp0(xN, X4)) | ~ aElementOf0(X4, szNzAzT0) | aElementOf0(X5, xS)), inference(subsumption_resolution, [], [f11401, f490])).
fof(f490, plain, aSet0(xS), inference(cnf_transformation, [], [f198])).
fof(f198, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xS)) & aSet0(xS)), inference(ennf_transformation, [], [f75])).
fof(f75, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, szNzAzT0)) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__3435)).
fof(f11401, plain, ! [X4, X5] : (~ aElementOf0(X4, szNzAzT0) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X4)) | aElementOf0(X5, xS) | ~ aSet0(xS)), inference(resolution, [], [f3621, f352])).
fof(f352, plain, ! [X0, X3, X1] : (~ aSubsetOf0(X1, X0) | ~ aElementOf0(X3, X1) | aElementOf0(X3, X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (~ aElementOf0(sK11(X0, X1), X0) & aElementOf0(sK11(X0, X1), X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11])], [f233, f234])).
fof(f234, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) => (~ aElementOf0(sK11(X0, X1), X0) & aElementOf0(sK11(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f233, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(rectify, [], [f232])).
fof(f232, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(flattening, [], [f231])).
fof(f231, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1))) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(nnf_transformation, [], [f105])).
fof(f105, plain, ! [X0] : (! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1))) | ~ aSet0(X0)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, X0)) & aSet0(X1)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mDefSub)).
fof(f3621, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f3620, f398])).
fof(f398, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sdtlseqdt0(sz00, X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : (sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => sdtlseqdt0(sz00, X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mZeroLess)).
fof(f3620, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f3616, f391])).
fof(f391, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mZeroNum)).
fof(f3616, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(sz00, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(superposition, [], [f567, f558])).
fof(f558, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f333])).
fof(f333, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f225, f332])).
fof(f332, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f225, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f204, e224, e223])).
fof(f223, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e223])).
fof(e223, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f224, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e224])).
fof(e224, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f204, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f203])).
fof(f203, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f96])).
fof(f96, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__3623)).
fof(f567, plain, ! [X0, X1] : (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f206])).
fof(f206, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__3754)).
fof(f7529, plain, ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, sK41)) | ~ aElementOf0(X1, sK42)), inference(subsumption_resolution, [], [f7519, f1109])).
fof(f1109, plain, sP8(sK41), inference(resolution, [], [f1092, f547])).
fof(f547, plain, ! [X0] : (~ sP9(X0) | sP8(X0)), inference(cnf_transformation, [], [f328])).
fof(f328, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(rectify, [], [f327])).
fof(f327, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(nnf_transformation, [], [f224])).
fof(f1092, plain, sP9(sK41), inference(resolution, [], [f1051, f573])).
fof(f1051, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sP9(X0)), inference(subsumption_resolution, [], [f1050, f564])).
fof(f564, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aSet0(sdtlpdtrp0(xN, X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__3671)).
fof(f1050, plain, ! [X0] : (sP9(X0) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f561, f565])).
fof(f565, plain, ! [X0] : (isCountable0(sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f205])).
fof(f561, plain, ! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f333])).
fof(f7519, plain, ! [X1] : (~ aElementOf0(X1, sK42) | aElementOf0(X1, sdtlpdtrp0(xN, sK41)) | ~ sP8(sK41)), inference(resolution, [], [f582, f553])).
fof(f553, plain, ! [X0, X1] : (~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ sP8(X0)), inference(cnf_transformation, [], [f331])).
fof(f331, plain, ! [X0] : (! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ aElement0(X1)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(rectify, [], [f330])).
fof(f330, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(flattening, [], [f329])).
fof(f329, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(nnf_transformation, [], [f223])).
fof(f582, plain, ! [X5] : (aElementOf0(X5, sdtmndt0(sdtlpdtrp0(xN, sK41), szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ aElementOf0(X5, sK42)), inference(cnf_transformation, [], [f343])).
fof(f256874, plain, (spl44_2 | ~ spl44_5), inference(avatar_split_clause, [], [f256873, f675, f660])).
fof(f660, plain, (spl44_2 <=> (xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))))), introduced(avatar_definition, [new_symbols(naming, [spl44_2])])).
fof(f675, plain, (spl44_5 <=> aElement0(szmzizndt0(sdtlpdtrp0(xN, sK41)))), introduced(avatar_definition, [new_symbols(naming, [spl44_5])])).
fof(f256873, plain, ((xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | ~ spl44_5), inference(forward_demodulation, [], [f15380, f543])).
fof(f543, plain, (xK = szszuzczcdt0(xk)), inference(cnf_transformation, [], [f80])).
fof(f80, plain, ((xK = szszuzczcdt0(xk)) & aElementOf0(xk, szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', m__3533)).
fof(f15380, plain, ((szszuzczcdt0(xk) = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | ~ spl44_5), inference(forward_demodulation, [], [f15379, f584])).
fof(f584, plain, (xk = sbrdtbr0(sK42)), inference(cnf_transformation, [], [f343])).
fof(f15379, plain, ((sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) = szszuzczcdt0(sbrdtbr0(sK42))) | ~ spl44_5), inference(subsumption_resolution, [], [f15378, f574])).
fof(f574, plain, aSet0(sK42), inference(cnf_transformation, [], [f343])).
fof(f15378, plain, ((sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) = szszuzczcdt0(sbrdtbr0(sK42))) | ~ aSet0(sK42) | ~ spl44_5), inference(subsumption_resolution, [], [f15377, f785])).
fof(f785, plain, isFinite0(sK42), inference(subsumption_resolution, [], [f784, f574])).
fof(f784, plain, (isFinite0(sK42) | ~ aSet0(sK42)), inference(subsumption_resolution, [], [f782, f542])).
fof(f542, plain, aElementOf0(xk, szNzAzT0), inference(cnf_transformation, [], [f80])).
fof(f782, plain, (~ aElementOf0(xk, szNzAzT0) | isFinite0(sK42) | ~ aSet0(sK42)), inference(superposition, [], [f409, f584])).
fof(f409, plain, ! [X0] : (~ aElementOf0(sbrdtbr0(X0), szNzAzT0) | isFinite0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ! [X0] : (((aElementOf0(sbrdtbr0(X0), szNzAzT0) | ~ isFinite0(X0)) & (isFinite0(X0) | ~ aElementOf0(sbrdtbr0(X0), szNzAzT0))) | ~ aSet0(X0)), inference(nnf_transformation, [], [f148])).
fof(f148, plain, ! [X0] : ((aElementOf0(sbrdtbr0(X0), szNzAzT0) <=> isFinite0(X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (aSet0(X0) => (aElementOf0(sbrdtbr0(X0), szNzAzT0) <=> isFinite0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mCardNum)).
fof(f15377, plain, ((sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) = szszuzczcdt0(sbrdtbr0(sK42))) | ~ isFinite0(sK42) | ~ aSet0(sK42) | ~ spl44_5), inference(subsumption_resolution, [], [f15375, f676])).
fof(f676, plain, (aElement0(szmzizndt0(sdtlpdtrp0(xN, sK41))) | ~ spl44_5), inference(avatar_component_clause, [], [f675])).
fof(f15375, plain, ((sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) = szszuzczcdt0(sbrdtbr0(sK42))) | ~ aElement0(szmzizndt0(sdtlpdtrp0(xN, sK41))) | ~ isFinite0(sK42) | ~ aSet0(sK42)), inference(resolution, [], [f7540, f413])).
fof(f413, plain, ! [X0, X1] : (aElementOf0(X1, X0) | (sbrdtbr0(sdtpldt0(X0, X1)) = szszuzczcdt0(sbrdtbr0(X0))) | ~ aElement0(X1) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f151])).
fof(f151, plain, ! [X0] : (! [X1] : ((sbrdtbr0(sdtpldt0(X0, X1)) = szszuzczcdt0(sbrdtbr0(X0))) | aElementOf0(X1, X0) | ~ aElement0(X1)) | ~ isFinite0(X0) | ~ aSet0(X0)), inference(flattening, [], [f150])).
fof(f150, plain, ! [X0] : (! [X1] : (((sbrdtbr0(sdtpldt0(X0, X1)) = szszuzczcdt0(sbrdtbr0(X0))) | aElementOf0(X1, X0)) | ~ aElement0(X1)) | (~ isFinite0(X0) | ~ aSet0(X0))), inference(ennf_transformation, [], [f43])).
fof(f43, plain, ! [X0] : ((isFinite0(X0) & aSet0(X0)) => ! [X1] : (aElement0(X1) => (~ aElementOf0(X1, X0) => (sbrdtbr0(sdtpldt0(X0, X1)) = szszuzczcdt0(sbrdtbr0(X0)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mCardCons)).
fof(f7540, plain, ~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sK42), inference(subsumption_resolution, [], [f7525, f1109])).
fof(f7525, plain, (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sK42) | ~ sP8(sK41)), inference(resolution, [], [f582, f639])).
fof(f639, plain, ! [X0] : (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(equality_resolution, [], [f554])).
fof(f554, plain, ! [X0, X1] : (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(cnf_transformation, [], [f331])).
fof(f104664, plain, (spl44_3 | ~ spl44_158 | ~ spl44_1133), inference(avatar_contradiction_clause, [], [f104663])).
fof(f104663, plain, ($false | (spl44_3 | ~ spl44_158 | ~ spl44_1133)), inference(subsumption_resolution, [], [f104491, f667])).
fof(f667, plain, (~ aElementOf0(sK43, xS) | spl44_3), inference(avatar_component_clause, [], [f665])).
fof(f665, plain, (spl44_3 <=> aElementOf0(sK43, xS)), introduced(avatar_definition, [new_symbols(naming, [spl44_3])])).
fof(f104491, plain, (aElementOf0(sK43, xS) | (~ spl44_158 | ~ spl44_1133)), inference(resolution, [], [f34193, f4023])).
fof(f4023, plain, (aElementOf0(sK43, sK42) | ~ spl44_158), inference(avatar_component_clause, [], [f4021])).
fof(f4021, plain, (spl44_158 <=> aElementOf0(sK43, sK42)), introduced(avatar_definition, [new_symbols(naming, [spl44_158])])).
fof(f34193, plain, (! [X89] : (~ aElementOf0(X89, sK42) | aElementOf0(X89, xS)) | ~ spl44_1133), inference(avatar_component_clause, [], [f34192])).
fof(f14201, plain, (spl44_3 | ~ spl44_4 | spl44_158), inference(avatar_contradiction_clause, [], [f14200])).
fof(f14200, plain, ($false | (spl44_3 | ~ spl44_4 | spl44_158)), inference(subsumption_resolution, [], [f14199, f391])).
fof(f14199, plain, (~ aElementOf0(sz00, szNzAzT0) | (spl44_3 | ~ spl44_4 | spl44_158)), inference(subsumption_resolution, [], [f14198, f695])).
fof(f695, plain, sdtlseqdt0(sz00, sK41), inference(resolution, [], [f398, f573])).
fof(f14198, plain, (~ sdtlseqdt0(sz00, sK41) | ~ aElementOf0(sz00, szNzAzT0) | (spl44_3 | ~ spl44_4 | spl44_158)), inference(subsumption_resolution, [], [f14193, f667])).
fof(f14193, plain, (aElementOf0(sK43, xS) | ~ sdtlseqdt0(sz00, sK41) | ~ aElementOf0(sz00, szNzAzT0) | (~ spl44_4 | spl44_158)), inference(superposition, [], [f8081, f558])).
fof(f8081, plain, (! [X1] : (aElementOf0(sK43, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, sK41) | ~ aElementOf0(X1, szNzAzT0)) | (~ spl44_4 | spl44_158)), inference(subsumption_resolution, [], [f8066, f573])).
fof(f8066, plain, (! [X1] : (aElementOf0(sK43, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, sK41) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(sK41, szNzAzT0)) | (~ spl44_4 | spl44_158)), inference(resolution, [], [f7670, f566])).
fof(f566, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, sdtlpdtrp0(xN, X0)) | aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f207])).
fof(f7670, plain, (aElementOf0(sK43, sdtlpdtrp0(xN, sK41)) | (~ spl44_4 | spl44_158)), inference(backward_demodulation, [], [f586, f7662])).
fof(f7662, plain, ((szmzizndt0(sdtlpdtrp0(xN, sK41)) = sK43) | (~ spl44_4 | spl44_158)), inference(subsumption_resolution, [], [f7635, f4022])).
fof(f4022, plain, (~ aElementOf0(sK43, sK42) | spl44_158), inference(avatar_component_clause, [], [f4021])).
fof(f7635, plain, (aElementOf0(sK43, sK42) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = sK43) | ~ spl44_4), inference(resolution, [], [f590, f672])).
fof(f672, plain, (aElementOf0(sK43, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | ~ spl44_4), inference(avatar_component_clause, [], [f670])).
fof(f670, plain, (spl44_4 <=> aElementOf0(sK43, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))), introduced(avatar_definition, [new_symbols(naming, [spl44_4])])).
fof(f590, plain, ! [X3] : (~ aElementOf0(X3, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41)))) | aElementOf0(X3, sK42) | (szmzizndt0(sdtlpdtrp0(xN, sK41)) = X3)), inference(cnf_transformation, [], [f343])).
fof(f586, plain, aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK41)), sdtlpdtrp0(xN, sK41)), inference(cnf_transformation, [], [f343])).
fof(f1801, plain, spl44_11, inference(avatar_contradiction_clause, [], [f1800])).
fof(f1800, plain, ($false | spl44_11), inference(subsumption_resolution, [], [f1799, f573])).
fof(f1799, plain, (~ aElementOf0(sK41, szNzAzT0) | spl44_11), inference(resolution, [], [f779, f562])).
fof(f562, plain, ! [X0] : (aSet0(sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f205])).
fof(f779, plain, (~ aSet0(sdtlpdtrp0(xN, sK41)) | spl44_11), inference(avatar_component_clause, [], [f777])).
fof(f777, plain, (spl44_11 <=> aSet0(sdtlpdtrp0(xN, sK41))), introduced(avatar_definition, [new_symbols(naming, [spl44_11])])).
fof(f780, plain, (~ spl44_11 | spl44_5), inference(avatar_split_clause, [], [f769, f675, f777])).
fof(f769, plain, (aElement0(szmzizndt0(sdtlpdtrp0(xN, sK41))) | ~ aSet0(sdtlpdtrp0(xN, sK41))), inference(resolution, [], [f586, f344])).
fof(f344, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM579+3.p', mEOfElem)).
fof(f673, plain, (spl44_4 | ~ spl44_2), inference(avatar_split_clause, [], [f593, f660, f670])).
fof(f593, plain, (~ (xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | aElementOf0(sK43, sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))), inference(cnf_transformation, [], [f343])).
fof(f668, plain, (~ spl44_3 | ~ spl44_2), inference(avatar_split_clause, [], [f594, f660, f665])).
fof(f594, plain, (~ (xK = sbrdtbr0(sdtpldt0(sK42, szmzizndt0(sdtlpdtrp0(xN, sK41))))) | ~ aElementOf0(sK43, xS)), inference(cnf_transformation, [], [f343])).