fof(f175409, plain, $false, inference(avatar_sat_refutation, [], [f857, f3276, f8833, f10303, f13681, f13685, f108917, f110144, f140234, f140307, f142427, f175408])).
fof(f175408, plain, (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468 | ~ spl18_1585), inference(avatar_contradiction_clause, [], [f175407])).
fof(f175407, plain, ($false | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468 | ~ spl18_1585)), inference(subsumption_resolution, [], [f175406, f128410])).
fof(f128410, plain, (iLess0(sdtpldt0(xr, sdtpldt0(xn, xm)), sdtpldt0(xp, sdtpldt0(xn, xm))) | ~ spl18_1585), inference(avatar_component_clause, [], [f128408])).
fof(f128408, plain, (spl18_1585 <=> iLess0(sdtpldt0(xr, sdtpldt0(xn, xm)), sdtpldt0(xp, sdtpldt0(xn, xm)))), introduced(avatar_definition, [new_symbols(naming, [spl18_1585])])).
fof(f175406, plain, (~ iLess0(sdtpldt0(xr, sdtpldt0(xn, xm)), sdtpldt0(xp, sdtpldt0(xn, xm))) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468)), inference(forward_demodulation, [], [f175405, f163059])).
fof(f163059, plain, ((sdtpldt0(sdtpldt0(xn, xm), xr) = sdtpldt0(xr, sdtpldt0(xn, xm))) | ~ spl18_51), inference(resolution, [], [f3250, f530])).
fof(f530, plain, ! [X13] : (~ aNaturalNumber0(X13) | (sdtpldt0(X13, xr) = sdtpldt0(xr, X13))), inference(resolution, [], [f189, f298])).
fof(f298, plain, aNaturalNumber0(xr), inference(cnf_transformation, [], [f176])).
fof(f176, plain, (isPrime0(xr) & ! [X0] : ((xr = X0) | (sz10 = X0) | (~ doDivides0(X0, xr) & ! [X1] : (~ (sdtasdt0(X0, X1) = xr) | ~ aNaturalNumber0(X1))) | ~ aNaturalNumber0(X0)) & ~ (sz10 = xr) & ~ (sz00 = xr) & doDivides0(xr, xk) & ((xk = sdtasdt0(xr, sK13)) & aNaturalNumber0(sK13)) & aNaturalNumber0(xr)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13])], [f134, f175])).
fof(f175, plain, (? [X2] : ((xk = sdtasdt0(xr, X2)) & aNaturalNumber0(X2)) => ((xk = sdtasdt0(xr, sK13)) & aNaturalNumber0(sK13))), introduced(choice_axiom, [])).
fof(f134, plain, (isPrime0(xr) & ! [X0] : ((xr = X0) | (sz10 = X0) | (~ doDivides0(X0, xr) & ! [X1] : (~ (sdtasdt0(X0, X1) = xr) | ~ aNaturalNumber0(X1))) | ~ aNaturalNumber0(X0)) & ~ (sz10 = xr) & ~ (sz00 = xr) & doDivides0(xr, xk) & ? [X2] : ((xk = sdtasdt0(xr, X2)) & aNaturalNumber0(X2)) & aNaturalNumber0(xr)), inference(flattening, [], [f133])).
fof(f133, plain, (isPrime0(xr) & ! [X0] : (((xr = X0) | (sz10 = X0)) | ((~ doDivides0(X0, xr) & ! [X1] : (~ (sdtasdt0(X0, X1) = xr) | ~ aNaturalNumber0(X1))) | ~ aNaturalNumber0(X0))) & ~ (sz10 = xr) & ~ (sz00 = xr) & doDivides0(xr, xk) & ? [X2] : ((xk = sdtasdt0(xr, X2)) & aNaturalNumber0(X2)) & aNaturalNumber0(xr)), inference(ennf_transformation, [], [f59])).
fof(f59, plain, (isPrime0(xr) & ! [X0] : (((doDivides0(X0, xr) | ? [X1] : ((sdtasdt0(X0, X1) = xr) & aNaturalNumber0(X1))) & aNaturalNumber0(X0)) => ((xr = X0) | (sz10 = X0))) & ~ (sz10 = xr) & ~ (sz00 = xr) & doDivides0(xr, xk) & ? [X2] : ((xk = sdtasdt0(xr, X2)) & aNaturalNumber0(X2)) & aNaturalNumber0(xr)), inference(rectify, [], [f48])).
fof(f48, plain, (isPrime0(xr) & ! [X0] : (((doDivides0(X0, xr) | ? [X1] : ((sdtasdt0(X0, X1) = xr) & aNaturalNumber0(X1))) & aNaturalNumber0(X0)) => ((xr = X0) | (sz10 = X0))) & ~ (sz10 = xr) & ~ (sz00 = xr) & doDivides0(xr, xk) & ? [X0] : ((xk = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) & aNaturalNumber0(xr)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__2342)).
fof(f189, plain, ! [X0, X1] : (~ aNaturalNumber0(X1) | (sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f67])).
fof(f67, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f66])).
fof(f66, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtpldt0(X0, X1) = sdtpldt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', mAddComm)).
fof(f3250, plain, (aNaturalNumber0(sdtpldt0(xn, xm)) | ~ spl18_51), inference(avatar_component_clause, [], [f3249])).
fof(f3249, plain, (spl18_51 <=> aNaturalNumber0(sdtpldt0(xn, xm))), introduced(avatar_definition, [new_symbols(naming, [spl18_51])])).
fof(f175405, plain, (~ iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468)), inference(subsumption_resolution, [], [f175404, f252])).
fof(f252, plain, aNaturalNumber0(xn), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (aNaturalNumber0(xp) & aNaturalNumber0(xm) & aNaturalNumber0(xn)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__1837)).
fof(f175404, plain, (~ iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | ~ aNaturalNumber0(xn) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468)), inference(subsumption_resolution, [], [f175403, f253])).
fof(f253, plain, aNaturalNumber0(xm), inference(cnf_transformation, [], [f39])).
fof(f175403, plain, (~ iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | ~ aNaturalNumber0(xm) | ~ aNaturalNumber0(xn) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468)), inference(subsumption_resolution, [], [f175402, f323])).
fof(f323, plain, ~ doDivides0(xr, xm), inference(cnf_transformation, [], [f135])).
fof(f135, plain, (~ doDivides0(xr, xm) & ! [X0] : (~ (xm = sdtasdt0(xr, X0)) | ~ aNaturalNumber0(X0)) & ~ doDivides0(xr, xn) & ! [X1] : (~ (xn = sdtasdt0(xr, X1)) | ~ aNaturalNumber0(X1))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ~ (doDivides0(xr, xm) | ? [X0] : ((xm = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) | doDivides0(xr, xn) | ? [X1] : ((xn = sdtasdt0(xr, X1)) & aNaturalNumber0(X1))), inference(rectify, [], [f53])).
fof(f53, plain, ~ (doDivides0(xr, xm) | ? [X0] : ((xm = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) | doDivides0(xr, xn) | ? [X0] : ((xn = sdtasdt0(xr, X0)) & aNaturalNumber0(X0))), inference(negated_conjecture, [], [f52])).
fof(f52, plain, ~ (doDivides0(xr, xm) | ? [X0] : ((xm = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) | doDivides0(xr, xn) | ? [X0] : ((xn = sdtasdt0(xr, X0)) & aNaturalNumber0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__)).
fof(f175402, plain, (~ iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | doDivides0(xr, xm) | ~ aNaturalNumber0(xm) | ~ aNaturalNumber0(xn) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166 | spl18_468)), inference(subsumption_resolution, [], [f175332, f36751])).
fof(f36751, plain, (~ sP1(xn, xr) | spl18_468), inference(avatar_component_clause, [], [f36750])).
fof(f36750, plain, (spl18_468 <=> sP1(xn, xr)), introduced(avatar_definition, [new_symbols(naming, [spl18_468])])).
fof(f175332, plain, (~ iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(xn, xr) | doDivides0(xr, xm) | ~ aNaturalNumber0(xm) | ~ aNaturalNumber0(xn) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166)), inference(equality_resolution, [], [f128951])).
fof(f128951, plain, (! [X78, X79] : (~ (sdtasdt0(xn, xm) = sdtasdt0(X78, X79)) | ~ iLess0(sdtpldt0(sdtpldt0(X78, X79), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(X78, xr) | doDivides0(xr, X79) | ~ aNaturalNumber0(X79) | ~ aNaturalNumber0(X78)) | (spl18_18 | ~ spl18_51 | ~ spl18_139 | ~ spl18_166)), inference(subsumption_resolution, [], [f109822, f816])).
fof(f816, plain, (~ sP0(xr) | spl18_18), inference(avatar_component_clause, [], [f814])).
fof(f814, plain, (spl18_18 <=> sP0(xr)), introduced(avatar_definition, [new_symbols(naming, [spl18_18])])).
fof(f109822, plain, (! [X78, X79] : (~ (sdtasdt0(xn, xm) = sdtasdt0(X78, X79)) | ~ iLess0(sdtpldt0(sdtpldt0(X78, X79), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(X78, xr) | doDivides0(xr, X79) | sP0(xr) | ~ aNaturalNumber0(X79) | ~ aNaturalNumber0(X78)) | (~ spl18_51 | ~ spl18_139 | ~ spl18_166)), inference(subsumption_resolution, [], [f109821, f8829])).
fof(f8829, plain, (aNaturalNumber0(xr) | ~ spl18_139), inference(avatar_component_clause, [], [f8828])).
fof(f8828, plain, (spl18_139 <=> aNaturalNumber0(xr)), introduced(avatar_definition, [new_symbols(naming, [spl18_139])])).
fof(f109821, plain, (! [X78, X79] : (~ (sdtasdt0(xn, xm) = sdtasdt0(X78, X79)) | ~ iLess0(sdtpldt0(sdtpldt0(X78, X79), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(X78, xr) | doDivides0(xr, X79) | sP0(xr) | ~ aNaturalNumber0(xr) | ~ aNaturalNumber0(X79) | ~ aNaturalNumber0(X78)) | (~ spl18_51 | ~ spl18_166)), inference(subsumption_resolution, [], [f16962, f10288])).
fof(f10288, plain, (aNaturalNumber0(sK14) | ~ spl18_166), inference(avatar_component_clause, [], [f10287])).
fof(f10287, plain, (spl18_166 <=> aNaturalNumber0(sK14)), introduced(avatar_definition, [new_symbols(naming, [spl18_166])])).
fof(f16962, plain, (! [X78, X79] : (~ (sdtasdt0(xn, xm) = sdtasdt0(X78, X79)) | ~ iLess0(sdtpldt0(sdtpldt0(X78, X79), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(X78, xr) | doDivides0(xr, X79) | ~ aNaturalNumber0(sK14) | sP0(xr) | ~ aNaturalNumber0(xr) | ~ aNaturalNumber0(X79) | ~ aNaturalNumber0(X78)) | ~ spl18_51), inference(superposition, [], [f16123, f310])).
fof(f310, plain, (sdtasdt0(xn, xm) = sdtasdt0(xr, sK14)), inference(cnf_transformation, [], [f179])).
fof(f179, plain, (doDivides0(xr, sdtasdt0(xn, xm)) & ((sdtasdt0(xn, xm) = sdtasdt0(xr, sK14)) & aNaturalNumber0(sK14)) & ((xk = sdtpldt0(xr, sK15)) & aNaturalNumber0(sK15))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15])], [f60, f178, f177])).
fof(f177, plain, (? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) => ((sdtasdt0(xn, xm) = sdtasdt0(xr, sK14)) & aNaturalNumber0(sK14))), introduced(choice_axiom, [])).
fof(f178, plain, (? [X1] : ((xk = sdtpldt0(xr, X1)) & aNaturalNumber0(X1)) => ((xk = sdtpldt0(xr, sK15)) & aNaturalNumber0(sK15))), introduced(choice_axiom, [])).
fof(f60, plain, (doDivides0(xr, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) & ? [X1] : ((xk = sdtpldt0(xr, X1)) & aNaturalNumber0(X1))), inference(rectify, [], [f49])).
fof(f49, plain, (doDivides0(xr, sdtasdt0(xn, xm)) & ? [X0] : ((sdtasdt0(xn, xm) = sdtasdt0(xr, X0)) & aNaturalNumber0(X0)) & ? [X0] : ((xk = sdtpldt0(xr, X0)) & aNaturalNumber0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__2362)).
fof(f16123, plain, (! [X4, X2, X0, X1] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(xp, sdtpldt0(xn, xm))) | sP1(X0, X2) | doDivides0(X2, X1) | ~ aNaturalNumber0(X4) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)) | ~ spl18_51), inference(backward_demodulation, [], [f269, f16087])).
fof(f16087, plain, ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(xp, sdtpldt0(xn, xm))) | ~ spl18_51), inference(resolution, [], [f528, f3250])).
fof(f528, plain, ! [X11] : (~ aNaturalNumber0(X11) | (sdtpldt0(X11, xp) = sdtpldt0(xp, X11))), inference(resolution, [], [f189, f254])).
fof(f254, plain, aNaturalNumber0(xp), inference(cnf_transformation, [], [f39])).
fof(f269, plain, ! [X4, X2, X0, X1] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | doDivides0(X2, X1) | ~ aNaturalNumber0(X4) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ((sdtasdt0(X2, sK9(X1, X2)) = X1) & aNaturalNumber0(sK9(X1, X2)))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X4] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | ~ aNaturalNumber0(X4))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9])], [f167, f168])).
fof(f168, plain, ! [X2, X1] : (? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3)) => ((sdtasdt0(X2, sK9(X1, X2)) = X1) & aNaturalNumber0(sK9(X1, X2)))), introduced(choice_axiom, [])).
fof(f167, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X4] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X4)) | ~ aNaturalNumber0(X4))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(rectify, [], [f138])).
fof(f138, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | sP1(X0, X2) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | sP0(X2) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(definition_folding, [], [f127, e137, e136])).
fof(f136, plain, ! [X2] : ((~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ sP0(X2)), inference(usedef, [], [e136])).
fof(e136, plain, ! [X2] : (sP0(X2) <=> (~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f137, plain, ! [X0, X2] : ((doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ sP1(X0, X2)), inference(usedef, [], [e137])).
fof(e137, plain, ! [X0, X2] : (sP1(X0, X2) <=> (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f127, plain, ! [X0, X1, X2] : ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) | (~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | (~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f126])).
fof(f126, plain, ! [X0, X1, X2] : (((((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7)))) | ~ iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp))) | ((~ doDivides0(X2, sdtasdt0(X0, X1)) & ! [X3] : (~ (sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) | ~ aNaturalNumber0(X3))) | (~ isPrime0(X2) & (? [X4] : ((~ (X2 = X4) & ~ (sz10 = X4)) & (doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4))) | (sz10 = X2) | (sz00 = X2))))) | (~ aNaturalNumber0(X2) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ! [X0, X1, X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (((doDivides0(X2, sdtasdt0(X0, X1)) | ? [X3] : ((sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) & aNaturalNumber0(X3))) & (isPrime0(X2) | (! [X4] : ((doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) => ((X2 = X4) | (sz10 = X4))) & ~ (sz10 = X2) & ~ (sz00 = X2)))) => (iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) => ((doDivides0(X2, X1) & ? [X6] : ((sdtasdt0(X2, X6) = X1) & aNaturalNumber0(X6))) | (doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))))))), inference(rectify, [], [f40])).
fof(f40, plain, ! [X0, X1, X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (((doDivides0(X2, sdtasdt0(X0, X1)) | ? [X3] : ((sdtasdt0(X0, X1) = sdtasdt0(X2, X3)) & aNaturalNumber0(X3))) & (isPrime0(X2) | (! [X3] : ((doDivides0(X3, X2) & ? [X4] : ((sdtasdt0(X3, X4) = X2) & aNaturalNumber0(X4)) & aNaturalNumber0(X3)) => ((X2 = X3) | (sz10 = X3))) & ~ (sz10 = X2) & ~ (sz00 = X2)))) => (iLess0(sdtpldt0(sdtpldt0(X0, X1), X2), sdtpldt0(sdtpldt0(xn, xm), xp)) => ((doDivides0(X2, X1) & ? [X3] : ((sdtasdt0(X2, X3) = X1) & aNaturalNumber0(X3))) | (doDivides0(X2, X0) & ? [X3] : ((sdtasdt0(X2, X3) = X0) & aNaturalNumber0(X3))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__1799)).
fof(f142427, plain, (spl18_1585 | ~ spl18_51 | ~ spl18_200), inference(avatar_split_clause, [], [f142425, f13590, f3249, f128408])).
fof(f13590, plain, (spl18_200 <=> iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp))), introduced(avatar_definition, [new_symbols(naming, [spl18_200])])).
fof(f142425, plain, (iLess0(sdtpldt0(xr, sdtpldt0(xn, xm)), sdtpldt0(xp, sdtpldt0(xn, xm))) | (~ spl18_51 | ~ spl18_200)), inference(backward_demodulation, [], [f142189, f139251])).
fof(f139251, plain, ((sdtpldt0(sdtpldt0(xn, xm), xr) = sdtpldt0(xr, sdtpldt0(xn, xm))) | ~ spl18_51), inference(resolution, [], [f3250, f530])).
fof(f142189, plain, (iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(xp, sdtpldt0(xn, xm))) | (~ spl18_51 | ~ spl18_200)), inference(forward_demodulation, [], [f13592, f139249])).
fof(f139249, plain, ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(xp, sdtpldt0(xn, xm))) | ~ spl18_51), inference(resolution, [], [f3250, f528])).
fof(f13592, plain, (iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) | ~ spl18_200), inference(avatar_component_clause, [], [f13590])).
fof(f140307, plain, (~ spl18_1582 | spl18_19 | ~ spl18_51), inference(avatar_split_clause, [], [f140306, f3249, f834, f128396])).
fof(f128396, plain, (spl18_1582 <=> aNaturalNumber0(sdtpldt0(xr, sdtpldt0(xn, xm)))), introduced(avatar_definition, [new_symbols(naming, [spl18_1582])])).
fof(f834, plain, (spl18_19 <=> aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr))), introduced(avatar_definition, [new_symbols(naming, [spl18_19])])).
fof(f140306, plain, (~ aNaturalNumber0(sdtpldt0(xr, sdtpldt0(xn, xm))) | (spl18_19 | ~ spl18_51)), inference(forward_demodulation, [], [f836, f139251])).
fof(f836, plain, (~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr)) | spl18_19), inference(avatar_component_clause, [], [f834])).
fof(f140234, plain, (~ spl18_51 | ~ spl18_139 | spl18_1582), inference(avatar_contradiction_clause, [], [f140233])).
fof(f140233, plain, ($false | (~ spl18_51 | ~ spl18_139 | spl18_1582)), inference(subsumption_resolution, [], [f140232, f8829])).
fof(f140232, plain, (~ aNaturalNumber0(xr) | (~ spl18_51 | spl18_1582)), inference(subsumption_resolution, [], [f140231, f3250])).
fof(f140231, plain, (~ aNaturalNumber0(sdtpldt0(xn, xm)) | ~ aNaturalNumber0(xr) | spl18_1582), inference(resolution, [], [f128398, f187])).
fof(f187, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f62])).
fof(f62, plain, ! [X0, X1] : (aNaturalNumber0(sdtpldt0(X0, X1)) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f4])).
fof(f4, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtpldt0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', mSortsB)).
fof(f128398, plain, (~ aNaturalNumber0(sdtpldt0(xr, sdtpldt0(xn, xm))) | spl18_1582), inference(avatar_component_clause, [], [f128396])).
fof(f110144, plain, ~ spl18_18, inference(avatar_contradiction_clause, [], [f110143])).
fof(f110143, plain, ($false | ~ spl18_18), inference(subsumption_resolution, [], [f110113, f306])).
fof(f306, plain, isPrime0(xr), inference(cnf_transformation, [], [f176])).
fof(f110113, plain, (~ isPrime0(xr) | ~ spl18_18), inference(resolution, [], [f815, f264])).
fof(f264, plain, ! [X0] : (~ sP0(X0) | ~ isPrime0(X0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : ((~ isPrime0(X0) & ((~ (sK7(X0) = X0) & ~ (sz10 = sK7(X0)) & doDivides0(sK7(X0), X0) & ((sdtasdt0(sK7(X0), sK8(X0)) = X0) & aNaturalNumber0(sK8(X0))) & aNaturalNumber0(sK7(X0))) | (sz10 = X0) | (sz00 = X0))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f163, f165, f164])).
fof(f164, plain, ! [X0] : (? [X1] : (~ (X0 = X1) & ~ (sz10 = X1) & doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) => (~ (sK7(X0) = X0) & ~ (sz10 = sK7(X0)) & doDivides0(sK7(X0), X0) & ? [X2] : ((sdtasdt0(sK7(X0), X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(sK7(X0)))), introduced(choice_axiom, [])).
fof(f165, plain, ! [X0] : (? [X2] : ((sdtasdt0(sK7(X0), X2) = X0) & aNaturalNumber0(X2)) => ((sdtasdt0(sK7(X0), sK8(X0)) = X0) & aNaturalNumber0(sK8(X0)))), introduced(choice_axiom, [])).
fof(f163, plain, ! [X0] : ((~ isPrime0(X0) & (? [X1] : (~ (X0 = X1) & ~ (sz10 = X1) & doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) & aNaturalNumber0(X1)) | (sz10 = X0) | (sz00 = X0))) | ~ sP0(X0)), inference(rectify, [], [f162])).
fof(f162, plain, ! [X2] : ((~ isPrime0(X2) & (? [X4] : (~ (X2 = X4) & ~ (sz10 = X4) & doDivides0(X4, X2) & ? [X5] : ((sdtasdt0(X4, X5) = X2) & aNaturalNumber0(X5)) & aNaturalNumber0(X4)) | (sz10 = X2) | (sz00 = X2))) | ~ sP0(X2)), inference(nnf_transformation, [], [f136])).
fof(f815, plain, (sP0(xr) | ~ spl18_18), inference(avatar_component_clause, [], [f814])).
fof(f108917, plain, ~ spl18_468, inference(avatar_contradiction_clause, [], [f108916])).
fof(f108916, plain, ($false | ~ spl18_468), inference(subsumption_resolution, [], [f108901, f321])).
fof(f321, plain, ~ doDivides0(xr, xn), inference(cnf_transformation, [], [f135])).
fof(f108901, plain, (doDivides0(xr, xn) | ~ spl18_468), inference(resolution, [], [f36752, f257])).
fof(f257, plain, ! [X0, X1] : (~ sP1(X0, X1) | doDivides0(X1, X0)), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ! [X0, X1] : ((doDivides0(X1, X0) & ((sdtasdt0(X1, sK6(X0, X1)) = X0) & aNaturalNumber0(sK6(X0, X1)))) | ~ sP1(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f159, f160])).
fof(f160, plain, ! [X1, X0] : (? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2)) => ((sdtasdt0(X1, sK6(X0, X1)) = X0) & aNaturalNumber0(sK6(X0, X1)))), introduced(choice_axiom, [])).
fof(f159, plain, ! [X0, X1] : ((doDivides0(X1, X0) & ? [X2] : ((sdtasdt0(X1, X2) = X0) & aNaturalNumber0(X2))) | ~ sP1(X0, X1)), inference(rectify, [], [f158])).
fof(f158, plain, ! [X0, X2] : ((doDivides0(X2, X0) & ? [X7] : ((sdtasdt0(X2, X7) = X0) & aNaturalNumber0(X7))) | ~ sP1(X0, X2)), inference(nnf_transformation, [], [f137])).
fof(f36752, plain, (sP1(xn, xr) | ~ spl18_468), inference(avatar_component_clause, [], [f36750])).
fof(f13685, plain, ~ spl18_199, inference(avatar_split_clause, [], [f316, f13583])).
fof(f13583, plain, (spl18_199 <=> (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr))), introduced(avatar_definition, [new_symbols(naming, [spl18_199])])).
fof(f316, plain, ~ (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr)), inference(cnf_transformation, [], [f183])).
fof(f183, plain, (sdtlseqdt0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) & ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(sdtpldt0(xn, xm), xr), sK17)) & aNaturalNumber0(sK17)) & ~ (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17])], [f51, f182])).
fof(f182, plain, (? [X0] : ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(sdtpldt0(xn, xm), xr), X0)) & aNaturalNumber0(X0)) => ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(sdtpldt0(xn, xm), xr), sK17)) & aNaturalNumber0(sK17))), introduced(choice_axiom, [])).
fof(f51, plain, (sdtlseqdt0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) & ? [X0] : ((sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(sdtpldt0(xn, xm), xr), X0)) & aNaturalNumber0(X0)) & ~ (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', m__2478)).
fof(f13681, plain, (spl18_199 | spl18_200 | ~ spl18_19 | ~ spl18_23), inference(avatar_split_clause, [], [f13680, f854, f834, f13590, f13583])).
fof(f854, plain, (spl18_23 <=> aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xp))), introduced(avatar_definition, [new_symbols(naming, [spl18_23])])).
fof(f13680, plain, (iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) | (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr)) | (~ spl18_19 | ~ spl18_23)), inference(subsumption_resolution, [], [f13679, f835])).
fof(f835, plain, (aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr)) | ~ spl18_19), inference(avatar_component_clause, [], [f834])).
fof(f13679, plain, (iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) | (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr)) | ~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr)) | ~ spl18_23), inference(subsumption_resolution, [], [f10431, f856])).
fof(f856, plain, (aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xp)) | ~ spl18_23), inference(avatar_component_clause, [], [f854])).
fof(f10431, plain, (iLess0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)) | (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(xn, xm), xr)) | ~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xp)) | ~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr))), inference(resolution, [], [f319, f230])).
fof(f230, plain, ! [X0, X1] : (~ sdtlseqdt0(X0, X1) | iLess0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ! [X0, X1] : (iLess0(X0, X1) | ~ sdtlseqdt0(X0, X1) | (X0 = X1) | ~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0)), inference(flattening, [], [f106])).
fof(f106, plain, ! [X0, X1] : ((iLess0(X0, X1) | (~ sdtlseqdt0(X0, X1) | (X0 = X1))) | (~ aNaturalNumber0(X1) | ~ aNaturalNumber0(X0))), inference(ennf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X0, X1) & ~ (X0 = X1)) => iLess0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM508+3.p', mIH_03)).
fof(f319, plain, sdtlseqdt0(sdtpldt0(sdtpldt0(xn, xm), xr), sdtpldt0(sdtpldt0(xn, xm), xp)), inference(cnf_transformation, [], [f183])).
fof(f10303, plain, spl18_166, inference(avatar_split_clause, [], [f309, f10287])).
fof(f309, plain, aNaturalNumber0(sK14), inference(cnf_transformation, [], [f179])).
fof(f8833, plain, spl18_139, inference(avatar_split_clause, [], [f298, f8828])).
fof(f3276, plain, spl18_51, inference(avatar_contradiction_clause, [], [f3275])).
fof(f3275, plain, ($false | spl18_51), inference(subsumption_resolution, [], [f3274, f252])).
fof(f3274, plain, (~ aNaturalNumber0(xn) | spl18_51), inference(subsumption_resolution, [], [f3273, f253])).
fof(f3273, plain, (~ aNaturalNumber0(xm) | ~ aNaturalNumber0(xn) | spl18_51), inference(resolution, [], [f3251, f187])).
fof(f3251, plain, (~ aNaturalNumber0(sdtpldt0(xn, xm)) | spl18_51), inference(avatar_component_clause, [], [f3249])).
fof(f857, plain, (~ spl18_19 | spl18_23), inference(avatar_split_clause, [], [f852, f854, f834])).
fof(f852, plain, (aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xp)) | ~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr))), inference(subsumption_resolution, [], [f831, f317])).
fof(f317, plain, aNaturalNumber0(sK17), inference(cnf_transformation, [], [f183])).
fof(f831, plain, (aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xp)) | ~ aNaturalNumber0(sK17) | ~ aNaturalNumber0(sdtpldt0(sdtpldt0(xn, xm), xr))), inference(superposition, [], [f187, f318])).
fof(f318, plain, (sdtpldt0(sdtpldt0(xn, xm), xp) = sdtpldt0(sdtpldt0(sdtpldt0(xn, xm), xr), sK17)), inference(cnf_transformation, [], [f183])).