fof(f155999, plain, $false, inference(avatar_sat_refutation, [], [f673, f678, f719, f1286, f3673, f10090, f11156, f11235, f17451, f155918, f155993])).
fof(f155993, plain, (~ spl43_5 | ~ spl43_9 | ~ spl43_494), inference(avatar_contradiction_clause, [], [f155992])).
fof(f155992, plain, ($false | (~ spl43_5 | ~ spl43_9 | ~ spl43_494)), inference(subsumption_resolution, [], [f155991, f1170])).
fof(f1170, plain, sP8(xn), inference(resolution, [], [f1117, f548])).
fof(f548, plain, ! [X0] : (~ sP9(X0) | sP8(X0)), inference(cnf_transformation, [], [f332])).
fof(f332, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(rectify, [], [f331])).
fof(f331, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(nnf_transformation, [], [f225])).
fof(f225, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e225])).
fof(e225, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f1117, plain, sP9(xn), inference(resolution, [], [f1099, f572])).
fof(f572, plain, aElementOf0(xn, szNzAzT0), inference(cnf_transformation, [], [f86])).
fof(f86, plain, (aElementOf0(xm, szNzAzT0) & aElementOf0(xn, szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', m__3904)).
fof(f1099, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sP9(X0)), inference(subsumption_resolution, [], [f1098, f565])).
fof(f565, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f206])).
fof(f206, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aSet0(sdtlpdtrp0(xN, X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', m__3671)).
fof(f1098, plain, ! [X0] : (sP9(X0) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f562, f566])).
fof(f566, plain, ! [X0] : (isCountable0(sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f206])).
fof(f562, plain, ! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f337])).
fof(f337, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK41(X0), szNzAzT0) & aElementOf0(sK41(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f226, f336])).
fof(f336, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK41(X0), szNzAzT0) & aElementOf0(sK41(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f226, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f205, e225, e224])).
fof(f224, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e224])).
fof(e224, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f205, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f204])).
fof(f204, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', m__3623)).
fof(f155991, plain, (~ sP8(xn) | (~ spl43_5 | ~ spl43_9 | ~ spl43_494)), inference(subsumption_resolution, [], [f155966, f10796])).
fof(f10796, plain, (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xm)) | ~ spl43_494), inference(avatar_component_clause, [], [f10795])).
fof(f10795, plain, (spl43_494 <=> aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xm))), introduced(avatar_definition, [new_symbols(naming, [spl43_494])])).
fof(f155966, plain, (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xm)) | ~ sP8(xn) | (~ spl43_5 | ~ spl43_9)), inference(resolution, [], [f155946, f633])).
fof(f633, plain, ! [X0] : (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(equality_resolution, [], [f555])).
fof(f555, plain, ! [X0, X1] : (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(cnf_transformation, [], [f335])).
fof(f335, plain, ! [X0] : (! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ aElement0(X1)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(rectify, [], [f334])).
fof(f334, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(flattening, [], [f333])).
fof(f333, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(nnf_transformation, [], [f224])).
fof(f155946, plain, (! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xm))) | (~ spl43_5 | ~ spl43_9)), inference(subsumption_resolution, [], [f155932, f687])).
fof(f687, plain, (aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ~ spl43_9), inference(avatar_component_clause, [], [f685])).
fof(f685, plain, (spl43_9 <=> aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))), introduced(avatar_definition, [new_symbols(naming, [spl43_9])])).
fof(f155932, plain, (! [X1] : (~ aElementOf0(X1, sdtlpdtrp0(xN, xm)) | aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ~ aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))) | ~ spl43_5), inference(resolution, [], [f666, f353])).
fof(f353, plain, ! [X0, X3, X1] : (~ aSubsetOf0(X1, X0) | ~ aElementOf0(X3, X1) | aElementOf0(X3, X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (~ aElementOf0(sK13(X0, X1), X0) & aElementOf0(sK13(X0, X1), X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13])], [f237, f238])).
fof(f238, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) => (~ aElementOf0(sK13(X0, X1), X0) & aElementOf0(sK13(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f237, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(rectify, [], [f236])).
fof(f236, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(flattening, [], [f235])).
fof(f235, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1))) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(nnf_transformation, [], [f106])).
fof(f106, plain, ! [X0] : (! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1))) | ~ aSet0(X0)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, X0)) & aSet0(X1)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', mDefSub)).
fof(f666, plain, (aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ~ spl43_5), inference(avatar_component_clause, [], [f665])).
fof(f665, plain, (spl43_5 <=> aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))), introduced(avatar_definition, [new_symbols(naming, [spl43_5])])).
fof(f155918, plain, (spl43_5 | ~ spl43_21 | ~ spl43_170), inference(avatar_contradiction_clause, [], [f155917])).
fof(f155917, plain, ($false | (spl43_5 | ~ spl43_21 | ~ spl43_170)), inference(subsumption_resolution, [], [f155916, f573])).
fof(f573, plain, aElementOf0(xm, szNzAzT0), inference(cnf_transformation, [], [f86])).
fof(f155916, plain, (~ aElementOf0(xm, szNzAzT0) | (spl43_5 | ~ spl43_21 | ~ spl43_170)), inference(subsumption_resolution, [], [f155915, f2966])).
fof(f2966, plain, (aElementOf0(szszuzczcdt0(xn), szNzAzT0) | ~ spl43_170), inference(avatar_component_clause, [], [f2965])).
fof(f2965, plain, (spl43_170 <=> aElementOf0(szszuzczcdt0(xn), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl43_170])])).
fof(f155915, plain, (~ aElementOf0(szszuzczcdt0(xn), szNzAzT0) | ~ aElementOf0(xm, szNzAzT0) | (spl43_5 | ~ spl43_21)), inference(subsumption_resolution, [], [f155910, f585])).
fof(f585, plain, sdtlseqdt0(szszuzczcdt0(xn), xm), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ((((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xm))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) | sP11) & sdtlseqdt0(szszuzczcdt0(xn), xm)), inference(definition_folding, [], [f212, e228, e227])).
fof(f227, plain, (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, xn)) & aElement0(X3))) | ~ sP10), inference(usedef, [], [e227])).
fof(e227, plain, (sP10 <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, xn)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f228, plain, ((~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X4, sdtlpdtrp0(xN, xm))) & sP10 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) | ~ sP11), inference(usedef, [], [e228])).
fof(e228, plain, (sP11 <=> (~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X4, sdtlpdtrp0(xN, xm))) & sP10 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f212, plain, ((((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xm))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) | (~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X4, sdtlpdtrp0(xN, xm))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, xn)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn)))) & sdtlseqdt0(szszuzczcdt0(xn), xm)), inference(flattening, [], [f211])).
fof(f211, plain, ((((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xm))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) | (((~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X4, sdtlpdtrp0(xN, xm)))) & (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, xn)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))))) & (! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))))) & sdtlseqdt0(szszuzczcdt0(xn), xm)), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ~ (sdtlseqdt0(szszuzczcdt0(xn), xm) => (~ ((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xm)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) & ((! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) => ((! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, xn)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))) => (aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, xm)) => aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))))))))), inference(rectify, [], [f88])).
fof(f88, plain, ~ (sdtlseqdt0(szszuzczcdt0(xn), xm) => (~ ((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xm)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) & ((! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) => ((! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) & aElementOf0(X0, sdtlpdtrp0(xN, xn)) & aElement0(X0))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))) => (aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xm)) => aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))))))))), inference(negated_conjecture, [], [f87])).
fof(f87, plain, ~ (sdtlseqdt0(szszuzczcdt0(xn), xm) => (~ ((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xn)) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xm)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xm)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))) & ((! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) => ((! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) & aElementOf0(X0, sdtlpdtrp0(xN, xn)) & aElement0(X0))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))) => (aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xm)) => aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', m__)).
fof(f155910, plain, (~ sdtlseqdt0(szszuzczcdt0(xn), xm) | ~ aElementOf0(szszuzczcdt0(xn), szNzAzT0) | ~ aElementOf0(xm, szNzAzT0) | (spl43_5 | ~ spl43_21)), inference(resolution, [], [f40141, f568])).
fof(f568, plain, ! [X0, X1] : (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f208])).
fof(f208, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f207])).
fof(f207, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', m__3754)).
fof(f40141, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtlpdtrp0(xN, szszuzczcdt0(xn))) | (spl43_5 | ~ spl43_21)), inference(subsumption_resolution, [], [f40140, f795])).
fof(f795, plain, (aSet0(sdtlpdtrp0(xN, xm)) | ~ spl43_21), inference(avatar_component_clause, [], [f794])).
fof(f794, plain, (spl43_21 <=> aSet0(sdtlpdtrp0(xN, xm))), introduced(avatar_definition, [new_symbols(naming, [spl43_21])])).
fof(f40140, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtlpdtrp0(xN, szszuzczcdt0(xn))) | ~ aSet0(sdtlpdtrp0(xN, xm)) | spl43_5), inference(subsumption_resolution, [], [f40133, f1117])).
fof(f40133, plain, (~ sP9(xn) | ~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtlpdtrp0(xN, szszuzczcdt0(xn))) | ~ aSet0(sdtlpdtrp0(xN, xm)) | spl43_5), inference(resolution, [], [f3090, f667])).
fof(f667, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | spl43_5), inference(avatar_component_clause, [], [f665])).
fof(f3090, plain, ! [X0, X1] : (aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP9(X0) | ~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ aSet0(X1)), inference(subsumption_resolution, [], [f3081, f547])).
fof(f547, plain, ! [X0] : (aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP9(X0)), inference(cnf_transformation, [], [f332])).
fof(f3081, plain, ! [X0, X1] : (~ sP9(X0) | aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aSet0(X1)), inference(resolution, [], [f551, f638])).
fof(f638, plain, ! [X2, X0, X1] : (~ aSubsetOf0(X1, X2) | aSubsetOf0(X0, X2) | ~ aSubsetOf0(X0, X1) | ~ aSet0(X2) | ~ aSet0(X0)), inference(subsumption_resolution, [], [f359, f352])).
fof(f352, plain, ! [X0, X1] : (~ aSubsetOf0(X1, X0) | aSet0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f239])).
fof(f359, plain, ! [X2, X0, X1] : (aSubsetOf0(X0, X2) | ~ aSubsetOf0(X1, X2) | ~ aSubsetOf0(X0, X1) | ~ aSet0(X2) | ~ aSet0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0, X1, X2] : (aSubsetOf0(X0, X2) | ~ aSubsetOf0(X1, X2) | ~ aSubsetOf0(X0, X1) | ~ aSet0(X2) | ~ aSet0(X1) | ~ aSet0(X0)), inference(flattening, [], [f112])).
fof(f112, plain, ! [X0, X1, X2] : ((aSubsetOf0(X0, X2) | (~ aSubsetOf0(X1, X2) | ~ aSubsetOf0(X0, X1))) | (~ aSet0(X2) | ~ aSet0(X1) | ~ aSet0(X0))), inference(ennf_transformation, [], [f14])).
fof(f14, plain, ! [X0, X1, X2] : ((aSet0(X2) & aSet0(X1) & aSet0(X0)) => ((aSubsetOf0(X1, X2) & aSubsetOf0(X0, X1)) => aSubsetOf0(X0, X2))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', mSubTrans)).
fof(f551, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP9(X0)), inference(cnf_transformation, [], [f332])).
fof(f17451, plain, spl43_9, inference(avatar_contradiction_clause, [], [f17450])).
fof(f17450, plain, ($false | spl43_9), inference(subsumption_resolution, [], [f17448, f1117])).
fof(f17448, plain, (~ sP9(xn) | spl43_9), inference(resolution, [], [f686, f547])).
fof(f686, plain, (~ aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | spl43_9), inference(avatar_component_clause, [], [f685])).
fof(f11235, plain, (spl43_494 | ~ spl43_16 | ~ spl43_20), inference(avatar_split_clause, [], [f11191, f734, f716, f10795])).
fof(f716, plain, (spl43_16 <=> (szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm)))), introduced(avatar_definition, [new_symbols(naming, [spl43_16])])).
fof(f734, plain, (spl43_20 <=> aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm))), introduced(avatar_definition, [new_symbols(naming, [spl43_20])])).
fof(f11191, plain, (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xm)) | (~ spl43_16 | ~ spl43_20)), inference(backward_demodulation, [], [f736, f718])).
fof(f718, plain, ((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) | ~ spl43_16), inference(avatar_component_clause, [], [f716])).
fof(f736, plain, (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm)) | ~ spl43_20), inference(avatar_component_clause, [], [f734])).
fof(f11156, plain, (spl43_6 | ~ spl43_7 | ~ spl43_170), inference(avatar_contradiction_clause, [], [f11155])).
fof(f11155, plain, ($false | (spl43_6 | ~ spl43_7 | ~ spl43_170)), inference(subsumption_resolution, [], [f11154, f2966])).
fof(f11154, plain, (~ aElementOf0(szszuzczcdt0(xn), szNzAzT0) | (spl43_6 | ~ spl43_7)), inference(subsumption_resolution, [], [f11151, f585])).
fof(f11151, plain, (~ sdtlseqdt0(szszuzczcdt0(xn), xm) | ~ aElementOf0(szszuzczcdt0(xn), szNzAzT0) | (spl43_6 | ~ spl43_7)), inference(resolution, [], [f3844, f10017])).
fof(f10017, plain, (! [X0] : (aElementOf0(sK42, sdtlpdtrp0(xN, X0)) | ~ sdtlseqdt0(X0, xm) | ~ aElementOf0(X0, szNzAzT0)) | ~ spl43_7), inference(subsumption_resolution, [], [f10012, f573])).
fof(f10012, plain, (! [X0] : (aElementOf0(sK42, sdtlpdtrp0(xN, X0)) | ~ sdtlseqdt0(X0, xm) | ~ aElementOf0(X0, szNzAzT0) | ~ aElementOf0(xm, szNzAzT0)) | ~ spl43_7), inference(resolution, [], [f677, f567])).
fof(f567, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, sdtlpdtrp0(xN, X0)) | aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f208])).
fof(f677, plain, (aElementOf0(sK42, sdtlpdtrp0(xN, xm)) | ~ spl43_7), inference(avatar_component_clause, [], [f675])).
fof(f675, plain, (spl43_7 <=> aElementOf0(sK42, sdtlpdtrp0(xN, xm))), introduced(avatar_definition, [new_symbols(naming, [spl43_7])])).
fof(f3844, plain, (~ aElementOf0(sK42, sdtlpdtrp0(xN, szszuzczcdt0(xn))) | spl43_6), inference(subsumption_resolution, [], [f3833, f1117])).
fof(f3833, plain, (~ aElementOf0(sK42, sdtlpdtrp0(xN, szszuzczcdt0(xn))) | ~ sP9(xn) | spl43_6), inference(resolution, [], [f550, f672])).
fof(f672, plain, (~ aElementOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | spl43_6), inference(avatar_component_clause, [], [f670])).
fof(f670, plain, (spl43_6 <=> aElementOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn))))), introduced(avatar_definition, [new_symbols(naming, [spl43_6])])).
fof(f550, plain, ! [X0, X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ sP9(X0)), inference(cnf_transformation, [], [f332])).
fof(f10090, plain, spl43_20, inference(avatar_contradiction_clause, [], [f10089])).
fof(f10089, plain, ($false | spl43_20), inference(subsumption_resolution, [], [f10085, f1118])).
fof(f1118, plain, sP9(xm), inference(resolution, [], [f1099, f573])).
fof(f10085, plain, (~ sP9(xm) | spl43_20), inference(resolution, [], [f735, f545])).
fof(f545, plain, ! [X0] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f332])).
fof(f735, plain, (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xm)), sdtlpdtrp0(xN, xm)) | spl43_20), inference(avatar_component_clause, [], [f734])).
fof(f3673, plain, spl43_170, inference(avatar_contradiction_clause, [], [f3672])).
fof(f3672, plain, ($false | spl43_170), inference(subsumption_resolution, [], [f3669, f572])).
fof(f3669, plain, (~ aElementOf0(xn, szNzAzT0) | spl43_170), inference(resolution, [], [f2967, f393])).
fof(f393, plain, ! [X0] : (aElementOf0(szszuzczcdt0(X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ! [X0] : ((~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM577+3.p', mSuccNum)).
fof(f2967, plain, (~ aElementOf0(szszuzczcdt0(xn), szNzAzT0) | spl43_170), inference(avatar_component_clause, [], [f2965])).
fof(f1286, plain, spl43_21, inference(avatar_contradiction_clause, [], [f1285])).
fof(f1285, plain, ($false | spl43_21), inference(subsumption_resolution, [], [f1284, f573])).
fof(f1284, plain, (~ aElementOf0(xm, szNzAzT0) | spl43_21), inference(resolution, [], [f796, f563])).
fof(f563, plain, ! [X0] : (aSet0(sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f206])).
fof(f796, plain, (~ aSet0(sdtlpdtrp0(xN, xm)) | spl43_21), inference(avatar_component_clause, [], [f794])).
fof(f719, plain, (spl43_4 | spl43_16), inference(avatar_split_clause, [], [f590, f716, f661])).
fof(f661, plain, (spl43_4 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl43_4])])).
fof(f590, plain, ((szmzizndt0(sdtlpdtrp0(xN, xn)) = szmzizndt0(sdtlpdtrp0(xN, xm))) | sP11), inference(cnf_transformation, [], [f229])).
fof(f678, plain, (~ spl43_4 | spl43_7), inference(avatar_split_clause, [], [f578, f675, f661])).
fof(f578, plain, (aElementOf0(sK42, sdtlpdtrp0(xN, xm)) | ~ sP11), inference(cnf_transformation, [], [f341])).
fof(f341, plain, ((~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & (~ aElementOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(sK42, sdtlpdtrp0(xN, xm))) & sP10 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) | ~ sP11), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f339, f340])).
fof(f340, plain, (? [X0] : (~ aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X0, sdtlpdtrp0(xN, xm))) => (~ aElementOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(sK42, sdtlpdtrp0(xN, xm)))), introduced(choice_axiom, [])).
fof(f339, plain, ((~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X0] : (~ aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X0, sdtlpdtrp0(xN, xm))) & sP10 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) | ~ sP11), inference(rectify, [], [f338])).
fof(f338, plain, ((~ aSubsetOf0(sdtlpdtrp0(xN, xm), sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & aElementOf0(X4, sdtlpdtrp0(xN, xm))) & sP10 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xn))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xn)), sdtlpdtrp0(xN, xn))) | ~ sP11), inference(nnf_transformation, [], [f228])).
fof(f673, plain, (~ spl43_4 | ~ spl43_6), inference(avatar_split_clause, [], [f579, f670, f661])).
fof(f579, plain, (~ aElementOf0(sK42, sdtmndt0(sdtlpdtrp0(xN, xn), szmzizndt0(sdtlpdtrp0(xN, xn)))) | ~ sP11), inference(cnf_transformation, [], [f341])).