fof(f453, plain, $false, inference(avatar_sat_refutation, [], [f358, f388, f393, f407, f451])).
fof(f451, plain, (~ spl15_2 | ~ spl15_9), inference(avatar_contradiction_clause, [], [f450])).
fof(f450, plain, ($false | (~ spl15_2 | ~ spl15_9)), inference(subsumption_resolution, [], [f443, f414])).
fof(f414, plain, (aElementOf0(szmzazxdt0(xS), szNzAzT0) | ~ spl15_9), inference(resolution, [], [f302, f387])).
fof(f387, plain, (aElementOf0(szmzazxdt0(xS), xS) | ~ spl15_9), inference(avatar_component_clause, [], [f385])).
fof(f385, plain, (spl15_9 <=> aElementOf0(szmzazxdt0(xS), xS)), introduced(avatar_definition, [new_symbols(naming, [spl15_9])])).
fof(f302, plain, ! [X0] : (~ aElementOf0(X0, xS) | aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, (isFinite0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xS)) & aSet0(xS)), inference(ennf_transformation, [], [f55])).
fof(f55, plain, (isFinite0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, szNzAzT0)) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', m__1986)).
fof(f443, plain, (~ aElementOf0(szmzazxdt0(xS), szNzAzT0) | ~ spl15_2), inference(resolution, [], [f253, f428])).
fof(f428, plain, (~ aElementOf0(szszuzczcdt0(szmzazxdt0(xS)), szNzAzT0) | ~ spl15_2), inference(resolution, [], [f321, f357])).
fof(f357, plain, (aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ spl15_2), inference(avatar_component_clause, [], [f355])).
fof(f355, plain, (spl15_2 <=> aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS))))), introduced(avatar_definition, [new_symbols(naming, [spl15_2])])).
fof(f321, plain, ! [X0] : (~ aSubsetOf0(xS, slbdtrb0(X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f204])).
fof(f204, plain, ! [X0] : ((~ aSubsetOf0(xS, slbdtrb0(X0)) & (~ aElementOf0(sK14(X0), slbdtrb0(X0)) & aElementOf0(sK14(X0), xS)) & ! [X2] : ((aElementOf0(X2, slbdtrb0(X0)) | ~ sdtlseqdt0(szszuzczcdt0(X2), X0) | ~ aElementOf0(X2, szNzAzT0)) & ((sdtlseqdt0(szszuzczcdt0(X2), X0) & aElementOf0(X2, szNzAzT0)) | ~ aElementOf0(X2, slbdtrb0(X0)))) & aSet0(slbdtrb0(X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f202, f203])).
fof(f203, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, slbdtrb0(X0)) & aElementOf0(X1, xS)) => (~ aElementOf0(sK14(X0), slbdtrb0(X0)) & aElementOf0(sK14(X0), xS))), introduced(choice_axiom, [])).
fof(f202, plain, ! [X0] : ((~ aSubsetOf0(xS, slbdtrb0(X0)) & ? [X1] : (~ aElementOf0(X1, slbdtrb0(X0)) & aElementOf0(X1, xS)) & ! [X2] : ((aElementOf0(X2, slbdtrb0(X0)) | ~ sdtlseqdt0(szszuzczcdt0(X2), X0) | ~ aElementOf0(X2, szNzAzT0)) & ((sdtlseqdt0(szszuzczcdt0(X2), X0) & aElementOf0(X2, szNzAzT0)) | ~ aElementOf0(X2, slbdtrb0(X0)))) & aSet0(slbdtrb0(X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(rectify, [], [f201])).
fof(f201, plain, ! [X0] : ((~ aSubsetOf0(xS, slbdtrb0(X0)) & ? [X2] : (~ aElementOf0(X2, slbdtrb0(X0)) & aElementOf0(X2, xS)) & ! [X1] : ((aElementOf0(X1, slbdtrb0(X0)) | ~ sdtlseqdt0(szszuzczcdt0(X1), X0) | ~ aElementOf0(X1, szNzAzT0)) & ((sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0)) | ~ aElementOf0(X1, slbdtrb0(X0)))) & aSet0(slbdtrb0(X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f200])).
fof(f200, plain, ! [X0] : ((~ aSubsetOf0(xS, slbdtrb0(X0)) & ? [X2] : (~ aElementOf0(X2, slbdtrb0(X0)) & aElementOf0(X2, xS)) & ! [X1] : ((aElementOf0(X1, slbdtrb0(X0)) | (~ sdtlseqdt0(szszuzczcdt0(X1), X0) | ~ aElementOf0(X1, szNzAzT0))) & ((sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0)) | ~ aElementOf0(X1, slbdtrb0(X0)))) & aSet0(slbdtrb0(X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(nnf_transformation, [], [f141])).
fof(f141, plain, ! [X0] : ((~ aSubsetOf0(xS, slbdtrb0(X0)) & ? [X2] : (~ aElementOf0(X2, slbdtrb0(X0)) & aElementOf0(X2, xS)) & ! [X1] : (aElementOf0(X1, slbdtrb0(X0)) <=> (sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0))) & aSet0(slbdtrb0(X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f140])).
fof(f140, plain, ! [X0] : (((~ aSubsetOf0(xS, slbdtrb0(X0)) & ? [X2] : (~ aElementOf0(X2, slbdtrb0(X0)) & aElementOf0(X2, xS))) & (! [X1] : (aElementOf0(X1, slbdtrb0(X0)) <=> (sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0))) & aSet0(slbdtrb0(X0)))) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f66])).
fof(f66, plain, ~ ? [X0] : (((! [X1] : (aElementOf0(X1, slbdtrb0(X0)) <=> (sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0))) & aSet0(slbdtrb0(X0))) => (aSubsetOf0(xS, slbdtrb0(X0)) | ! [X2] : (aElementOf0(X2, xS) => aElementOf0(X2, slbdtrb0(X0))))) & aElementOf0(X0, szNzAzT0)), inference(rectify, [], [f58])).
fof(f58, plain, ~ ? [X0] : (((! [X1] : (aElementOf0(X1, slbdtrb0(X0)) <=> (sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0))) & aSet0(slbdtrb0(X0))) => (aSubsetOf0(xS, slbdtrb0(X0)) | ! [X1] : (aElementOf0(X1, xS) => aElementOf0(X1, slbdtrb0(X0))))) & aElementOf0(X0, szNzAzT0)), inference(negated_conjecture, [], [f57])).
fof(f57, plain, ~ ? [X0] : (((! [X1] : (aElementOf0(X1, slbdtrb0(X0)) <=> (sdtlseqdt0(szszuzczcdt0(X1), X0) & aElementOf0(X1, szNzAzT0))) & aSet0(slbdtrb0(X0))) => (aSubsetOf0(xS, slbdtrb0(X0)) | ! [X1] : (aElementOf0(X1, xS) => aElementOf0(X1, slbdtrb0(X0))))) & aElementOf0(X0, szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', m__)).
fof(f253, plain, ! [X0] : (aElementOf0(szszuzczcdt0(X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ! [X0] : ((~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', mSuccNum)).
fof(f407, plain, ~ spl15_10, inference(avatar_contradiction_clause, [], [f406])).
fof(f406, plain, ($false | ~ spl15_10), inference(resolution, [], [f252, f405])).
fof(f405, plain, (! [X0] : ~ aElementOf0(X0, szNzAzT0) | ~ spl15_10), inference(subsumption_resolution, [], [f402, f322])).
fof(f322, plain, ! [X2] : ~ aElementOf0(X2, slcrc0), inference(equality_resolution, [], [f207])).
fof(f207, plain, ! [X2, X0] : (~ aElementOf0(X2, X0) | ~ (slcrc0 = X0)), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ! [X0] : (((slcrc0 = X0) | aElementOf0(sK5(X0), X0) | ~ aSet0(X0)) & ((! [X2] : ~ aElementOf0(X2, X0) & aSet0(X0)) | ~ (slcrc0 = X0))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5])], [f152, f153])).
fof(f153, plain, ! [X0] : (? [X1] : aElementOf0(X1, X0) => aElementOf0(sK5(X0), X0)), introduced(choice_axiom, [])).
fof(f152, plain, ! [X0] : (((slcrc0 = X0) | ? [X1] : aElementOf0(X1, X0) | ~ aSet0(X0)) & ((! [X2] : ~ aElementOf0(X2, X0) & aSet0(X0)) | ~ (slcrc0 = X0))), inference(rectify, [], [f151])).
fof(f151, plain, ! [X0] : (((slcrc0 = X0) | ? [X1] : aElementOf0(X1, X0) | ~ aSet0(X0)) & ((! [X1] : ~ aElementOf0(X1, X0) & aSet0(X0)) | ~ (slcrc0 = X0))), inference(flattening, [], [f150])).
fof(f150, plain, ! [X0] : (((slcrc0 = X0) | (? [X1] : aElementOf0(X1, X0) | ~ aSet0(X0))) & ((! [X1] : ~ aElementOf0(X1, X0) & aSet0(X0)) | ~ (slcrc0 = X0))), inference(nnf_transformation, [], [f69])).
fof(f69, plain, ! [X0] : ((slcrc0 = X0) <=> (! [X1] : ~ aElementOf0(X1, X0) & aSet0(X0))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0] : ((slcrc0 = X0) <=> (~ ? [X1] : aElementOf0(X1, X0) & aSet0(X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', mDefEmp)).
fof(f402, plain, (! [X0] : (aElementOf0(sK14(X0), slcrc0) | ~ aElementOf0(X0, szNzAzT0)) | ~ spl15_10), inference(backward_demodulation, [], [f319, f392])).
fof(f392, plain, ((slcrc0 = xS) | ~ spl15_10), inference(avatar_component_clause, [], [f390])).
fof(f390, plain, (spl15_10 <=> (slcrc0 = xS)), introduced(avatar_definition, [new_symbols(naming, [spl15_10])])).
fof(f319, plain, ! [X0] : (aElementOf0(sK14(X0), xS) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f204])).
fof(f252, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', mZeroNum)).
fof(f393, plain, (spl15_10 | spl15_1), inference(avatar_split_clause, [], [f314, f351, f390])).
fof(f351, plain, (spl15_1 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl15_1])])).
fof(f314, plain, (sP4 | (slcrc0 = xS)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, (sP4 | ((slcrc0 = xS) & ! [X0] : ~ aElementOf0(X0, xS))), inference(definition_folding, [], [f139, e148])).
fof(f148, plain, ((aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X1, xS)) & ! [X2] : (aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) <=> (sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (sdtlseqdt0(X3, szmzazxdt0(xS)) | ~ aElementOf0(X3, xS)) & aElementOf0(szmzazxdt0(xS), xS)) | ~ sP4), inference(usedef, [], [e148])).
fof(e148, plain, (sP4 <=> (aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X1, xS)) & ! [X2] : (aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) <=> (sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (sdtlseqdt0(X3, szmzazxdt0(xS)) | ~ aElementOf0(X3, xS)) & aElementOf0(szmzazxdt0(xS), xS))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f139, plain, ((aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X1, xS)) & ! [X2] : (aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) <=> (sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (sdtlseqdt0(X3, szmzazxdt0(xS)) | ~ aElementOf0(X3, xS)) & aElementOf0(szmzazxdt0(xS), xS)) | ((slcrc0 = xS) & ! [X0] : ~ aElementOf0(X0, xS))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, (~ ((slcrc0 = xS) & ~ ? [X0] : aElementOf0(X0, xS)) => (aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, xS) => aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS))))) & ! [X2] : (aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) <=> (sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (aElementOf0(X3, xS) => sdtlseqdt0(X3, szmzazxdt0(xS))) & aElementOf0(szmzazxdt0(xS), xS))), inference(rectify, [], [f56])).
fof(f56, plain, (~ ((slcrc0 = xS) & ~ ? [X0] : aElementOf0(X0, xS)) => (aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS))))) & ! [X0] : (aElementOf0(X0, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) <=> (sdtlseqdt0(szszuzczcdt0(X0), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X0, szNzAzT0))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X0] : (aElementOf0(X0, xS) => sdtlseqdt0(X0, szmzazxdt0(xS))) & aElementOf0(szmzazxdt0(xS), xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM545+2.p', m__2035)).
fof(f388, plain, (~ spl15_1 | spl15_9), inference(avatar_split_clause, [], [f305, f385, f351])).
fof(f305, plain, (aElementOf0(szmzazxdt0(xS), xS) | ~ sP4), inference(cnf_transformation, [], [f199])).
fof(f199, plain, ((aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X0] : (aElementOf0(X0, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X0, xS)) & ! [X1] : ((aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ sdtlseqdt0(szszuzczcdt0(X1), szszuzczcdt0(szmzazxdt0(xS))) | ~ aElementOf0(X1, szNzAzT0)) & ((sdtlseqdt0(szszuzczcdt0(X1), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X1, szNzAzT0)) | ~ aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X2] : (sdtlseqdt0(X2, szmzazxdt0(xS)) | ~ aElementOf0(X2, xS)) & aElementOf0(szmzazxdt0(xS), xS)) | ~ sP4), inference(rectify, [], [f198])).
fof(f198, plain, ((aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X1, xS)) & ! [X2] : ((aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) | ~ aElementOf0(X2, szNzAzT0)) & ((sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0)) | ~ aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (sdtlseqdt0(X3, szmzazxdt0(xS)) | ~ aElementOf0(X3, xS)) & aElementOf0(szmzazxdt0(xS), xS)) | ~ sP4), inference(flattening, [], [f197])).
fof(f197, plain, ((aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X1] : (aElementOf0(X1, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ aElementOf0(X1, xS)) & ! [X2] : ((aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | (~ sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) | ~ aElementOf0(X2, szNzAzT0))) & ((sdtlseqdt0(szszuzczcdt0(X2), szszuzczcdt0(szmzazxdt0(xS))) & aElementOf0(X2, szNzAzT0)) | ~ aElementOf0(X2, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))))) & aSet0(slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) & ! [X3] : (sdtlseqdt0(X3, szmzazxdt0(xS)) | ~ aElementOf0(X3, xS)) & aElementOf0(szmzazxdt0(xS), xS)) | ~ sP4), inference(nnf_transformation, [], [f148])).
fof(f358, plain, (~ spl15_1 | spl15_2), inference(avatar_split_clause, [], [f312, f355, f351])).
fof(f312, plain, (aSubsetOf0(xS, slbdtrb0(szszuzczcdt0(szmzazxdt0(xS)))) | ~ sP4), inference(cnf_transformation, [], [f199])).