fof(f5575, plain, $false, inference(avatar_sat_refutation, [], [f4951, f5574])).
fof(f5574, plain, ~ spl55_183, inference(avatar_contradiction_clause, [], [f5573])).
fof(f5573, plain, ($false | ~ spl55_183), inference(subsumption_resolution, [], [f5572, f690])).
fof(f690, plain, aElementOf0(sK53, szDzozmdt0(sdtlpdtrp0(xC, xi))), inference(cnf_transformation, [], [f389])).
fof(f389, plain, (aElementOf0(xx, sdtlcdtrc0(sdtlpdtrp0(xC, xi), szDzozmdt0(sdtlpdtrp0(xC, xi)))) & ((xx = sdtlpdtrp0(sdtlpdtrp0(xC, xi), sK53)) & aElementOf0(sK53, szDzozmdt0(sdtlpdtrp0(xC, xi))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK53])], [f88, f388])).
fof(f388, plain, (? [X0] : ((sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) & aElementOf0(X0, szDzozmdt0(sdtlpdtrp0(xC, xi)))) => ((xx = sdtlpdtrp0(sdtlpdtrp0(xC, xi), sK53)) & aElementOf0(sK53, szDzozmdt0(sdtlpdtrp0(xC, xi))))), introduced(choice_axiom, [])).
fof(f88, plain, (aElementOf0(xx, sdtlcdtrc0(sdtlpdtrp0(xC, xi), szDzozmdt0(sdtlpdtrp0(xC, xi)))) & ? [X0] : ((sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) & aElementOf0(X0, szDzozmdt0(sdtlpdtrp0(xC, xi))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM586+3.p', m__4200_02)).
fof(f5572, plain, (~ aElementOf0(sK53, szDzozmdt0(sdtlpdtrp0(xC, xi))) | ~ spl55_183), inference(trivial_inequality_removal, [], [f5571])).
fof(f5571, plain, (~ (xx = xx) | ~ aElementOf0(sK53, szDzozmdt0(sdtlpdtrp0(xC, xi))) | ~ spl55_183), inference(superposition, [], [f5570, f691])).
fof(f691, plain, (xx = sdtlpdtrp0(sdtlpdtrp0(xC, xi), sK53)), inference(cnf_transformation, [], [f389])).
fof(f5570, plain, (! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | ~ aElementOf0(X0, szDzozmdt0(sdtlpdtrp0(xC, xi)))) | ~ spl55_183), inference(backward_demodulation, [], [f704, f5559])).
fof(f5559, plain, ((szDzozmdt0(sdtlpdtrp0(xC, xi)) = slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) | ~ spl55_183), inference(resolution, [], [f4947, f687])).
fof(f687, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0)))), inference(cnf_transformation, [], [f387])).
fof(f387, plain, (! [X0] : ((sP16(X0) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & sP15(X0) & sP17(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), inference(rectify, [], [f245])).
fof(f245, plain, (! [X0] : ((sP16(X0) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & sP15(X0) & sP17(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X11] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X11) | ~ aElementOf0(X11, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), inference(definition_folding, [], [f219, e244, e243, e242, e241, e240])).
fof(f240, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP13(X0)), inference(usedef, [], [e240])).
fof(e240, plain, ! [X0] : (sP13(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f241, plain, ! [X0, X1] : ((~ aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X4, X1)))) & sP13(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP14(X0, X1)), inference(usedef, [], [e241])).
fof(e241, plain, ! [X0, X1] : (sP14(X0, X1) <=> (~ aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X4, X1)))) & sP13(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f242, plain, ! [X0] : (! [X7] : ((aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))) | ~ (xk = sbrdtbr0(X7)) | (~ aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & (? [X8] : (~ aElementOf0(X8, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X8, X7)) | ~ aSet0(X7)))) & (((xk = sbrdtbr0(X7)) & aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X9] : (aElementOf0(X9, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X9, X7)) & aSet0(X7)) | ~ aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))))) | ~ sP15(X0)), inference(usedef, [], [e242])).
fof(e242, plain, ! [X0] : (sP15(X0) <=> ! [X7] : ((aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))) | ~ (xk = sbrdtbr0(X7)) | (~ aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & (? [X8] : (~ aElementOf0(X8, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X8, X7)) | ~ aSet0(X7)))) & (((xk = sbrdtbr0(X7)) & aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X9] : (aElementOf0(X9, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X9, X7)) & aSet0(X7)) | ~ aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f243, plain, ! [X0] : (! [X1] : (((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X5] : (aElementOf0(X5, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X5) | aElementOf0(X5, X1)) & aElement0(X5))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | sP14(X0, X1) | ~ aSet0(X1)) | ~ sP16(X0)), inference(usedef, [], [e243])).
fof(e243, plain, ! [X0] : (sP16(X0) <=> ! [X1] : (((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X5] : (aElementOf0(X5, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X5) | aElementOf0(X5, X1)) & aElement0(X5))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | sP14(X0, X1) | ~ aSet0(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f244, plain, ! [X0] : (! [X10] : (aElementOf0(X10, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X10) & aElementOf0(X10, sdtlpdtrp0(xN, X0)) & aElement0(X10))) | ~ sP17(X0)), inference(usedef, [], [e244])).
fof(e244, plain, ! [X0] : (sP17(X0) <=> ! [X10] : (aElementOf0(X10, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X10) & aElementOf0(X10, sdtlpdtrp0(xN, X0)) & aElement0(X10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f219, plain, (! [X0] : ((! [X1] : (((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X5] : (aElementOf0(X5, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X5) | aElementOf0(X5, X1)) & aElement0(X5))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X4, X1)))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ aSet0(X1)) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & ! [X7] : ((aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))) | ~ (xk = sbrdtbr0(X7)) | (~ aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & (? [X8] : (~ aElementOf0(X8, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X8, X7)) | ~ aSet0(X7)))) & (((xk = sbrdtbr0(X7)) & aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X9] : (aElementOf0(X9, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X9, X7)) & aSet0(X7)) | ~ aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))))) & ! [X10] : (aElementOf0(X10, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X10) & aElementOf0(X10, sdtlpdtrp0(xN, X0)) & aElement0(X10))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X11] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X11) | ~ aElementOf0(X11, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), inference(flattening, [], [f218])).
fof(f218, plain, (! [X0] : ((! [X1] : (((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X5] : (aElementOf0(X5, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X5) | aElementOf0(X5, X1)) & aElement0(X5))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X6) | ~ aElementOf0(X6, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ((((~ aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ? [X4] : (~ aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X4, X1))))) & (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) & (! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))) | ~ aSet0(X1))) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & ! [X7] : ((aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))) | (~ (xk = sbrdtbr0(X7)) | (~ aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & (? [X8] : (~ aElementOf0(X8, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & aElementOf0(X8, X7)) | ~ aSet0(X7))))) & (((xk = sbrdtbr0(X7)) & aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X9] : (aElementOf0(X9, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X9, X7)) & aSet0(X7)) | ~ aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))))) & ! [X10] : (aElementOf0(X10, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X10) & aElementOf0(X10, sdtlpdtrp0(xN, X0)) & aElement0(X10))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X11] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X11) | ~ aElementOf0(X11, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), inference(ennf_transformation, [], [f103])).
fof(f103, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => (! [X1] : ((((! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) => ((! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) => (aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) | ((sbrdtbr0(X1) = xk) & (aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ! [X4] : (aElementOf0(X4, X1) => aElementOf0(X4, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))))))) & aSet0(X1)) => ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X5] : (aElementOf0(X5, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X5) | aElementOf0(X5, X1)) & aElement0(X5))) & ! [X6] : (aElementOf0(X6, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X6)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & ! [X7] : ((((xk = sbrdtbr0(X7)) & (aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (! [X8] : (aElementOf0(X8, X7) => aElementOf0(X8, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(X7)))) => aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0)))) & (aElementOf0(X7, szDzozmdt0(sdtlpdtrp0(xC, X0))) => ((xk = sbrdtbr0(X7)) & aSubsetOf0(X7, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X9] : (aElementOf0(X9, X7) => aElementOf0(X9, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(X7)))) & ! [X10] : (aElementOf0(X10, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X10) & aElementOf0(X10, sdtlpdtrp0(xN, X0)) & aElement0(X10))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X11] : (aElementOf0(X11, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X11)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0)))) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), inference(rectify, [], [f86])).
fof(f86, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => (! [X1] : ((((! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) => ((! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, X0)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) => (aElementOf0(X1, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk)) | ((sbrdtbr0(X1) = xk) & (aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))))))) & aSet0(X1)) => ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xc, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0))))) & ! [X2] : (aElementOf0(X2, sdtpldt0(X1, szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, X0)) = X2) | aElementOf0(X2, X1)) & aElement0(X2))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))) & (slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))), xk) = szDzozmdt0(sdtlpdtrp0(xC, X0))) & ! [X1] : ((((sbrdtbr0(X1) = xk) & (aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(X1)))) => aElementOf0(X1, szDzozmdt0(sdtlpdtrp0(xC, X0)))) & (aElementOf0(X1, szDzozmdt0(sdtlpdtrp0(xC, X0))) => ((sbrdtbr0(X1) = xk) & aSubsetOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(X1)))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & aFunction0(sdtlpdtrp0(xC, X0)))) & (szNzAzT0 = szDzozmdt0(xC)) & aFunction0(xC)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM586+3.p', m__4151)).
fof(f4947, plain, (aElementOf0(xi, szNzAzT0) | ~ spl55_183), inference(avatar_component_clause, [], [f4946])).
fof(f4946, plain, (spl55_183 <=> aElementOf0(xi, szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl55_183])])).
fof(f704, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | ~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk))), inference(cnf_transformation, [], [f395])).
fof(f395, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | (~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (~ (sbrdtbr0(X0) = xk) | (~ aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ((~ aElementOf0(sK54(X0), sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(sK54(X0), X0)) | ~ aSet0(X0)))) & sP18 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f393, f394])).
fof(f394, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(X1, X0)) => (~ aElementOf0(sK54(X0), sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(sK54(X0), X0))), introduced(choice_axiom, [])).
fof(f393, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | (~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (~ (sbrdtbr0(X0) = xk) | (~ aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & (? [X1] : (~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(X1, X0)) | ~ aSet0(X0)))) & sP18 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)))), inference(rectify, [], [f247])).
fof(f247, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | (~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (~ (sbrdtbr0(X0) = xk) | (~ aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & (? [X3] : (~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(X3, X0)) | ~ aSet0(X0)))) & sP18 & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)))), inference(definition_folding, [], [f221, e246])).
fof(f246, plain, (! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, xi)) & aElement0(X2))) | ~ sP18), inference(usedef, [], [e246])).
fof(e246, plain, (sP18 <=> ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, xi)) & aElement0(X2)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f221, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | (~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (~ (sbrdtbr0(X0) = xk) | (~ aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & (? [X3] : (~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(X3, X0)) | ~ aSet0(X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, xi)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)))), inference(flattening, [], [f220])).
fof(f220, plain, ! [X0] : (~ (sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) | (((~ aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (~ (sbrdtbr0(X0) = xk) | (~ aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & (? [X3] : (~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & aElementOf0(X3, X0)) | ~ aSet0(X0))))) & (! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, xi)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))))) & (! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))))), inference(ennf_transformation, [], [f104])).
fof(f104, plain, ~ ? [X0] : ((sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) & ((! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X2) & aElementOf0(X2, sdtlpdtrp0(xN, xi)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) | ((sbrdtbr0(X0) = xk) & (aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | (! [X3] : (aElementOf0(X3, X0) => aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) & aSet0(X0)))))))), inference(rectify, [], [f90])).
fof(f90, plain, ~ ? [X0] : ((sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) & ((! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) | ((sbrdtbr0(X0) = xk) & (aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | (! [X1] : (aElementOf0(X1, X0) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) & aSet0(X0)))))))), inference(negated_conjecture, [], [f89])).
fof(f89, plain, ~ ? [X0] : ((sdtlpdtrp0(sdtlpdtrp0(xC, xi), X0) = xx) & ((! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aElementOf0(X0, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) | ((sbrdtbr0(X0) = xk) & (aSubsetOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | (! [X1] : (aElementOf0(X1, X0) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) & aSet0(X0)))))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM586+3.p', m__)).
fof(f4951, plain, spl55_183, inference(avatar_split_clause, [], [f689, f4946])).
fof(f689, plain, aElementOf0(xi, szNzAzT0), inference(cnf_transformation, [], [f87])).
fof(f87, plain, aElementOf0(xi, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM586+3.p', m__4200)).