fof(f210820, plain, $false, inference(avatar_sat_refutation, [], [f2987, f20326, f20479, f65337, f210819])).
fof(f210819, plain, (~ spl42_119 | spl42_695 | ~ spl42_1459), inference(avatar_contradiction_clause, [], [f210818])).
fof(f210818, plain, ($false | (~ spl42_119 | spl42_695 | ~ spl42_1459)), inference(subsumption_resolution, [], [f210817, f556])).
fof(f556, plain, aElementOf0(sK40, szNzAzT0), inference(cnf_transformation, [], [f331])).
fof(f331, plain, ((szmzizndt0(sdtlpdtrp0(xN, sK40)) = szmzizndt0(sdtlpdtrp0(xN, sK41))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK40)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, sK41)) & ! [X3] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK40)), X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, sK40))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, sK40)) & ~ (sK40 = sK41) & aElementOf0(sK41, szNzAzT0) & aElementOf0(sK40, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40, sK41])], [f207, f330])).
fof(f330, plain, (? [X0, X1] : ((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X3] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & ~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ((szmzizndt0(sdtlpdtrp0(xN, sK40)) = szmzizndt0(sdtlpdtrp0(xN, sK41))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK40)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, sK41))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, sK41)) & ! [X3] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, sK40)), X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, sK40))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, sK40)) & ~ (sK40 = sK41) & aElementOf0(sK41, szNzAzT0) & aElementOf0(sK40, szNzAzT0))), introduced(choice_axiom, [])).
fof(f207, plain, ? [X0, X1] : ((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X3] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) & ~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f206])).
fof(f206, plain, ? [X0, X1] : (((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X3] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) & (~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ~ ! [X0, X1] : ((~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ~ ((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X3] : (aElementOf0(X3, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X3)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), inference(rectify, [], [f85])).
fof(f85, plain, ~ ! [X0, X1] : ((~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ~ ((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), inference(negated_conjecture, [], [f84])).
fof(f84, plain, ~ ! [X0, X1] : ((~ (X0 = X1) & aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ~ ((szmzizndt0(sdtlpdtrp0(xN, X0)) = szmzizndt0(sdtlpdtrp0(xN, X1))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', m__)).
fof(f210817, plain, (~ aElementOf0(sK40, szNzAzT0) | (~ spl42_119 | spl42_695 | ~ spl42_1459)), inference(subsumption_resolution, [], [f210816, f2974])).
fof(f2974, plain, (aElementOf0(sK41, szNzAzT0) | ~ spl42_119), inference(avatar_component_clause, [], [f2973])).
fof(f2973, plain, (spl42_119 <=> aElementOf0(sK41, szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_119])])).
fof(f210816, plain, (~ aElementOf0(sK41, szNzAzT0) | ~ aElementOf0(sK40, szNzAzT0) | (~ spl42_119 | spl42_695 | ~ spl42_1459)), inference(subsumption_resolution, [], [f210810, f21173])).
fof(f21173, plain, (~ sdtlseqdt0(sK40, sK41) | (~ spl42_119 | spl42_695)), inference(subsumption_resolution, [], [f21172, f556])).
fof(f21172, plain, (~ sdtlseqdt0(sK40, sK41) | ~ aElementOf0(sK40, szNzAzT0) | (~ spl42_119 | spl42_695)), inference(subsumption_resolution, [], [f21171, f2974])).
fof(f21171, plain, (~ sdtlseqdt0(sK40, sK41) | ~ aElementOf0(sK41, szNzAzT0) | ~ aElementOf0(sK40, szNzAzT0) | (~ spl42_119 | spl42_695)), inference(subsumption_resolution, [], [f21168, f558])).
fof(f558, plain, ~ (sK40 = sK41), inference(cnf_transformation, [], [f331])).
fof(f21168, plain, ((sK40 = sK41) | ~ sdtlseqdt0(sK40, sK41) | ~ aElementOf0(sK41, szNzAzT0) | ~ aElementOf0(sK40, szNzAzT0) | (~ spl42_119 | spl42_695)), inference(resolution, [], [f20848, f392])).
fof(f392, plain, ! [X0, X1] : (~ sdtlseqdt0(X1, X0) | (X0 = X1) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0, X1] : ((X0 = X1) | ~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f138])).
fof(f138, plain, ! [X0, X1] : (((X0 = X1) | (~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1))) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ((sdtlseqdt0(X1, X0) & sdtlseqdt0(X0, X1)) => (X0 = X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', mLessASymm)).
fof(f20848, plain, (sdtlseqdt0(sK41, sK40) | (~ spl42_119 | spl42_695)), inference(subsumption_resolution, [], [f20847, f2974])).
fof(f20847, plain, (sdtlseqdt0(sK41, sK40) | ~ aElementOf0(sK41, szNzAzT0) | spl42_695), inference(subsumption_resolution, [], [f20842, f556])).
fof(f20842, plain, (sdtlseqdt0(sK41, sK40) | ~ aElementOf0(sK40, szNzAzT0) | ~ aElementOf0(sK41, szNzAzT0) | spl42_695), inference(resolution, [], [f20325, f394])).
fof(f394, plain, ! [X0, X1] : (sdtlseqdt0(szszuzczcdt0(X1), X0) | sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0, X1] : (sdtlseqdt0(szszuzczcdt0(X1), X0) | sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f142])).
fof(f142, plain, ! [X0, X1] : ((sdtlseqdt0(szszuzczcdt0(X1), X0) | sdtlseqdt0(X0, X1)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f37])).
fof(f37, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(szszuzczcdt0(X1), X0) | sdtlseqdt0(X0, X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', mLessTotal)).
fof(f20325, plain, (~ sdtlseqdt0(szszuzczcdt0(sK40), sK41) | spl42_695), inference(avatar_component_clause, [], [f20323])).
fof(f20323, plain, (spl42_695 <=> sdtlseqdt0(szszuzczcdt0(sK40), sK41)), introduced(avatar_definition, [new_symbols(naming, [spl42_695])])).
fof(f210810, plain, (sdtlseqdt0(sK40, sK41) | ~ aElementOf0(sK41, szNzAzT0) | ~ aElementOf0(sK40, szNzAzT0) | ~ spl42_1459), inference(resolution, [], [f210786, f394])).
fof(f210786, plain, (~ sdtlseqdt0(szszuzczcdt0(sK41), sK40) | ~ spl42_1459), inference(subsumption_resolution, [], [f210785, f556])).
fof(f210785, plain, (~ sdtlseqdt0(szszuzczcdt0(sK41), sK40) | ~ aElementOf0(sK40, szNzAzT0) | ~ spl42_1459), inference(subsumption_resolution, [], [f210780, f65193])).
fof(f65193, plain, (aElementOf0(szszuzczcdt0(sK41), szNzAzT0) | ~ spl42_1459), inference(avatar_component_clause, [], [f65192])).
fof(f65192, plain, (spl42_1459 <=> aElementOf0(szszuzczcdt0(sK41), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_1459])])).
fof(f210780, plain, (~ sdtlseqdt0(szszuzczcdt0(sK41), sK40) | ~ aElementOf0(szszuzczcdt0(sK41), szNzAzT0) | ~ aElementOf0(sK40, szNzAzT0)), inference(resolution, [], [f34933, f3588])).
fof(f3588, plain, ! [X0, X1] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f3574, f966])).
fof(f966, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sP9(X0)), inference(subsumption_resolution, [], [f965, f552])).
fof(f552, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f203])).
fof(f203, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aSet0(sdtlpdtrp0(xN, X0))) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f82])).
fof(f82, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (isCountable0(sdtlpdtrp0(xN, X0)) & aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', m__3671)).
fof(f965, plain, ! [X0] : (sP9(X0) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f549, f553])).
fof(f553, plain, ! [X0] : (isCountable0(sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f203])).
fof(f549, plain, ! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | ~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f329])).
fof(f329, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f221, f328])).
fof(f328, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f221, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f202, e220, e219])).
fof(f219, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e219])).
fof(e219, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f220, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e220])).
fof(e220, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f202, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f201])).
fof(f201, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f95])).
fof(f95, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', m__3623)).
fof(f3574, plain, ! [X0, X1] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0) | ~ sP9(X0)), inference(resolution, [], [f554, f532])).
fof(f532, plain, ! [X0] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)) | ~ sP9(X0)), inference(cnf_transformation, [], [f324])).
fof(f324, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(rectify, [], [f323])).
fof(f323, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(nnf_transformation, [], [f220])).
fof(f554, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, sdtlpdtrp0(xN, X0)) | aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f205])).
fof(f205, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f204])).
fof(f204, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', m__3754)).
fof(f34933, plain, ~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, szszuzczcdt0(sK41))), inference(subsumption_resolution, [], [f34924, f1005])).
fof(f1005, plain, sP9(sK41), inference(resolution, [], [f966, f557])).
fof(f557, plain, aElementOf0(sK41, szNzAzT0), inference(cnf_transformation, [], [f331])).
fof(f34924, plain, (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, szszuzczcdt0(sK41))) | ~ sP9(sK41)), inference(superposition, [], [f3549, f563])).
fof(f563, plain, (szmzizndt0(sdtlpdtrp0(xN, sK40)) = szmzizndt0(sdtlpdtrp0(xN, sK41))), inference(cnf_transformation, [], [f331])).
fof(f3549, plain, ! [X10] : (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X10)), sdtlpdtrp0(xN, szszuzczcdt0(X10))) | ~ sP9(X10)), inference(subsumption_resolution, [], [f3542, f535])).
fof(f535, plain, ! [X0] : (~ sP9(X0) | sP8(X0)), inference(cnf_transformation, [], [f324])).
fof(f3542, plain, ! [X10] : (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X10)), sdtlpdtrp0(xN, szszuzczcdt0(X10))) | ~ sP9(X10) | ~ sP8(X10)), inference(resolution, [], [f537, f606])).
fof(f606, plain, ! [X0] : (~ aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(equality_resolution, [], [f542])).
fof(f542, plain, ! [X0, X1] : (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ sP8(X0)), inference(cnf_transformation, [], [f327])).
fof(f327, plain, ! [X0] : (! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0)) | ~ aElement0(X1)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(rectify, [], [f326])).
fof(f326, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(flattening, [], [f325])).
fof(f325, plain, ! [X0] : (! [X3] : ((aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ((szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) | ~ aElementOf0(X3, sdtlpdtrp0(xN, X0)) | ~ aElement0(X3))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)) | ~ aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))))) | ~ sP8(X0)), inference(nnf_transformation, [], [f219])).
fof(f537, plain, ! [X0, X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ sP9(X0)), inference(cnf_transformation, [], [f324])).
fof(f65337, plain, (~ spl42_119 | spl42_1459), inference(avatar_contradiction_clause, [], [f65336])).
fof(f65336, plain, ($false | (~ spl42_119 | spl42_1459)), inference(subsumption_resolution, [], [f65332, f2974])).
fof(f65332, plain, (~ aElementOf0(sK41, szNzAzT0) | spl42_1459), inference(resolution, [], [f65194, f380])).
fof(f380, plain, ! [X0] : (aElementOf0(szszuzczcdt0(X0), szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ! [X0] : ((~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f25])).
fof(f25, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (~ (sz00 = szszuzczcdt0(X0)) & aElementOf0(szszuzczcdt0(X0), szNzAzT0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM575+3.p', mSuccNum)).
fof(f65194, plain, (~ aElementOf0(szszuzczcdt0(sK41), szNzAzT0) | spl42_1459), inference(avatar_component_clause, [], [f65192])).
fof(f20479, plain, spl42_694, inference(avatar_contradiction_clause, [], [f20478])).
fof(f20478, plain, ($false | spl42_694), inference(subsumption_resolution, [], [f20474, f556])).
fof(f20474, plain, (~ aElementOf0(sK40, szNzAzT0) | spl42_694), inference(resolution, [], [f20321, f380])).
fof(f20321, plain, (~ aElementOf0(szszuzczcdt0(sK40), szNzAzT0) | spl42_694), inference(avatar_component_clause, [], [f20319])).
fof(f20319, plain, (spl42_694 <=> aElementOf0(szszuzczcdt0(sK40), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_694])])).
fof(f20326, plain, (~ spl42_694 | ~ spl42_695 | ~ spl42_119), inference(avatar_split_clause, [], [f20317, f2973, f20323, f20319])).
fof(f20317, plain, (~ sdtlseqdt0(szszuzczcdt0(sK40), sK41) | ~ aElementOf0(szszuzczcdt0(sK40), szNzAzT0) | ~ spl42_119), inference(subsumption_resolution, [], [f20305, f1004])).
fof(f1004, plain, sP9(sK40), inference(resolution, [], [f966, f556])).
fof(f20305, plain, (~ sdtlseqdt0(szszuzczcdt0(sK40), sK41) | ~ aElementOf0(szszuzczcdt0(sK40), szNzAzT0) | ~ sP9(sK40) | ~ spl42_119), inference(resolution, [], [f3589, f3549])).
fof(f3589, plain, (! [X3] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, X3)) | ~ sdtlseqdt0(X3, sK41) | ~ aElementOf0(X3, szNzAzT0)) | ~ spl42_119), inference(subsumption_resolution, [], [f3576, f2974])).
fof(f3576, plain, ! [X3] : (aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, X3)) | ~ sdtlseqdt0(X3, sK41) | ~ aElementOf0(X3, szNzAzT0) | ~ aElementOf0(sK41, szNzAzT0)), inference(resolution, [], [f554, f561])).
fof(f561, plain, aElementOf0(szmzizndt0(sdtlpdtrp0(xN, sK40)), sdtlpdtrp0(xN, sK41)), inference(cnf_transformation, [], [f331])).
fof(f2987, plain, spl42_119, inference(avatar_split_clause, [], [f557, f2973])).