fof(f684, plain, $false, inference(avatar_sat_refutation, [], [f629, f634, f654, f673, f681])).
fof(f681, plain, (~ spl42_1 | spl42_3), inference(avatar_contradiction_clause, [], [f680])).
fof(f680, plain, ($false | (~ spl42_1 | spl42_3)), inference(subsumption_resolution, [], [f679, f482])).
fof(f482, plain, aSubsetOf0(xS, szNzAzT0), inference(cnf_transformation, [], [f197])).
fof(f197, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xS)) & aSet0(xS)), inference(ennf_transformation, [], [f75])).
fof(f75, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, szNzAzT0)) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM571+3.p', m__3435)).
fof(f679, plain, (~ aSubsetOf0(xS, szNzAzT0) | (~ spl42_1 | spl42_3)), inference(forward_demodulation, [], [f678, f548])).
fof(f548, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f329])).
fof(f329, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f221, f328])).
fof(f328, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f221, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f203, e220, e219])).
fof(f219, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e219])).
fof(e219, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f220, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e220])).
fof(e220, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f203, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f202])).
fof(f202, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f96])).
fof(f96, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM571+3.p', m__3623)).
fof(f678, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, sz00), szNzAzT0) | (~ spl42_1 | spl42_3)), inference(forward_demodulation, [], [f632, f624])).
fof(f624, plain, ((sz00 = xi) | ~ spl42_1), inference(avatar_component_clause, [], [f622])).
fof(f622, plain, (spl42_1 <=> (sz00 = xi)), introduced(avatar_definition, [new_symbols(naming, [spl42_1])])).
fof(f632, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) | spl42_3), inference(avatar_component_clause, [], [f631])).
fof(f631, plain, (spl42_3 <=> aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_3])])).
fof(f673, plain, (~ spl42_1 | spl42_2), inference(avatar_contradiction_clause, [], [f672])).
fof(f672, plain, ($false | (~ spl42_1 | spl42_2)), inference(subsumption_resolution, [], [f671, f483])).
fof(f483, plain, isCountable0(xS), inference(cnf_transformation, [], [f197])).
fof(f671, plain, (~ isCountable0(xS) | (~ spl42_1 | spl42_2)), inference(forward_demodulation, [], [f665, f548])).
fof(f665, plain, (~ isCountable0(sdtlpdtrp0(xN, sz00)) | (~ spl42_1 | spl42_2)), inference(backward_demodulation, [], [f627, f624])).
fof(f627, plain, (~ isCountable0(sdtlpdtrp0(xN, xi)) | spl42_2), inference(avatar_component_clause, [], [f626])).
fof(f626, plain, (spl42_2 <=> isCountable0(sdtlpdtrp0(xN, xi))), introduced(avatar_definition, [new_symbols(naming, [spl42_2])])).
fof(f654, plain, (~ spl42_3 | ~ spl42_2), inference(avatar_split_clause, [], [f565, f626, f631])).
fof(f565, plain, (~ isCountable0(sdtlpdtrp0(xN, xi)) | ~ aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0)), inference(cnf_transformation, [], [f333])).
fof(f333, plain, (~ isCountable0(sdtlpdtrp0(xN, xi)) | (~ aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & ((~ aElementOf0(sK41, szNzAzT0) & aElementOf0(sK41, sdtlpdtrp0(xN, xi))) | ~ aSet0(sdtlpdtrp0(xN, xi))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f207, f332])).
fof(f332, plain, (? [X0] : (~ aElementOf0(X0, szNzAzT0) & aElementOf0(X0, sdtlpdtrp0(xN, xi))) => (~ aElementOf0(sK41, szNzAzT0) & aElementOf0(sK41, sdtlpdtrp0(xN, xi)))), introduced(choice_axiom, [])).
fof(f207, plain, (~ isCountable0(sdtlpdtrp0(xN, xi)) | (~ aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & (? [X0] : (~ aElementOf0(X0, szNzAzT0) & aElementOf0(X0, sdtlpdtrp0(xN, xi))) | ~ aSet0(sdtlpdtrp0(xN, xi))))), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ (isCountable0(sdtlpdtrp0(xN, xi)) & (aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) | (! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, xi))))), inference(negated_conjecture, [], [f85])).
fof(f85, plain, ~ (isCountable0(sdtlpdtrp0(xN, xi)) & (aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) | (! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, xi))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM571+3.p', m__)).
fof(f634, plain, (spl42_1 | spl42_3), inference(avatar_split_clause, [], [f561, f631, f622])).
fof(f561, plain, (aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) | (sz00 = xi)), inference(cnf_transformation, [], [f331])).
fof(f331, plain, ((isCountable0(sdtlpdtrp0(xN, xi)) & aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aSet0(sdtlpdtrp0(xN, xi)) & ((xi = szszuzczcdt0(sK40)) & aElementOf0(sK40, szNzAzT0))) | (sz00 = xi)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40])], [f206, f330])).
fof(f330, plain, (? [X1] : ((szszuzczcdt0(X1) = xi) & aElementOf0(X1, szNzAzT0)) => ((xi = szszuzczcdt0(sK40)) & aElementOf0(sK40, szNzAzT0))), introduced(choice_axiom, [])).
fof(f206, plain, ((isCountable0(sdtlpdtrp0(xN, xi)) & aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aSet0(sdtlpdtrp0(xN, xi)) & ? [X1] : ((szszuzczcdt0(X1) = xi) & aElementOf0(X1, szNzAzT0))) | (sz00 = xi)), inference(ennf_transformation, [], [f97])).
fof(f97, plain, (~ (sz00 = xi) => (isCountable0(sdtlpdtrp0(xN, xi)) & aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, xi)) & ? [X1] : ((szszuzczcdt0(X1) = xi) & aElementOf0(X1, szNzAzT0)))), inference(rectify, [], [f84])).
fof(f84, plain, (~ (sz00 = xi) => (isCountable0(sdtlpdtrp0(xN, xi)) & aSubsetOf0(sdtlpdtrp0(xN, xi), szNzAzT0) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, xi)) & ? [X0] : ((szszuzczcdt0(X0) = xi) & aElementOf0(X0, szNzAzT0)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM571+3.p', m__3702_02)).
fof(f629, plain, (spl42_1 | spl42_2), inference(avatar_split_clause, [], [f562, f626, f622])).
fof(f562, plain, (isCountable0(sdtlpdtrp0(xN, xi)) | (sz00 = xi)), inference(cnf_transformation, [], [f331])).