fof(f3663, plain, $false, inference(avatar_sat_refutation, [], [f637, f642, f643, f3379, f3524, f3580, f3618, f3655])).
fof(f3655, plain, (spl41_144 | ~ spl41_2 | ~ spl41_34), inference(avatar_split_clause, [], [f3654, f1209, f630, f3376])).
fof(f3376, plain, (spl41_144 <=> sdtlseqdt0(xj, sz00)), introduced(avatar_definition, [new_symbols(naming, [spl41_144])])).
fof(f630, plain, (spl41_2 <=> sdtlseqdt0(xj, xi)), introduced(avatar_definition, [new_symbols(naming, [spl41_2])])).
fof(f1209, plain, (spl41_34 <=> (sz00 = xi)), introduced(avatar_definition, [new_symbols(naming, [spl41_34])])).
fof(f3654, plain, (sdtlseqdt0(xj, sz00) | (~ spl41_2 | ~ spl41_34)), inference(forward_demodulation, [], [f631, f1210])).
fof(f1210, plain, ((sz00 = xi) | ~ spl41_34), inference(avatar_component_clause, [], [f1209])).
fof(f631, plain, (sdtlseqdt0(xj, xi) | ~ spl41_2), inference(avatar_component_clause, [], [f630])).
fof(f3618, plain, (~ spl41_30 | ~ spl41_34), inference(avatar_contradiction_clause, [], [f3617])).
fof(f3617, plain, ($false | (~ spl41_30 | ~ spl41_34)), inference(subsumption_resolution, [], [f3616, f3414])).
fof(f3414, plain, (~ aElementOf0(sK40, xS) | ~ spl41_30), inference(forward_demodulation, [], [f3386, f551])).
fof(f551, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f333])).
fof(f333, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f225, f332])).
fof(f332, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f225, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f204, e224, e223])).
fof(f223, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e223])).
fof(e223, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f224, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e224])).
fof(e224, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f204, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f203])).
fof(f203, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f97])).
fof(f97, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', m__3623)).
fof(f3386, plain, (~ aElementOf0(sK40, sdtlpdtrp0(xN, sz00)) | ~ spl41_30), inference(backward_demodulation, [], [f567, f1183])).
fof(f1183, plain, ((sz00 = xj) | ~ spl41_30), inference(avatar_component_clause, [], [f1182])).
fof(f1182, plain, (spl41_30 <=> (sz00 = xj)), introduced(avatar_definition, [new_symbols(naming, [spl41_30])])).
fof(f567, plain, ~ aElementOf0(sK40, sdtlpdtrp0(xN, xj)), inference(cnf_transformation, [], [f336])).
fof(f336, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & (~ aElementOf0(sK40, sdtlpdtrp0(xN, xj)) & aElementOf0(sK40, sdtlpdtrp0(xN, xi))) & sdtlseqdt0(xj, xi)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40])], [f211, f335])).
fof(f335, plain, (? [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, xj)) & aElementOf0(X0, sdtlpdtrp0(xN, xi))) => (~ aElementOf0(sK40, sdtlpdtrp0(xN, xj)) & aElementOf0(sK40, sdtlpdtrp0(xN, xi)))), introduced(choice_axiom, [])).
fof(f211, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ? [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, xj)) & aElementOf0(X0, sdtlpdtrp0(xN, xi))) & sdtlseqdt0(xj, xi)), inference(flattening, [], [f210])).
fof(f210, plain, ((~ aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ? [X0] : (~ aElementOf0(X0, sdtlpdtrp0(xN, xj)) & aElementOf0(X0, sdtlpdtrp0(xN, xi)))) & sdtlseqdt0(xj, xi)), inference(ennf_transformation, [], [f87])).
fof(f87, plain, ~ (sdtlseqdt0(xj, xi) => (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, sdtlpdtrp0(xN, xj))))), inference(negated_conjecture, [], [f86])).
fof(f86, plain, ~ (sdtlseqdt0(xj, xi) => (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, sdtlpdtrp0(xN, xj))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', m__)).
fof(f3616, plain, (aElementOf0(sK40, xS) | ~ spl41_34), inference(forward_demodulation, [], [f3589, f551])).
fof(f3589, plain, (aElementOf0(sK40, sdtlpdtrp0(xN, sz00)) | ~ spl41_34), inference(backward_demodulation, [], [f566, f1210])).
fof(f566, plain, aElementOf0(sK40, sdtlpdtrp0(xN, xi)), inference(cnf_transformation, [], [f336])).
fof(f3580, plain, (spl41_34 | spl41_146), inference(avatar_split_clause, [], [f3579, f3455, f1209])).
fof(f3455, plain, (spl41_146 <=> aElementOf0(sK14(xi), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl41_146])])).
fof(f3579, plain, ((sz00 = xi) | spl41_146), inference(subsumption_resolution, [], [f3563, f560])).
fof(f560, plain, aElementOf0(xi, szNzAzT0), inference(cnf_transformation, [], [f83])).
fof(f83, plain, (aElementOf0(xi, szNzAzT0) & aElementOf0(xj, szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', m__3786)).
fof(f3563, plain, ((sz00 = xi) | ~ aElementOf0(xi, szNzAzT0) | spl41_146), inference(resolution, [], [f3457, f388])).
fof(f388, plain, ! [X0] : (aElementOf0(sK14(X0), szNzAzT0) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ! [X0] : (((szszuzczcdt0(sK14(X0)) = X0) & aElementOf0(sK14(X0), szNzAzT0)) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14])], [f132, f248])).
fof(f248, plain, ! [X0] : (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) => ((szszuzczcdt0(sK14(X0)) = X0) & aElementOf0(sK14(X0), szNzAzT0))), introduced(choice_axiom, [])).
fof(f132, plain, ! [X0] : (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f131])).
fof(f131, plain, ! [X0] : ((? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => (? [X1] : ((szszuzczcdt0(X1) = X0) & aElementOf0(X1, szNzAzT0)) | (sz00 = X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', mNatExtra)).
fof(f3457, plain, (~ aElementOf0(sK14(xi), szNzAzT0) | spl41_146), inference(avatar_component_clause, [], [f3455])).
fof(f3524, plain, (~ spl41_146 | ~ spl41_1 | spl41_34), inference(avatar_split_clause, [], [f3453, f1209, f627, f3455])).
fof(f627, plain, (spl41_1 <=> ! [X1] : (~ (szszuzczcdt0(X1) = xi) | ~ aElementOf0(X1, szNzAzT0))), introduced(avatar_definition, [new_symbols(naming, [spl41_1])])).
fof(f3453, plain, (~ aElementOf0(sK14(xi), szNzAzT0) | (~ spl41_1 | spl41_34)), inference(trivial_inequality_removal, [], [f3452])).
fof(f3452, plain, (~ (xi = xi) | ~ aElementOf0(sK14(xi), szNzAzT0) | (~ spl41_1 | spl41_34)), inference(superposition, [], [f628, f1483])).
fof(f1483, plain, ((xi = szszuzczcdt0(sK14(xi))) | spl41_34), inference(subsumption_resolution, [], [f1473, f1211])).
fof(f1211, plain, (~ (sz00 = xi) | spl41_34), inference(avatar_component_clause, [], [f1209])).
fof(f1473, plain, ((sz00 = xi) | (xi = szszuzczcdt0(sK14(xi)))), inference(resolution, [], [f389, f560])).
fof(f389, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | (sz00 = X0) | (szszuzczcdt0(sK14(X0)) = X0)), inference(cnf_transformation, [], [f249])).
fof(f628, plain, (! [X1] : (~ (szszuzczcdt0(X1) = xi) | ~ aElementOf0(X1, szNzAzT0)) | ~ spl41_1), inference(avatar_component_clause, [], [f627])).
fof(f3379, plain, (~ spl41_144 | spl41_30), inference(avatar_split_clause, [], [f3374, f1182, f3376])).
fof(f3374, plain, ((sz00 = xj) | ~ sdtlseqdt0(xj, sz00)), inference(subsumption_resolution, [], [f3373, f559])).
fof(f559, plain, aElementOf0(xj, szNzAzT0), inference(cnf_transformation, [], [f83])).
fof(f3373, plain, ((sz00 = xj) | ~ sdtlseqdt0(xj, sz00) | ~ aElementOf0(xj, szNzAzT0)), inference(subsumption_resolution, [], [f2464, f384])).
fof(f384, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', mZeroNum)).
fof(f2464, plain, ((sz00 = xj) | ~ sdtlseqdt0(xj, sz00) | ~ aElementOf0(sz00, szNzAzT0) | ~ aElementOf0(xj, szNzAzT0)), inference(resolution, [], [f397, f654])).
fof(f654, plain, sdtlseqdt0(sz00, xj), inference(resolution, [], [f391, f559])).
fof(f391, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sdtlseqdt0(sz00, X0)), inference(cnf_transformation, [], [f134])).
fof(f134, plain, ! [X0] : (sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => sdtlseqdt0(sz00, X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', mZeroLess)).
fof(f397, plain, ! [X0, X1] : (~ sdtlseqdt0(X1, X0) | (X0 = X1) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ! [X0, X1] : ((X0 = X1) | ~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f140])).
fof(f140, plain, ! [X0, X1] : (((X0 = X1) | (~ sdtlseqdt0(X1, X0) | ~ sdtlseqdt0(X0, X1))) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => ((sdtlseqdt0(X1, X0) & sdtlseqdt0(X0, X1)) => (X0 = X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', mLessASymm)).
fof(f643, plain, spl41_2, inference(avatar_split_clause, [], [f565, f630])).
fof(f565, plain, sdtlseqdt0(xj, xi), inference(cnf_transformation, [], [f336])).
fof(f642, plain, ~ spl41_3, inference(avatar_split_clause, [], [f568, f634])).
fof(f634, plain, (spl41_3 <=> aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj))), introduced(avatar_definition, [new_symbols(naming, [spl41_3])])).
fof(f568, plain, ~ aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)), inference(cnf_transformation, [], [f336])).
fof(f637, plain, (spl41_1 | ~ spl41_2 | spl41_3), inference(avatar_split_clause, [], [f613, f634, f630, f627])).
fof(f613, plain, ! [X1] : (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) | ~ sdtlseqdt0(xj, xi) | ~ (szszuzczcdt0(X1) = xi) | ~ aElementOf0(X1, szNzAzT0)), inference(duplicate_literal_removal, [], [f564])).
fof(f564, plain, ! [X1] : (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) | ~ sdtlseqdt0(xj, xi) | ~ (szszuzczcdt0(X1) = xi) | ~ aElementOf0(X1, szNzAzT0) | ~ sdtlseqdt0(xj, xi)), inference(cnf_transformation, [], [f334])).
fof(f334, plain, ((aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xj)) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi)))) | ~ sdtlseqdt0(xj, xi) | ! [X1] : (~ (szszuzczcdt0(X1) = xi) | ~ aElementOf0(X1, szNzAzT0)) | ~ sdtlseqdt0(xj, xi)), inference(rectify, [], [f209])).
fof(f209, plain, ((aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xj)) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi)))) | ~ sdtlseqdt0(xj, xi) | ! [X0] : (~ (szszuzczcdt0(X0) = xi) | ~ aElementOf0(X0, szNzAzT0)) | ~ sdtlseqdt0(xj, xi)), inference(flattening, [], [f208])).
fof(f208, plain, (((aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xj)) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi)))) | ~ sdtlseqdt0(xj, xi)) | (! [X0] : (~ (szszuzczcdt0(X0) = xi) | ~ aElementOf0(X0, szNzAzT0)) | ~ sdtlseqdt0(xj, xi))), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ((? [X0] : ((szszuzczcdt0(X0) = xi) & aElementOf0(X0, szNzAzT0)) & sdtlseqdt0(xj, xi)) => (sdtlseqdt0(xj, xi) => (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xi)) => aElementOf0(X1, sdtlpdtrp0(xN, xj)))))), inference(rectify, [], [f85])).
fof(f85, plain, ((? [X0] : ((szszuzczcdt0(X0) = xi) & aElementOf0(X0, szNzAzT0)) & sdtlseqdt0(xj, xi)) => (sdtlseqdt0(xj, xi) => (aSubsetOf0(sdtlpdtrp0(xN, xi), sdtlpdtrp0(xN, xj)) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, sdtlpdtrp0(xN, xj)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM574+3.p', m__3786_02)).