fof(f294667, plain, $false, inference(avatar_sat_refutation, [], [f19018, f80421, f293116, f293143, f294177])).
fof(f294177, plain, (~ spl42_456 | spl42_1980), inference(avatar_contradiction_clause, [], [f294176])).
fof(f294176, plain, ($false | (~ spl42_456 | spl42_1980)), inference(subsumption_resolution, [], [f293144, f80075])).
fof(f80075, plain, (~ aElementOf0(sK41, sdtlpdtrp0(xN, xi)) | spl42_1980), inference(avatar_component_clause, [], [f80073])).
fof(f80073, plain, (spl42_1980 <=> aElementOf0(sK41, sdtlpdtrp0(xN, xi))), introduced(avatar_definition, [new_symbols(naming, [spl42_1980])])).
fof(f293144, plain, (aElementOf0(sK41, sdtlpdtrp0(xN, xi)) | ~ spl42_456), inference(backward_demodulation, [], [f583, f10961])).
fof(f10961, plain, ((szmzizndt0(sdtlpdtrp0(xN, xi)) = sK41) | ~ spl42_456), inference(avatar_component_clause, [], [f10959])).
fof(f10959, plain, (spl42_456 <=> (szmzizndt0(sdtlpdtrp0(xN, xi)) = sK41)), introduced(avatar_definition, [new_symbols(naming, [spl42_456])])).
fof(f583, plain, aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)), inference(cnf_transformation, [], [f345])).
fof(f345, plain, (aElementOf0(xQ, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (xk = sbrdtbr0(xQ)) & aSubsetOf0(xQ, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | ~ aElementOf0(X0, xQ)) & aSet0(xQ) & ! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi)) | ~ aElement0(X1)) & ((~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(flattening, [], [f344])).
fof(f344, plain, (aElementOf0(xQ, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (xk = sbrdtbr0(xQ)) & aSubsetOf0(xQ, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | ~ aElementOf0(X0, xQ)) & aSet0(xQ) & ! [X1] : ((aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | ((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi)) | ~ aElement0(X1))) & ((~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1)) | ~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(nnf_transformation, [], [f215])).
fof(f215, plain, (aElementOf0(xQ, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (xk = sbrdtbr0(xQ)) & aSubsetOf0(xQ, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | ~ aElementOf0(X0, xQ)) & aSet0(xQ) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, (aElementOf0(xQ, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (xk = sbrdtbr0(xQ)) & aSubsetOf0(xQ, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, xQ) => aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) & aSet0(xQ) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, xi)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(rectify, [], [f86])).
fof(f86, plain, (aElementOf0(xQ, slbdtsldtrb0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))), xk)) & (xk = sbrdtbr0(xQ)) & aSubsetOf0(xQ, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, xQ) => aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi))))) & aSet0(xQ) & ! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) & aElementOf0(X0, sdtlpdtrp0(xN, xi)) & aElement0(X0))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__3989_02)).
fof(f293143, plain, (spl42_456 | spl42_457), inference(avatar_split_clause, [], [f11734, f10963, f10959])).
fof(f10963, plain, (spl42_457 <=> aElementOf0(sK41, xQ)), introduced(avatar_definition, [new_symbols(naming, [spl42_457])])).
fof(f11734, plain, (aElementOf0(sK41, xQ) | (szmzizndt0(sdtlpdtrp0(xN, xi)) = sK41)), inference(resolution, [], [f599, f610])).
fof(f610, plain, aElementOf0(sK41, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))), inference(cnf_transformation, [], [f352])).
fof(f352, plain, (~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & (~ aElementOf0(sK41, xS) & aElementOf0(sK41, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X1] : ((aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & ~ aElementOf0(X1, xQ)) | ~ aElement0(X1)) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1)) | ~ aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41])], [f350, f351])).
fof(f351, plain, (? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) => (~ aElementOf0(sK41, xS) & aElementOf0(sK41, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))), introduced(choice_axiom, [])).
fof(f350, plain, (~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & ? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X1] : ((aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & ~ aElementOf0(X1, xQ)) | ~ aElement0(X1)) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1)) | ~ aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X2) | ~ aElementOf0(X2, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(rectify, [], [f349])).
fof(f349, plain, (~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X1] : ((aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & ~ aElementOf0(X1, xQ)) | ~ aElement0(X1)) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1)) | ~ aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(flattening, [], [f348])).
fof(f348, plain, (~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X1] : ((aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | ((~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) & ~ aElementOf0(X1, xQ)) | ~ aElement0(X1))) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1)) | ~ aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(nnf_transformation, [], [f218])).
fof(f218, plain, (~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X1] : (aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(flattening, [], [f217])).
fof(f217, plain, (((~ aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) & ? [X2] : (~ aElementOf0(X2, xS) & aElementOf0(X2, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & (! [X1] : (aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & (! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0) | ~ aElementOf0(X0, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi)))), inference(ennf_transformation, [], [f103])).
fof(f103, plain, ~ ((! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X1] : (aElementOf0(X1, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X1) | aElementOf0(X1, xQ)) & aElement0(X1))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) | ! [X2] : (aElementOf0(X2, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) => aElementOf0(X2, xS))))), inference(rectify, [], [f89])).
fof(f89, plain, ~ ((! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) | ! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) => aElementOf0(X0, xS))))), inference(negated_conjecture, [], [f88])).
fof(f88, plain, ~ ((! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))) => ((! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) => (aSubsetOf0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))), xS) | ! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) => aElementOf0(X0, xS))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__)).
fof(f599, plain, ! [X0] : (~ aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | aElementOf0(X0, xQ) | (szmzizndt0(sdtlpdtrp0(xN, xi)) = X0)), inference(cnf_transformation, [], [f347])).
fof(f347, plain, ((xK = sbrdtbr0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X0] : ((aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) & ~ aElementOf0(X0, xQ)) | ~ aElement0(X0)) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0)) | ~ aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(flattening, [], [f346])).
fof(f346, plain, ((xK = sbrdtbr0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X0] : ((aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) | ((~ (szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) & ~ aElementOf0(X0, xQ)) | ~ aElement0(X0))) & ((((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0)) | ~ aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(nnf_transformation, [], [f216])).
fof(f216, plain, ((xK = sbrdtbr0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(ennf_transformation, [], [f102])).
fof(f102, plain, ((xK = sbrdtbr0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), inference(rectify, [], [f87])).
fof(f87, plain, ((xK = sbrdtbr0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi))))) & ! [X0] : (aElementOf0(X0, sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xi)) = X0) | aElementOf0(X0, xQ)) & aElement0(X0))) & aSet0(sdtpldt0(xQ, szmzizndt0(sdtlpdtrp0(xN, xi)))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xi)), X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, xi)), sdtlpdtrp0(xN, xi))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__4007)).
fof(f293116, plain, (~ spl42_819 | ~ spl42_1980), inference(avatar_contradiction_clause, [], [f293115])).
fof(f293115, plain, ($false | (~ spl42_819 | ~ spl42_1980)), inference(subsumption_resolution, [], [f293114, f611])).
fof(f611, plain, ~ aElementOf0(sK41, xS), inference(cnf_transformation, [], [f352])).
fof(f293114, plain, (aElementOf0(sK41, xS) | (~ spl42_819 | ~ spl42_1980)), inference(subsumption_resolution, [], [f293025, f18916])).
fof(f18916, plain, (aElementOf0(xi, szNzAzT0) | ~ spl42_819), inference(avatar_component_clause, [], [f18915])).
fof(f18915, plain, (spl42_819 <=> aElementOf0(xi, szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl42_819])])).
fof(f293025, plain, (~ aElementOf0(xi, szNzAzT0) | aElementOf0(sK41, xS) | ~ spl42_1980), inference(resolution, [], [f7932, f80074])).
fof(f80074, plain, (aElementOf0(sK41, sdtlpdtrp0(xN, xi)) | ~ spl42_1980), inference(avatar_component_clause, [], [f80073])).
fof(f7932, plain, ! [X4, X5] : (~ aElementOf0(X5, sdtlpdtrp0(xN, X4)) | ~ aElementOf0(X4, szNzAzT0) | aElementOf0(X5, xS)), inference(subsumption_resolution, [], [f7926, f499])).
fof(f499, plain, aSet0(xS), inference(cnf_transformation, [], [f203])).
fof(f203, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xS)) & aSet0(xS)), inference(ennf_transformation, [], [f75])).
fof(f75, plain, (isCountable0(xS) & aSubsetOf0(xS, szNzAzT0) & ! [X0] : (aElementOf0(X0, xS) => aElementOf0(X0, szNzAzT0)) & aSet0(xS)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__3435)).
fof(f7926, plain, ! [X4, X5] : (~ aElementOf0(X4, szNzAzT0) | ~ aElementOf0(X5, sdtlpdtrp0(xN, X4)) | aElementOf0(X5, xS) | ~ aSet0(xS)), inference(resolution, [], [f4044, f361])).
fof(f361, plain, ! [X0, X3, X1] : (~ aSubsetOf0(X1, X0) | ~ aElementOf0(X3, X1) | aElementOf0(X3, X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f242])).
fof(f242, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (~ aElementOf0(sK11(X0, X1), X0) & aElementOf0(sK11(X0, X1), X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11])], [f240, f241])).
fof(f241, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) => (~ aElementOf0(sK11(X0, X1), X0) & aElementOf0(sK11(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f240, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X3] : (aElementOf0(X3, X0) | ~ aElementOf0(X3, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(rectify, [], [f239])).
fof(f239, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | ? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1)) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(flattening, [], [f238])).
fof(f238, plain, ! [X0] : (! [X1] : ((aSubsetOf0(X1, X0) | (? [X2] : (~ aElementOf0(X2, X0) & aElementOf0(X2, X1)) | ~ aSet0(X1))) & ((! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1)) | ~ aSubsetOf0(X1, X0))) | ~ aSet0(X0)), inference(nnf_transformation, [], [f110])).
fof(f110, plain, ! [X0] : (! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X0) | ~ aElementOf0(X2, X1)) & aSet0(X1))) | ~ aSet0(X0)), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aSubsetOf0(X1, X0) <=> (! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, X0)) & aSet0(X1)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', mDefSub)).
fof(f4044, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f4043, f407])).
fof(f407, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sdtlseqdt0(sz00, X0)), inference(cnf_transformation, [], [f139])).
fof(f139, plain, ! [X0] : (sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => sdtlseqdt0(sz00, X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', mZeroLess)).
fof(f4043, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f4039, f400])).
fof(f400, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', mZeroNum)).
fof(f4039, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(sz00, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(superposition, [], [f576, f567])).
fof(f567, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f340])).
fof(f340, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK39])], [f232, f339])).
fof(f339, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK39(X0), szNzAzT0) & aElementOf0(sK39(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f232, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f209, e231, e230])).
fof(f230, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e230])).
fof(e230, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f231, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e231])).
fof(e231, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f209, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f208])).
fof(f208, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f99])).
fof(f99, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__3623)).
fof(f576, plain, ! [X0, X1] : (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f211])).
fof(f211, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__3754)).
fof(f80421, plain, (~ spl42_457 | spl42_1980), inference(avatar_contradiction_clause, [], [f80420])).
fof(f80420, plain, ($false | (~ spl42_457 | spl42_1980)), inference(subsumption_resolution, [], [f80416, f10965])).
fof(f10965, plain, (aElementOf0(sK41, xQ) | ~ spl42_457), inference(avatar_component_clause, [], [f10963])).
fof(f80416, plain, (~ aElementOf0(sK41, xQ) | spl42_1980), inference(resolution, [], [f80075, f2885])).
fof(f2885, plain, ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) | ~ aElementOf0(X0, xQ)), inference(resolution, [], [f587, f591])).
fof(f591, plain, ! [X0] : (aElementOf0(X0, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | ~ aElementOf0(X0, xQ)), inference(cnf_transformation, [], [f345])).
fof(f587, plain, ! [X1] : (~ aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, xi), szmzizndt0(sdtlpdtrp0(xN, xi)))) | aElementOf0(X1, sdtlpdtrp0(xN, xi))), inference(cnf_transformation, [], [f345])).
fof(f19018, plain, spl42_819, inference(avatar_split_clause, [], [f582, f18915])).
fof(f582, plain, aElementOf0(xi, szNzAzT0), inference(cnf_transformation, [], [f85])).
fof(f85, plain, aElementOf0(xi, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM581+3.p', m__3989)).