fof(f556039, plain, $false, inference(avatar_sat_refutation, [], [f35125, f38835, f364024, f556038])).
fof(f556038, plain, (~ spl72_47 | ~ spl72_1116 | ~ spl72_1313), inference(avatar_contradiction_clause, [], [f556037])).
fof(f556037, plain, ($false | (~ spl72_47 | ~ spl72_1116 | ~ spl72_1313)), inference(subsumption_resolution, [], [f556036, f1931])).
fof(f1931, plain, (aElementOf0(sK69(sK71), szNzAzT0) | ~ spl72_47), inference(avatar_component_clause, [], [f1930])).
fof(f1930, plain, (spl72_47 <=> aElementOf0(sK69(sK71), szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl72_47])])).
fof(f556036, plain, (~ aElementOf0(sK69(sK71), szNzAzT0) | (~ spl72_1116 | ~ spl72_1313)), inference(subsumption_resolution, [], [f556029, f864])).
fof(f864, plain, ~ aElementOf0(sK71, xS), inference(cnf_transformation, [], [f472])).
fof(f472, plain, (~ aSubsetOf0(xO, xS) & (~ aElementOf0(sK71, xS) & aElementOf0(sK71, xO))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK71])], [f246, f471])).
fof(f471, plain, (? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, xO)) => (~ aElementOf0(sK71, xS) & aElementOf0(sK71, xO))), introduced(choice_axiom, [])).
fof(f246, plain, (~ aSubsetOf0(xO, xS) & ? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, xO))), inference(ennf_transformation, [], [f99])).
fof(f99, plain, ~ (aSubsetOf0(xO, xS) | ! [X0] : (aElementOf0(X0, xO) => aElementOf0(X0, xS))), inference(negated_conjecture, [], [f98])).
fof(f98, plain, ~ (aSubsetOf0(xO, xS) | ! [X0] : (aElementOf0(X0, xO) => aElementOf0(X0, xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__)).
fof(f556029, plain, (aElementOf0(sK71, xS) | ~ aElementOf0(sK69(sK71), szNzAzT0) | (~ spl72_1116 | ~ spl72_1313)), inference(superposition, [], [f42422, f38822])).
fof(f38822, plain, ((sK71 = sdtlpdtrp0(xe, sK69(sK71))) | ~ spl72_1313), inference(avatar_component_clause, [], [f38820])).
fof(f38820, plain, (spl72_1313 <=> (sK71 = sdtlpdtrp0(xe, sK69(sK71)))), introduced(avatar_definition, [new_symbols(naming, [spl72_1313])])).
fof(f42422, plain, (! [X0] : (aElementOf0(sdtlpdtrp0(xe, X0), xS) | ~ aElementOf0(X0, szNzAzT0)) | ~ spl72_1116), inference(subsumption_resolution, [], [f42421, f527])).
fof(f527, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sdtlseqdt0(sz00, X0)), inference(cnf_transformation, [], [f154])).
fof(f154, plain, ! [X0] : (sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => sdtlseqdt0(sz00, X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', mZeroLess)).
fof(f42421, plain, (! [X0] : (aElementOf0(sdtlpdtrp0(xe, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)) | ~ spl72_1116), inference(subsumption_resolution, [], [f42414, f34741])).
fof(f34741, plain, (aElementOf0(sz00, szNzAzT0) | ~ spl72_1116), inference(avatar_component_clause, [], [f34740])).
fof(f34740, plain, (spl72_1116 <=> aElementOf0(sz00, szNzAzT0)), introduced(avatar_definition, [new_symbols(naming, [spl72_1116])])).
fof(f42414, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xe, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(sz00, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(superposition, [], [f6787, f687])).
fof(f687, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f387])).
fof(f387, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK54(X0), szNzAzT0) & aElementOf0(sK54(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f260, f386])).
fof(f386, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK54(X0), szNzAzT0) & aElementOf0(sK54(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f260, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f224, e259, e258])).
fof(f258, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e258])).
fof(e258, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f259, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e259])).
fof(e259, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f224, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f223])).
fof(f223, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f109])).
fof(f109, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__3623)).
fof(f6787, plain, ! [X17, X18] : (aElementOf0(sdtlpdtrp0(xe, X17), sdtlpdtrp0(xN, X18)) | ~ sdtlseqdt0(X18, X17) | ~ aElementOf0(X18, szNzAzT0) | ~ aElementOf0(X17, szNzAzT0)), inference(duplicate_literal_removal, [], [f6776])).
fof(f6776, plain, ! [X17, X18] : (aElementOf0(sdtlpdtrp0(xe, X17), sdtlpdtrp0(xN, X18)) | ~ sdtlseqdt0(X18, X17) | ~ aElementOf0(X18, szNzAzT0) | ~ aElementOf0(X17, szNzAzT0) | ~ aElementOf0(X17, szNzAzT0)), inference(resolution, [], [f695, f824])).
fof(f824, plain, ! [X0] : (aElementOf0(sdtlpdtrp0(xe, X0), sdtlpdtrp0(xN, X0)) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f241])).
fof(f241, plain, (! [X0] : (((szmzizndt0(sdtlpdtrp0(xN, X0)) = sdtlpdtrp0(xe, X0)) & ! [X1] : (sdtlseqdt0(sdtlpdtrp0(xe, X0), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aElementOf0(sdtlpdtrp0(xe, X0), sdtlpdtrp0(xN, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xe)) & aFunction0(xe)), inference(ennf_transformation, [], [f91])).
fof(f91, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((szmzizndt0(sdtlpdtrp0(xN, X0)) = sdtlpdtrp0(xe, X0)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(sdtlpdtrp0(xe, X0), X1)) & aElementOf0(sdtlpdtrp0(xe, X0), sdtlpdtrp0(xN, X0)))) & (szNzAzT0 = szDzozmdt0(xe)) & aFunction0(xe)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__4660)).
fof(f695, plain, ! [X2, X0, X1] : (~ aElementOf0(X2, sdtlpdtrp0(xN, X0)) | aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f226])).
fof(f226, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__3754)).
fof(f364024, plain, spl72_47, inference(avatar_contradiction_clause, [], [f364023])).
fof(f364023, plain, ($false | spl72_47), inference(subsumption_resolution, [], [f364018, f863])).
fof(f863, plain, aElementOf0(sK71, xO), inference(cnf_transformation, [], [f472])).
fof(f364018, plain, (~ aElementOf0(sK71, xO) | spl72_47), inference(resolution, [], [f1932, f1948])).
fof(f1948, plain, ! [X0] : (aElementOf0(sK69(X0), szNzAzT0) | ~ aElementOf0(X0, xO)), inference(resolution, [], [f849, f962])).
fof(f962, plain, ! [X0] : (~ aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))) | aElementOf0(X0, szNzAzT0)), inference(forward_demodulation, [], [f841, f828])).
fof(f828, plain, (szNzAzT0 = szDzozmdt0(xd)), inference(cnf_transformation, [], [f456])).
fof(f456, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | (~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & (~ aElementOf0(sK67(X0, X1), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK67(X0, X1), X1))))) | ~ aSet0(X1)) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK67])], [f243, f455])).
fof(f455, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1)) => (~ aElementOf0(sK67(X0, X1), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK67(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f243, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | (~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1))))) | ~ aSet0(X1)) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(flattening, [], [f242])).
fof(f242, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | ((~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1))))) | ~ aSet0(X1))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(ennf_transformation, [], [f92])).
fof(f92, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ! [X1] : (((aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) | ((sbrdtbr0(X1) = xk) & (aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))))))) & aSet0(X1)) => (sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)))) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__4730)).
fof(f841, plain, ! [X0] : (aElementOf0(X0, szDzozmdt0(xd)) | ~ aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd)))), inference(cnf_transformation, [], [f462])).
fof(f462, plain, (! [X0] : ((aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))) | ~ (sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) | ~ aElementOf0(X0, szDzozmdt0(xd))) & (((sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) & aElementOf0(X0, szDzozmdt0(xd))) | ~ aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aElementOf0(szDzizrdt0(xd), xT)), inference(flattening, [], [f461])).
fof(f461, plain, (! [X0] : ((aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))) | (~ (sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) | ~ aElementOf0(X0, szDzozmdt0(xd)))) & (((sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) & aElementOf0(X0, szDzozmdt0(xd))) | ~ aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aElementOf0(szDzizrdt0(xd), xT)), inference(nnf_transformation, [], [f94])).
fof(f94, plain, (! [X0] : (aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))) <=> ((sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) & aElementOf0(X0, szDzozmdt0(xd)))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aElementOf0(szDzizrdt0(xd), xT)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__4854)).
fof(f849, plain, ! [X0] : (aElementOf0(sK69(X0), sdtlbdtrb0(xd, szDzizrdt0(xd))) | ~ aElementOf0(X0, xO)), inference(cnf_transformation, [], [f467])).
fof(f467, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : ((aElementOf0(X0, xO) | ! [X1] : (~ (sdtlpdtrp0(xe, X1) = X0) | ~ aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & (((sdtlpdtrp0(xe, sK69(X0)) = X0) & aElementOf0(sK69(X0), sdtlbdtrb0(xd, szDzizrdt0(xd)))) | ~ aElementOf0(X0, xO))) & ! [X3] : ((aElementOf0(X3, sdtlbdtrb0(xd, szDzizrdt0(xd))) | ~ (szDzizrdt0(xd) = sdtlpdtrp0(xd, X3)) | ~ aElementOf0(X3, szDzozmdt0(xd))) & (((szDzizrdt0(xd) = sdtlpdtrp0(xd, X3)) & aElementOf0(X3, szDzozmdt0(xd))) | ~ aElementOf0(X3, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK69])], [f465, f466])).
fof(f466, plain, ! [X0] : (? [X2] : ((sdtlpdtrp0(xe, X2) = X0) & aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd)))) => ((sdtlpdtrp0(xe, sK69(X0)) = X0) & aElementOf0(sK69(X0), sdtlbdtrb0(xd, szDzizrdt0(xd))))), introduced(choice_axiom, [])).
fof(f465, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : ((aElementOf0(X0, xO) | ! [X1] : (~ (sdtlpdtrp0(xe, X1) = X0) | ~ aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & (? [X2] : ((sdtlpdtrp0(xe, X2) = X0) & aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd)))) | ~ aElementOf0(X0, xO))) & ! [X3] : ((aElementOf0(X3, sdtlbdtrb0(xd, szDzizrdt0(xd))) | ~ (szDzizrdt0(xd) = sdtlpdtrp0(xd, X3)) | ~ aElementOf0(X3, szDzozmdt0(xd))) & (((szDzizrdt0(xd) = sdtlpdtrp0(xd, X3)) & aElementOf0(X3, szDzozmdt0(xd))) | ~ aElementOf0(X3, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), inference(rectify, [], [f464])).
fof(f464, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : ((aElementOf0(X0, xO) | ! [X1] : (~ (sdtlpdtrp0(xe, X1) = X0) | ~ aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & (? [X1] : ((sdtlpdtrp0(xe, X1) = X0) & aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd)))) | ~ aElementOf0(X0, xO))) & ! [X2] : ((aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd))) | ~ (szDzizrdt0(xd) = sdtlpdtrp0(xd, X2)) | ~ aElementOf0(X2, szDzozmdt0(xd))) & (((szDzizrdt0(xd) = sdtlpdtrp0(xd, X2)) & aElementOf0(X2, szDzozmdt0(xd))) | ~ aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), inference(flattening, [], [f463])).
fof(f463, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : ((aElementOf0(X0, xO) | ! [X1] : (~ (sdtlpdtrp0(xe, X1) = X0) | ~ aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & (? [X1] : ((sdtlpdtrp0(xe, X1) = X0) & aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd)))) | ~ aElementOf0(X0, xO))) & ! [X2] : ((aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd))) | (~ (szDzizrdt0(xd) = sdtlpdtrp0(xd, X2)) | ~ aElementOf0(X2, szDzozmdt0(xd)))) & (((szDzizrdt0(xd) = sdtlpdtrp0(xd, X2)) & aElementOf0(X2, szDzozmdt0(xd))) | ~ aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : (aElementOf0(X0, xO) <=> ? [X1] : ((sdtlpdtrp0(xe, X1) = X0) & aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & ! [X2] : (aElementOf0(X2, sdtlbdtrb0(xd, szDzizrdt0(xd))) <=> ((szDzizrdt0(xd) = sdtlpdtrp0(xd, X2)) & aElementOf0(X2, szDzozmdt0(xd)))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), inference(rectify, [], [f95])).
fof(f95, plain, ((xO = sdtlcdtrc0(xe, sdtlbdtrb0(xd, szDzizrdt0(xd)))) & ! [X0] : (aElementOf0(X0, xO) <=> ? [X1] : ((sdtlpdtrp0(xe, X1) = X0) & aElementOf0(X1, sdtlbdtrb0(xd, szDzizrdt0(xd))))) & ! [X0] : (aElementOf0(X0, sdtlbdtrb0(xd, szDzizrdt0(xd))) <=> ((sdtlpdtrp0(xd, X0) = szDzizrdt0(xd)) & aElementOf0(X0, szDzozmdt0(xd)))) & aSet0(sdtlbdtrb0(xd, szDzizrdt0(xd))) & aSet0(xO)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', m__4891)).
fof(f1932, plain, (~ aElementOf0(sK69(sK71), szNzAzT0) | spl72_47), inference(avatar_component_clause, [], [f1930])).
fof(f38835, plain, spl72_1313, inference(avatar_split_clause, [], [f1156, f38820])).
fof(f1156, plain, (sK71 = sdtlpdtrp0(xe, sK69(sK71))), inference(resolution, [], [f850, f863])).
fof(f850, plain, ! [X0] : (~ aElementOf0(X0, xO) | (sdtlpdtrp0(xe, sK69(X0)) = X0)), inference(cnf_transformation, [], [f467])).
fof(f35125, plain, spl72_1116, inference(avatar_split_clause, [], [f520, f34740])).
fof(f520, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM601+3.p', mZeroNum)).