fof(f14136, plain, $false, inference(subsumption_resolution, [], [f14128, f870])).
fof(f870, plain, aElementOf0(xi, szNzAzT0), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ((xx = sdtlpdtrp0(xe, xi)) & aElementOf0(xi, szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', m__5034)).
fof(f14128, plain, ~ aElementOf0(xi, szNzAzT0), inference(resolution, [], [f4896, f874])).
fof(f874, plain, ~ aSubsetOf0(sdtlpdtrp0(xN, xi), xS), inference(cnf_transformation, [], [f476])).
fof(f476, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xi), xS) & (~ aElementOf0(sK72, xS) & aElementOf0(sK72, sdtlpdtrp0(xN, xi)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK72])], [f248, f475])).
fof(f475, plain, (? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, sdtlpdtrp0(xN, xi))) => (~ aElementOf0(sK72, xS) & aElementOf0(sK72, sdtlpdtrp0(xN, xi)))), introduced(choice_axiom, [])).
fof(f248, plain, (~ aSubsetOf0(sdtlpdtrp0(xN, xi), xS) & ? [X0] : (~ aElementOf0(X0, xS) & aElementOf0(X0, sdtlpdtrp0(xN, xi)))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, ~ (aSubsetOf0(sdtlpdtrp0(xN, xi), xS) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, xS))), inference(negated_conjecture, [], [f100])).
fof(f100, plain, ~ (aSubsetOf0(sdtlpdtrp0(xN, xi), xS) | ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xi)) => aElementOf0(X0, xS))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', m__)).
fof(f4896, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f4895, f531])).
fof(f531, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | sdtlseqdt0(sz00, X0)), inference(cnf_transformation, [], [f156])).
fof(f156, plain, ! [X0] : (sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => sdtlseqdt0(sz00, X0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', mZeroLess)).
fof(f4895, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(X0, szNzAzT0)), inference(subsumption_resolution, [], [f4891, f524])).
fof(f524, plain, aElementOf0(sz00, szNzAzT0), inference(cnf_transformation, [], [f24])).
fof(f24, plain, aElementOf0(sz00, szNzAzT0), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', mZeroNum)).
fof(f4891, plain, ! [X0] : (aSubsetOf0(sdtlpdtrp0(xN, X0), xS) | ~ sdtlseqdt0(sz00, X0) | ~ aElementOf0(sz00, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(superposition, [], [f700, f691])).
fof(f691, plain, (xS = sdtlpdtrp0(xN, sz00)), inference(cnf_transformation, [], [f389])).
fof(f389, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & ((~ aElementOf0(sK54(X0), szNzAzT0) & aElementOf0(sK54(X0), sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f262, f388])).
fof(f388, plain, ! [X0] : (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) => (~ aElementOf0(sK54(X0), szNzAzT0) & aElementOf0(sK54(X0), sdtlpdtrp0(xN, X0)))), introduced(choice_axiom, [])).
fof(f262, plain, (! [X0] : (sP9(X0) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(definition_folding, [], [f226, e261, e260])).
fof(f260, plain, ! [X0] : (! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) | ~ sP8(X0)), inference(usedef, [], [e260])).
fof(e260, plain, ! [X0] : (sP8(X0) <=> ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f261, plain, ! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ sP9(X0)), inference(usedef, [], [e261])).
fof(e261, plain, ! [X0] : (sP9(X0) <=> (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & sP8(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f226, plain, (! [X0] : ((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | ~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(flattening, [], [f225])).
fof(f225, plain, (! [X0] : (((isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) | ~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0)))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4) | ~ aElementOf0(X4, sdtlpdtrp0(xN, X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))) | (~ isCountable0(sdtlpdtrp0(xN, X0)) | (~ aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) & (? [X1] : (~ aElementOf0(X1, szNzAzT0) & aElementOf0(X1, sdtlpdtrp0(xN, X0))) | ~ aSet0(sdtlpdtrp0(xN, X0)))))) | ~ aElementOf0(X0, szNzAzT0)) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(ennf_transformation, [], [f111])).
fof(f111, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X2, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X3] : (aElementOf0(X3, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X3) & aElementOf0(X3, sdtlpdtrp0(xN, X0)) & aElement0(X3))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X4] : (aElementOf0(X4, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X4)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), inference(rectify, [], [f81])).
fof(f81, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((isCountable0(sdtlpdtrp0(xN, X0)) & (aSubsetOf0(sdtlpdtrp0(xN, X0), szNzAzT0) | (! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => aElementOf0(X1, szNzAzT0)) & aSet0(sdtlpdtrp0(xN, X0))))) => (isCountable0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aSubsetOf0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) => aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0))))) & aSet0(sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ! [X1] : (aElementOf0(X1, sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) <=> (~ (szmzizndt0(sdtlpdtrp0(xN, X0)) = X1) & aElementOf0(X1, sdtlpdtrp0(xN, X0)) & aElement0(X1))) & aSet0(sdtmndt0(sdtlpdtrp0(xN, X0), szmzizndt0(sdtlpdtrp0(xN, X0)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, X0)), X1)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN, X0)), sdtlpdtrp0(xN, X0))))) & (xS = sdtlpdtrp0(xN, sz00)) & (szNzAzT0 = szDzozmdt0(xN)) & aFunction0(xN)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', m__3623)).
fof(f700, plain, ! [X0, X1] : (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ! [X0, X1] : ((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0) | ~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f228])).
fof(f228, plain, ! [X0, X1] : (((aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X1)) | ~ aElementOf0(X2, sdtlpdtrp0(xN, X0)))) | ~ sdtlseqdt0(X1, X0)) | (~ aElementOf0(X1, szNzAzT0) | ~ aElementOf0(X0, szNzAzT0))), inference(ennf_transformation, [], [f83])).
fof(f83, plain, ! [X0, X1] : ((aElementOf0(X1, szNzAzT0) & aElementOf0(X0, szNzAzT0)) => (sdtlseqdt0(X1, X0) => (aSubsetOf0(sdtlpdtrp0(xN, X0), sdtlpdtrp0(xN, X1)) & ! [X2] : (aElementOf0(X2, sdtlpdtrp0(xN, X0)) => aElementOf0(X2, sdtlpdtrp0(xN, X1)))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM603+3.p', m__3754)).