fof(f106911, plain, $false, inference(avatar_sat_refutation, [], [f2331, f106896])).
fof(f106896, plain, ~ spl73_55, inference(avatar_contradiction_clause, [], [f106895])).
fof(f106895, plain, ($false | ~ spl73_55), inference(subsumption_resolution, [], [f106894, f1088])).
fof(f1088, plain, ~ (szDzizrdt0(xd) = sdtlpdtrp0(xc, xQ)), inference(forward_demodulation, [], [f970, f946])).
fof(f946, plain, (szDzizrdt0(xd) = sdtlpdtrp0(xd, xn)), inference(cnf_transformation, [], [f111])).
fof(f111, plain, ((xp = sdtlpdtrp0(xe, xn)) & aElementOf0(xn, szNzAzT0) & aElementOf0(xn, sdtlbdtrb0(xd, szDzizrdt0(xd))) & (szDzizrdt0(xd) = sdtlpdtrp0(xd, xn)) & aElementOf0(xn, szDzozmdt0(xd))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5309)).
fof(f970, plain, ~ (sdtlpdtrp0(xd, xn) = sdtlpdtrp0(xc, xQ)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ~ (sdtlpdtrp0(xd, xn) = sdtlpdtrp0(xc, xQ)), inference(flattening, [], [f118])).
fof(f118, plain, ~ (sdtlpdtrp0(xd, xn) = sdtlpdtrp0(xc, xQ)), inference(negated_conjecture, [], [f117])).
fof(f117, plain, ~ (sdtlpdtrp0(xd, xn) = sdtlpdtrp0(xc, xQ)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__)).
fof(f106894, plain, ((szDzizrdt0(xd) = sdtlpdtrp0(xc, xQ)) | ~ spl73_55), inference(forward_demodulation, [], [f4162, f106893])).
fof(f106893, plain, (szDzizrdt0(xd) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)), inference(backward_demodulation, [], [f106833, f106877])).
fof(f106877, plain, (szDzizrdt0(xd) = sK65(xn)), inference(forward_demodulation, [], [f106876, f946])).
fof(f106876, plain, (sdtlpdtrp0(xd, xn) = sK65(xn)), inference(forward_demodulation, [], [f106828, f106833])).
fof(f106828, plain, (sdtlpdtrp0(xd, xn) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)), inference(subsumption_resolution, [], [f106819, f948])).
fof(f948, plain, aElementOf0(xn, szNzAzT0), inference(cnf_transformation, [], [f111])).
fof(f106819, plain, ((sdtlpdtrp0(xd, xn) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) | ~ aElementOf0(xn, szNzAzT0)), inference(resolution, [], [f952, f9443])).
fof(f9443, plain, ! [X6] : (~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | (sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sdtlpdtrp0(xd, X6)) | ~ aElementOf0(X6, szNzAzT0)), inference(subsumption_resolution, [], [f9440, f929])).
fof(f929, plain, aSet0(xP), inference(cnf_transformation, [], [f512])).
fof(f512, plain, ((xP = sdtmndt0(xQ, szmzizndt0(xQ))) & ! [X0] : ((aElementOf0(X0, xP) | (szmzizndt0(xQ) = X0) | ~ aElementOf0(X0, xQ) | ~ aElement0(X0)) & ((~ (szmzizndt0(xQ) = X0) & aElementOf0(X0, xQ) & aElement0(X0)) | ~ aElementOf0(X0, xP))) & ! [X1] : (sdtlseqdt0(szmzizndt0(xQ), X1) | ~ aElementOf0(X1, xQ)) & aSet0(xP)), inference(flattening, [], [f511])).
fof(f511, plain, ((xP = sdtmndt0(xQ, szmzizndt0(xQ))) & ! [X0] : ((aElementOf0(X0, xP) | ((szmzizndt0(xQ) = X0) | ~ aElementOf0(X0, xQ) | ~ aElement0(X0))) & ((~ (szmzizndt0(xQ) = X0) & aElementOf0(X0, xQ) & aElement0(X0)) | ~ aElementOf0(X0, xP))) & ! [X1] : (sdtlseqdt0(szmzizndt0(xQ), X1) | ~ aElementOf0(X1, xQ)) & aSet0(xP)), inference(nnf_transformation, [], [f278])).
fof(f278, plain, ((xP = sdtmndt0(xQ, szmzizndt0(xQ))) & ! [X0] : (aElementOf0(X0, xP) <=> (~ (szmzizndt0(xQ) = X0) & aElementOf0(X0, xQ) & aElement0(X0))) & ! [X1] : (sdtlseqdt0(szmzizndt0(xQ), X1) | ~ aElementOf0(X1, xQ)) & aSet0(xP)), inference(ennf_transformation, [], [f140])).
fof(f140, plain, ((xP = sdtmndt0(xQ, szmzizndt0(xQ))) & ! [X0] : (aElementOf0(X0, xP) <=> (~ (szmzizndt0(xQ) = X0) & aElementOf0(X0, xQ) & aElement0(X0))) & ! [X1] : (aElementOf0(X1, xQ) => sdtlseqdt0(szmzizndt0(xQ), X1)) & aSet0(xP)), inference(rectify, [], [f104])).
fof(f104, plain, ((xP = sdtmndt0(xQ, szmzizndt0(xQ))) & ! [X0] : (aElementOf0(X0, xP) <=> (~ (szmzizndt0(xQ) = X0) & aElementOf0(X0, xQ) & aElement0(X0))) & ! [X0] : (aElementOf0(X0, xQ) => sdtlseqdt0(szmzizndt0(xQ), X0)) & aSet0(xP)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5164)).
fof(f9440, plain, ! [X6] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sdtlpdtrp0(xd, X6)) | ~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | ~ aSet0(xP) | ~ aElementOf0(X6, szNzAzT0)), inference(trivial_inequality_removal, [], [f9439])).
fof(f9439, plain, ! [X6] : (~ (xk = xk) | (sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sdtlpdtrp0(xd, X6)) | ~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | ~ aSet0(xP) | ~ aElementOf0(X6, szNzAzT0)), inference(superposition, [], [f877, f943])).
fof(f943, plain, (xk = sbrdtbr0(xP)), inference(cnf_transformation, [], [f109])).
fof(f109, plain, (xk = sbrdtbr0(xP)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5217)).
fof(f877, plain, ! [X0, X1] : (~ (sbrdtbr0(X1) = xk) | (sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | ~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ aSet0(X1) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f494])).
fof(f494, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | (~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & (~ aElementOf0(sK67(X0, X1), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK67(X0, X1), X1))))) | ~ aSet0(X1)) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK67])], [f268, f493])).
fof(f493, plain, ! [X1, X0] : (? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1)) => (~ aElementOf0(sK67(X0, X1), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK67(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f268, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | (~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1))))) | ~ aSet0(X1)) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(flattening, [], [f267])).
fof(f267, plain, (! [X0] : (! [X1] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)) | ((~ aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X1) = xk) | (~ aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X2] : (~ aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X2, X1))))) | ~ aSet0(X1))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), inference(ennf_transformation, [], [f92])).
fof(f92, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ! [X1] : (((aElementOf0(X1, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) | ((sbrdtbr0(X1) = xk) & (aSubsetOf0(X1, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ! [X2] : (aElementOf0(X2, X1) => aElementOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))))))) & aSet0(X1)) => (sdtlpdtrp0(sdtlpdtrp0(xC, X0), X1) = sdtlpdtrp0(xd, X0)))) & (szNzAzT0 = szDzozmdt0(xd)) & aFunction0(xd)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__4730)).
fof(f952, plain, aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(xn))), inference(cnf_transformation, [], [f281])).
fof(f281, plain, (aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(xn))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, szszuzczcdt0(xn))) | ~ aElementOf0(X0, xP))), inference(ennf_transformation, [], [f113])).
fof(f113, plain, (aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(xn))) & ! [X0] : (aElementOf0(X0, xP) => aElementOf0(X0, sdtlpdtrp0(xN, szszuzczcdt0(xn))))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5334)).
fof(f106833, plain, (sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP) = sK65(xn)), inference(subsumption_resolution, [], [f106820, f948])).
fof(f106820, plain, ((sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP) = sK65(xn)) | ~ aElementOf0(xn, szNzAzT0)), inference(resolution, [], [f952, f9295])).
fof(f9295, plain, ! [X6] : (~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | (sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sK65(X6)) | ~ aElementOf0(X6, szNzAzT0)), inference(subsumption_resolution, [], [f9292, f929])).
fof(f9292, plain, ! [X6] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sK65(X6)) | ~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | ~ aSet0(xP) | ~ aElementOf0(X6, szNzAzT0)), inference(trivial_inequality_removal, [], [f9291])).
fof(f9291, plain, ! [X6] : (~ (xk = xk) | (sdtlpdtrp0(sdtlpdtrp0(xC, X6), xP) = sK65(X6)) | ~ aSubsetOf0(xP, sdtlpdtrp0(xN, szszuzczcdt0(X6))) | ~ aSet0(xP) | ~ aElementOf0(X6, szNzAzT0)), inference(superposition, [], [f866, f943])).
fof(f866, plain, ! [X2, X0] : (~ (sbrdtbr0(X2) = xk) | (sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = sK65(X0)) | ~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ~ aSet0(X2) | ~ aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f492])).
fof(f492, plain, ! [X0] : ((! [X2] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = sK65(X0)) | (~ aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X2) = xk) | (~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & (~ aElementOf0(sK66(X0, X2), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK66(X0, X2), X2))))) | ~ aSet0(X2)) & aElementOf0(sK65(X0), xT)) | ~ aElementOf0(X0, szNzAzT0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK65, sK66])], [f265, f491, f490])).
fof(f490, plain, ! [X0] : (? [X1] : (! [X2] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = X1) | (~ aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X2) = xk) | (~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X3, X2))))) | ~ aSet0(X2)) & aElementOf0(X1, xT)) => (! [X2] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = sK65(X0)) | (~ aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X2) = xk) | (~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X3, X2))))) | ~ aSet0(X2)) & aElementOf0(sK65(X0), xT))), introduced(choice_axiom, [])).
fof(f491, plain, ! [X2, X0] : (? [X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X3, X2)) => (~ aElementOf0(sK66(X0, X2), sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(sK66(X0, X2), X2))), introduced(choice_axiom, [])).
fof(f265, plain, ! [X0] : (? [X1] : (! [X2] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = X1) | (~ aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X2) = xk) | (~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X3, X2))))) | ~ aSet0(X2)) & aElementOf0(X1, xT)) | ~ aElementOf0(X0, szNzAzT0)), inference(flattening, [], [f264])).
fof(f264, plain, ! [X0] : (? [X1] : (! [X2] : ((sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = X1) | ((~ aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) & (~ (sbrdtbr0(X2) = xk) | (~ aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & ? [X3] : (~ aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))) & aElementOf0(X3, X2))))) | ~ aSet0(X2))) & aElementOf0(X1, xT)) | ~ aElementOf0(X0, szNzAzT0)), inference(ennf_transformation, [], [f90])).
fof(f90, plain, ! [X0] : (aElementOf0(X0, szNzAzT0) => ? [X1] : (! [X2] : (((aElementOf0(X2, slbdtsldtrb0(sdtlpdtrp0(xN, szszuzczcdt0(X0)), xk)) | ((sbrdtbr0(X2) = xk) & (aSubsetOf0(X2, sdtlpdtrp0(xN, szszuzczcdt0(X0))) | ! [X3] : (aElementOf0(X3, X2) => aElementOf0(X3, sdtlpdtrp0(xN, szszuzczcdt0(X0))))))) & aSet0(X2)) => (sdtlpdtrp0(sdtlpdtrp0(xC, X0), X2) = X1)) & aElementOf0(X1, xT))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__4618)).
fof(f4162, plain, ((sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP) = sdtlpdtrp0(xc, xQ)) | ~ spl73_55), inference(backward_demodulation, [], [f3144, f4159])).
fof(f4159, plain, ((xQ = sdtpldt0(xP, xp)) | ~ spl73_55), inference(forward_demodulation, [], [f4158, f4061])).
fof(f4061, plain, ((xP = sdtmndt0(xQ, xp)) | ~ spl73_55), inference(subsumption_resolution, [], [f4055, f2321])).
fof(f2321, plain, (sP3(xQ, xp) | ~ spl73_55), inference(avatar_component_clause, [], [f2320])).
fof(f2320, plain, (spl73_55 <=> sP3(xQ, xp)), introduced(avatar_definition, [new_symbols(naming, [spl73_55])])).
fof(f4055, plain, ((xP = sdtmndt0(xQ, xp)) | ~ sP3(xQ, xp) | ~ spl73_55), inference(resolution, [], [f547, f2373])).
fof(f2373, plain, (sP2(xp, xQ, xP) | ~ spl73_55), inference(subsumption_resolution, [], [f2372, f2321])).
fof(f2372, plain, (sP2(xp, xQ, xP) | ~ sP3(xQ, xp)), inference(superposition, [], [f976, f1073])).
fof(f1073, plain, (xP = sdtmndt0(xQ, xp)), inference(forward_demodulation, [], [f935, f928])).
fof(f928, plain, (xp = szmzizndt0(xQ)), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((xp = szmzizndt0(xQ)) & ! [X0] : (sdtlseqdt0(xp, X0) | ~ aElementOf0(X0, xQ)) & aElementOf0(xp, xQ)), inference(ennf_transformation, [], [f103])).
fof(f103, plain, ((xp = szmzizndt0(xQ)) & ! [X0] : (aElementOf0(X0, xQ) => sdtlseqdt0(xp, X0)) & aElementOf0(xp, xQ)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5147)).
fof(f935, plain, (xP = sdtmndt0(xQ, szmzizndt0(xQ))), inference(cnf_transformation, [], [f512])).
fof(f976, plain, ! [X0, X1] : (sP2(X1, X0, sdtmndt0(X0, X1)) | ~ sP3(X0, X1)), inference(equality_resolution, [], [f546])).
fof(f546, plain, ! [X2, X0, X1] : (sP2(X1, X0, X2) | ~ (sdtmndt0(X0, X1) = X2) | ~ sP3(X0, X1)), inference(cnf_transformation, [], [f334])).
fof(f334, plain, ! [X0, X1] : (! [X2] : (((sdtmndt0(X0, X1) = X2) | ~ sP2(X1, X0, X2)) & (sP2(X1, X0, X2) | ~ (sdtmndt0(X0, X1) = X2))) | ~ sP3(X0, X1)), inference(nnf_transformation, [], [f289])).
fof(f289, plain, ! [X0, X1] : (! [X2] : ((sdtmndt0(X0, X1) = X2) <=> sP2(X1, X0, X2)) | ~ sP3(X0, X1)), inference(usedef, [], [e289])).
fof(e289, plain, ! [X0, X1] : (sP3(X0, X1) <=> ! [X2] : ((sdtmndt0(X0, X1) = X2) <=> sP2(X1, X0, X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f547, plain, ! [X2, X0, X1] : (~ sP2(X1, X0, X2) | (sdtmndt0(X0, X1) = X2) | ~ sP3(X0, X1)), inference(cnf_transformation, [], [f334])).
fof(f4158, plain, (xQ = sdtpldt0(sdtmndt0(xQ, xp), xp)), inference(subsumption_resolution, [], [f4112, f911])).
fof(f911, plain, aSet0(xQ), inference(cnf_transformation, [], [f272])).
fof(f272, plain, (aElementOf0(xQ, slbdtsldtrb0(xO, xK)) & (xK = sbrdtbr0(xQ)) & aSubsetOf0(xQ, xO) & ! [X0] : (aElementOf0(X0, xO) | ~ aElementOf0(X0, xQ)) & aSet0(xQ)), inference(ennf_transformation, [], [f99])).
fof(f99, plain, (aElementOf0(xQ, slbdtsldtrb0(xO, xK)) & (xK = sbrdtbr0(xQ)) & aSubsetOf0(xQ, xO) & ! [X0] : (aElementOf0(X0, xQ) => aElementOf0(X0, xO)) & aSet0(xQ)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5078)).
fof(f4112, plain, ((xQ = sdtpldt0(sdtmndt0(xQ, xp), xp)) | ~ aSet0(xQ)), inference(resolution, [], [f558, f926])).
fof(f926, plain, aElementOf0(xp, xQ), inference(cnf_transformation, [], [f277])).
fof(f558, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | (sdtpldt0(sdtmndt0(X0, X1), X1) = X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : (! [X1] : ((sdtpldt0(sdtmndt0(X0, X1), X1) = X0) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => (sdtpldt0(sdtmndt0(X0, X1), X1) = X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', mConsDiff)).
fof(f3144, plain, (sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP) = sdtlpdtrp0(xc, sdtpldt0(xP, xp))), inference(backward_demodulation, [], [f969, f3137])).
fof(f3137, plain, (xp = szmzizndt0(sdtlpdtrp0(xN, xn))), inference(forward_demodulation, [], [f3124, f949])).
fof(f949, plain, (xp = sdtlpdtrp0(xe, xn)), inference(cnf_transformation, [], [f111])).
fof(f3124, plain, (sdtlpdtrp0(xe, xn) = szmzizndt0(sdtlpdtrp0(xN, xn))), inference(resolution, [], [f872, f948])).
fof(f872, plain, ! [X0] : (~ aElementOf0(X0, szNzAzT0) | (szmzizndt0(sdtlpdtrp0(xN, X0)) = sdtlpdtrp0(xe, X0))), inference(cnf_transformation, [], [f266])).
fof(f266, plain, (! [X0] : (((szmzizndt0(sdtlpdtrp0(xN, X0)) = sdtlpdtrp0(xe, X0)) & ! [X1] : (sdtlseqdt0(sdtlpdtrp0(xe, X0), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, X0))) & aElementOf0(sdtlpdtrp0(xe, X0), sdtlpdtrp0(xN, X0))) | ~ aElementOf0(X0, szNzAzT0)) & (szNzAzT0 = szDzozmdt0(xe)) & aFunction0(xe)), inference(ennf_transformation, [], [f91])).
fof(f91, plain, (! [X0] : (aElementOf0(X0, szNzAzT0) => ((szmzizndt0(sdtlpdtrp0(xN, X0)) = sdtlpdtrp0(xe, X0)) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, X0)) => sdtlseqdt0(sdtlpdtrp0(xe, X0), X1)) & aElementOf0(sdtlpdtrp0(xe, X0), sdtlpdtrp0(xN, X0)))) & (szNzAzT0 = szDzozmdt0(xe)) & aFunction0(xe)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__4660)).
fof(f969, plain, (sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)), inference(cnf_transformation, [], [f518])).
fof(f518, plain, ((sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) & ! [X0] : ((aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) | (~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) & ~ aElementOf0(X0, xP)) | ~ aElement0(X0)) & ((((szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) | aElementOf0(X0, xP)) & aElement0(X0)) | ~ aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))))) & aSet0(sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xn)))), inference(flattening, [], [f517])).
fof(f517, plain, ((sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) & ! [X0] : ((aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) | ((~ (szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) & ~ aElementOf0(X0, xP)) | ~ aElement0(X0))) & ((((szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) | aElementOf0(X0, xP)) & aElement0(X0)) | ~ aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))))) & aSet0(sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xn)))), inference(nnf_transformation, [], [f284])).
fof(f284, plain, ((sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) & ! [X0] : (aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) | aElementOf0(X0, xP)) & aElement0(X0))) & aSet0(sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1) | ~ aElementOf0(X1, sdtlpdtrp0(xN, xn)))), inference(ennf_transformation, [], [f142])).
fof(f142, plain, ((sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) & ! [X0] : (aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) | aElementOf0(X0, xP)) & aElement0(X0))) & aSet0(sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X1] : (aElementOf0(X1, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X1))), inference(rectify, [], [f116])).
fof(f116, plain, ((sdtlpdtrp0(xc, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) = sdtlpdtrp0(sdtlpdtrp0(xC, xn), xP)) & ! [X0] : (aElementOf0(X0, sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) <=> (((szmzizndt0(sdtlpdtrp0(xN, xn)) = X0) | aElementOf0(X0, xP)) & aElement0(X0))) & aSet0(sdtpldt0(xP, szmzizndt0(sdtlpdtrp0(xN, xn)))) & ! [X0] : (aElementOf0(X0, sdtlpdtrp0(xN, xn)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN, xn)), X0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5632)).
fof(f2331, plain, spl73_55, inference(avatar_contradiction_clause, [], [f2330])).
fof(f2330, plain, ($false | spl73_55), inference(subsumption_resolution, [], [f2329, f911])).
fof(f2329, plain, (~ aSet0(xQ) | spl73_55), inference(subsumption_resolution, [], [f2328, f1195])).
fof(f1195, plain, aElement0(xp), inference(subsumption_resolution, [], [f1177, f564])).
fof(f564, plain, aSet0(szNzAzT0), inference(cnf_transformation, [], [f23])).
fof(f23, plain, (isCountable0(szNzAzT0) & aSet0(szNzAzT0)), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', mNATSet)).
fof(f1177, plain, (aElement0(xp) | ~ aSet0(szNzAzT0)), inference(resolution, [], [f519, f1115])).
fof(f1115, plain, aElementOf0(xp, szNzAzT0), inference(resolution, [], [f919, f926])).
fof(f919, plain, ! [X0] : (~ aElementOf0(X0, xQ) | aElementOf0(X0, szNzAzT0)), inference(cnf_transformation, [], [f275])).
fof(f275, plain, (aSubsetOf0(xQ, szNzAzT0) & ! [X0] : (aElementOf0(X0, szNzAzT0) | ~ aElementOf0(X0, xQ))), inference(ennf_transformation, [], [f101])).
fof(f101, plain, (aSubsetOf0(xQ, szNzAzT0) & ! [X0] : (aElementOf0(X0, xQ) => aElementOf0(X0, szNzAzT0))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', m__5106)).
fof(f519, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f144])).
fof(f144, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', mEOfElem)).
fof(f2328, plain, (~ aElement0(xp) | ~ aSet0(xQ) | spl73_55), inference(resolution, [], [f2322, f557])).
fof(f557, plain, ! [X0, X1] : (sP3(X0, X1) | ~ aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ! [X0, X1] : (sP3(X0, X1) | ~ aElement0(X1) | ~ aSet0(X0)), inference(definition_folding, [], [f161, e289, e288])).
fof(f288, plain, ! [X1, X0, X2] : (sP2(X1, X0, X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> (~ (X1 = X3) & aElementOf0(X3, X0) & aElement0(X3))) & aSet0(X2))), inference(usedef, [], [e288])).
fof(e288, plain, ! [X1, X0, X2] : (sP2(X1, X0, X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> (~ (X1 = X3) & aElementOf0(X3, X0) & aElement0(X3))) & aSet0(X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f161, plain, ! [X0, X1] : (! [X2] : ((sdtmndt0(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> (~ (X1 = X3) & aElementOf0(X3, X0) & aElement0(X3))) & aSet0(X2))) | ~ aElement0(X1) | ~ aSet0(X0)), inference(flattening, [], [f160])).
fof(f160, plain, ! [X0, X1] : (! [X2] : ((sdtmndt0(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> (~ (X1 = X3) & aElementOf0(X3, X0) & aElement0(X3))) & aSet0(X2))) | (~ aElement0(X1) | ~ aSet0(X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0, X1] : ((aElement0(X1) & aSet0(X0)) => ! [X2] : ((sdtmndt0(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> (~ (X1 = X3) & aElementOf0(X3, X0) & aElement0(X3))) & aSet0(X2)))), file('/home/ubuntu/library/tptp/Problems/NUM/NUM631+3.p', mDefDiff)).
fof(f2322, plain, (~ sP3(xQ, xp) | spl73_55), inference(avatar_component_clause, [], [f2320])).