fof(f665, plain, $false, inference(resolution, [], [f662, f278])).
fof(f278, plain, ~ arboreal(sK19), inference(subsumption_resolution, [], [f276, f254])).
fof(f254, plain, ~ atomic(tptp0), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ~ atomic(tptp0), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', sos_37)).
fof(f276, plain, (~ arboreal(sK19) | atomic(tptp0)), inference(resolution, [], [f192, f268])).
fof(f268, plain, occurrence_of(sK19, tptp0), inference(cnf_transformation, [], [f182])).
fof(f182, plain, (! [X1, X2] : (((min_precedes(X1, sK20(X1), tptp0) & occurrence_of(sK20(X1), tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK19)) & occurrence_of(sK19, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19, sK20])], [f138, f181, f180])).
fof(f180, plain, (? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)) => (! [X2, X1] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK19)) & occurrence_of(sK19, tptp0))), introduced(choice_axiom, [])).
fof(f181, plain, ! [X1] : (? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) => (min_precedes(X1, sK20(X1), tptp0) & occurrence_of(sK20(X1), tptp1))), introduced(choice_axiom, [])).
fof(f138, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(definition_folding, [], [f136, e137])).
fof(f137, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(usedef, [], [e137])).
fof(e137, plain, ! [X1, X2] : (sP0(X1, X2) <=> (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f136, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2] : ((occurrence_of(X2, tptp2) => ~ ? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1))) & (occurrence_of(X2, tptp1) => ~ ? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2))) & leaf_occ(X2, X0))), inference(rectify, [], [f50])).
fof(f50, plain, ~ ! [X109] : (occurrence_of(X109, tptp0) => ? [X110, X111] : ((occurrence_of(X111, tptp2) => ~ ? [X113] : (min_precedes(X110, X113, tptp0) & occurrence_of(X113, tptp1))) & (occurrence_of(X111, tptp1) => ~ ? [X112] : (min_precedes(X110, X112, tptp0) & occurrence_of(X112, tptp2))) & leaf_occ(X111, X109))), inference(negated_conjecture, [], [f49])).
fof(f49, plain, ~ ! [X109] : (occurrence_of(X109, tptp0) => ? [X110, X111] : ((occurrence_of(X111, tptp2) => ~ ? [X113] : (min_precedes(X110, X113, tptp0) & occurrence_of(X113, tptp1))) & (occurrence_of(X111, tptp1) => ~ ? [X112] : (min_precedes(X110, X112, tptp0) & occurrence_of(X112, tptp2))) & leaf_occ(X111, X109))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', goals)).
fof(f192, plain, ! [X0, X1] : (~ occurrence_of(X0, X1) | ~ arboreal(X0) | atomic(X1)), inference(cnf_transformation, [], [f141])).
fof(f141, plain, ! [X0, X1] : (((arboreal(X0) | ~ atomic(X1)) & (atomic(X1) | ~ arboreal(X0))) | ~ occurrence_of(X0, X1)), inference(nnf_transformation, [], [f99])).
fof(f99, plain, ! [X0, X1] : ((arboreal(X0) <=> atomic(X1)) | ~ occurrence_of(X0, X1)), inference(ennf_transformation, [], [f57])).
fof(f57, plain, ! [X0, X1] : (occurrence_of(X0, X1) => (arboreal(X0) <=> atomic(X1))), inference(rectify, [], [f8])).
fof(f8, plain, ! [X16, X17] : (occurrence_of(X16, X17) => (arboreal(X16) <=> atomic(X17))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', sos_07)).
fof(f662, plain, ! [X0] : arboreal(X0), inference(subsumption_resolution, [], [f659, f374])).
fof(f374, plain, ! [X2, X0, X1] : (~ min_precedes(X0, X1, X2) | arboreal(X0)), inference(resolution, [], [f201, f318])).
fof(f318, plain, ! [X0, X1] : (~ atocc(X0, X1) | arboreal(X0)), inference(subsumption_resolution, [], [f314, f221])).
fof(f221, plain, ! [X0, X1] : (atomic(sK8(X0, X1)) | ~ atocc(X0, X1)), inference(cnf_transformation, [], [f159])).
fof(f159, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & ((occurrence_of(X0, sK8(X0, X1)) & atomic(sK8(X0, X1)) & subactivity(X1, sK8(X0, X1))) | ~ atocc(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8])], [f157, f158])).
fof(f158, plain, ! [X1, X0] : (? [X3] : (occurrence_of(X0, X3) & atomic(X3) & subactivity(X1, X3)) => (occurrence_of(X0, sK8(X0, X1)) & atomic(sK8(X0, X1)) & subactivity(X1, sK8(X0, X1)))), introduced(choice_axiom, [])).
fof(f157, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & (? [X3] : (occurrence_of(X0, X3) & atomic(X3) & subactivity(X1, X3)) | ~ atocc(X0, X1))), inference(rectify, [], [f156])).
fof(f156, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & (? [X2] : (occurrence_of(X0, X2) & atomic(X2) & subactivity(X1, X2)) | ~ atocc(X0, X1))), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ! [X0, X1] : (atocc(X0, X1) <=> ? [X2] : (occurrence_of(X0, X2) & atomic(X2) & subactivity(X1, X2))), inference(rectify, [], [f24])).
fof(f24, plain, ! [X64, X65] : (atocc(X64, X65) <=> ? [X66] : (occurrence_of(X64, X66) & atomic(X66) & subactivity(X65, X66))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', sos_23)).
fof(f314, plain, ! [X0, X1] : (~ atocc(X0, X1) | ~ atomic(sK8(X0, X1)) | arboreal(X0)), inference(resolution, [], [f222, f193])).
fof(f193, plain, ! [X0, X1] : (~ occurrence_of(X0, X1) | ~ atomic(X1) | arboreal(X0)), inference(cnf_transformation, [], [f141])).
fof(f222, plain, ! [X0, X1] : (occurrence_of(X0, sK8(X0, X1)) | ~ atocc(X0, X1)), inference(cnf_transformation, [], [f159])).
fof(f201, plain, ! [X2, X0, X1] : (atocc(X1, sK2(X0, X1, X2)) | ~ min_precedes(X1, X2, X0)), inference(cnf_transformation, [], [f145])).
fof(f145, plain, ! [X0, X1, X2] : ((atocc(X2, sK3(X0, X1, X2)) & atocc(X1, sK2(X0, X1, X2)) & subactivity(sK3(X0, X1, X2), X0) & subactivity(sK2(X0, X1, X2), X0)) | ~ min_precedes(X1, X2, X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3])], [f103, f144])).
fof(f144, plain, ! [X2, X1, X0] : (? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0)) => (atocc(X2, sK3(X0, X1, X2)) & atocc(X1, sK2(X0, X1, X2)) & subactivity(sK3(X0, X1, X2), X0) & subactivity(sK2(X0, X1, X2), X0))), introduced(choice_axiom, [])).
fof(f103, plain, ! [X0, X1, X2] : (? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0)) | ~ min_precedes(X1, X2, X0)), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0, X1, X2] : (min_precedes(X1, X2, X0) => ? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0))), inference(rectify, [], [f12])).
fof(f12, plain, ! [X23, X24, X25] : (min_precedes(X24, X25, X23) => ? [X26, X27] : (atocc(X25, X27) & atocc(X24, X26) & subactivity(X27, X23) & subactivity(X26, X23))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', sos_11)).
fof(f659, plain, ! [X0] : (arboreal(X0) | min_precedes(X0, sK18(X0), tptp0)), inference(resolution, [], [f645, f267])).
fof(f267, plain, ! [X0, X1] : (~ sP0(X0, X1) | min_precedes(X0, sK18(X0), tptp0)), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ! [X0, X1] : (((min_precedes(X0, sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18])], [f177, f178])).
fof(f178, plain, ! [X0] : (? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) => (min_precedes(X0, sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp2))), introduced(choice_axiom, [])).
fof(f177, plain, ! [X0, X1] : ((? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(rectify, [], [f176])).
fof(f176, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f137])).
fof(f645, plain, ! [X13] : (sP0(X13, sK17(sK19)) | arboreal(X13)), inference(resolution, [], [f438, f374])).
fof(f438, plain, ! [X0] : (min_precedes(X0, sK20(X0), tptp0) | sP0(X0, sK17(sK19))), inference(subsumption_resolution, [], [f437, f268])).
fof(f437, plain, ! [X0] : (sP0(X0, sK17(sK19)) | min_precedes(X0, sK20(X0), tptp0) | ~ occurrence_of(sK19, tptp0)), inference(resolution, [], [f271, f252])).
fof(f252, plain, ! [X0] : (leaf_occ(sK17(X0), X0) | ~ occurrence_of(X0, tptp0)), inference(cnf_transformation, [], [f175])).
fof(f175, plain, ! [X0] : ((leaf_occ(sK17(X0), X0) & next_subocc(sK16(X0), sK17(X0), tptp0) & (occurrence_of(sK17(X0), tptp2) | occurrence_of(sK17(X0), tptp1)) & next_subocc(sK15(X0), sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp4) & root_occ(sK15(X0), X0) & occurrence_of(sK15(X0), tptp3)) | ~ occurrence_of(X0, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17])], [f135, f174])).
fof(f174, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) => (leaf_occ(sK17(X0), X0) & next_subocc(sK16(X0), sK17(X0), tptp0) & (occurrence_of(sK17(X0), tptp2) | occurrence_of(sK17(X0), tptp1)) & next_subocc(sK15(X0), sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp4) & root_occ(sK15(X0), X0) & occurrence_of(sK15(X0), tptp3))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) | ~ occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f85])).
fof(f85, plain, ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3))), inference(rectify, [], [f36])).
fof(f36, plain, ! [X105] : (occurrence_of(X105, tptp0) => ? [X106, X107, X108] : (leaf_occ(X108, X105) & next_subocc(X107, X108, tptp0) & (occurrence_of(X108, tptp2) | occurrence_of(X108, tptp1)) & next_subocc(X106, X107, tptp0) & occurrence_of(X107, tptp4) & root_occ(X106, X105) & occurrence_of(X106, tptp3))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+1.p', sos_35)).
fof(f271, plain, ! [X2, X1] : (~ leaf_occ(X2, sK19) | sP0(X1, X2) | min_precedes(X1, sK20(X1), tptp0)), inference(cnf_transformation, [], [f182])).