fof(f619, plain, $false, inference(subsumption_resolution, [], [f618, f254])).
fof(f254, plain, occurrence_of(sK17, tptp0), inference(cnf_transformation, [], [f177])).
fof(f177, plain, (! [X1, X2] : (((min_precedes(X1, sK18(X1), tptp0) & occurrence_of(sK18(X1), tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK17)) & occurrence_of(sK17, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18])], [f132, f176, f175])).
fof(f175, plain, (? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)) => (! [X2, X1] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK17)) & occurrence_of(sK17, tptp0))), introduced(choice_axiom, [])).
fof(f176, plain, ! [X1] : (? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) => (min_precedes(X1, sK18(X1), tptp0) & occurrence_of(sK18(X1), tptp1))), introduced(choice_axiom, [])).
fof(f132, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(definition_folding, [], [f130, e131])).
fof(f131, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(usedef, [], [e131])).
fof(e131, plain, ! [X1, X2] : (sP0(X1, X2) <=> (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f130, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ~ ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2] : ((occurrence_of(X2, tptp2) => ~ ? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1))) & (occurrence_of(X2, tptp1) => ~ ? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2))) & leaf_occ(X2, X0))), inference(rectify, [], [f47])).
fof(f47, plain, ~ ! [X105] : (occurrence_of(X105, tptp0) => ? [X106, X107] : ((occurrence_of(X107, tptp2) => ~ ? [X109] : (min_precedes(X106, X109, tptp0) & occurrence_of(X109, tptp1))) & (occurrence_of(X107, tptp1) => ~ ? [X108] : (min_precedes(X106, X108, tptp0) & occurrence_of(X108, tptp2))) & leaf_occ(X107, X105))), inference(negated_conjecture, [], [f46])).
fof(f46, plain, ~ ! [X105] : (occurrence_of(X105, tptp0) => ? [X106, X107] : ((occurrence_of(X107, tptp2) => ~ ? [X109] : (min_precedes(X106, X109, tptp0) & occurrence_of(X109, tptp1))) & (occurrence_of(X107, tptp1) => ~ ? [X108] : (min_precedes(X106, X108, tptp0) & occurrence_of(X108, tptp2))) & leaf_occ(X107, X105))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+4.p', goals)).
fof(f618, plain, ~ occurrence_of(sK17, tptp0), inference(resolution, [], [f615, f239])).
fof(f239, plain, ! [X0] : (leaf_occ(sK15(X0), X0) | ~ occurrence_of(X0, tptp0)), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ! [X0] : ((leaf_occ(sK15(X0), X0) & next_subocc(sK14(X0), sK15(X0), tptp0) & (occurrence_of(sK15(X0), tptp2) | occurrence_of(sK15(X0), tptp1)) & next_subocc(sK13(X0), sK14(X0), tptp0) & occurrence_of(sK14(X0), tptp4) & root_occ(sK13(X0), X0) & occurrence_of(sK13(X0), tptp3)) | ~ occurrence_of(X0, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK13, sK14, sK15])], [f129, f169])).
fof(f169, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) => (leaf_occ(sK15(X0), X0) & next_subocc(sK14(X0), sK15(X0), tptp0) & (occurrence_of(sK15(X0), tptp2) | occurrence_of(sK15(X0), tptp1)) & next_subocc(sK13(X0), sK14(X0), tptp0) & occurrence_of(sK14(X0), tptp4) & root_occ(sK13(X0), X0) & occurrence_of(sK13(X0), tptp3))), introduced(choice_axiom, [])).
fof(f129, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) | ~ occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3))), inference(rectify, [], [f33])).
fof(f33, plain, ! [X101] : (occurrence_of(X101, tptp0) => ? [X102, X103, X104] : (leaf_occ(X104, X101) & next_subocc(X103, X104, tptp0) & (occurrence_of(X104, tptp2) | occurrence_of(X104, tptp1)) & next_subocc(X102, X103, tptp0) & occurrence_of(X103, tptp4) & root_occ(X102, X101) & occurrence_of(X102, tptp3))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+4.p', sos_32)).
fof(f615, plain, ~ leaf_occ(sK15(sK17), sK17), inference(subsumption_resolution, [], [f613, f589])).
fof(f589, plain, ! [X0] : ~ leaf(X0, tptp0), inference(subsumption_resolution, [], [f586, f200])).
fof(f200, plain, ! [X4, X0, X1] : (~ min_precedes(X0, X4, X1) | ~ leaf(X0, X1)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0, X1] : ((leaf(X0, X1) | min_precedes(X0, sK7(X0, X1), X1) | (! [X3] : ~ min_precedes(X3, X0, X1) & ~ root(X0, X1))) & ((! [X4] : ~ min_precedes(X0, X4, X1) & (min_precedes(sK8(X0, X1), X0, X1) | root(X0, X1))) | ~ leaf(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f147, f149, f148])).
fof(f148, plain, ! [X1, X0] : (? [X2] : min_precedes(X0, X2, X1) => min_precedes(X0, sK7(X0, X1), X1)), introduced(choice_axiom, [])).
fof(f149, plain, ! [X1, X0] : (? [X5] : min_precedes(X5, X0, X1) => min_precedes(sK8(X0, X1), X0, X1)), introduced(choice_axiom, [])).
fof(f147, plain, ! [X0, X1] : ((leaf(X0, X1) | ? [X2] : min_precedes(X0, X2, X1) | (! [X3] : ~ min_precedes(X3, X0, X1) & ~ root(X0, X1))) & ((! [X4] : ~ min_precedes(X0, X4, X1) & (? [X5] : min_precedes(X5, X0, X1) | root(X0, X1))) | ~ leaf(X0, X1))), inference(rectify, [], [f146])).
fof(f146, plain, ! [X0, X1] : ((leaf(X0, X1) | ? [X2] : min_precedes(X0, X2, X1) | (! [X3] : ~ min_precedes(X3, X0, X1) & ~ root(X0, X1))) & ((! [X2] : ~ min_precedes(X0, X2, X1) & (? [X3] : min_precedes(X3, X0, X1) | root(X0, X1))) | ~ leaf(X0, X1))), inference(flattening, [], [f145])).
fof(f145, plain, ! [X0, X1] : ((leaf(X0, X1) | (? [X2] : min_precedes(X0, X2, X1) | (! [X3] : ~ min_precedes(X3, X0, X1) & ~ root(X0, X1)))) & ((! [X2] : ~ min_precedes(X0, X2, X1) & (? [X3] : min_precedes(X3, X0, X1) | root(X0, X1))) | ~ leaf(X0, X1))), inference(nnf_transformation, [], [f110])).
fof(f110, plain, ! [X0, X1] : (leaf(X0, X1) <=> (! [X2] : ~ min_precedes(X0, X2, X1) & (? [X3] : min_precedes(X3, X0, X1) | root(X0, X1)))), inference(ennf_transformation, [], [f62])).
fof(f62, plain, ! [X0, X1] : (leaf(X0, X1) <=> (~ ? [X2] : min_precedes(X0, X2, X1) & (? [X3] : min_precedes(X3, X0, X1) | root(X0, X1)))), inference(rectify, [], [f16])).
fof(f16, plain, ! [X47, X48] : (leaf(X47, X48) <=> (~ ? [X50] : min_precedes(X47, X50, X48) & (? [X49] : min_precedes(X49, X47, X48) | root(X47, X48)))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+4.p', sos_15)).
fof(f586, plain, ! [X0] : (~ leaf(X0, tptp0) | min_precedes(X0, sK16(X0), tptp0)), inference(resolution, [], [f532, f253])).
fof(f253, plain, ! [X0, X1] : (~ sP0(X0, X1) | min_precedes(X0, sK16(X0), tptp0)), inference(cnf_transformation, [], [f174])).
fof(f174, plain, ! [X0, X1] : (((min_precedes(X0, sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16])], [f172, f173])).
fof(f173, plain, ! [X0] : (? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) => (min_precedes(X0, sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp2))), introduced(choice_axiom, [])).
fof(f172, plain, ! [X0, X1] : ((? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(rectify, [], [f171])).
fof(f171, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f131])).
fof(f532, plain, ! [X7] : (sP0(X7, sK15(sK17)) | ~ leaf(X7, tptp0)), inference(resolution, [], [f431, f200])).
fof(f431, plain, ! [X0] : (min_precedes(X0, sK18(X0), tptp0) | sP0(X0, sK15(sK17))), inference(subsumption_resolution, [], [f376, f254])).
fof(f376, plain, ! [X0] : (sP0(X0, sK15(sK17)) | min_precedes(X0, sK18(X0), tptp0) | ~ occurrence_of(sK17, tptp0)), inference(resolution, [], [f257, f239])).
fof(f257, plain, ! [X2, X1] : (~ leaf_occ(X2, sK17) | sP0(X1, X2) | min_precedes(X1, sK18(X1), tptp0)), inference(cnf_transformation, [], [f177])).
fof(f613, plain, (leaf(sK15(sK17), tptp0) | ~ leaf_occ(sK15(sK17), sK17)), inference(superposition, [], [f208, f606])).
fof(f606, plain, (tptp0 = sK9(sK15(sK17), sK17)), inference(subsumption_resolution, [], [f605, f254])).
fof(f605, plain, ((tptp0 = sK9(sK15(sK17), sK17)) | ~ occurrence_of(sK17, tptp0)), inference(resolution, [], [f367, f239])).
fof(f367, plain, ! [X0] : (~ leaf_occ(X0, sK17) | (tptp0 = sK9(X0, sK17))), inference(resolution, [], [f308, f206])).
fof(f206, plain, ! [X0, X1] : (occurrence_of(X1, sK9(X0, X1)) | ~ leaf_occ(X0, X1)), inference(cnf_transformation, [], [f155])).
fof(f155, plain, ! [X0, X1] : ((leaf_occ(X0, X1) | ! [X2] : (~ leaf(X0, X2) | ~ subactivity_occurrence(X0, X1) | ~ occurrence_of(X1, X2))) & ((leaf(X0, sK9(X0, X1)) & subactivity_occurrence(X0, X1) & occurrence_of(X1, sK9(X0, X1))) | ~ leaf_occ(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9])], [f153, f154])).
fof(f154, plain, ! [X1, X0] : (? [X3] : (leaf(X0, X3) & subactivity_occurrence(X0, X1) & occurrence_of(X1, X3)) => (leaf(X0, sK9(X0, X1)) & subactivity_occurrence(X0, X1) & occurrence_of(X1, sK9(X0, X1)))), introduced(choice_axiom, [])).
fof(f153, plain, ! [X0, X1] : ((leaf_occ(X0, X1) | ! [X2] : (~ leaf(X0, X2) | ~ subactivity_occurrence(X0, X1) | ~ occurrence_of(X1, X2))) & (? [X3] : (leaf(X0, X3) & subactivity_occurrence(X0, X1) & occurrence_of(X1, X3)) | ~ leaf_occ(X0, X1))), inference(rectify, [], [f152])).
fof(f152, plain, ! [X0, X1] : ((leaf_occ(X0, X1) | ! [X2] : (~ leaf(X0, X2) | ~ subactivity_occurrence(X0, X1) | ~ occurrence_of(X1, X2))) & (? [X2] : (leaf(X0, X2) & subactivity_occurrence(X0, X1) & occurrence_of(X1, X2)) | ~ leaf_occ(X0, X1))), inference(nnf_transformation, [], [f65])).
fof(f65, plain, ! [X0, X1] : (leaf_occ(X0, X1) <=> ? [X2] : (leaf(X0, X2) & subactivity_occurrence(X0, X1) & occurrence_of(X1, X2))), inference(rectify, [], [f19])).
fof(f19, plain, ! [X55, X56] : (leaf_occ(X55, X56) <=> ? [X57] : (leaf(X55, X57) & subactivity_occurrence(X55, X56) & occurrence_of(X56, X57))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+4.p', sos_18)).
fof(f308, plain, ! [X12] : (~ occurrence_of(sK17, X12) | (tptp0 = X12)), inference(resolution, [], [f190, f254])).
fof(f190, plain, ! [X2, X0, X1] : (~ occurrence_of(X0, X2) | (X1 = X2) | ~ occurrence_of(X0, X1)), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ! [X0, X1, X2] : ((X1 = X2) | ~ occurrence_of(X0, X2) | ~ occurrence_of(X0, X1)), inference(flattening, [], [f100])).
fof(f100, plain, ! [X0, X1, X2] : ((X1 = X2) | (~ occurrence_of(X0, X2) | ~ occurrence_of(X0, X1))), inference(ennf_transformation, [], [f55])).
fof(f55, plain, ! [X0, X1, X2] : ((occurrence_of(X0, X2) & occurrence_of(X0, X1)) => (X1 = X2)), inference(rectify, [], [f9])).
fof(f9, plain, ! [X28, X29, X30] : ((occurrence_of(X28, X30) & occurrence_of(X28, X29)) => (X29 = X30)), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+4.p', sos_08)).
fof(f208, plain, ! [X0, X1] : (leaf(X0, sK9(X0, X1)) | ~ leaf_occ(X0, X1)), inference(cnf_transformation, [], [f155])).