fof(f839, plain, $false, inference(resolution, [], [f838, f357])).
fof(f357, plain, ~ arboreal(sK21), inference(subsumption_resolution, [], [f355, f333])).
fof(f333, plain, ~ atomic(tptp0), inference(cnf_transformation, [], [f52])).
fof(f52, plain, ~ atomic(tptp0), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+3.p', sos_51)).
fof(f355, plain, (~ arboreal(sK21) | atomic(tptp0)), inference(resolution, [], [f252, f347])).
fof(f347, plain, occurrence_of(sK21, tptp0), inference(cnf_transformation, [], [f242])).
fof(f242, plain, (! [X1, X2] : (((min_precedes(X1, sK22(X1), tptp0) & occurrence_of(sK22(X1), tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK21)) & occurrence_of(sK21, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22])], [f189, f241, f240])).
fof(f240, plain, (? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)) => (! [X2, X1] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK21)) & occurrence_of(sK21, tptp0))), introduced(choice_axiom, [])).
fof(f241, plain, ! [X1] : (? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) => (min_precedes(X1, sK22(X1), tptp0) & occurrence_of(sK22(X1), tptp1))), introduced(choice_axiom, [])).
fof(f189, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(definition_folding, [], [f187, e188])).
fof(f188, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(usedef, [], [e188])).
fof(e188, plain, ! [X1, X2] : (sP0(X1, X2) <=> (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f187, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ~ ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2] : ((occurrence_of(X2, tptp2) => ~ ? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1))) & (occurrence_of(X2, tptp1) => ~ ? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2))) & leaf_occ(X2, X0))), inference(rectify, [], [f64])).
fof(f64, plain, ~ ! [X165] : (occurrence_of(X165, tptp0) => ? [X166, X167] : ((occurrence_of(X167, tptp2) => ~ ? [X169] : (min_precedes(X166, X169, tptp0) & occurrence_of(X169, tptp1))) & (occurrence_of(X167, tptp1) => ~ ? [X168] : (min_precedes(X166, X168, tptp0) & occurrence_of(X168, tptp2))) & leaf_occ(X167, X165))), inference(negated_conjecture, [], [f63])).
fof(f63, plain, ~ ! [X165] : (occurrence_of(X165, tptp0) => ? [X166, X167] : ((occurrence_of(X167, tptp2) => ~ ? [X169] : (min_precedes(X166, X169, tptp0) & occurrence_of(X169, tptp1))) & (occurrence_of(X167, tptp1) => ~ ? [X168] : (min_precedes(X166, X168, tptp0) & occurrence_of(X168, tptp2))) & leaf_occ(X167, X165))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+3.p', goals)).
fof(f252, plain, ! [X0, X1] : (~ occurrence_of(X0, X1) | ~ arboreal(X0) | atomic(X1)), inference(cnf_transformation, [], [f192])).
fof(f192, plain, ! [X0, X1] : (((arboreal(X0) | ~ atomic(X1)) & (atomic(X1) | ~ arboreal(X0))) | ~ occurrence_of(X0, X1)), inference(nnf_transformation, [], [f125])).
fof(f125, plain, ! [X0, X1] : ((arboreal(X0) <=> atomic(X1)) | ~ occurrence_of(X0, X1)), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ! [X0, X1] : (occurrence_of(X0, X1) => (arboreal(X0) <=> atomic(X1))), inference(rectify, [], [f8])).
fof(f8, plain, ! [X16, X17] : (occurrence_of(X16, X17) => (arboreal(X16) <=> atomic(X17))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+3.p', sos_07)).
fof(f838, plain, ! [X0] : arboreal(X0), inference(subsumption_resolution, [], [f835, f322])).
fof(f322, plain, ! [X2, X0, X1] : (~ min_precedes(X0, X1, X2) | arboreal(X0)), inference(cnf_transformation, [], [f181])).
fof(f181, plain, ! [X0, X1, X2] : (arboreal(X0) | ~ min_precedes(X0, X1, X2)), inference(ennf_transformation, [], [f110])).
fof(f110, plain, ! [X0, X1, X2] : (min_precedes(X0, X1, X2) => arboreal(X0)), inference(rectify, [], [f47])).
fof(f47, plain, ! [X148, X149, X150] : (min_precedes(X148, X149, X150) => arboreal(X148)), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+3.p', sos_46)).
fof(f835, plain, ! [X0] : (arboreal(X0) | min_precedes(X0, sK20(X0), tptp0)), inference(resolution, [], [f828, f346])).
fof(f346, plain, ! [X0, X1] : (~ sP0(X0, X1) | min_precedes(X0, sK20(X0), tptp0)), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ! [X0, X1] : (((min_precedes(X0, sK20(X0), tptp0) & occurrence_of(sK20(X0), tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK20])], [f237, f238])).
fof(f238, plain, ! [X0] : (? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) => (min_precedes(X0, sK20(X0), tptp0) & occurrence_of(sK20(X0), tptp2))), introduced(choice_axiom, [])).
fof(f237, plain, ! [X0, X1] : ((? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(rectify, [], [f236])).
fof(f236, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f188])).
fof(f828, plain, ! [X17] : (sP0(X17, sK19(sK21)) | arboreal(X17)), inference(resolution, [], [f554, f322])).
fof(f554, plain, ! [X0] : (min_precedes(X0, sK22(X0), tptp0) | sP0(X0, sK19(sK21))), inference(subsumption_resolution, [], [f553, f347])).
fof(f553, plain, ! [X0] : (sP0(X0, sK19(sK21)) | min_precedes(X0, sK22(X0), tptp0) | ~ occurrence_of(sK21, tptp0)), inference(resolution, [], [f350, f331])).
fof(f331, plain, ! [X0] : (leaf_occ(sK19(X0), X0) | ~ occurrence_of(X0, tptp0)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0] : ((leaf_occ(sK19(X0), X0) & next_subocc(sK18(X0), sK19(X0), tptp0) & (occurrence_of(sK19(X0), tptp2) | occurrence_of(sK19(X0), tptp1)) & next_subocc(sK17(X0), sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp4) & root_occ(sK17(X0), X0) & occurrence_of(sK17(X0), tptp3)) | ~ occurrence_of(X0, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18, sK19])], [f186, f234])).
fof(f234, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) => (leaf_occ(sK19(X0), X0) & next_subocc(sK18(X0), sK19(X0), tptp0) & (occurrence_of(sK19(X0), tptp2) | occurrence_of(sK19(X0), tptp1)) & next_subocc(sK17(X0), sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp4) & root_occ(sK17(X0), X0) & occurrence_of(sK17(X0), tptp3))), introduced(choice_axiom, [])).
fof(f186, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) | ~ occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f113])).
fof(f113, plain, ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3))), inference(rectify, [], [f50])).
fof(f50, plain, ! [X161] : (occurrence_of(X161, tptp0) => ? [X162, X163, X164] : (leaf_occ(X164, X161) & next_subocc(X163, X164, tptp0) & (occurrence_of(X164, tptp2) | occurrence_of(X164, tptp1)) & next_subocc(X162, X163, tptp0) & occurrence_of(X163, tptp4) & root_occ(X162, X161) & occurrence_of(X162, tptp3))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+3.p', sos_49)).
fof(f350, plain, ! [X2, X1] : (~ leaf_occ(X2, sK21) | sP0(X1, X2) | min_precedes(X1, sK22(X1), tptp0)), inference(cnf_transformation, [], [f242])).