fof(f738, plain, $false, inference(resolution, [], [f737, f268])).
fof(f268, plain, ~ arboreal(sK19), inference(subsumption_resolution, [], [f266, f244])).
fof(f244, plain, ~ atomic(tptp0), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ~ atomic(tptp0), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', sos_34)).
fof(f266, plain, (~ arboreal(sK19) | atomic(tptp0)), inference(resolution, [], [f198, f258])).
fof(f258, plain, occurrence_of(sK19, tptp0), inference(cnf_transformation, [], [f173])).
fof(f173, plain, (! [X1, X2] : (((min_precedes(X1, sK20(X1), tptp0) & occurrence_of(sK20(X1), tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK19)) & occurrence_of(sK19, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19, sK20])], [f126, f172, f171])).
fof(f171, plain, (? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)) => (! [X2, X1] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, sK19)) & occurrence_of(sK19, tptp0))), introduced(choice_axiom, [])).
fof(f172, plain, ! [X1] : (? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) => (min_precedes(X1, sK20(X1), tptp0) & occurrence_of(sK20(X1), tptp1))), introduced(choice_axiom, [])).
fof(f126, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | sP0(X1, X2) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(definition_folding, [], [f124, e125])).
fof(f125, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(usedef, [], [e125])).
fof(e125, plain, ! [X1, X2] : (sP0(X1, X2) <=> (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f124, plain, ? [X0] : (! [X1, X2] : ((? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1)) & occurrence_of(X2, tptp2)) | (? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ leaf_occ(X2, X0)) & occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ~ ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2] : ((occurrence_of(X2, tptp2) => ~ ? [X3] : (min_precedes(X1, X3, tptp0) & occurrence_of(X3, tptp1))) & (occurrence_of(X2, tptp1) => ~ ? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2))) & leaf_occ(X2, X0))), inference(rectify, [], [f47])).
fof(f47, plain, ~ ! [X99] : (occurrence_of(X99, tptp0) => ? [X100, X101] : ((occurrence_of(X101, tptp2) => ~ ? [X103] : (min_precedes(X100, X103, tptp0) & occurrence_of(X103, tptp1))) & (occurrence_of(X101, tptp1) => ~ ? [X102] : (min_precedes(X100, X102, tptp0) & occurrence_of(X102, tptp2))) & leaf_occ(X101, X99))), inference(negated_conjecture, [], [f46])).
fof(f46, plain, ~ ! [X99] : (occurrence_of(X99, tptp0) => ? [X100, X101] : ((occurrence_of(X101, tptp2) => ~ ? [X103] : (min_precedes(X100, X103, tptp0) & occurrence_of(X103, tptp1))) & (occurrence_of(X101, tptp1) => ~ ? [X102] : (min_precedes(X100, X102, tptp0) & occurrence_of(X102, tptp2))) & leaf_occ(X101, X99))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', goals)).
fof(f198, plain, ! [X0, X1] : (~ occurrence_of(X0, X1) | ~ arboreal(X0) | atomic(X1)), inference(cnf_transformation, [], [f140])).
fof(f140, plain, ! [X0, X1] : (((arboreal(X0) | ~ atomic(X1)) & (atomic(X1) | ~ arboreal(X0))) | ~ occurrence_of(X0, X1)), inference(nnf_transformation, [], [f97])).
fof(f97, plain, ! [X0, X1] : ((arboreal(X0) <=> atomic(X1)) | ~ occurrence_of(X0, X1)), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ! [X0, X1] : (occurrence_of(X0, X1) => (arboreal(X0) <=> atomic(X1))), inference(rectify, [], [f14])).
fof(f14, plain, ! [X40, X41] : (occurrence_of(X40, X41) => (arboreal(X40) <=> atomic(X41))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', sos_13)).
fof(f737, plain, ! [X0] : arboreal(X0), inference(subsumption_resolution, [], [f734, f411])).
fof(f411, plain, ! [X2, X0, X1] : (~ min_precedes(X0, X1, X2) | arboreal(X0)), inference(resolution, [], [f226, f316])).
fof(f316, plain, ! [X0, X1] : (~ atocc(X0, X1) | arboreal(X0)), inference(subsumption_resolution, [], [f312, f205])).
fof(f205, plain, ! [X0, X1] : (atomic(sK6(X0, X1)) | ~ atocc(X0, X1)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & ((occurrence_of(X0, sK6(X0, X1)) & atomic(sK6(X0, X1)) & subactivity(X1, sK6(X0, X1))) | ~ atocc(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f148, f149])).
fof(f149, plain, ! [X1, X0] : (? [X3] : (occurrence_of(X0, X3) & atomic(X3) & subactivity(X1, X3)) => (occurrence_of(X0, sK6(X0, X1)) & atomic(sK6(X0, X1)) & subactivity(X1, sK6(X0, X1)))), introduced(choice_axiom, [])).
fof(f148, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & (? [X3] : (occurrence_of(X0, X3) & atomic(X3) & subactivity(X1, X3)) | ~ atocc(X0, X1))), inference(rectify, [], [f147])).
fof(f147, plain, ! [X0, X1] : ((atocc(X0, X1) | ! [X2] : (~ occurrence_of(X0, X2) | ~ atomic(X2) | ~ subactivity(X1, X2))) & (? [X2] : (occurrence_of(X0, X2) & atomic(X2) & subactivity(X1, X2)) | ~ atocc(X0, X1))), inference(nnf_transformation, [], [f62])).
fof(f62, plain, ! [X0, X1] : (atocc(X0, X1) <=> ? [X2] : (occurrence_of(X0, X2) & atomic(X2) & subactivity(X1, X2))), inference(rectify, [], [f16])).
fof(f16, plain, ! [X46, X47] : (atocc(X46, X47) <=> ? [X48] : (occurrence_of(X46, X48) & atomic(X48) & subactivity(X47, X48))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', sos_15)).
fof(f312, plain, ! [X0, X1] : (~ atocc(X0, X1) | ~ atomic(sK6(X0, X1)) | arboreal(X0)), inference(resolution, [], [f206, f199])).
fof(f199, plain, ! [X0, X1] : (~ occurrence_of(X0, X1) | ~ atomic(X1) | arboreal(X0)), inference(cnf_transformation, [], [f140])).
fof(f206, plain, ! [X0, X1] : (occurrence_of(X0, sK6(X0, X1)) | ~ atocc(X0, X1)), inference(cnf_transformation, [], [f150])).
fof(f226, plain, ! [X2, X0, X1] : (atocc(X1, sK11(X0, X1, X2)) | ~ min_precedes(X1, X2, X0)), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ! [X0, X1, X2] : ((atocc(X2, sK12(X0, X1, X2)) & atocc(X1, sK11(X0, X1, X2)) & subactivity(sK12(X0, X1, X2), X0) & subactivity(sK11(X0, X1, X2), X0)) | ~ min_precedes(X1, X2, X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12])], [f115, f159])).
fof(f159, plain, ! [X2, X1, X0] : (? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0)) => (atocc(X2, sK12(X0, X1, X2)) & atocc(X1, sK11(X0, X1, X2)) & subactivity(sK12(X0, X1, X2), X0) & subactivity(sK11(X0, X1, X2), X0))), introduced(choice_axiom, [])).
fof(f115, plain, ! [X0, X1, X2] : (? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0)) | ~ min_precedes(X1, X2, X0)), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ! [X0, X1, X2] : (min_precedes(X1, X2, X0) => ? [X3, X4] : (atocc(X2, X4) & atocc(X1, X3) & subactivity(X4, X0) & subactivity(X3, X0))), inference(rectify, [], [f27])).
fof(f27, plain, ! [X77, X78, X79] : (min_precedes(X78, X79, X77) => ? [X80, X81] : (atocc(X79, X81) & atocc(X78, X80) & subactivity(X81, X77) & subactivity(X80, X77))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', sos_26)).
fof(f734, plain, ! [X0] : (arboreal(X0) | min_precedes(X0, sK18(X0), tptp0)), inference(resolution, [], [f677, f257])).
fof(f257, plain, ! [X0, X1] : (~ sP0(X0, X1) | min_precedes(X0, sK18(X0), tptp0)), inference(cnf_transformation, [], [f170])).
fof(f170, plain, ! [X0, X1] : (((min_precedes(X0, sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18])], [f168, f169])).
fof(f169, plain, ! [X0] : (? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) => (min_precedes(X0, sK18(X0), tptp0) & occurrence_of(sK18(X0), tptp2))), introduced(choice_axiom, [])).
fof(f168, plain, ! [X0, X1] : ((? [X2] : (min_precedes(X0, X2, tptp0) & occurrence_of(X2, tptp2)) & occurrence_of(X1, tptp1)) | ~ sP0(X0, X1)), inference(rectify, [], [f167])).
fof(f167, plain, ! [X1, X2] : ((? [X4] : (min_precedes(X1, X4, tptp0) & occurrence_of(X4, tptp2)) & occurrence_of(X2, tptp1)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f125])).
fof(f677, plain, ! [X13] : (sP0(X13, sK17(sK19)) | arboreal(X13)), inference(resolution, [], [f521, f411])).
fof(f521, plain, ! [X0] : (min_precedes(X0, sK20(X0), tptp0) | sP0(X0, sK17(sK19))), inference(subsumption_resolution, [], [f432, f258])).
fof(f432, plain, ! [X0] : (sP0(X0, sK17(sK19)) | min_precedes(X0, sK20(X0), tptp0) | ~ occurrence_of(sK19, tptp0)), inference(resolution, [], [f261, f242])).
fof(f242, plain, ! [X0] : (leaf_occ(sK17(X0), X0) | ~ occurrence_of(X0, tptp0)), inference(cnf_transformation, [], [f166])).
fof(f166, plain, ! [X0] : ((leaf_occ(sK17(X0), X0) & next_subocc(sK16(X0), sK17(X0), tptp0) & (occurrence_of(sK17(X0), tptp2) | occurrence_of(sK17(X0), tptp1)) & next_subocc(sK15(X0), sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp4) & root_occ(sK15(X0), X0) & occurrence_of(sK15(X0), tptp3)) | ~ occurrence_of(X0, tptp0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17])], [f123, f165])).
fof(f165, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) => (leaf_occ(sK17(X0), X0) & next_subocc(sK16(X0), sK17(X0), tptp0) & (occurrence_of(sK17(X0), tptp2) | occurrence_of(sK17(X0), tptp1)) & next_subocc(sK15(X0), sK16(X0), tptp0) & occurrence_of(sK16(X0), tptp4) & root_occ(sK15(X0), X0) & occurrence_of(sK15(X0), tptp3))), introduced(choice_axiom, [])).
fof(f123, plain, ! [X0] : (? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3)) | ~ occurrence_of(X0, tptp0)), inference(ennf_transformation, [], [f79])).
fof(f79, plain, ! [X0] : (occurrence_of(X0, tptp0) => ? [X1, X2, X3] : (leaf_occ(X3, X0) & next_subocc(X2, X3, tptp0) & (occurrence_of(X3, tptp2) | occurrence_of(X3, tptp1)) & next_subocc(X1, X2, tptp0) & occurrence_of(X2, tptp4) & root_occ(X1, X0) & occurrence_of(X1, tptp3))), inference(rectify, [], [f33])).
fof(f33, plain, ! [X95] : (occurrence_of(X95, tptp0) => ? [X96, X97, X98] : (leaf_occ(X98, X95) & next_subocc(X97, X98, tptp0) & (occurrence_of(X98, tptp2) | occurrence_of(X98, tptp1)) & next_subocc(X96, X97, tptp0) & occurrence_of(X97, tptp4) & root_occ(X96, X95) & occurrence_of(X96, tptp3))), file('/home/ubuntu/library/tptp/Problems/PRO/PRO010+2.p', sos_32)).
fof(f261, plain, ! [X2, X1] : (~ leaf_occ(X2, sK19) | sP0(X1, X2) | min_precedes(X1, sK20(X1), tptp0)), inference(cnf_transformation, [], [f173])).