fof(f15014, plain, $false, inference(avatar_sat_refutation, [], [f10661, f11374, f11382, f11407, f11413, f11419, f11425, f11431, f11437, f11443, f11449, f11455, f11461, f11467, f11473, f11479, f11485, f11491, f11497, f11503, f11509, f11515, f11521, f11527, f11533, f11539, f11545, f11551, f11557, f11563, f11569, f11575, f11581, f11587, f11593, f11599, f11605, f11611, f11617, f11623, f11629, f11635, f11641, f11647, f11653, f11659, f11665, f11671, f11677, f11683, f11689, f11695, f11701, f11707, f11713, f11719, f11725, f11731, f11737, f11743, f11749, f11755, f11761, f11767, f11773, f11779, f11785, f11791, f11797, f11803, f11809, f11815, f11821, f11827, f11833, f11839, f11845, f11851, f11857, f11863, f11869, f11875, f11881, f11887, f11893, f11899, f11905, f11909, f11913, f11919, f11925, f11931, f11937, f11943, f11949, f11955, f11961, f11967, f11973, f11979, f11985, f11991, f11997, f12003, f12009, f12015, f12021, f12027, f12033, f12039, f12045, f12051, f12057, f12063, f12069, f12075, f12081, f12087, f12093, f12099, f12105, f12111, f12117, f12123, f12129, f12135, f12141, f12147, f12153, f12159, f12165, f12171, f12177, f12183, f12189, f12195, f12201, f12207, f12213, f12219, f12225, f12231, f12237, f12243, f12249, f12255, f12261, f12267, f12273, f12279, f12285, f12291, f12297, f12303, f12309, f12315, f12321, f12327, f12333, f12339, f12345, f12351, f12357, f12363, f12369, f12375, f12381, f12387, f12393, f12399, f12405, f12411, f12417, f12423, f12429, f13205, f13211, f13249, f13255, f13261, f13267, f13273, f13279, f13285, f13291, f13297, f13303, f13309, f13315, f13321, f13327, f13333, f13339, f13345, f13351, f13357, f13363, f13369, f13375, f13381, f13387, f13393, f13399, f13405, f13411, f13417, f13423, f13429, f13435, f13441, f13447, f13453, f13459, f13465, f13471, f13477, f13483, f13489, f13495, f13501, f13507, f13513, f13519, f13525, f13531, f13537, f13543, f13549, f13555, f13561, f13567, f13573, f13579, f13585, f13591, f13597, f13603, f13609, f13615, f13621, f13627, f13633, f13639, f13645, f13651, f13657, f13663, f13669, f13675, f13681, f13687, f13693, f13699, f13705, f13711, f13717, f13723, f13729, f13735, f13741, f13747, f13753, f13759, f13765, f13771, f13777, f13783, f13789, f13795, f13801, f13807, f13813, f13819, f13825, f13831, f13837, f13843, f13849, f13855, f13861, f13867, f13873, f13879, f13885, f13891, f13897, f13903, f13909, f13915, f13921, f13927, f13933, f13939, f13945, f13951, f13957, f13963, f13969, f13975, f13981, f13987, f13993, f13999, f14005, f14011, f14017, f14023, f14029, f14035, f14041, f14047, f14053, f14059, f14065, f14071, f14077, f14083, f14089, f14095, f14101, f14107, f14113, f14119, f14125, f14131, f14137, f14143, f14149, f14155, f14161, f14167, f14173, f14179, f14185, f14191, f14197, f14203, f14209, f14215, f14221, f14227, f14233, f14239, f14245, f14251, f14257, f14263, f14269, f14297, f14303, f15007, f15013])).
fof(f15013, plain, ~ spl177_260, inference(avatar_contradiction_clause, [], [f15012])).
fof(f15012, plain, ($false | ~ spl177_260), inference(subsumption_resolution, [], [f15009, f6714])).
fof(f6714, plain, ~ (cell89 = cell100), inference(cnf_transformation, [], [f1036])).
fof(f1036, plain, ~ (cell89 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26025)).
fof(f15009, plain, ((cell89 = cell100) | ~ spl177_260), inference(resolution, [], [f12764, f5510])).
fof(f5510, plain, ! [X0, X1] : (~ sP80(X0, X1) | (cell89 = X1)), inference(cnf_transformation, [], [f5237])).
fof(f5237, plain, ! [X0, X1] : (((cell90 = X0) & (cell89 = X1)) | ~ sP80(X0, X1)), inference(nnf_transformation, [], [f5043])).
fof(f5043, plain, ! [X0, X1] : (((cell90 = X0) & (cell89 = X1)) | ~ sP80(X0, X1)), inference(usedef, [], [e5043])).
fof(e5043, plain, ! [X0, X1] : (sP80(X0, X1) <=> ((cell90 = X0) & (cell89 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP80])])).
fof(f12764, plain, (sP80(cell1, cell100) | ~ spl177_260), inference(avatar_component_clause, [], [f12762])).
fof(f12762, plain, (spl177_260 <=> sP80(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_260])])).
fof(f15007, plain, (spl177_259 | spl177_260 | ~ spl177_1 | spl177_180 | spl177_181 | spl177_182 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353), inference(avatar_split_clause, [], [f15006, f13134, f13130, f13126, f13122, f13118, f13114, f13110, f13106, f13102, f13098, f13094, f13090, f13086, f13082, f13078, f13074, f13070, f13066, f13062, f13058, f13054, f13050, f13046, f13042, f13038, f13034, f13030, f13026, f13022, f13018, f13014, f13010, f13006, f13002, f12998, f12994, f12990, f12986, f12982, f12978, f12974, f12970, f12966, f12962, f12958, f12954, f12950, f12946, f12942, f12938, f12934, f12930, f12926, f12922, f12918, f12914, f12910, f12906, f12902, f12898, f12894, f12890, f12886, f12882, f12878, f12874, f12870, f12866, f12862, f12858, f12854, f12850, f12846, f12842, f12838, f12834, f12830, f12826, f12822, f12818, f12814, f12810, f12806, f12802, f12798, f12794, f12790, f12786, f12782, f12778, f12774, f12770, f12766, f12754, f12750, f12746, f12742, f12738, f12734, f12730, f12726, f12722, f12718, f12714, f12710, f12706, f12702, f12698, f12694, f12690, f12686, f12682, f12678, f12674, f12670, f12666, f12662, f12658, f12654, f12650, f12646, f12642, f12638, f12634, f12630, f12626, f12622, f12618, f12614, f12610, f12606, f12602, f12598, f12594, f12590, f12586, f12582, f12578, f12574, f12570, f12566, f12562, f12558, f12554, f12550, f12546, f12542, f12538, f12534, f12530, f12526, f12522, f12518, f12514, f12510, f12506, f12502, f12498, f12494, f12490, f12486, f12482, f12478, f12474, f12470, f12466, f12462, f12458, f12454, f12450, f12446, f12442, f10654, f12762, f12758])).
fof(f12758, plain, (spl177_259 <=> sP79(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_259])])).
fof(f10654, plain, (spl177_1 <=> edge(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_1])])).
fof(f12442, plain, (spl177_180 <=> sP0(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_180])])).
fof(f12446, plain, (spl177_181 <=> sP1(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_181])])).
fof(f12450, plain, (spl177_182 <=> sP2(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_182])])).
fof(f12454, plain, (spl177_183 <=> sP3(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_183])])).
fof(f12458, plain, (spl177_184 <=> sP4(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_184])])).
fof(f12462, plain, (spl177_185 <=> sP5(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_185])])).
fof(f12466, plain, (spl177_186 <=> sP6(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_186])])).
fof(f12470, plain, (spl177_187 <=> sP7(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_187])])).
fof(f12474, plain, (spl177_188 <=> sP8(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_188])])).
fof(f12478, plain, (spl177_189 <=> sP9(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_189])])).
fof(f12482, plain, (spl177_190 <=> sP10(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_190])])).
fof(f12486, plain, (spl177_191 <=> sP11(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_191])])).
fof(f12490, plain, (spl177_192 <=> sP12(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_192])])).
fof(f12494, plain, (spl177_193 <=> sP13(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_193])])).
fof(f12498, plain, (spl177_194 <=> sP14(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_194])])).
fof(f12502, plain, (spl177_195 <=> sP15(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_195])])).
fof(f12506, plain, (spl177_196 <=> sP16(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_196])])).
fof(f12510, plain, (spl177_197 <=> sP17(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_197])])).
fof(f12514, plain, (spl177_198 <=> sP18(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_198])])).
fof(f12518, plain, (spl177_199 <=> sP19(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_199])])).
fof(f12522, plain, (spl177_200 <=> sP20(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_200])])).
fof(f12526, plain, (spl177_201 <=> sP21(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_201])])).
fof(f12530, plain, (spl177_202 <=> sP22(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_202])])).
fof(f12534, plain, (spl177_203 <=> sP23(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_203])])).
fof(f12538, plain, (spl177_204 <=> sP24(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_204])])).
fof(f12542, plain, (spl177_205 <=> sP25(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_205])])).
fof(f12546, plain, (spl177_206 <=> sP26(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_206])])).
fof(f12550, plain, (spl177_207 <=> sP27(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_207])])).
fof(f12554, plain, (spl177_208 <=> sP28(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_208])])).
fof(f12558, plain, (spl177_209 <=> sP29(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_209])])).
fof(f12562, plain, (spl177_210 <=> sP30(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_210])])).
fof(f12566, plain, (spl177_211 <=> sP31(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_211])])).
fof(f12570, plain, (spl177_212 <=> sP32(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_212])])).
fof(f12574, plain, (spl177_213 <=> sP33(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_213])])).
fof(f12578, plain, (spl177_214 <=> sP34(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_214])])).
fof(f12582, plain, (spl177_215 <=> sP35(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_215])])).
fof(f12586, plain, (spl177_216 <=> sP36(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_216])])).
fof(f12590, plain, (spl177_217 <=> sP37(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_217])])).
fof(f12594, plain, (spl177_218 <=> sP38(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_218])])).
fof(f12598, plain, (spl177_219 <=> sP39(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_219])])).
fof(f12602, plain, (spl177_220 <=> sP40(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_220])])).
fof(f12606, plain, (spl177_221 <=> sP41(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_221])])).
fof(f12610, plain, (spl177_222 <=> sP42(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_222])])).
fof(f12614, plain, (spl177_223 <=> sP43(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_223])])).
fof(f12618, plain, (spl177_224 <=> sP44(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_224])])).
fof(f12622, plain, (spl177_225 <=> sP45(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_225])])).
fof(f12626, plain, (spl177_226 <=> sP46(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_226])])).
fof(f12630, plain, (spl177_227 <=> sP47(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_227])])).
fof(f12634, plain, (spl177_228 <=> sP48(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_228])])).
fof(f12638, plain, (spl177_229 <=> sP49(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_229])])).
fof(f12642, plain, (spl177_230 <=> sP50(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_230])])).
fof(f12646, plain, (spl177_231 <=> sP51(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_231])])).
fof(f12650, plain, (spl177_232 <=> sP52(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_232])])).
fof(f12654, plain, (spl177_233 <=> sP53(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_233])])).
fof(f12658, plain, (spl177_234 <=> sP54(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_234])])).
fof(f12662, plain, (spl177_235 <=> sP55(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_235])])).
fof(f12666, plain, (spl177_236 <=> sP56(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_236])])).
fof(f12670, plain, (spl177_237 <=> sP57(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_237])])).
fof(f12674, plain, (spl177_238 <=> sP58(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_238])])).
fof(f12678, plain, (spl177_239 <=> sP59(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_239])])).
fof(f12682, plain, (spl177_240 <=> sP60(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_240])])).
fof(f12686, plain, (spl177_241 <=> sP61(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_241])])).
fof(f12690, plain, (spl177_242 <=> sP62(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_242])])).
fof(f12694, plain, (spl177_243 <=> sP63(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_243])])).
fof(f12698, plain, (spl177_244 <=> sP64(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_244])])).
fof(f12702, plain, (spl177_245 <=> sP65(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_245])])).
fof(f12706, plain, (spl177_246 <=> sP66(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_246])])).
fof(f12710, plain, (spl177_247 <=> sP67(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_247])])).
fof(f12714, plain, (spl177_248 <=> sP68(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_248])])).
fof(f12718, plain, (spl177_249 <=> sP69(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_249])])).
fof(f12722, plain, (spl177_250 <=> sP70(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_250])])).
fof(f12726, plain, (spl177_251 <=> sP71(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_251])])).
fof(f12730, plain, (spl177_252 <=> sP72(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_252])])).
fof(f12734, plain, (spl177_253 <=> sP73(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_253])])).
fof(f12738, plain, (spl177_254 <=> sP74(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_254])])).
fof(f12742, plain, (spl177_255 <=> sP75(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_255])])).
fof(f12746, plain, (spl177_256 <=> sP76(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_256])])).
fof(f12750, plain, (spl177_257 <=> sP77(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_257])])).
fof(f12754, plain, (spl177_258 <=> sP78(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_258])])).
fof(f12766, plain, (spl177_261 <=> sP81(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_261])])).
fof(f12770, plain, (spl177_262 <=> sP82(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_262])])).
fof(f12774, plain, (spl177_263 <=> sP83(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_263])])).
fof(f12778, plain, (spl177_264 <=> sP84(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_264])])).
fof(f12782, plain, (spl177_265 <=> sP85(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_265])])).
fof(f12786, plain, (spl177_266 <=> sP86(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_266])])).
fof(f12790, plain, (spl177_267 <=> sP90(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_267])])).
fof(f12794, plain, (spl177_268 <=> sP91(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_268])])).
fof(f12798, plain, (spl177_269 <=> sP92(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_269])])).
fof(f12802, plain, (spl177_270 <=> sP93(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_270])])).
fof(f12806, plain, (spl177_271 <=> sP94(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_271])])).
fof(f12810, plain, (spl177_272 <=> sP95(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_272])])).
fof(f12814, plain, (spl177_273 <=> sP96(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_273])])).
fof(f12818, plain, (spl177_274 <=> sP97(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_274])])).
fof(f12822, plain, (spl177_275 <=> sP98(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_275])])).
fof(f12826, plain, (spl177_276 <=> sP99(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_276])])).
fof(f12830, plain, (spl177_277 <=> sP100(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_277])])).
fof(f12834, plain, (spl177_278 <=> sP101(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_278])])).
fof(f12838, plain, (spl177_279 <=> sP102(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_279])])).
fof(f12842, plain, (spl177_280 <=> sP103(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_280])])).
fof(f12846, plain, (spl177_281 <=> sP104(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_281])])).
fof(f12850, plain, (spl177_282 <=> sP105(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_282])])).
fof(f12854, plain, (spl177_283 <=> sP106(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_283])])).
fof(f12858, plain, (spl177_284 <=> sP107(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_284])])).
fof(f12862, plain, (spl177_285 <=> sP108(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_285])])).
fof(f12866, plain, (spl177_286 <=> sP109(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_286])])).
fof(f12870, plain, (spl177_287 <=> sP110(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_287])])).
fof(f12874, plain, (spl177_288 <=> sP111(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_288])])).
fof(f12878, plain, (spl177_289 <=> sP112(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_289])])).
fof(f12882, plain, (spl177_290 <=> sP113(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_290])])).
fof(f12886, plain, (spl177_291 <=> sP114(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_291])])).
fof(f12890, plain, (spl177_292 <=> sP115(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_292])])).
fof(f12894, plain, (spl177_293 <=> sP116(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_293])])).
fof(f12898, plain, (spl177_294 <=> sP117(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_294])])).
fof(f12902, plain, (spl177_295 <=> sP118(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_295])])).
fof(f12906, plain, (spl177_296 <=> sP119(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_296])])).
fof(f12910, plain, (spl177_297 <=> sP120(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_297])])).
fof(f12914, plain, (spl177_298 <=> sP121(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_298])])).
fof(f12918, plain, (spl177_299 <=> sP122(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_299])])).
fof(f12922, plain, (spl177_300 <=> sP123(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_300])])).
fof(f12926, plain, (spl177_301 <=> sP124(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_301])])).
fof(f12930, plain, (spl177_302 <=> sP125(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_302])])).
fof(f12934, plain, (spl177_303 <=> sP126(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_303])])).
fof(f12938, plain, (spl177_304 <=> sP127(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_304])])).
fof(f12942, plain, (spl177_305 <=> sP128(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_305])])).
fof(f12946, plain, (spl177_306 <=> sP129(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_306])])).
fof(f12950, plain, (spl177_307 <=> sP130(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_307])])).
fof(f12954, plain, (spl177_308 <=> sP131(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_308])])).
fof(f12958, plain, (spl177_309 <=> sP132(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_309])])).
fof(f12962, plain, (spl177_310 <=> sP133(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_310])])).
fof(f12966, plain, (spl177_311 <=> sP134(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_311])])).
fof(f12970, plain, (spl177_312 <=> sP135(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_312])])).
fof(f12974, plain, (spl177_313 <=> sP136(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_313])])).
fof(f12978, plain, (spl177_314 <=> sP137(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_314])])).
fof(f12982, plain, (spl177_315 <=> sP138(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_315])])).
fof(f12986, plain, (spl177_316 <=> sP139(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_316])])).
fof(f12990, plain, (spl177_317 <=> sP140(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_317])])).
fof(f12994, plain, (spl177_318 <=> sP141(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_318])])).
fof(f12998, plain, (spl177_319 <=> sP142(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_319])])).
fof(f13002, plain, (spl177_320 <=> sP143(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_320])])).
fof(f13006, plain, (spl177_321 <=> sP144(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_321])])).
fof(f13010, plain, (spl177_322 <=> sP145(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_322])])).
fof(f13014, plain, (spl177_323 <=> sP146(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_323])])).
fof(f13018, plain, (spl177_324 <=> sP147(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_324])])).
fof(f13022, plain, (spl177_325 <=> sP148(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_325])])).
fof(f13026, plain, (spl177_326 <=> sP149(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_326])])).
fof(f13030, plain, (spl177_327 <=> sP150(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_327])])).
fof(f13034, plain, (spl177_328 <=> sP151(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_328])])).
fof(f13038, plain, (spl177_329 <=> sP152(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_329])])).
fof(f13042, plain, (spl177_330 <=> sP153(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_330])])).
fof(f13046, plain, (spl177_331 <=> sP154(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_331])])).
fof(f13050, plain, (spl177_332 <=> sP155(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_332])])).
fof(f13054, plain, (spl177_333 <=> sP156(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_333])])).
fof(f13058, plain, (spl177_334 <=> sP157(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_334])])).
fof(f13062, plain, (spl177_335 <=> sP158(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_335])])).
fof(f13066, plain, (spl177_336 <=> sP159(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_336])])).
fof(f13070, plain, (spl177_337 <=> sP160(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_337])])).
fof(f13074, plain, (spl177_338 <=> sP161(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_338])])).
fof(f13078, plain, (spl177_339 <=> sP162(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_339])])).
fof(f13082, plain, (spl177_340 <=> sP163(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_340])])).
fof(f13086, plain, (spl177_341 <=> sP164(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_341])])).
fof(f13090, plain, (spl177_342 <=> sP165(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_342])])).
fof(f13094, plain, (spl177_343 <=> sP166(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_343])])).
fof(f13098, plain, (spl177_344 <=> sP167(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_344])])).
fof(f13102, plain, (spl177_345 <=> sP168(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_345])])).
fof(f13106, plain, (spl177_346 <=> sP169(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_346])])).
fof(f13110, plain, (spl177_347 <=> sP170(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_347])])).
fof(f13114, plain, (spl177_348 <=> sP171(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_348])])).
fof(f13118, plain, (spl177_349 <=> sP172(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_349])])).
fof(f13122, plain, (spl177_350 <=> sP173(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_350])])).
fof(f13126, plain, (spl177_351 <=> sP174(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_351])])).
fof(f13130, plain, (spl177_352 <=> sP175(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_352])])).
fof(f13134, plain, (spl177_353 <=> sP176(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_353])])).
fof(f15006, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | (~ spl177_1 | spl177_180 | spl177_181 | spl177_182 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15005, f10630])).
fof(f10630, plain, ~ (cell1 = cell100), inference(cnf_transformation, [], [f4952])).
fof(f4952, plain, ~ (cell1 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29941)).
fof(f15005, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_180 | spl177_181 | spl177_182 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15004, f12443])).
fof(f12443, plain, (~ sP0(cell1, cell100) | spl177_180), inference(avatar_component_clause, [], [f12442])).
fof(f15004, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_181 | spl177_182 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15003, f12447])).
fof(f12447, plain, (~ sP1(cell1, cell100) | spl177_181), inference(avatar_component_clause, [], [f12446])).
fof(f15003, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_182 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15002, f12451])).
fof(f12451, plain, (~ sP2(cell1, cell100) | spl177_182), inference(avatar_component_clause, [], [f12450])).
fof(f15002, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_183 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15001, f12455])).
fof(f12455, plain, (~ sP3(cell1, cell100) | spl177_183), inference(avatar_component_clause, [], [f12454])).
fof(f15001, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_184 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f15000, f12459])).
fof(f12459, plain, (~ sP4(cell1, cell100) | spl177_184), inference(avatar_component_clause, [], [f12458])).
fof(f15000, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_185 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14999, f12463])).
fof(f12463, plain, (~ sP5(cell1, cell100) | spl177_185), inference(avatar_component_clause, [], [f12462])).
fof(f14999, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_186 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14998, f12467])).
fof(f12467, plain, (~ sP6(cell1, cell100) | spl177_186), inference(avatar_component_clause, [], [f12466])).
fof(f14998, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_187 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14997, f12471])).
fof(f12471, plain, (~ sP7(cell1, cell100) | spl177_187), inference(avatar_component_clause, [], [f12470])).
fof(f14997, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_188 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14996, f12475])).
fof(f12475, plain, (~ sP8(cell1, cell100) | spl177_188), inference(avatar_component_clause, [], [f12474])).
fof(f14996, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_189 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14995, f12479])).
fof(f12479, plain, (~ sP9(cell1, cell100) | spl177_189), inference(avatar_component_clause, [], [f12478])).
fof(f14995, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_190 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14994, f12483])).
fof(f12483, plain, (~ sP10(cell1, cell100) | spl177_190), inference(avatar_component_clause, [], [f12482])).
fof(f14994, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_191 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14993, f12487])).
fof(f12487, plain, (~ sP11(cell1, cell100) | spl177_191), inference(avatar_component_clause, [], [f12486])).
fof(f14993, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_192 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14992, f12491])).
fof(f12491, plain, (~ sP12(cell1, cell100) | spl177_192), inference(avatar_component_clause, [], [f12490])).
fof(f14992, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_193 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14991, f12495])).
fof(f12495, plain, (~ sP13(cell1, cell100) | spl177_193), inference(avatar_component_clause, [], [f12494])).
fof(f14991, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_194 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14990, f12499])).
fof(f12499, plain, (~ sP14(cell1, cell100) | spl177_194), inference(avatar_component_clause, [], [f12498])).
fof(f14990, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_195 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14989, f12503])).
fof(f12503, plain, (~ sP15(cell1, cell100) | spl177_195), inference(avatar_component_clause, [], [f12502])).
fof(f14989, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_196 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14988, f12507])).
fof(f12507, plain, (~ sP16(cell1, cell100) | spl177_196), inference(avatar_component_clause, [], [f12506])).
fof(f14988, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_197 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14987, f12511])).
fof(f12511, plain, (~ sP17(cell1, cell100) | spl177_197), inference(avatar_component_clause, [], [f12510])).
fof(f14987, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_198 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14986, f12515])).
fof(f12515, plain, (~ sP18(cell1, cell100) | spl177_198), inference(avatar_component_clause, [], [f12514])).
fof(f14986, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_199 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14985, f12519])).
fof(f12519, plain, (~ sP19(cell1, cell100) | spl177_199), inference(avatar_component_clause, [], [f12518])).
fof(f14985, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_200 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14984, f12523])).
fof(f12523, plain, (~ sP20(cell1, cell100) | spl177_200), inference(avatar_component_clause, [], [f12522])).
fof(f14984, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_201 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14983, f12527])).
fof(f12527, plain, (~ sP21(cell1, cell100) | spl177_201), inference(avatar_component_clause, [], [f12526])).
fof(f14983, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_202 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14982, f12531])).
fof(f12531, plain, (~ sP22(cell1, cell100) | spl177_202), inference(avatar_component_clause, [], [f12530])).
fof(f14982, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_203 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14981, f12535])).
fof(f12535, plain, (~ sP23(cell1, cell100) | spl177_203), inference(avatar_component_clause, [], [f12534])).
fof(f14981, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_204 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14980, f12539])).
fof(f12539, plain, (~ sP24(cell1, cell100) | spl177_204), inference(avatar_component_clause, [], [f12538])).
fof(f14980, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_205 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14979, f12543])).
fof(f12543, plain, (~ sP25(cell1, cell100) | spl177_205), inference(avatar_component_clause, [], [f12542])).
fof(f14979, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_206 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14978, f12547])).
fof(f12547, plain, (~ sP26(cell1, cell100) | spl177_206), inference(avatar_component_clause, [], [f12546])).
fof(f14978, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_207 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14977, f12551])).
fof(f12551, plain, (~ sP27(cell1, cell100) | spl177_207), inference(avatar_component_clause, [], [f12550])).
fof(f14977, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_208 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14976, f12555])).
fof(f12555, plain, (~ sP28(cell1, cell100) | spl177_208), inference(avatar_component_clause, [], [f12554])).
fof(f14976, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_209 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14975, f12559])).
fof(f12559, plain, (~ sP29(cell1, cell100) | spl177_209), inference(avatar_component_clause, [], [f12558])).
fof(f14975, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_210 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14974, f12563])).
fof(f12563, plain, (~ sP30(cell1, cell100) | spl177_210), inference(avatar_component_clause, [], [f12562])).
fof(f14974, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_211 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14973, f12567])).
fof(f12567, plain, (~ sP31(cell1, cell100) | spl177_211), inference(avatar_component_clause, [], [f12566])).
fof(f14973, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_212 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14972, f12571])).
fof(f12571, plain, (~ sP32(cell1, cell100) | spl177_212), inference(avatar_component_clause, [], [f12570])).
fof(f14972, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_213 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14971, f12575])).
fof(f12575, plain, (~ sP33(cell1, cell100) | spl177_213), inference(avatar_component_clause, [], [f12574])).
fof(f14971, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_214 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14970, f12579])).
fof(f12579, plain, (~ sP34(cell1, cell100) | spl177_214), inference(avatar_component_clause, [], [f12578])).
fof(f14970, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_215 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14969, f12583])).
fof(f12583, plain, (~ sP35(cell1, cell100) | spl177_215), inference(avatar_component_clause, [], [f12582])).
fof(f14969, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_216 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14968, f12587])).
fof(f12587, plain, (~ sP36(cell1, cell100) | spl177_216), inference(avatar_component_clause, [], [f12586])).
fof(f14968, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_217 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14967, f12591])).
fof(f12591, plain, (~ sP37(cell1, cell100) | spl177_217), inference(avatar_component_clause, [], [f12590])).
fof(f14967, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_218 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14966, f12595])).
fof(f12595, plain, (~ sP38(cell1, cell100) | spl177_218), inference(avatar_component_clause, [], [f12594])).
fof(f14966, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_219 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14965, f12599])).
fof(f12599, plain, (~ sP39(cell1, cell100) | spl177_219), inference(avatar_component_clause, [], [f12598])).
fof(f14965, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_220 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14964, f12603])).
fof(f12603, plain, (~ sP40(cell1, cell100) | spl177_220), inference(avatar_component_clause, [], [f12602])).
fof(f14964, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_221 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14963, f12607])).
fof(f12607, plain, (~ sP41(cell1, cell100) | spl177_221), inference(avatar_component_clause, [], [f12606])).
fof(f14963, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_222 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14962, f12611])).
fof(f12611, plain, (~ sP42(cell1, cell100) | spl177_222), inference(avatar_component_clause, [], [f12610])).
fof(f14962, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_223 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14961, f12615])).
fof(f12615, plain, (~ sP43(cell1, cell100) | spl177_223), inference(avatar_component_clause, [], [f12614])).
fof(f14961, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_224 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14960, f12619])).
fof(f12619, plain, (~ sP44(cell1, cell100) | spl177_224), inference(avatar_component_clause, [], [f12618])).
fof(f14960, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_225 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14959, f12623])).
fof(f12623, plain, (~ sP45(cell1, cell100) | spl177_225), inference(avatar_component_clause, [], [f12622])).
fof(f14959, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_226 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14958, f12627])).
fof(f12627, plain, (~ sP46(cell1, cell100) | spl177_226), inference(avatar_component_clause, [], [f12626])).
fof(f14958, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_227 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14957, f12631])).
fof(f12631, plain, (~ sP47(cell1, cell100) | spl177_227), inference(avatar_component_clause, [], [f12630])).
fof(f14957, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_228 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14956, f12635])).
fof(f12635, plain, (~ sP48(cell1, cell100) | spl177_228), inference(avatar_component_clause, [], [f12634])).
fof(f14956, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_229 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14955, f12639])).
fof(f12639, plain, (~ sP49(cell1, cell100) | spl177_229), inference(avatar_component_clause, [], [f12638])).
fof(f14955, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_230 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14954, f12643])).
fof(f12643, plain, (~ sP50(cell1, cell100) | spl177_230), inference(avatar_component_clause, [], [f12642])).
fof(f14954, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_231 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14953, f12647])).
fof(f12647, plain, (~ sP51(cell1, cell100) | spl177_231), inference(avatar_component_clause, [], [f12646])).
fof(f14953, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_232 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14952, f12651])).
fof(f12651, plain, (~ sP52(cell1, cell100) | spl177_232), inference(avatar_component_clause, [], [f12650])).
fof(f14952, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_233 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14951, f12655])).
fof(f12655, plain, (~ sP53(cell1, cell100) | spl177_233), inference(avatar_component_clause, [], [f12654])).
fof(f14951, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_234 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14950, f12659])).
fof(f12659, plain, (~ sP54(cell1, cell100) | spl177_234), inference(avatar_component_clause, [], [f12658])).
fof(f14950, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_235 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14949, f12663])).
fof(f12663, plain, (~ sP55(cell1, cell100) | spl177_235), inference(avatar_component_clause, [], [f12662])).
fof(f14949, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_236 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14948, f12667])).
fof(f12667, plain, (~ sP56(cell1, cell100) | spl177_236), inference(avatar_component_clause, [], [f12666])).
fof(f14948, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_237 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14947, f12671])).
fof(f12671, plain, (~ sP57(cell1, cell100) | spl177_237), inference(avatar_component_clause, [], [f12670])).
fof(f14947, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_238 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14946, f12675])).
fof(f12675, plain, (~ sP58(cell1, cell100) | spl177_238), inference(avatar_component_clause, [], [f12674])).
fof(f14946, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_239 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14945, f12679])).
fof(f12679, plain, (~ sP59(cell1, cell100) | spl177_239), inference(avatar_component_clause, [], [f12678])).
fof(f14945, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_240 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14944, f12683])).
fof(f12683, plain, (~ sP60(cell1, cell100) | spl177_240), inference(avatar_component_clause, [], [f12682])).
fof(f14944, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_241 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14943, f12687])).
fof(f12687, plain, (~ sP61(cell1, cell100) | spl177_241), inference(avatar_component_clause, [], [f12686])).
fof(f14943, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_242 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14942, f12691])).
fof(f12691, plain, (~ sP62(cell1, cell100) | spl177_242), inference(avatar_component_clause, [], [f12690])).
fof(f14942, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_243 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14941, f12695])).
fof(f12695, plain, (~ sP63(cell1, cell100) | spl177_243), inference(avatar_component_clause, [], [f12694])).
fof(f14941, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_244 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14940, f12699])).
fof(f12699, plain, (~ sP64(cell1, cell100) | spl177_244), inference(avatar_component_clause, [], [f12698])).
fof(f14940, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_245 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14939, f12703])).
fof(f12703, plain, (~ sP65(cell1, cell100) | spl177_245), inference(avatar_component_clause, [], [f12702])).
fof(f14939, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_246 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14938, f12707])).
fof(f12707, plain, (~ sP66(cell1, cell100) | spl177_246), inference(avatar_component_clause, [], [f12706])).
fof(f14938, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_247 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14937, f12711])).
fof(f12711, plain, (~ sP67(cell1, cell100) | spl177_247), inference(avatar_component_clause, [], [f12710])).
fof(f14937, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_248 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14936, f12715])).
fof(f12715, plain, (~ sP68(cell1, cell100) | spl177_248), inference(avatar_component_clause, [], [f12714])).
fof(f14936, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_249 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14935, f12719])).
fof(f12719, plain, (~ sP69(cell1, cell100) | spl177_249), inference(avatar_component_clause, [], [f12718])).
fof(f14935, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_250 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14934, f12723])).
fof(f12723, plain, (~ sP70(cell1, cell100) | spl177_250), inference(avatar_component_clause, [], [f12722])).
fof(f14934, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_251 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14933, f12727])).
fof(f12727, plain, (~ sP71(cell1, cell100) | spl177_251), inference(avatar_component_clause, [], [f12726])).
fof(f14933, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_252 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14932, f12731])).
fof(f12731, plain, (~ sP72(cell1, cell100) | spl177_252), inference(avatar_component_clause, [], [f12730])).
fof(f14932, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_253 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14931, f12735])).
fof(f12735, plain, (~ sP73(cell1, cell100) | spl177_253), inference(avatar_component_clause, [], [f12734])).
fof(f14931, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_254 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14930, f12739])).
fof(f12739, plain, (~ sP74(cell1, cell100) | spl177_254), inference(avatar_component_clause, [], [f12738])).
fof(f14930, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_255 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14929, f12743])).
fof(f12743, plain, (~ sP75(cell1, cell100) | spl177_255), inference(avatar_component_clause, [], [f12742])).
fof(f14929, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_256 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14928, f12747])).
fof(f12747, plain, (~ sP76(cell1, cell100) | spl177_256), inference(avatar_component_clause, [], [f12746])).
fof(f14928, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_257 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14927, f12751])).
fof(f12751, plain, (~ sP77(cell1, cell100) | spl177_257), inference(avatar_component_clause, [], [f12750])).
fof(f14927, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_258 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14926, f12755])).
fof(f12755, plain, (~ sP78(cell1, cell100) | spl177_258), inference(avatar_component_clause, [], [f12754])).
fof(f14926, plain, (sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_261 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14925, f12767])).
fof(f12767, plain, (~ sP81(cell1, cell100) | spl177_261), inference(avatar_component_clause, [], [f12766])).
fof(f14925, plain, (sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_262 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14924, f12771])).
fof(f12771, plain, (~ sP82(cell1, cell100) | spl177_262), inference(avatar_component_clause, [], [f12770])).
fof(f14924, plain, (sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_263 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14923, f12775])).
fof(f12775, plain, (~ sP83(cell1, cell100) | spl177_263), inference(avatar_component_clause, [], [f12774])).
fof(f14923, plain, (sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_264 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14922, f12779])).
fof(f12779, plain, (~ sP84(cell1, cell100) | spl177_264), inference(avatar_component_clause, [], [f12778])).
fof(f14922, plain, (sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_265 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14921, f12783])).
fof(f12783, plain, (~ sP85(cell1, cell100) | spl177_265), inference(avatar_component_clause, [], [f12782])).
fof(f14921, plain, (sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_266 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14920, f12787])).
fof(f12787, plain, (~ sP86(cell1, cell100) | spl177_266), inference(avatar_component_clause, [], [f12786])).
fof(f14920, plain, (sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_267 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14919, f12791])).
fof(f12791, plain, (~ sP90(cell1, cell100) | spl177_267), inference(avatar_component_clause, [], [f12790])).
fof(f14919, plain, (sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_268 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14918, f12795])).
fof(f12795, plain, (~ sP91(cell1, cell100) | spl177_268), inference(avatar_component_clause, [], [f12794])).
fof(f14918, plain, (sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_269 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14917, f12799])).
fof(f12799, plain, (~ sP92(cell1, cell100) | spl177_269), inference(avatar_component_clause, [], [f12798])).
fof(f14917, plain, (sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_270 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14916, f12803])).
fof(f12803, plain, (~ sP93(cell1, cell100) | spl177_270), inference(avatar_component_clause, [], [f12802])).
fof(f14916, plain, (sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_271 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14915, f12807])).
fof(f12807, plain, (~ sP94(cell1, cell100) | spl177_271), inference(avatar_component_clause, [], [f12806])).
fof(f14915, plain, (sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_272 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14914, f12811])).
fof(f12811, plain, (~ sP95(cell1, cell100) | spl177_272), inference(avatar_component_clause, [], [f12810])).
fof(f14914, plain, (sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_273 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14913, f12815])).
fof(f12815, plain, (~ sP96(cell1, cell100) | spl177_273), inference(avatar_component_clause, [], [f12814])).
fof(f14913, plain, (sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_274 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14912, f12819])).
fof(f12819, plain, (~ sP97(cell1, cell100) | spl177_274), inference(avatar_component_clause, [], [f12818])).
fof(f14912, plain, (sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_275 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14911, f12823])).
fof(f12823, plain, (~ sP98(cell1, cell100) | spl177_275), inference(avatar_component_clause, [], [f12822])).
fof(f14911, plain, (sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_276 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14910, f12827])).
fof(f12827, plain, (~ sP99(cell1, cell100) | spl177_276), inference(avatar_component_clause, [], [f12826])).
fof(f14910, plain, (sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_277 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14909, f12831])).
fof(f12831, plain, (~ sP100(cell1, cell100) | spl177_277), inference(avatar_component_clause, [], [f12830])).
fof(f14909, plain, (sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_278 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14908, f12835])).
fof(f12835, plain, (~ sP101(cell1, cell100) | spl177_278), inference(avatar_component_clause, [], [f12834])).
fof(f14908, plain, (sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_279 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14907, f12839])).
fof(f12839, plain, (~ sP102(cell1, cell100) | spl177_279), inference(avatar_component_clause, [], [f12838])).
fof(f14907, plain, (sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_280 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14906, f12843])).
fof(f12843, plain, (~ sP103(cell1, cell100) | spl177_280), inference(avatar_component_clause, [], [f12842])).
fof(f14906, plain, (sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_281 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14905, f12847])).
fof(f12847, plain, (~ sP104(cell1, cell100) | spl177_281), inference(avatar_component_clause, [], [f12846])).
fof(f14905, plain, (sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_282 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14904, f12851])).
fof(f12851, plain, (~ sP105(cell1, cell100) | spl177_282), inference(avatar_component_clause, [], [f12850])).
fof(f14904, plain, (sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_283 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14903, f12855])).
fof(f12855, plain, (~ sP106(cell1, cell100) | spl177_283), inference(avatar_component_clause, [], [f12854])).
fof(f14903, plain, (sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_284 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14902, f12859])).
fof(f12859, plain, (~ sP107(cell1, cell100) | spl177_284), inference(avatar_component_clause, [], [f12858])).
fof(f14902, plain, (sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_285 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14901, f12863])).
fof(f12863, plain, (~ sP108(cell1, cell100) | spl177_285), inference(avatar_component_clause, [], [f12862])).
fof(f14901, plain, (sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_286 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14900, f12867])).
fof(f12867, plain, (~ sP109(cell1, cell100) | spl177_286), inference(avatar_component_clause, [], [f12866])).
fof(f14900, plain, (sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_287 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14899, f12871])).
fof(f12871, plain, (~ sP110(cell1, cell100) | spl177_287), inference(avatar_component_clause, [], [f12870])).
fof(f14899, plain, (sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_288 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14898, f12875])).
fof(f12875, plain, (~ sP111(cell1, cell100) | spl177_288), inference(avatar_component_clause, [], [f12874])).
fof(f14898, plain, (sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_289 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14897, f12879])).
fof(f12879, plain, (~ sP112(cell1, cell100) | spl177_289), inference(avatar_component_clause, [], [f12878])).
fof(f14897, plain, (sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_290 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14896, f12883])).
fof(f12883, plain, (~ sP113(cell1, cell100) | spl177_290), inference(avatar_component_clause, [], [f12882])).
fof(f14896, plain, (sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_291 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14895, f12887])).
fof(f12887, plain, (~ sP114(cell1, cell100) | spl177_291), inference(avatar_component_clause, [], [f12886])).
fof(f14895, plain, (sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_292 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14894, f12891])).
fof(f12891, plain, (~ sP115(cell1, cell100) | spl177_292), inference(avatar_component_clause, [], [f12890])).
fof(f14894, plain, (sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_293 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14893, f12895])).
fof(f12895, plain, (~ sP116(cell1, cell100) | spl177_293), inference(avatar_component_clause, [], [f12894])).
fof(f14893, plain, (sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_294 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14892, f12899])).
fof(f12899, plain, (~ sP117(cell1, cell100) | spl177_294), inference(avatar_component_clause, [], [f12898])).
fof(f14892, plain, (sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_295 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14891, f12903])).
fof(f12903, plain, (~ sP118(cell1, cell100) | spl177_295), inference(avatar_component_clause, [], [f12902])).
fof(f14891, plain, (sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_296 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14890, f12907])).
fof(f12907, plain, (~ sP119(cell1, cell100) | spl177_296), inference(avatar_component_clause, [], [f12906])).
fof(f14890, plain, (sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_297 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14889, f12911])).
fof(f12911, plain, (~ sP120(cell1, cell100) | spl177_297), inference(avatar_component_clause, [], [f12910])).
fof(f14889, plain, (sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_298 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14888, f12915])).
fof(f12915, plain, (~ sP121(cell1, cell100) | spl177_298), inference(avatar_component_clause, [], [f12914])).
fof(f14888, plain, (sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_299 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14887, f12919])).
fof(f12919, plain, (~ sP122(cell1, cell100) | spl177_299), inference(avatar_component_clause, [], [f12918])).
fof(f14887, plain, (sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_300 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14886, f12923])).
fof(f12923, plain, (~ sP123(cell1, cell100) | spl177_300), inference(avatar_component_clause, [], [f12922])).
fof(f14886, plain, (sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_301 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14885, f12927])).
fof(f12927, plain, (~ sP124(cell1, cell100) | spl177_301), inference(avatar_component_clause, [], [f12926])).
fof(f14885, plain, (sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_302 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14884, f12931])).
fof(f12931, plain, (~ sP125(cell1, cell100) | spl177_302), inference(avatar_component_clause, [], [f12930])).
fof(f14884, plain, (sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_303 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14883, f12935])).
fof(f12935, plain, (~ sP126(cell1, cell100) | spl177_303), inference(avatar_component_clause, [], [f12934])).
fof(f14883, plain, (sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_304 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14882, f12939])).
fof(f12939, plain, (~ sP127(cell1, cell100) | spl177_304), inference(avatar_component_clause, [], [f12938])).
fof(f14882, plain, (sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_305 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14881, f12943])).
fof(f12943, plain, (~ sP128(cell1, cell100) | spl177_305), inference(avatar_component_clause, [], [f12942])).
fof(f14881, plain, (sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_306 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14880, f12947])).
fof(f12947, plain, (~ sP129(cell1, cell100) | spl177_306), inference(avatar_component_clause, [], [f12946])).
fof(f14880, plain, (sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_307 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14879, f12951])).
fof(f12951, plain, (~ sP130(cell1, cell100) | spl177_307), inference(avatar_component_clause, [], [f12950])).
fof(f14879, plain, (sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_308 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14878, f12955])).
fof(f12955, plain, (~ sP131(cell1, cell100) | spl177_308), inference(avatar_component_clause, [], [f12954])).
fof(f14878, plain, (sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_309 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14877, f12959])).
fof(f12959, plain, (~ sP132(cell1, cell100) | spl177_309), inference(avatar_component_clause, [], [f12958])).
fof(f14877, plain, (sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_310 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14876, f12963])).
fof(f12963, plain, (~ sP133(cell1, cell100) | spl177_310), inference(avatar_component_clause, [], [f12962])).
fof(f14876, plain, (sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_311 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14875, f12967])).
fof(f12967, plain, (~ sP134(cell1, cell100) | spl177_311), inference(avatar_component_clause, [], [f12966])).
fof(f14875, plain, (sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_312 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14874, f12971])).
fof(f12971, plain, (~ sP135(cell1, cell100) | spl177_312), inference(avatar_component_clause, [], [f12970])).
fof(f14874, plain, (sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_313 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14873, f12975])).
fof(f12975, plain, (~ sP136(cell1, cell100) | spl177_313), inference(avatar_component_clause, [], [f12974])).
fof(f14873, plain, (sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_314 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14872, f12979])).
fof(f12979, plain, (~ sP137(cell1, cell100) | spl177_314), inference(avatar_component_clause, [], [f12978])).
fof(f14872, plain, (sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_315 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14871, f12983])).
fof(f12983, plain, (~ sP138(cell1, cell100) | spl177_315), inference(avatar_component_clause, [], [f12982])).
fof(f14871, plain, (sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_316 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14870, f12987])).
fof(f12987, plain, (~ sP139(cell1, cell100) | spl177_316), inference(avatar_component_clause, [], [f12986])).
fof(f14870, plain, (sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_317 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14869, f12991])).
fof(f12991, plain, (~ sP140(cell1, cell100) | spl177_317), inference(avatar_component_clause, [], [f12990])).
fof(f14869, plain, (sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_318 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14868, f12995])).
fof(f12995, plain, (~ sP141(cell1, cell100) | spl177_318), inference(avatar_component_clause, [], [f12994])).
fof(f14868, plain, (sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_319 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14867, f12999])).
fof(f12999, plain, (~ sP142(cell1, cell100) | spl177_319), inference(avatar_component_clause, [], [f12998])).
fof(f14867, plain, (sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_320 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14866, f13003])).
fof(f13003, plain, (~ sP143(cell1, cell100) | spl177_320), inference(avatar_component_clause, [], [f13002])).
fof(f14866, plain, (sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_321 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14865, f13007])).
fof(f13007, plain, (~ sP144(cell1, cell100) | spl177_321), inference(avatar_component_clause, [], [f13006])).
fof(f14865, plain, (sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_322 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14864, f13011])).
fof(f13011, plain, (~ sP145(cell1, cell100) | spl177_322), inference(avatar_component_clause, [], [f13010])).
fof(f14864, plain, (sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_323 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14863, f13015])).
fof(f13015, plain, (~ sP146(cell1, cell100) | spl177_323), inference(avatar_component_clause, [], [f13014])).
fof(f14863, plain, (sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_324 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14862, f13019])).
fof(f13019, plain, (~ sP147(cell1, cell100) | spl177_324), inference(avatar_component_clause, [], [f13018])).
fof(f14862, plain, (sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_325 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14861, f13023])).
fof(f13023, plain, (~ sP148(cell1, cell100) | spl177_325), inference(avatar_component_clause, [], [f13022])).
fof(f14861, plain, (sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_326 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14860, f13027])).
fof(f13027, plain, (~ sP149(cell1, cell100) | spl177_326), inference(avatar_component_clause, [], [f13026])).
fof(f14860, plain, (sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_327 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14859, f13031])).
fof(f13031, plain, (~ sP150(cell1, cell100) | spl177_327), inference(avatar_component_clause, [], [f13030])).
fof(f14859, plain, (sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_328 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14858, f13035])).
fof(f13035, plain, (~ sP151(cell1, cell100) | spl177_328), inference(avatar_component_clause, [], [f13034])).
fof(f14858, plain, (sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_329 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14857, f13039])).
fof(f13039, plain, (~ sP152(cell1, cell100) | spl177_329), inference(avatar_component_clause, [], [f13038])).
fof(f14857, plain, (sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_330 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14856, f13043])).
fof(f13043, plain, (~ sP153(cell1, cell100) | spl177_330), inference(avatar_component_clause, [], [f13042])).
fof(f14856, plain, (sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_331 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14855, f13047])).
fof(f13047, plain, (~ sP154(cell1, cell100) | spl177_331), inference(avatar_component_clause, [], [f13046])).
fof(f14855, plain, (sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_332 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14854, f13051])).
fof(f13051, plain, (~ sP155(cell1, cell100) | spl177_332), inference(avatar_component_clause, [], [f13050])).
fof(f14854, plain, (sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_333 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14853, f13055])).
fof(f13055, plain, (~ sP156(cell1, cell100) | spl177_333), inference(avatar_component_clause, [], [f13054])).
fof(f14853, plain, (sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_334 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14852, f13059])).
fof(f13059, plain, (~ sP157(cell1, cell100) | spl177_334), inference(avatar_component_clause, [], [f13058])).
fof(f14852, plain, (sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_335 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14851, f13063])).
fof(f13063, plain, (~ sP158(cell1, cell100) | spl177_335), inference(avatar_component_clause, [], [f13062])).
fof(f14851, plain, (sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_336 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14850, f13067])).
fof(f13067, plain, (~ sP159(cell1, cell100) | spl177_336), inference(avatar_component_clause, [], [f13066])).
fof(f14850, plain, (sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_337 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14849, f13071])).
fof(f13071, plain, (~ sP160(cell1, cell100) | spl177_337), inference(avatar_component_clause, [], [f13070])).
fof(f14849, plain, (sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_338 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14848, f13075])).
fof(f13075, plain, (~ sP161(cell1, cell100) | spl177_338), inference(avatar_component_clause, [], [f13074])).
fof(f14848, plain, (sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_339 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14847, f13079])).
fof(f13079, plain, (~ sP162(cell1, cell100) | spl177_339), inference(avatar_component_clause, [], [f13078])).
fof(f14847, plain, (sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_340 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14846, f13083])).
fof(f13083, plain, (~ sP163(cell1, cell100) | spl177_340), inference(avatar_component_clause, [], [f13082])).
fof(f14846, plain, (sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_341 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14845, f13087])).
fof(f13087, plain, (~ sP164(cell1, cell100) | spl177_341), inference(avatar_component_clause, [], [f13086])).
fof(f14845, plain, (sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_342 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14844, f13091])).
fof(f13091, plain, (~ sP165(cell1, cell100) | spl177_342), inference(avatar_component_clause, [], [f13090])).
fof(f14844, plain, (sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_343 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14843, f13095])).
fof(f13095, plain, (~ sP166(cell1, cell100) | spl177_343), inference(avatar_component_clause, [], [f13094])).
fof(f14843, plain, (sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_344 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14842, f13099])).
fof(f13099, plain, (~ sP167(cell1, cell100) | spl177_344), inference(avatar_component_clause, [], [f13098])).
fof(f14842, plain, (sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_345 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14841, f13103])).
fof(f13103, plain, (~ sP168(cell1, cell100) | spl177_345), inference(avatar_component_clause, [], [f13102])).
fof(f14841, plain, (sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_346 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14840, f13107])).
fof(f13107, plain, (~ sP169(cell1, cell100) | spl177_346), inference(avatar_component_clause, [], [f13106])).
fof(f14840, plain, (sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_347 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14839, f13111])).
fof(f13111, plain, (~ sP170(cell1, cell100) | spl177_347), inference(avatar_component_clause, [], [f13110])).
fof(f14839, plain, (sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_348 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14838, f13115])).
fof(f13115, plain, (~ sP171(cell1, cell100) | spl177_348), inference(avatar_component_clause, [], [f13114])).
fof(f14838, plain, (sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_349 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14837, f13119])).
fof(f13119, plain, (~ sP172(cell1, cell100) | spl177_349), inference(avatar_component_clause, [], [f13118])).
fof(f14837, plain, (sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_350 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14836, f13123])).
fof(f13123, plain, (~ sP173(cell1, cell100) | spl177_350), inference(avatar_component_clause, [], [f13122])).
fof(f14836, plain, (sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_351 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14835, f13127])).
fof(f13127, plain, (~ sP174(cell1, cell100) | spl177_351), inference(avatar_component_clause, [], [f13126])).
fof(f14835, plain, (sP174(cell1, cell100) | sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_352 | spl177_353)), inference(subsumption_resolution, [], [f14834, f13131])).
fof(f13131, plain, (~ sP175(cell1, cell100) | spl177_352), inference(avatar_component_clause, [], [f13130])).
fof(f14834, plain, (sP175(cell1, cell100) | sP174(cell1, cell100) | sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | (~ spl177_1 | spl177_353)), inference(subsumption_resolution, [], [f14833, f13135])).
fof(f13135, plain, (~ sP176(cell1, cell100) | spl177_353), inference(avatar_component_clause, [], [f13134])).
fof(f14833, plain, (sP176(cell1, cell100) | sP175(cell1, cell100) | sP174(cell1, cell100) | sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | ~ spl177_1), inference(subsumption_resolution, [], [f14832, f5876])).
fof(f5876, plain, ~ (cell1 = cell98), inference(cnf_transformation, [], [f198])).
fof(f198, plain, ~ (cell1 = cell98), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25187)).
fof(f14832, plain, ((cell1 = cell98) | sP176(cell1, cell100) | sP175(cell1, cell100) | sP174(cell1, cell100) | sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | ~ spl177_1), inference(subsumption_resolution, [], [f14284, f5778])).
fof(f5778, plain, ~ (cell1 = cell99), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ~ (cell1 = cell99), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25089)).
fof(f14284, plain, ((cell1 = cell99) | (cell1 = cell98) | sP176(cell1, cell100) | sP175(cell1, cell100) | sP174(cell1, cell100) | sP173(cell1, cell100) | sP172(cell1, cell100) | sP171(cell1, cell100) | sP170(cell1, cell100) | sP169(cell1, cell100) | sP168(cell1, cell100) | sP167(cell1, cell100) | sP166(cell1, cell100) | sP165(cell1, cell100) | sP164(cell1, cell100) | sP163(cell1, cell100) | sP162(cell1, cell100) | sP161(cell1, cell100) | sP160(cell1, cell100) | sP159(cell1, cell100) | sP158(cell1, cell100) | sP157(cell1, cell100) | sP156(cell1, cell100) | sP155(cell1, cell100) | sP154(cell1, cell100) | sP153(cell1, cell100) | sP152(cell1, cell100) | sP151(cell1, cell100) | sP150(cell1, cell100) | sP149(cell1, cell100) | sP148(cell1, cell100) | sP147(cell1, cell100) | sP146(cell1, cell100) | sP145(cell1, cell100) | sP144(cell1, cell100) | sP143(cell1, cell100) | sP142(cell1, cell100) | sP141(cell1, cell100) | sP140(cell1, cell100) | sP139(cell1, cell100) | sP138(cell1, cell100) | sP137(cell1, cell100) | sP136(cell1, cell100) | sP135(cell1, cell100) | sP134(cell1, cell100) | sP133(cell1, cell100) | sP132(cell1, cell100) | sP131(cell1, cell100) | sP130(cell1, cell100) | sP129(cell1, cell100) | sP128(cell1, cell100) | sP127(cell1, cell100) | sP126(cell1, cell100) | sP125(cell1, cell100) | sP124(cell1, cell100) | sP123(cell1, cell100) | sP122(cell1, cell100) | sP121(cell1, cell100) | sP120(cell1, cell100) | sP119(cell1, cell100) | sP118(cell1, cell100) | sP117(cell1, cell100) | sP116(cell1, cell100) | sP115(cell1, cell100) | sP114(cell1, cell100) | sP113(cell1, cell100) | sP112(cell1, cell100) | sP111(cell1, cell100) | sP110(cell1, cell100) | sP109(cell1, cell100) | sP108(cell1, cell100) | sP107(cell1, cell100) | sP106(cell1, cell100) | sP105(cell1, cell100) | sP104(cell1, cell100) | sP103(cell1, cell100) | sP102(cell1, cell100) | sP101(cell1, cell100) | sP100(cell1, cell100) | sP99(cell1, cell100) | sP98(cell1, cell100) | sP97(cell1, cell100) | sP96(cell1, cell100) | sP95(cell1, cell100) | sP94(cell1, cell100) | sP93(cell1, cell100) | sP92(cell1, cell100) | sP91(cell1, cell100) | sP90(cell1, cell100) | sP86(cell1, cell100) | sP85(cell1, cell100) | sP84(cell1, cell100) | sP83(cell1, cell100) | sP82(cell1, cell100) | sP81(cell1, cell100) | sP80(cell1, cell100) | sP79(cell1, cell100) | sP78(cell1, cell100) | sP77(cell1, cell100) | sP76(cell1, cell100) | sP75(cell1, cell100) | sP74(cell1, cell100) | sP73(cell1, cell100) | sP72(cell1, cell100) | sP71(cell1, cell100) | sP70(cell1, cell100) | sP69(cell1, cell100) | sP68(cell1, cell100) | sP67(cell1, cell100) | sP66(cell1, cell100) | sP65(cell1, cell100) | sP64(cell1, cell100) | sP63(cell1, cell100) | sP62(cell1, cell100) | sP61(cell1, cell100) | sP60(cell1, cell100) | sP59(cell1, cell100) | sP58(cell1, cell100) | sP57(cell1, cell100) | sP56(cell1, cell100) | sP55(cell1, cell100) | sP54(cell1, cell100) | sP53(cell1, cell100) | sP52(cell1, cell100) | sP51(cell1, cell100) | sP50(cell1, cell100) | sP49(cell1, cell100) | sP48(cell1, cell100) | sP47(cell1, cell100) | sP46(cell1, cell100) | sP45(cell1, cell100) | sP44(cell1, cell100) | sP43(cell1, cell100) | sP42(cell1, cell100) | sP41(cell1, cell100) | sP40(cell1, cell100) | sP39(cell1, cell100) | sP38(cell1, cell100) | sP37(cell1, cell100) | sP36(cell1, cell100) | sP35(cell1, cell100) | sP34(cell1, cell100) | sP33(cell1, cell100) | sP32(cell1, cell100) | sP31(cell1, cell100) | sP30(cell1, cell100) | sP29(cell1, cell100) | sP28(cell1, cell100) | sP27(cell1, cell100) | sP26(cell1, cell100) | sP25(cell1, cell100) | sP24(cell1, cell100) | sP23(cell1, cell100) | sP22(cell1, cell100) | sP21(cell1, cell100) | sP20(cell1, cell100) | sP19(cell1, cell100) | sP18(cell1, cell100) | sP17(cell1, cell100) | sP16(cell1, cell100) | sP15(cell1, cell100) | sP14(cell1, cell100) | sP13(cell1, cell100) | sP12(cell1, cell100) | sP11(cell1, cell100) | sP10(cell1, cell100) | sP9(cell1, cell100) | sP8(cell1, cell100) | sP7(cell1, cell100) | sP6(cell1, cell100) | sP5(cell1, cell100) | sP4(cell1, cell100) | sP3(cell1, cell100) | sP2(cell1, cell100) | sP1(cell1, cell100) | sP0(cell1, cell100) | (cell1 = cell100) | ~ spl177_1), inference(resolution, [], [f10656, f10634])).
fof(f10634, plain, ! [X0, X1] : (~ edge(X1, X0) | (cell99 = X0) | (cell98 = X0) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | (cell100 = X0)), inference(subsumption_resolution, [], [f10633, f5493])).
fof(f5493, plain, ! [X0, X1] : (~ sP89(X0, X1) | (cell100 = X0)), inference(cnf_transformation, [], [f5228])).
fof(f5228, plain, ! [X0, X1] : (((cell100 = X0) & (cell99 = X1)) | ~ sP89(X0, X1)), inference(nnf_transformation, [], [f5052])).
fof(f5052, plain, ! [X0, X1] : (((cell100 = X0) & (cell99 = X1)) | ~ sP89(X0, X1)), inference(usedef, [], [e5052])).
fof(e5052, plain, ! [X0, X1] : (sP89(X0, X1) <=> ((cell100 = X0) & (cell99 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP89])])).
fof(f10633, plain, ! [X0, X1] : ((cell100 = X0) | (cell99 = X0) | (cell98 = X0) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(subsumption_resolution, [], [f10632, f5495])).
fof(f5495, plain, ! [X0, X1] : (~ sP88(X0, X1) | (cell99 = X0)), inference(cnf_transformation, [], [f5229])).
fof(f5229, plain, ! [X0, X1] : (((cell99 = X0) & (cell98 = X1)) | ~ sP88(X0, X1)), inference(nnf_transformation, [], [f5051])).
fof(f5051, plain, ! [X0, X1] : (((cell99 = X0) & (cell98 = X1)) | ~ sP88(X0, X1)), inference(usedef, [], [e5051])).
fof(e5051, plain, ! [X0, X1] : (sP88(X0, X1) <=> ((cell99 = X0) & (cell98 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP88])])).
fof(f10632, plain, ! [X0, X1] : ((cell100 = X0) | (cell99 = X0) | (cell98 = X0) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(subsumption_resolution, [], [f5679, f5497])).
fof(f5497, plain, ! [X0, X1] : (~ sP87(X0, X1) | (cell98 = X0)), inference(cnf_transformation, [], [f5230])).
fof(f5230, plain, ! [X0, X1] : (((cell98 = X0) & (cell97 = X1)) | ~ sP87(X0, X1)), inference(nnf_transformation, [], [f5050])).
fof(f5050, plain, ! [X0, X1] : (((cell98 = X0) & (cell97 = X1)) | ~ sP87(X0, X1)), inference(usedef, [], [e5050])).
fof(e5050, plain, ! [X0, X1] : (sP87(X0, X1) <=> ((cell98 = X0) & (cell97 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP87])])).
fof(f5679, plain, ! [X0, X1] : ((cell100 = X0) | (cell99 = X0) | (cell98 = X0) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP87(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(cnf_transformation, [], [f5140])).
fof(f5140, plain, ! [X0, X1] : (((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP87(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(definition_folding, [], [f4960, e5139, e5138, e5137, e5136, e5135, e5134, e5133, e5132, e5131, e5130, e5129, e5128, e5127, e5126, e5125, e5124, e5123, e5122, e5121, e5120, e5119, e5118, e5117, e5116, e5115, e5114, e5113, e5112, e5111, e5110, e5109, e5108, e5107, e5106, e5105, e5104, e5103, e5102, e5101, e5100, e5099, e5098, e5097, e5096, e5095, e5094, e5093, e5092, e5091, e5090, e5089, e5088, e5087, e5086, e5085, e5084, e5083, e5082, e5081, e5080, e5079, e5078, e5077, e5076, e5075, e5074, e5073, e5072, e5071, e5070, e5069, e5068, e5067, e5066, e5065, e5064, e5063, e5062, e5061, e5060, e5059, e5058, e5057, e5056, e5055, e5054, e5053, e5052, e5051, e5050, e5049, e5048, e5047, e5046, e5045, e5044, e5043, e5042, e5041, e5040, e5039, e5038, e5037, e5036, e5035, e5034, e5033, e5032, e5031, e5030, e5029, e5028, e5027, e5026, e5025, e5024, e5023, e5022, e5021, e5020, e5019, e5018, e5017, e5016, e5015, e5014, e5013, e5012, e5011, e5010, e5009, e5008, e5007, e5006, e5005, e5004, e5003, e5002, e5001, e5000, e4999, e4998, e4997, e4996, e4995, e4994, e4993, e4992, e4991, e4990, e4989, e4988, e4987, e4986, e4985, e4984, e4983, e4982, e4981, e4980, e4979, e4978, e4977, e4976, e4975, e4974, e4973, e4972, e4971, e4970, e4969, e4968, e4967, e4966, e4965, e4964, e4963])).
fof(f4963, plain, ! [X0, X1] : (((cell2 = X0) & (cell1 = X1)) | ~ sP0(X0, X1)), inference(usedef, [], [e4963])).
fof(e4963, plain, ! [X0, X1] : (sP0(X0, X1) <=> ((cell2 = X0) & (cell1 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f4964, plain, ! [X0, X1] : (((cell3 = X0) & (cell2 = X1)) | ~ sP1(X0, X1)), inference(usedef, [], [e4964])).
fof(e4964, plain, ! [X0, X1] : (sP1(X0, X1) <=> ((cell3 = X0) & (cell2 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f4965, plain, ! [X0, X1] : (((cell4 = X0) & (cell3 = X1)) | ~ sP2(X0, X1)), inference(usedef, [], [e4965])).
fof(e4965, plain, ! [X0, X1] : (sP2(X0, X1) <=> ((cell4 = X0) & (cell3 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f4966, plain, ! [X0, X1] : (((cell5 = X0) & (cell4 = X1)) | ~ sP3(X0, X1)), inference(usedef, [], [e4966])).
fof(e4966, plain, ! [X0, X1] : (sP3(X0, X1) <=> ((cell5 = X0) & (cell4 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f4967, plain, ! [X0, X1] : (((cell6 = X0) & (cell5 = X1)) | ~ sP4(X0, X1)), inference(usedef, [], [e4967])).
fof(e4967, plain, ! [X0, X1] : (sP4(X0, X1) <=> ((cell6 = X0) & (cell5 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f4968, plain, ! [X0, X1] : (((cell7 = X0) & (cell6 = X1)) | ~ sP5(X0, X1)), inference(usedef, [], [e4968])).
fof(e4968, plain, ! [X0, X1] : (sP5(X0, X1) <=> ((cell7 = X0) & (cell6 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f4969, plain, ! [X0, X1] : (((cell8 = X0) & (cell7 = X1)) | ~ sP6(X0, X1)), inference(usedef, [], [e4969])).
fof(e4969, plain, ! [X0, X1] : (sP6(X0, X1) <=> ((cell8 = X0) & (cell7 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f4970, plain, ! [X0, X1] : (((cell9 = X0) & (cell8 = X1)) | ~ sP7(X0, X1)), inference(usedef, [], [e4970])).
fof(e4970, plain, ! [X0, X1] : (sP7(X0, X1) <=> ((cell9 = X0) & (cell8 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f4971, plain, ! [X0, X1] : (((cell10 = X0) & (cell9 = X1)) | ~ sP8(X0, X1)), inference(usedef, [], [e4971])).
fof(e4971, plain, ! [X0, X1] : (sP8(X0, X1) <=> ((cell10 = X0) & (cell9 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f4972, plain, ! [X0, X1] : (((cell12 = X0) & (cell11 = X1)) | ~ sP9(X0, X1)), inference(usedef, [], [e4972])).
fof(e4972, plain, ! [X0, X1] : (sP9(X0, X1) <=> ((cell12 = X0) & (cell11 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f4973, plain, ! [X0, X1] : (((cell13 = X0) & (cell12 = X1)) | ~ sP10(X0, X1)), inference(usedef, [], [e4973])).
fof(e4973, plain, ! [X0, X1] : (sP10(X0, X1) <=> ((cell13 = X0) & (cell12 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f4974, plain, ! [X0, X1] : (((cell14 = X0) & (cell13 = X1)) | ~ sP11(X0, X1)), inference(usedef, [], [e4974])).
fof(e4974, plain, ! [X0, X1] : (sP11(X0, X1) <=> ((cell14 = X0) & (cell13 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f4975, plain, ! [X0, X1] : (((cell15 = X0) & (cell14 = X1)) | ~ sP12(X0, X1)), inference(usedef, [], [e4975])).
fof(e4975, plain, ! [X0, X1] : (sP12(X0, X1) <=> ((cell15 = X0) & (cell14 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f4976, plain, ! [X0, X1] : (((cell16 = X0) & (cell15 = X1)) | ~ sP13(X0, X1)), inference(usedef, [], [e4976])).
fof(e4976, plain, ! [X0, X1] : (sP13(X0, X1) <=> ((cell16 = X0) & (cell15 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f4977, plain, ! [X0, X1] : (((cell17 = X0) & (cell16 = X1)) | ~ sP14(X0, X1)), inference(usedef, [], [e4977])).
fof(e4977, plain, ! [X0, X1] : (sP14(X0, X1) <=> ((cell17 = X0) & (cell16 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f4978, plain, ! [X0, X1] : (((cell18 = X0) & (cell17 = X1)) | ~ sP15(X0, X1)), inference(usedef, [], [e4978])).
fof(e4978, plain, ! [X0, X1] : (sP15(X0, X1) <=> ((cell18 = X0) & (cell17 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f4979, plain, ! [X0, X1] : (((cell19 = X0) & (cell18 = X1)) | ~ sP16(X0, X1)), inference(usedef, [], [e4979])).
fof(e4979, plain, ! [X0, X1] : (sP16(X0, X1) <=> ((cell19 = X0) & (cell18 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f4980, plain, ! [X0, X1] : (((cell20 = X0) & (cell19 = X1)) | ~ sP17(X0, X1)), inference(usedef, [], [e4980])).
fof(e4980, plain, ! [X0, X1] : (sP17(X0, X1) <=> ((cell20 = X0) & (cell19 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f4981, plain, ! [X0, X1] : (((cell22 = X0) & (cell21 = X1)) | ~ sP18(X0, X1)), inference(usedef, [], [e4981])).
fof(e4981, plain, ! [X0, X1] : (sP18(X0, X1) <=> ((cell22 = X0) & (cell21 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f4982, plain, ! [X0, X1] : (((cell23 = X0) & (cell22 = X1)) | ~ sP19(X0, X1)), inference(usedef, [], [e4982])).
fof(e4982, plain, ! [X0, X1] : (sP19(X0, X1) <=> ((cell23 = X0) & (cell22 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f4983, plain, ! [X0, X1] : (((cell24 = X0) & (cell23 = X1)) | ~ sP20(X0, X1)), inference(usedef, [], [e4983])).
fof(e4983, plain, ! [X0, X1] : (sP20(X0, X1) <=> ((cell24 = X0) & (cell23 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f4984, plain, ! [X0, X1] : (((cell25 = X0) & (cell24 = X1)) | ~ sP21(X0, X1)), inference(usedef, [], [e4984])).
fof(e4984, plain, ! [X0, X1] : (sP21(X0, X1) <=> ((cell25 = X0) & (cell24 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f4985, plain, ! [X0, X1] : (((cell26 = X0) & (cell25 = X1)) | ~ sP22(X0, X1)), inference(usedef, [], [e4985])).
fof(e4985, plain, ! [X0, X1] : (sP22(X0, X1) <=> ((cell26 = X0) & (cell25 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f4986, plain, ! [X0, X1] : (((cell27 = X0) & (cell26 = X1)) | ~ sP23(X0, X1)), inference(usedef, [], [e4986])).
fof(e4986, plain, ! [X0, X1] : (sP23(X0, X1) <=> ((cell27 = X0) & (cell26 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f4987, plain, ! [X0, X1] : (((cell28 = X0) & (cell27 = X1)) | ~ sP24(X0, X1)), inference(usedef, [], [e4987])).
fof(e4987, plain, ! [X0, X1] : (sP24(X0, X1) <=> ((cell28 = X0) & (cell27 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f4988, plain, ! [X0, X1] : (((cell29 = X0) & (cell28 = X1)) | ~ sP25(X0, X1)), inference(usedef, [], [e4988])).
fof(e4988, plain, ! [X0, X1] : (sP25(X0, X1) <=> ((cell29 = X0) & (cell28 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f4989, plain, ! [X0, X1] : (((cell30 = X0) & (cell29 = X1)) | ~ sP26(X0, X1)), inference(usedef, [], [e4989])).
fof(e4989, plain, ! [X0, X1] : (sP26(X0, X1) <=> ((cell30 = X0) & (cell29 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f4990, plain, ! [X0, X1] : (((cell32 = X0) & (cell31 = X1)) | ~ sP27(X0, X1)), inference(usedef, [], [e4990])).
fof(e4990, plain, ! [X0, X1] : (sP27(X0, X1) <=> ((cell32 = X0) & (cell31 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f4991, plain, ! [X0, X1] : (((cell33 = X0) & (cell32 = X1)) | ~ sP28(X0, X1)), inference(usedef, [], [e4991])).
fof(e4991, plain, ! [X0, X1] : (sP28(X0, X1) <=> ((cell33 = X0) & (cell32 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f4992, plain, ! [X0, X1] : (((cell34 = X0) & (cell33 = X1)) | ~ sP29(X0, X1)), inference(usedef, [], [e4992])).
fof(e4992, plain, ! [X0, X1] : (sP29(X0, X1) <=> ((cell34 = X0) & (cell33 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f4993, plain, ! [X0, X1] : (((cell35 = X0) & (cell34 = X1)) | ~ sP30(X0, X1)), inference(usedef, [], [e4993])).
fof(e4993, plain, ! [X0, X1] : (sP30(X0, X1) <=> ((cell35 = X0) & (cell34 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f4994, plain, ! [X0, X1] : (((cell36 = X0) & (cell35 = X1)) | ~ sP31(X0, X1)), inference(usedef, [], [e4994])).
fof(e4994, plain, ! [X0, X1] : (sP31(X0, X1) <=> ((cell36 = X0) & (cell35 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f4995, plain, ! [X0, X1] : (((cell37 = X0) & (cell36 = X1)) | ~ sP32(X0, X1)), inference(usedef, [], [e4995])).
fof(e4995, plain, ! [X0, X1] : (sP32(X0, X1) <=> ((cell37 = X0) & (cell36 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f4996, plain, ! [X0, X1] : (((cell38 = X0) & (cell37 = X1)) | ~ sP33(X0, X1)), inference(usedef, [], [e4996])).
fof(e4996, plain, ! [X0, X1] : (sP33(X0, X1) <=> ((cell38 = X0) & (cell37 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f4997, plain, ! [X0, X1] : (((cell39 = X0) & (cell38 = X1)) | ~ sP34(X0, X1)), inference(usedef, [], [e4997])).
fof(e4997, plain, ! [X0, X1] : (sP34(X0, X1) <=> ((cell39 = X0) & (cell38 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f4998, plain, ! [X0, X1] : (((cell40 = X0) & (cell39 = X1)) | ~ sP35(X0, X1)), inference(usedef, [], [e4998])).
fof(e4998, plain, ! [X0, X1] : (sP35(X0, X1) <=> ((cell40 = X0) & (cell39 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f4999, plain, ! [X0, X1] : (((cell42 = X0) & (cell41 = X1)) | ~ sP36(X0, X1)), inference(usedef, [], [e4999])).
fof(e4999, plain, ! [X0, X1] : (sP36(X0, X1) <=> ((cell42 = X0) & (cell41 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f5000, plain, ! [X0, X1] : (((cell43 = X0) & (cell42 = X1)) | ~ sP37(X0, X1)), inference(usedef, [], [e5000])).
fof(e5000, plain, ! [X0, X1] : (sP37(X0, X1) <=> ((cell43 = X0) & (cell42 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f5001, plain, ! [X0, X1] : (((cell44 = X0) & (cell43 = X1)) | ~ sP38(X0, X1)), inference(usedef, [], [e5001])).
fof(e5001, plain, ! [X0, X1] : (sP38(X0, X1) <=> ((cell44 = X0) & (cell43 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f5002, plain, ! [X0, X1] : (((cell45 = X0) & (cell44 = X1)) | ~ sP39(X0, X1)), inference(usedef, [], [e5002])).
fof(e5002, plain, ! [X0, X1] : (sP39(X0, X1) <=> ((cell45 = X0) & (cell44 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f5003, plain, ! [X0, X1] : (((cell46 = X0) & (cell45 = X1)) | ~ sP40(X0, X1)), inference(usedef, [], [e5003])).
fof(e5003, plain, ! [X0, X1] : (sP40(X0, X1) <=> ((cell46 = X0) & (cell45 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f5004, plain, ! [X0, X1] : (((cell47 = X0) & (cell46 = X1)) | ~ sP41(X0, X1)), inference(usedef, [], [e5004])).
fof(e5004, plain, ! [X0, X1] : (sP41(X0, X1) <=> ((cell47 = X0) & (cell46 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f5005, plain, ! [X0, X1] : (((cell48 = X0) & (cell47 = X1)) | ~ sP42(X0, X1)), inference(usedef, [], [e5005])).
fof(e5005, plain, ! [X0, X1] : (sP42(X0, X1) <=> ((cell48 = X0) & (cell47 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f5006, plain, ! [X0, X1] : (((cell49 = X0) & (cell48 = X1)) | ~ sP43(X0, X1)), inference(usedef, [], [e5006])).
fof(e5006, plain, ! [X0, X1] : (sP43(X0, X1) <=> ((cell49 = X0) & (cell48 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f5007, plain, ! [X0, X1] : (((cell50 = X0) & (cell49 = X1)) | ~ sP44(X0, X1)), inference(usedef, [], [e5007])).
fof(e5007, plain, ! [X0, X1] : (sP44(X0, X1) <=> ((cell50 = X0) & (cell49 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f5008, plain, ! [X0, X1] : (((cell52 = X0) & (cell51 = X1)) | ~ sP45(X0, X1)), inference(usedef, [], [e5008])).
fof(e5008, plain, ! [X0, X1] : (sP45(X0, X1) <=> ((cell52 = X0) & (cell51 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f5009, plain, ! [X0, X1] : (((cell53 = X0) & (cell52 = X1)) | ~ sP46(X0, X1)), inference(usedef, [], [e5009])).
fof(e5009, plain, ! [X0, X1] : (sP46(X0, X1) <=> ((cell53 = X0) & (cell52 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f5010, plain, ! [X0, X1] : (((cell54 = X0) & (cell53 = X1)) | ~ sP47(X0, X1)), inference(usedef, [], [e5010])).
fof(e5010, plain, ! [X0, X1] : (sP47(X0, X1) <=> ((cell54 = X0) & (cell53 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f5011, plain, ! [X0, X1] : (((cell55 = X0) & (cell54 = X1)) | ~ sP48(X0, X1)), inference(usedef, [], [e5011])).
fof(e5011, plain, ! [X0, X1] : (sP48(X0, X1) <=> ((cell55 = X0) & (cell54 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f5012, plain, ! [X0, X1] : (((cell56 = X0) & (cell55 = X1)) | ~ sP49(X0, X1)), inference(usedef, [], [e5012])).
fof(e5012, plain, ! [X0, X1] : (sP49(X0, X1) <=> ((cell56 = X0) & (cell55 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f5013, plain, ! [X0, X1] : (((cell57 = X0) & (cell56 = X1)) | ~ sP50(X0, X1)), inference(usedef, [], [e5013])).
fof(e5013, plain, ! [X0, X1] : (sP50(X0, X1) <=> ((cell57 = X0) & (cell56 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f5014, plain, ! [X0, X1] : (((cell58 = X0) & (cell57 = X1)) | ~ sP51(X0, X1)), inference(usedef, [], [e5014])).
fof(e5014, plain, ! [X0, X1] : (sP51(X0, X1) <=> ((cell58 = X0) & (cell57 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f5015, plain, ! [X0, X1] : (((cell59 = X0) & (cell58 = X1)) | ~ sP52(X0, X1)), inference(usedef, [], [e5015])).
fof(e5015, plain, ! [X0, X1] : (sP52(X0, X1) <=> ((cell59 = X0) & (cell58 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f5016, plain, ! [X0, X1] : (((cell60 = X0) & (cell59 = X1)) | ~ sP53(X0, X1)), inference(usedef, [], [e5016])).
fof(e5016, plain, ! [X0, X1] : (sP53(X0, X1) <=> ((cell60 = X0) & (cell59 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f5017, plain, ! [X0, X1] : (((cell62 = X0) & (cell61 = X1)) | ~ sP54(X0, X1)), inference(usedef, [], [e5017])).
fof(e5017, plain, ! [X0, X1] : (sP54(X0, X1) <=> ((cell62 = X0) & (cell61 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f5018, plain, ! [X0, X1] : (((cell63 = X0) & (cell62 = X1)) | ~ sP55(X0, X1)), inference(usedef, [], [e5018])).
fof(e5018, plain, ! [X0, X1] : (sP55(X0, X1) <=> ((cell63 = X0) & (cell62 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f5019, plain, ! [X0, X1] : (((cell64 = X0) & (cell63 = X1)) | ~ sP56(X0, X1)), inference(usedef, [], [e5019])).
fof(e5019, plain, ! [X0, X1] : (sP56(X0, X1) <=> ((cell64 = X0) & (cell63 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f5020, plain, ! [X0, X1] : (((cell65 = X0) & (cell64 = X1)) | ~ sP57(X0, X1)), inference(usedef, [], [e5020])).
fof(e5020, plain, ! [X0, X1] : (sP57(X0, X1) <=> ((cell65 = X0) & (cell64 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f5021, plain, ! [X0, X1] : (((cell66 = X0) & (cell65 = X1)) | ~ sP58(X0, X1)), inference(usedef, [], [e5021])).
fof(e5021, plain, ! [X0, X1] : (sP58(X0, X1) <=> ((cell66 = X0) & (cell65 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f5022, plain, ! [X0, X1] : (((cell67 = X0) & (cell66 = X1)) | ~ sP59(X0, X1)), inference(usedef, [], [e5022])).
fof(e5022, plain, ! [X0, X1] : (sP59(X0, X1) <=> ((cell67 = X0) & (cell66 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f5023, plain, ! [X0, X1] : (((cell68 = X0) & (cell67 = X1)) | ~ sP60(X0, X1)), inference(usedef, [], [e5023])).
fof(e5023, plain, ! [X0, X1] : (sP60(X0, X1) <=> ((cell68 = X0) & (cell67 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f5024, plain, ! [X0, X1] : (((cell69 = X0) & (cell68 = X1)) | ~ sP61(X0, X1)), inference(usedef, [], [e5024])).
fof(e5024, plain, ! [X0, X1] : (sP61(X0, X1) <=> ((cell69 = X0) & (cell68 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f5025, plain, ! [X0, X1] : (((cell70 = X0) & (cell69 = X1)) | ~ sP62(X0, X1)), inference(usedef, [], [e5025])).
fof(e5025, plain, ! [X0, X1] : (sP62(X0, X1) <=> ((cell70 = X0) & (cell69 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f5026, plain, ! [X0, X1] : (((cell72 = X0) & (cell71 = X1)) | ~ sP63(X0, X1)), inference(usedef, [], [e5026])).
fof(e5026, plain, ! [X0, X1] : (sP63(X0, X1) <=> ((cell72 = X0) & (cell71 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f5027, plain, ! [X0, X1] : (((cell73 = X0) & (cell72 = X1)) | ~ sP64(X0, X1)), inference(usedef, [], [e5027])).
fof(e5027, plain, ! [X0, X1] : (sP64(X0, X1) <=> ((cell73 = X0) & (cell72 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f5028, plain, ! [X0, X1] : (((cell74 = X0) & (cell73 = X1)) | ~ sP65(X0, X1)), inference(usedef, [], [e5028])).
fof(e5028, plain, ! [X0, X1] : (sP65(X0, X1) <=> ((cell74 = X0) & (cell73 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f5029, plain, ! [X0, X1] : (((cell75 = X0) & (cell74 = X1)) | ~ sP66(X0, X1)), inference(usedef, [], [e5029])).
fof(e5029, plain, ! [X0, X1] : (sP66(X0, X1) <=> ((cell75 = X0) & (cell74 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f5030, plain, ! [X0, X1] : (((cell76 = X0) & (cell75 = X1)) | ~ sP67(X0, X1)), inference(usedef, [], [e5030])).
fof(e5030, plain, ! [X0, X1] : (sP67(X0, X1) <=> ((cell76 = X0) & (cell75 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f5031, plain, ! [X0, X1] : (((cell77 = X0) & (cell76 = X1)) | ~ sP68(X0, X1)), inference(usedef, [], [e5031])).
fof(e5031, plain, ! [X0, X1] : (sP68(X0, X1) <=> ((cell77 = X0) & (cell76 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f5032, plain, ! [X0, X1] : (((cell78 = X0) & (cell77 = X1)) | ~ sP69(X0, X1)), inference(usedef, [], [e5032])).
fof(e5032, plain, ! [X0, X1] : (sP69(X0, X1) <=> ((cell78 = X0) & (cell77 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f5033, plain, ! [X0, X1] : (((cell79 = X0) & (cell78 = X1)) | ~ sP70(X0, X1)), inference(usedef, [], [e5033])).
fof(e5033, plain, ! [X0, X1] : (sP70(X0, X1) <=> ((cell79 = X0) & (cell78 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f5034, plain, ! [X0, X1] : (((cell80 = X0) & (cell79 = X1)) | ~ sP71(X0, X1)), inference(usedef, [], [e5034])).
fof(e5034, plain, ! [X0, X1] : (sP71(X0, X1) <=> ((cell80 = X0) & (cell79 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f5035, plain, ! [X0, X1] : (((cell82 = X0) & (cell81 = X1)) | ~ sP72(X0, X1)), inference(usedef, [], [e5035])).
fof(e5035, plain, ! [X0, X1] : (sP72(X0, X1) <=> ((cell82 = X0) & (cell81 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f5036, plain, ! [X0, X1] : (((cell83 = X0) & (cell82 = X1)) | ~ sP73(X0, X1)), inference(usedef, [], [e5036])).
fof(e5036, plain, ! [X0, X1] : (sP73(X0, X1) <=> ((cell83 = X0) & (cell82 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f5037, plain, ! [X0, X1] : (((cell84 = X0) & (cell83 = X1)) | ~ sP74(X0, X1)), inference(usedef, [], [e5037])).
fof(e5037, plain, ! [X0, X1] : (sP74(X0, X1) <=> ((cell84 = X0) & (cell83 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f5038, plain, ! [X0, X1] : (((cell85 = X0) & (cell84 = X1)) | ~ sP75(X0, X1)), inference(usedef, [], [e5038])).
fof(e5038, plain, ! [X0, X1] : (sP75(X0, X1) <=> ((cell85 = X0) & (cell84 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f5039, plain, ! [X0, X1] : (((cell86 = X0) & (cell85 = X1)) | ~ sP76(X0, X1)), inference(usedef, [], [e5039])).
fof(e5039, plain, ! [X0, X1] : (sP76(X0, X1) <=> ((cell86 = X0) & (cell85 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f5040, plain, ! [X0, X1] : (((cell87 = X0) & (cell86 = X1)) | ~ sP77(X0, X1)), inference(usedef, [], [e5040])).
fof(e5040, plain, ! [X0, X1] : (sP77(X0, X1) <=> ((cell87 = X0) & (cell86 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f5041, plain, ! [X0, X1] : (((cell88 = X0) & (cell87 = X1)) | ~ sP78(X0, X1)), inference(usedef, [], [e5041])).
fof(e5041, plain, ! [X0, X1] : (sP78(X0, X1) <=> ((cell88 = X0) & (cell87 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP78])])).
fof(f5042, plain, ! [X0, X1] : (((cell89 = X0) & (cell88 = X1)) | ~ sP79(X0, X1)), inference(usedef, [], [e5042])).
fof(e5042, plain, ! [X0, X1] : (sP79(X0, X1) <=> ((cell89 = X0) & (cell88 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP79])])).
fof(f5044, plain, ! [X0, X1] : (((cell92 = X0) & (cell91 = X1)) | ~ sP81(X0, X1)), inference(usedef, [], [e5044])).
fof(e5044, plain, ! [X0, X1] : (sP81(X0, X1) <=> ((cell92 = X0) & (cell91 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP81])])).
fof(f5045, plain, ! [X0, X1] : (((cell93 = X0) & (cell92 = X1)) | ~ sP82(X0, X1)), inference(usedef, [], [e5045])).
fof(e5045, plain, ! [X0, X1] : (sP82(X0, X1) <=> ((cell93 = X0) & (cell92 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP82])])).
fof(f5046, plain, ! [X0, X1] : (((cell94 = X0) & (cell93 = X1)) | ~ sP83(X0, X1)), inference(usedef, [], [e5046])).
fof(e5046, plain, ! [X0, X1] : (sP83(X0, X1) <=> ((cell94 = X0) & (cell93 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP83])])).
fof(f5047, plain, ! [X0, X1] : (((cell95 = X0) & (cell94 = X1)) | ~ sP84(X0, X1)), inference(usedef, [], [e5047])).
fof(e5047, plain, ! [X0, X1] : (sP84(X0, X1) <=> ((cell95 = X0) & (cell94 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP84])])).
fof(f5048, plain, ! [X0, X1] : (((cell96 = X0) & (cell95 = X1)) | ~ sP85(X0, X1)), inference(usedef, [], [e5048])).
fof(e5048, plain, ! [X0, X1] : (sP85(X0, X1) <=> ((cell96 = X0) & (cell95 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP85])])).
fof(f5049, plain, ! [X0, X1] : (((cell97 = X0) & (cell96 = X1)) | ~ sP86(X0, X1)), inference(usedef, [], [e5049])).
fof(e5049, plain, ! [X0, X1] : (sP86(X0, X1) <=> ((cell97 = X0) & (cell96 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP86])])).
fof(f5053, plain, ! [X0, X1] : (((cell11 = X0) & (cell1 = X1)) | ~ sP90(X0, X1)), inference(usedef, [], [e5053])).
fof(e5053, plain, ! [X0, X1] : (sP90(X0, X1) <=> ((cell11 = X0) & (cell1 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP90])])).
fof(f5054, plain, ! [X0, X1] : (((cell12 = X0) & (cell2 = X1)) | ~ sP91(X0, X1)), inference(usedef, [], [e5054])).
fof(e5054, plain, ! [X0, X1] : (sP91(X0, X1) <=> ((cell12 = X0) & (cell2 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP91])])).
fof(f5055, plain, ! [X0, X1] : (((cell13 = X0) & (cell3 = X1)) | ~ sP92(X0, X1)), inference(usedef, [], [e5055])).
fof(e5055, plain, ! [X0, X1] : (sP92(X0, X1) <=> ((cell13 = X0) & (cell3 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP92])])).
fof(f5056, plain, ! [X0, X1] : (((cell14 = X0) & (cell4 = X1)) | ~ sP93(X0, X1)), inference(usedef, [], [e5056])).
fof(e5056, plain, ! [X0, X1] : (sP93(X0, X1) <=> ((cell14 = X0) & (cell4 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP93])])).
fof(f5057, plain, ! [X0, X1] : (((cell15 = X0) & (cell5 = X1)) | ~ sP94(X0, X1)), inference(usedef, [], [e5057])).
fof(e5057, plain, ! [X0, X1] : (sP94(X0, X1) <=> ((cell15 = X0) & (cell5 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP94])])).
fof(f5058, plain, ! [X0, X1] : (((cell16 = X0) & (cell6 = X1)) | ~ sP95(X0, X1)), inference(usedef, [], [e5058])).
fof(e5058, plain, ! [X0, X1] : (sP95(X0, X1) <=> ((cell16 = X0) & (cell6 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP95])])).
fof(f5059, plain, ! [X0, X1] : (((cell17 = X0) & (cell7 = X1)) | ~ sP96(X0, X1)), inference(usedef, [], [e5059])).
fof(e5059, plain, ! [X0, X1] : (sP96(X0, X1) <=> ((cell17 = X0) & (cell7 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP96])])).
fof(f5060, plain, ! [X0, X1] : (((cell18 = X0) & (cell8 = X1)) | ~ sP97(X0, X1)), inference(usedef, [], [e5060])).
fof(e5060, plain, ! [X0, X1] : (sP97(X0, X1) <=> ((cell18 = X0) & (cell8 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP97])])).
fof(f5061, plain, ! [X0, X1] : (((cell19 = X0) & (cell9 = X1)) | ~ sP98(X0, X1)), inference(usedef, [], [e5061])).
fof(e5061, plain, ! [X0, X1] : (sP98(X0, X1) <=> ((cell19 = X0) & (cell9 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP98])])).
fof(f5062, plain, ! [X0, X1] : (((cell20 = X0) & (cell10 = X1)) | ~ sP99(X0, X1)), inference(usedef, [], [e5062])).
fof(e5062, plain, ! [X0, X1] : (sP99(X0, X1) <=> ((cell20 = X0) & (cell10 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP99])])).
fof(f5063, plain, ! [X0, X1] : (((cell21 = X0) & (cell11 = X1)) | ~ sP100(X0, X1)), inference(usedef, [], [e5063])).
fof(e5063, plain, ! [X0, X1] : (sP100(X0, X1) <=> ((cell21 = X0) & (cell11 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP100])])).
fof(f5064, plain, ! [X0, X1] : (((cell22 = X0) & (cell12 = X1)) | ~ sP101(X0, X1)), inference(usedef, [], [e5064])).
fof(e5064, plain, ! [X0, X1] : (sP101(X0, X1) <=> ((cell22 = X0) & (cell12 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP101])])).
fof(f5065, plain, ! [X0, X1] : (((cell23 = X0) & (cell13 = X1)) | ~ sP102(X0, X1)), inference(usedef, [], [e5065])).
fof(e5065, plain, ! [X0, X1] : (sP102(X0, X1) <=> ((cell23 = X0) & (cell13 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP102])])).
fof(f5066, plain, ! [X0, X1] : (((cell24 = X0) & (cell14 = X1)) | ~ sP103(X0, X1)), inference(usedef, [], [e5066])).
fof(e5066, plain, ! [X0, X1] : (sP103(X0, X1) <=> ((cell24 = X0) & (cell14 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP103])])).
fof(f5067, plain, ! [X0, X1] : (((cell25 = X0) & (cell15 = X1)) | ~ sP104(X0, X1)), inference(usedef, [], [e5067])).
fof(e5067, plain, ! [X0, X1] : (sP104(X0, X1) <=> ((cell25 = X0) & (cell15 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP104])])).
fof(f5068, plain, ! [X0, X1] : (((cell26 = X0) & (cell16 = X1)) | ~ sP105(X0, X1)), inference(usedef, [], [e5068])).
fof(e5068, plain, ! [X0, X1] : (sP105(X0, X1) <=> ((cell26 = X0) & (cell16 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP105])])).
fof(f5069, plain, ! [X0, X1] : (((cell27 = X0) & (cell17 = X1)) | ~ sP106(X0, X1)), inference(usedef, [], [e5069])).
fof(e5069, plain, ! [X0, X1] : (sP106(X0, X1) <=> ((cell27 = X0) & (cell17 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP106])])).
fof(f5070, plain, ! [X0, X1] : (((cell28 = X0) & (cell18 = X1)) | ~ sP107(X0, X1)), inference(usedef, [], [e5070])).
fof(e5070, plain, ! [X0, X1] : (sP107(X0, X1) <=> ((cell28 = X0) & (cell18 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP107])])).
fof(f5071, plain, ! [X0, X1] : (((cell29 = X0) & (cell19 = X1)) | ~ sP108(X0, X1)), inference(usedef, [], [e5071])).
fof(e5071, plain, ! [X0, X1] : (sP108(X0, X1) <=> ((cell29 = X0) & (cell19 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP108])])).
fof(f5072, plain, ! [X0, X1] : (((cell30 = X0) & (cell20 = X1)) | ~ sP109(X0, X1)), inference(usedef, [], [e5072])).
fof(e5072, plain, ! [X0, X1] : (sP109(X0, X1) <=> ((cell30 = X0) & (cell20 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP109])])).
fof(f5073, plain, ! [X0, X1] : (((cell31 = X0) & (cell21 = X1)) | ~ sP110(X0, X1)), inference(usedef, [], [e5073])).
fof(e5073, plain, ! [X0, X1] : (sP110(X0, X1) <=> ((cell31 = X0) & (cell21 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP110])])).
fof(f5074, plain, ! [X0, X1] : (((cell32 = X0) & (cell22 = X1)) | ~ sP111(X0, X1)), inference(usedef, [], [e5074])).
fof(e5074, plain, ! [X0, X1] : (sP111(X0, X1) <=> ((cell32 = X0) & (cell22 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP111])])).
fof(f5075, plain, ! [X0, X1] : (((cell33 = X0) & (cell23 = X1)) | ~ sP112(X0, X1)), inference(usedef, [], [e5075])).
fof(e5075, plain, ! [X0, X1] : (sP112(X0, X1) <=> ((cell33 = X0) & (cell23 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP112])])).
fof(f5076, plain, ! [X0, X1] : (((cell34 = X0) & (cell24 = X1)) | ~ sP113(X0, X1)), inference(usedef, [], [e5076])).
fof(e5076, plain, ! [X0, X1] : (sP113(X0, X1) <=> ((cell34 = X0) & (cell24 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP113])])).
fof(f5077, plain, ! [X0, X1] : (((cell35 = X0) & (cell25 = X1)) | ~ sP114(X0, X1)), inference(usedef, [], [e5077])).
fof(e5077, plain, ! [X0, X1] : (sP114(X0, X1) <=> ((cell35 = X0) & (cell25 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP114])])).
fof(f5078, plain, ! [X0, X1] : (((cell36 = X0) & (cell26 = X1)) | ~ sP115(X0, X1)), inference(usedef, [], [e5078])).
fof(e5078, plain, ! [X0, X1] : (sP115(X0, X1) <=> ((cell36 = X0) & (cell26 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP115])])).
fof(f5079, plain, ! [X0, X1] : (((cell37 = X0) & (cell27 = X1)) | ~ sP116(X0, X1)), inference(usedef, [], [e5079])).
fof(e5079, plain, ! [X0, X1] : (sP116(X0, X1) <=> ((cell37 = X0) & (cell27 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP116])])).
fof(f5080, plain, ! [X0, X1] : (((cell38 = X0) & (cell28 = X1)) | ~ sP117(X0, X1)), inference(usedef, [], [e5080])).
fof(e5080, plain, ! [X0, X1] : (sP117(X0, X1) <=> ((cell38 = X0) & (cell28 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP117])])).
fof(f5081, plain, ! [X0, X1] : (((cell39 = X0) & (cell29 = X1)) | ~ sP118(X0, X1)), inference(usedef, [], [e5081])).
fof(e5081, plain, ! [X0, X1] : (sP118(X0, X1) <=> ((cell39 = X0) & (cell29 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP118])])).
fof(f5082, plain, ! [X0, X1] : (((cell40 = X0) & (cell30 = X1)) | ~ sP119(X0, X1)), inference(usedef, [], [e5082])).
fof(e5082, plain, ! [X0, X1] : (sP119(X0, X1) <=> ((cell40 = X0) & (cell30 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP119])])).
fof(f5083, plain, ! [X0, X1] : (((cell41 = X0) & (cell31 = X1)) | ~ sP120(X0, X1)), inference(usedef, [], [e5083])).
fof(e5083, plain, ! [X0, X1] : (sP120(X0, X1) <=> ((cell41 = X0) & (cell31 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP120])])).
fof(f5084, plain, ! [X0, X1] : (((cell42 = X0) & (cell32 = X1)) | ~ sP121(X0, X1)), inference(usedef, [], [e5084])).
fof(e5084, plain, ! [X0, X1] : (sP121(X0, X1) <=> ((cell42 = X0) & (cell32 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP121])])).
fof(f5085, plain, ! [X0, X1] : (((cell43 = X0) & (cell33 = X1)) | ~ sP122(X0, X1)), inference(usedef, [], [e5085])).
fof(e5085, plain, ! [X0, X1] : (sP122(X0, X1) <=> ((cell43 = X0) & (cell33 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP122])])).
fof(f5086, plain, ! [X0, X1] : (((cell44 = X0) & (cell34 = X1)) | ~ sP123(X0, X1)), inference(usedef, [], [e5086])).
fof(e5086, plain, ! [X0, X1] : (sP123(X0, X1) <=> ((cell44 = X0) & (cell34 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP123])])).
fof(f5087, plain, ! [X0, X1] : (((cell45 = X0) & (cell35 = X1)) | ~ sP124(X0, X1)), inference(usedef, [], [e5087])).
fof(e5087, plain, ! [X0, X1] : (sP124(X0, X1) <=> ((cell45 = X0) & (cell35 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP124])])).
fof(f5088, plain, ! [X0, X1] : (((cell46 = X0) & (cell36 = X1)) | ~ sP125(X0, X1)), inference(usedef, [], [e5088])).
fof(e5088, plain, ! [X0, X1] : (sP125(X0, X1) <=> ((cell46 = X0) & (cell36 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP125])])).
fof(f5089, plain, ! [X0, X1] : (((cell47 = X0) & (cell37 = X1)) | ~ sP126(X0, X1)), inference(usedef, [], [e5089])).
fof(e5089, plain, ! [X0, X1] : (sP126(X0, X1) <=> ((cell47 = X0) & (cell37 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP126])])).
fof(f5090, plain, ! [X0, X1] : (((cell48 = X0) & (cell38 = X1)) | ~ sP127(X0, X1)), inference(usedef, [], [e5090])).
fof(e5090, plain, ! [X0, X1] : (sP127(X0, X1) <=> ((cell48 = X0) & (cell38 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP127])])).
fof(f5091, plain, ! [X0, X1] : (((cell49 = X0) & (cell39 = X1)) | ~ sP128(X0, X1)), inference(usedef, [], [e5091])).
fof(e5091, plain, ! [X0, X1] : (sP128(X0, X1) <=> ((cell49 = X0) & (cell39 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP128])])).
fof(f5092, plain, ! [X0, X1] : (((cell50 = X0) & (cell40 = X1)) | ~ sP129(X0, X1)), inference(usedef, [], [e5092])).
fof(e5092, plain, ! [X0, X1] : (sP129(X0, X1) <=> ((cell50 = X0) & (cell40 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP129])])).
fof(f5093, plain, ! [X0, X1] : (((cell51 = X0) & (cell41 = X1)) | ~ sP130(X0, X1)), inference(usedef, [], [e5093])).
fof(e5093, plain, ! [X0, X1] : (sP130(X0, X1) <=> ((cell51 = X0) & (cell41 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP130])])).
fof(f5094, plain, ! [X0, X1] : (((cell52 = X0) & (cell42 = X1)) | ~ sP131(X0, X1)), inference(usedef, [], [e5094])).
fof(e5094, plain, ! [X0, X1] : (sP131(X0, X1) <=> ((cell52 = X0) & (cell42 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP131])])).
fof(f5095, plain, ! [X0, X1] : (((cell53 = X0) & (cell43 = X1)) | ~ sP132(X0, X1)), inference(usedef, [], [e5095])).
fof(e5095, plain, ! [X0, X1] : (sP132(X0, X1) <=> ((cell53 = X0) & (cell43 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP132])])).
fof(f5096, plain, ! [X0, X1] : (((cell54 = X0) & (cell44 = X1)) | ~ sP133(X0, X1)), inference(usedef, [], [e5096])).
fof(e5096, plain, ! [X0, X1] : (sP133(X0, X1) <=> ((cell54 = X0) & (cell44 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP133])])).
fof(f5097, plain, ! [X0, X1] : (((cell55 = X0) & (cell45 = X1)) | ~ sP134(X0, X1)), inference(usedef, [], [e5097])).
fof(e5097, plain, ! [X0, X1] : (sP134(X0, X1) <=> ((cell55 = X0) & (cell45 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP134])])).
fof(f5098, plain, ! [X0, X1] : (((cell56 = X0) & (cell46 = X1)) | ~ sP135(X0, X1)), inference(usedef, [], [e5098])).
fof(e5098, plain, ! [X0, X1] : (sP135(X0, X1) <=> ((cell56 = X0) & (cell46 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP135])])).
fof(f5099, plain, ! [X0, X1] : (((cell57 = X0) & (cell47 = X1)) | ~ sP136(X0, X1)), inference(usedef, [], [e5099])).
fof(e5099, plain, ! [X0, X1] : (sP136(X0, X1) <=> ((cell57 = X0) & (cell47 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP136])])).
fof(f5100, plain, ! [X0, X1] : (((cell58 = X0) & (cell48 = X1)) | ~ sP137(X0, X1)), inference(usedef, [], [e5100])).
fof(e5100, plain, ! [X0, X1] : (sP137(X0, X1) <=> ((cell58 = X0) & (cell48 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP137])])).
fof(f5101, plain, ! [X0, X1] : (((cell59 = X0) & (cell49 = X1)) | ~ sP138(X0, X1)), inference(usedef, [], [e5101])).
fof(e5101, plain, ! [X0, X1] : (sP138(X0, X1) <=> ((cell59 = X0) & (cell49 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP138])])).
fof(f5102, plain, ! [X0, X1] : (((cell60 = X0) & (cell50 = X1)) | ~ sP139(X0, X1)), inference(usedef, [], [e5102])).
fof(e5102, plain, ! [X0, X1] : (sP139(X0, X1) <=> ((cell60 = X0) & (cell50 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP139])])).
fof(f5103, plain, ! [X0, X1] : (((cell61 = X0) & (cell51 = X1)) | ~ sP140(X0, X1)), inference(usedef, [], [e5103])).
fof(e5103, plain, ! [X0, X1] : (sP140(X0, X1) <=> ((cell61 = X0) & (cell51 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP140])])).
fof(f5104, plain, ! [X0, X1] : (((cell62 = X0) & (cell52 = X1)) | ~ sP141(X0, X1)), inference(usedef, [], [e5104])).
fof(e5104, plain, ! [X0, X1] : (sP141(X0, X1) <=> ((cell62 = X0) & (cell52 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP141])])).
fof(f5105, plain, ! [X0, X1] : (((cell63 = X0) & (cell53 = X1)) | ~ sP142(X0, X1)), inference(usedef, [], [e5105])).
fof(e5105, plain, ! [X0, X1] : (sP142(X0, X1) <=> ((cell63 = X0) & (cell53 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP142])])).
fof(f5106, plain, ! [X0, X1] : (((cell64 = X0) & (cell54 = X1)) | ~ sP143(X0, X1)), inference(usedef, [], [e5106])).
fof(e5106, plain, ! [X0, X1] : (sP143(X0, X1) <=> ((cell64 = X0) & (cell54 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP143])])).
fof(f5107, plain, ! [X0, X1] : (((cell65 = X0) & (cell55 = X1)) | ~ sP144(X0, X1)), inference(usedef, [], [e5107])).
fof(e5107, plain, ! [X0, X1] : (sP144(X0, X1) <=> ((cell65 = X0) & (cell55 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP144])])).
fof(f5108, plain, ! [X0, X1] : (((cell66 = X0) & (cell56 = X1)) | ~ sP145(X0, X1)), inference(usedef, [], [e5108])).
fof(e5108, plain, ! [X0, X1] : (sP145(X0, X1) <=> ((cell66 = X0) & (cell56 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP145])])).
fof(f5109, plain, ! [X0, X1] : (((cell67 = X0) & (cell57 = X1)) | ~ sP146(X0, X1)), inference(usedef, [], [e5109])).
fof(e5109, plain, ! [X0, X1] : (sP146(X0, X1) <=> ((cell67 = X0) & (cell57 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP146])])).
fof(f5110, plain, ! [X0, X1] : (((cell68 = X0) & (cell58 = X1)) | ~ sP147(X0, X1)), inference(usedef, [], [e5110])).
fof(e5110, plain, ! [X0, X1] : (sP147(X0, X1) <=> ((cell68 = X0) & (cell58 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP147])])).
fof(f5111, plain, ! [X0, X1] : (((cell69 = X0) & (cell59 = X1)) | ~ sP148(X0, X1)), inference(usedef, [], [e5111])).
fof(e5111, plain, ! [X0, X1] : (sP148(X0, X1) <=> ((cell69 = X0) & (cell59 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP148])])).
fof(f5112, plain, ! [X0, X1] : (((cell70 = X0) & (cell60 = X1)) | ~ sP149(X0, X1)), inference(usedef, [], [e5112])).
fof(e5112, plain, ! [X0, X1] : (sP149(X0, X1) <=> ((cell70 = X0) & (cell60 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP149])])).
fof(f5113, plain, ! [X0, X1] : (((cell71 = X0) & (cell61 = X1)) | ~ sP150(X0, X1)), inference(usedef, [], [e5113])).
fof(e5113, plain, ! [X0, X1] : (sP150(X0, X1) <=> ((cell71 = X0) & (cell61 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP150])])).
fof(f5114, plain, ! [X0, X1] : (((cell72 = X0) & (cell62 = X1)) | ~ sP151(X0, X1)), inference(usedef, [], [e5114])).
fof(e5114, plain, ! [X0, X1] : (sP151(X0, X1) <=> ((cell72 = X0) & (cell62 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP151])])).
fof(f5115, plain, ! [X0, X1] : (((cell73 = X0) & (cell63 = X1)) | ~ sP152(X0, X1)), inference(usedef, [], [e5115])).
fof(e5115, plain, ! [X0, X1] : (sP152(X0, X1) <=> ((cell73 = X0) & (cell63 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP152])])).
fof(f5116, plain, ! [X0, X1] : (((cell74 = X0) & (cell64 = X1)) | ~ sP153(X0, X1)), inference(usedef, [], [e5116])).
fof(e5116, plain, ! [X0, X1] : (sP153(X0, X1) <=> ((cell74 = X0) & (cell64 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP153])])).
fof(f5117, plain, ! [X0, X1] : (((cell75 = X0) & (cell65 = X1)) | ~ sP154(X0, X1)), inference(usedef, [], [e5117])).
fof(e5117, plain, ! [X0, X1] : (sP154(X0, X1) <=> ((cell75 = X0) & (cell65 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP154])])).
fof(f5118, plain, ! [X0, X1] : (((cell76 = X0) & (cell66 = X1)) | ~ sP155(X0, X1)), inference(usedef, [], [e5118])).
fof(e5118, plain, ! [X0, X1] : (sP155(X0, X1) <=> ((cell76 = X0) & (cell66 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP155])])).
fof(f5119, plain, ! [X0, X1] : (((cell77 = X0) & (cell67 = X1)) | ~ sP156(X0, X1)), inference(usedef, [], [e5119])).
fof(e5119, plain, ! [X0, X1] : (sP156(X0, X1) <=> ((cell77 = X0) & (cell67 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP156])])).
fof(f5120, plain, ! [X0, X1] : (((cell78 = X0) & (cell68 = X1)) | ~ sP157(X0, X1)), inference(usedef, [], [e5120])).
fof(e5120, plain, ! [X0, X1] : (sP157(X0, X1) <=> ((cell78 = X0) & (cell68 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP157])])).
fof(f5121, plain, ! [X0, X1] : (((cell79 = X0) & (cell69 = X1)) | ~ sP158(X0, X1)), inference(usedef, [], [e5121])).
fof(e5121, plain, ! [X0, X1] : (sP158(X0, X1) <=> ((cell79 = X0) & (cell69 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP158])])).
fof(f5122, plain, ! [X0, X1] : (((cell80 = X0) & (cell70 = X1)) | ~ sP159(X0, X1)), inference(usedef, [], [e5122])).
fof(e5122, plain, ! [X0, X1] : (sP159(X0, X1) <=> ((cell80 = X0) & (cell70 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP159])])).
fof(f5123, plain, ! [X0, X1] : (((cell81 = X0) & (cell71 = X1)) | ~ sP160(X0, X1)), inference(usedef, [], [e5123])).
fof(e5123, plain, ! [X0, X1] : (sP160(X0, X1) <=> ((cell81 = X0) & (cell71 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP160])])).
fof(f5124, plain, ! [X0, X1] : (((cell82 = X0) & (cell72 = X1)) | ~ sP161(X0, X1)), inference(usedef, [], [e5124])).
fof(e5124, plain, ! [X0, X1] : (sP161(X0, X1) <=> ((cell82 = X0) & (cell72 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP161])])).
fof(f5125, plain, ! [X0, X1] : (((cell83 = X0) & (cell73 = X1)) | ~ sP162(X0, X1)), inference(usedef, [], [e5125])).
fof(e5125, plain, ! [X0, X1] : (sP162(X0, X1) <=> ((cell83 = X0) & (cell73 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP162])])).
fof(f5126, plain, ! [X0, X1] : (((cell84 = X0) & (cell74 = X1)) | ~ sP163(X0, X1)), inference(usedef, [], [e5126])).
fof(e5126, plain, ! [X0, X1] : (sP163(X0, X1) <=> ((cell84 = X0) & (cell74 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP163])])).
fof(f5127, plain, ! [X0, X1] : (((cell85 = X0) & (cell75 = X1)) | ~ sP164(X0, X1)), inference(usedef, [], [e5127])).
fof(e5127, plain, ! [X0, X1] : (sP164(X0, X1) <=> ((cell85 = X0) & (cell75 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP164])])).
fof(f5128, plain, ! [X0, X1] : (((cell86 = X0) & (cell76 = X1)) | ~ sP165(X0, X1)), inference(usedef, [], [e5128])).
fof(e5128, plain, ! [X0, X1] : (sP165(X0, X1) <=> ((cell86 = X0) & (cell76 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP165])])).
fof(f5129, plain, ! [X0, X1] : (((cell87 = X0) & (cell77 = X1)) | ~ sP166(X0, X1)), inference(usedef, [], [e5129])).
fof(e5129, plain, ! [X0, X1] : (sP166(X0, X1) <=> ((cell87 = X0) & (cell77 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP166])])).
fof(f5130, plain, ! [X0, X1] : (((cell88 = X0) & (cell78 = X1)) | ~ sP167(X0, X1)), inference(usedef, [], [e5130])).
fof(e5130, plain, ! [X0, X1] : (sP167(X0, X1) <=> ((cell88 = X0) & (cell78 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP167])])).
fof(f5131, plain, ! [X0, X1] : (((cell89 = X0) & (cell79 = X1)) | ~ sP168(X0, X1)), inference(usedef, [], [e5131])).
fof(e5131, plain, ! [X0, X1] : (sP168(X0, X1) <=> ((cell89 = X0) & (cell79 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP168])])).
fof(f5132, plain, ! [X0, X1] : (((cell90 = X0) & (cell80 = X1)) | ~ sP169(X0, X1)), inference(usedef, [], [e5132])).
fof(e5132, plain, ! [X0, X1] : (sP169(X0, X1) <=> ((cell90 = X0) & (cell80 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP169])])).
fof(f5133, plain, ! [X0, X1] : (((cell91 = X0) & (cell81 = X1)) | ~ sP170(X0, X1)), inference(usedef, [], [e5133])).
fof(e5133, plain, ! [X0, X1] : (sP170(X0, X1) <=> ((cell91 = X0) & (cell81 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP170])])).
fof(f5134, plain, ! [X0, X1] : (((cell92 = X0) & (cell82 = X1)) | ~ sP171(X0, X1)), inference(usedef, [], [e5134])).
fof(e5134, plain, ! [X0, X1] : (sP171(X0, X1) <=> ((cell92 = X0) & (cell82 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP171])])).
fof(f5135, plain, ! [X0, X1] : (((cell93 = X0) & (cell83 = X1)) | ~ sP172(X0, X1)), inference(usedef, [], [e5135])).
fof(e5135, plain, ! [X0, X1] : (sP172(X0, X1) <=> ((cell93 = X0) & (cell83 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP172])])).
fof(f5136, plain, ! [X0, X1] : (((cell94 = X0) & (cell84 = X1)) | ~ sP173(X0, X1)), inference(usedef, [], [e5136])).
fof(e5136, plain, ! [X0, X1] : (sP173(X0, X1) <=> ((cell94 = X0) & (cell84 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP173])])).
fof(f5137, plain, ! [X0, X1] : (((cell95 = X0) & (cell85 = X1)) | ~ sP174(X0, X1)), inference(usedef, [], [e5137])).
fof(e5137, plain, ! [X0, X1] : (sP174(X0, X1) <=> ((cell95 = X0) & (cell85 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP174])])).
fof(f5138, plain, ! [X0, X1] : (((cell96 = X0) & (cell86 = X1)) | ~ sP175(X0, X1)), inference(usedef, [], [e5138])).
fof(e5138, plain, ! [X0, X1] : (sP175(X0, X1) <=> ((cell96 = X0) & (cell86 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP175])])).
fof(f5139, plain, ! [X0, X1] : (((cell97 = X0) & (cell87 = X1)) | ~ sP176(X0, X1)), inference(usedef, [], [e5139])).
fof(e5139, plain, ! [X0, X1] : (sP176(X0, X1) <=> ((cell97 = X0) & (cell87 = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP176])])).
fof(f4960, plain, ! [X0, X1] : (((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1)) | ((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1)) | ~ edge(X1, X0)), inference(flattening, [], [f4959])).
fof(f4959, plain, ! [X0, X1] : ((((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1)) | ((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1))) | ~ edge(X1, X0)), inference(ennf_transformation, [], [f4958])).
fof(f4958, plain, ! [X0, X1] : (edge(X1, X0) => (((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1)) | ((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1)))), inference(unused_predicate_definition_removal, [], [f1])).
fof(f1, plain, ! [X0, X1] : (edge(X1, X0) <=> (((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1)) | ((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1)))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof24990)).
fof(f10656, plain, (edge(cell100, cell1) | ~ spl177_1), inference(avatar_component_clause, [], [f10654])).
fof(f14303, plain, ~ spl177_259, inference(avatar_contradiction_clause, [], [f14302])).
fof(f14302, plain, ($false | ~ spl177_259), inference(subsumption_resolution, [], [f14299, f6802])).
fof(f6802, plain, ~ (cell88 = cell100), inference(cnf_transformation, [], [f1124])).
fof(f1124, plain, ~ (cell88 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26113)).
fof(f14299, plain, ((cell88 = cell100) | ~ spl177_259), inference(resolution, [], [f12760, f5512])).
fof(f5512, plain, ! [X0, X1] : (~ sP79(X0, X1) | (cell88 = X1)), inference(cnf_transformation, [], [f5238])).
fof(f5238, plain, ! [X0, X1] : (((cell89 = X0) & (cell88 = X1)) | ~ sP79(X0, X1)), inference(nnf_transformation, [], [f5042])).
fof(f12760, plain, (sP79(cell1, cell100) | ~ spl177_259), inference(avatar_component_clause, [], [f12758])).
fof(f14297, plain, ~ spl177_179, inference(avatar_contradiction_clause, [], [f14296])).
fof(f14296, plain, ($false | ~ spl177_179), inference(subsumption_resolution, [], [f14293, f5973])).
fof(f5973, plain, ~ (cell1 = cell97), inference(cnf_transformation, [], [f295])).
fof(f295, plain, ~ (cell1 = cell97), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25284)).
fof(f14293, plain, ((cell1 = cell97) | ~ spl177_179), inference(resolution, [], [f11399, f5496])).
fof(f5496, plain, ! [X0, X1] : (~ sP87(X0, X1) | (cell97 = X1)), inference(cnf_transformation, [], [f5230])).
fof(f11399, plain, (sP87(cell100, cell1) | ~ spl177_179), inference(avatar_component_clause, [], [f11397])).
fof(f11397, plain, (spl177_179 <=> sP87(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_179])])).
fof(f14269, plain, ~ spl177_353, inference(avatar_contradiction_clause, [], [f14268])).
fof(f14268, plain, ($false | ~ spl177_353), inference(subsumption_resolution, [], [f14265, f6889])).
fof(f6889, plain, ~ (cell87 = cell100), inference(cnf_transformation, [], [f1211])).
fof(f1211, plain, ~ (cell87 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26200)).
fof(f14265, plain, ((cell87 = cell100) | ~ spl177_353), inference(resolution, [], [f13136, f5318])).
fof(f5318, plain, ! [X0, X1] : (~ sP176(X0, X1) | (cell87 = X1)), inference(cnf_transformation, [], [f5141])).
fof(f5141, plain, ! [X0, X1] : (((cell97 = X0) & (cell87 = X1)) | ~ sP176(X0, X1)), inference(nnf_transformation, [], [f5139])).
fof(f13136, plain, (sP176(cell1, cell100) | ~ spl177_353), inference(avatar_component_clause, [], [f13134])).
fof(f14263, plain, ~ spl177_352, inference(avatar_contradiction_clause, [], [f14262])).
fof(f14262, plain, ($false | ~ spl177_352), inference(subsumption_resolution, [], [f14259, f6975])).
fof(f6975, plain, ~ (cell86 = cell100), inference(cnf_transformation, [], [f1297])).
fof(f1297, plain, ~ (cell86 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26286)).
fof(f14259, plain, ((cell86 = cell100) | ~ spl177_352), inference(resolution, [], [f13132, f5320])).
fof(f5320, plain, ! [X0, X1] : (~ sP175(X0, X1) | (cell86 = X1)), inference(cnf_transformation, [], [f5142])).
fof(f5142, plain, ! [X0, X1] : (((cell96 = X0) & (cell86 = X1)) | ~ sP175(X0, X1)), inference(nnf_transformation, [], [f5138])).
fof(f13132, plain, (sP175(cell1, cell100) | ~ spl177_352), inference(avatar_component_clause, [], [f13130])).
fof(f14257, plain, ~ spl177_351, inference(avatar_contradiction_clause, [], [f14256])).
fof(f14256, plain, ($false | ~ spl177_351), inference(subsumption_resolution, [], [f14253, f7060])).
fof(f7060, plain, ~ (cell85 = cell100), inference(cnf_transformation, [], [f1382])).
fof(f1382, plain, ~ (cell85 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26371)).
fof(f14253, plain, ((cell85 = cell100) | ~ spl177_351), inference(resolution, [], [f13128, f5322])).
fof(f5322, plain, ! [X0, X1] : (~ sP174(X0, X1) | (cell85 = X1)), inference(cnf_transformation, [], [f5143])).
fof(f5143, plain, ! [X0, X1] : (((cell95 = X0) & (cell85 = X1)) | ~ sP174(X0, X1)), inference(nnf_transformation, [], [f5137])).
fof(f13128, plain, (sP174(cell1, cell100) | ~ spl177_351), inference(avatar_component_clause, [], [f13126])).
fof(f14251, plain, ~ spl177_350, inference(avatar_contradiction_clause, [], [f14250])).
fof(f14250, plain, ($false | ~ spl177_350), inference(subsumption_resolution, [], [f14247, f7144])).
fof(f7144, plain, ~ (cell84 = cell100), inference(cnf_transformation, [], [f1466])).
fof(f1466, plain, ~ (cell84 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26455)).
fof(f14247, plain, ((cell84 = cell100) | ~ spl177_350), inference(resolution, [], [f13124, f5324])).
fof(f5324, plain, ! [X0, X1] : (~ sP173(X0, X1) | (cell84 = X1)), inference(cnf_transformation, [], [f5144])).
fof(f5144, plain, ! [X0, X1] : (((cell94 = X0) & (cell84 = X1)) | ~ sP173(X0, X1)), inference(nnf_transformation, [], [f5136])).
fof(f13124, plain, (sP173(cell1, cell100) | ~ spl177_350), inference(avatar_component_clause, [], [f13122])).
fof(f14245, plain, ~ spl177_349, inference(avatar_contradiction_clause, [], [f14244])).
fof(f14244, plain, ($false | ~ spl177_349), inference(subsumption_resolution, [], [f14241, f7227])).
fof(f7227, plain, ~ (cell83 = cell100), inference(cnf_transformation, [], [f1549])).
fof(f1549, plain, ~ (cell83 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26538)).
fof(f14241, plain, ((cell83 = cell100) | ~ spl177_349), inference(resolution, [], [f13120, f5326])).
fof(f5326, plain, ! [X0, X1] : (~ sP172(X0, X1) | (cell83 = X1)), inference(cnf_transformation, [], [f5145])).
fof(f5145, plain, ! [X0, X1] : (((cell93 = X0) & (cell83 = X1)) | ~ sP172(X0, X1)), inference(nnf_transformation, [], [f5135])).
fof(f13120, plain, (sP172(cell1, cell100) | ~ spl177_349), inference(avatar_component_clause, [], [f13118])).
fof(f14239, plain, ~ spl177_348, inference(avatar_contradiction_clause, [], [f14238])).
fof(f14238, plain, ($false | ~ spl177_348), inference(subsumption_resolution, [], [f14235, f7390])).
fof(f7390, plain, ~ (cell82 = cell100), inference(cnf_transformation, [], [f1712])).
fof(f1712, plain, ~ (cell82 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26701)).
fof(f14235, plain, ((cell82 = cell100) | ~ spl177_348), inference(resolution, [], [f13116, f5328])).
fof(f5328, plain, ! [X0, X1] : (~ sP171(X0, X1) | (cell82 = X1)), inference(cnf_transformation, [], [f5146])).
fof(f5146, plain, ! [X0, X1] : (((cell92 = X0) & (cell82 = X1)) | ~ sP171(X0, X1)), inference(nnf_transformation, [], [f5134])).
fof(f13116, plain, (sP171(cell1, cell100) | ~ spl177_348), inference(avatar_component_clause, [], [f13114])).
fof(f14233, plain, ~ spl177_347, inference(avatar_contradiction_clause, [], [f14232])).
fof(f14232, plain, ($false | ~ spl177_347), inference(subsumption_resolution, [], [f14229, f7309])).
fof(f7309, plain, ~ (cell81 = cell100), inference(cnf_transformation, [], [f1631])).
fof(f1631, plain, ~ (cell81 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26620)).
fof(f14229, plain, ((cell81 = cell100) | ~ spl177_347), inference(resolution, [], [f13112, f5330])).
fof(f5330, plain, ! [X0, X1] : (~ sP170(X0, X1) | (cell81 = X1)), inference(cnf_transformation, [], [f5147])).
fof(f5147, plain, ! [X0, X1] : (((cell91 = X0) & (cell81 = X1)) | ~ sP170(X0, X1)), inference(nnf_transformation, [], [f5133])).
fof(f13112, plain, (sP170(cell1, cell100) | ~ spl177_347), inference(avatar_component_clause, [], [f13110])).
fof(f14227, plain, ~ spl177_346, inference(avatar_contradiction_clause, [], [f14226])).
fof(f14226, plain, ($false | ~ spl177_346), inference(subsumption_resolution, [], [f14223, f7470])).
fof(f7470, plain, ~ (cell80 = cell100), inference(cnf_transformation, [], [f1792])).
fof(f1792, plain, ~ (cell80 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26781)).
fof(f14223, plain, ((cell80 = cell100) | ~ spl177_346), inference(resolution, [], [f13108, f5332])).
fof(f5332, plain, ! [X0, X1] : (~ sP169(X0, X1) | (cell80 = X1)), inference(cnf_transformation, [], [f5148])).
fof(f5148, plain, ! [X0, X1] : (((cell90 = X0) & (cell80 = X1)) | ~ sP169(X0, X1)), inference(nnf_transformation, [], [f5132])).
fof(f13108, plain, (sP169(cell1, cell100) | ~ spl177_346), inference(avatar_component_clause, [], [f13106])).
fof(f14221, plain, ~ spl177_345, inference(avatar_contradiction_clause, [], [f14220])).
fof(f14220, plain, ($false | ~ spl177_345), inference(subsumption_resolution, [], [f14217, f7549])).
fof(f7549, plain, ~ (cell79 = cell100), inference(cnf_transformation, [], [f1871])).
fof(f1871, plain, ~ (cell79 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26860)).
fof(f14217, plain, ((cell79 = cell100) | ~ spl177_345), inference(resolution, [], [f13104, f5334])).
fof(f5334, plain, ! [X0, X1] : (~ sP168(X0, X1) | (cell79 = X1)), inference(cnf_transformation, [], [f5149])).
fof(f5149, plain, ! [X0, X1] : (((cell89 = X0) & (cell79 = X1)) | ~ sP168(X0, X1)), inference(nnf_transformation, [], [f5131])).
fof(f13104, plain, (sP168(cell1, cell100) | ~ spl177_345), inference(avatar_component_clause, [], [f13102])).
fof(f14215, plain, ~ spl177_344, inference(avatar_contradiction_clause, [], [f14214])).
fof(f14214, plain, ($false | ~ spl177_344), inference(subsumption_resolution, [], [f14211, f7627])).
fof(f7627, plain, ~ (cell78 = cell100), inference(cnf_transformation, [], [f1949])).
fof(f1949, plain, ~ (cell78 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26938)).
fof(f14211, plain, ((cell78 = cell100) | ~ spl177_344), inference(resolution, [], [f13100, f5336])).
fof(f5336, plain, ! [X0, X1] : (~ sP167(X0, X1) | (cell78 = X1)), inference(cnf_transformation, [], [f5150])).
fof(f5150, plain, ! [X0, X1] : (((cell88 = X0) & (cell78 = X1)) | ~ sP167(X0, X1)), inference(nnf_transformation, [], [f5130])).
fof(f13100, plain, (sP167(cell1, cell100) | ~ spl177_344), inference(avatar_component_clause, [], [f13098])).
fof(f14209, plain, ~ spl177_343, inference(avatar_contradiction_clause, [], [f14208])).
fof(f14208, plain, ($false | ~ spl177_343), inference(subsumption_resolution, [], [f14205, f7704])).
fof(f7704, plain, ~ (cell77 = cell100), inference(cnf_transformation, [], [f2026])).
fof(f2026, plain, ~ (cell77 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27015)).
fof(f14205, plain, ((cell77 = cell100) | ~ spl177_343), inference(resolution, [], [f13096, f5338])).
fof(f5338, plain, ! [X0, X1] : (~ sP166(X0, X1) | (cell77 = X1)), inference(cnf_transformation, [], [f5151])).
fof(f5151, plain, ! [X0, X1] : (((cell87 = X0) & (cell77 = X1)) | ~ sP166(X0, X1)), inference(nnf_transformation, [], [f5129])).
fof(f13096, plain, (sP166(cell1, cell100) | ~ spl177_343), inference(avatar_component_clause, [], [f13094])).
fof(f14203, plain, ~ spl177_342, inference(avatar_contradiction_clause, [], [f14202])).
fof(f14202, plain, ($false | ~ spl177_342), inference(subsumption_resolution, [], [f14199, f7780])).
fof(f7780, plain, ~ (cell76 = cell100), inference(cnf_transformation, [], [f2102])).
fof(f2102, plain, ~ (cell76 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27091)).
fof(f14199, plain, ((cell76 = cell100) | ~ spl177_342), inference(resolution, [], [f13092, f5340])).
fof(f5340, plain, ! [X0, X1] : (~ sP165(X0, X1) | (cell76 = X1)), inference(cnf_transformation, [], [f5152])).
fof(f5152, plain, ! [X0, X1] : (((cell86 = X0) & (cell76 = X1)) | ~ sP165(X0, X1)), inference(nnf_transformation, [], [f5128])).
fof(f13092, plain, (sP165(cell1, cell100) | ~ spl177_342), inference(avatar_component_clause, [], [f13090])).
fof(f14197, plain, ~ spl177_341, inference(avatar_contradiction_clause, [], [f14196])).
fof(f14196, plain, ($false | ~ spl177_341), inference(subsumption_resolution, [], [f14193, f7855])).
fof(f7855, plain, ~ (cell75 = cell100), inference(cnf_transformation, [], [f2177])).
fof(f2177, plain, ~ (cell75 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27166)).
fof(f14193, plain, ((cell75 = cell100) | ~ spl177_341), inference(resolution, [], [f13088, f5342])).
fof(f5342, plain, ! [X0, X1] : (~ sP164(X0, X1) | (cell75 = X1)), inference(cnf_transformation, [], [f5153])).
fof(f5153, plain, ! [X0, X1] : (((cell85 = X0) & (cell75 = X1)) | ~ sP164(X0, X1)), inference(nnf_transformation, [], [f5127])).
fof(f13088, plain, (sP164(cell1, cell100) | ~ spl177_341), inference(avatar_component_clause, [], [f13086])).
fof(f14191, plain, ~ spl177_340, inference(avatar_contradiction_clause, [], [f14190])).
fof(f14190, plain, ($false | ~ spl177_340), inference(subsumption_resolution, [], [f14187, f7929])).
fof(f7929, plain, ~ (cell74 = cell100), inference(cnf_transformation, [], [f2251])).
fof(f2251, plain, ~ (cell74 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27240)).
fof(f14187, plain, ((cell74 = cell100) | ~ spl177_340), inference(resolution, [], [f13084, f5344])).
fof(f5344, plain, ! [X0, X1] : (~ sP163(X0, X1) | (cell74 = X1)), inference(cnf_transformation, [], [f5154])).
fof(f5154, plain, ! [X0, X1] : (((cell84 = X0) & (cell74 = X1)) | ~ sP163(X0, X1)), inference(nnf_transformation, [], [f5126])).
fof(f13084, plain, (sP163(cell1, cell100) | ~ spl177_340), inference(avatar_component_clause, [], [f13082])).
fof(f14185, plain, ~ spl177_339, inference(avatar_contradiction_clause, [], [f14184])).
fof(f14184, plain, ($false | ~ spl177_339), inference(subsumption_resolution, [], [f14181, f8002])).
fof(f8002, plain, ~ (cell73 = cell100), inference(cnf_transformation, [], [f2324])).
fof(f2324, plain, ~ (cell73 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27313)).
fof(f14181, plain, ((cell73 = cell100) | ~ spl177_339), inference(resolution, [], [f13080, f5346])).
fof(f5346, plain, ! [X0, X1] : (~ sP162(X0, X1) | (cell73 = X1)), inference(cnf_transformation, [], [f5155])).
fof(f5155, plain, ! [X0, X1] : (((cell83 = X0) & (cell73 = X1)) | ~ sP162(X0, X1)), inference(nnf_transformation, [], [f5125])).
fof(f13080, plain, (sP162(cell1, cell100) | ~ spl177_339), inference(avatar_component_clause, [], [f13078])).
fof(f14179, plain, ~ spl177_338, inference(avatar_contradiction_clause, [], [f14178])).
fof(f14178, plain, ($false | ~ spl177_338), inference(subsumption_resolution, [], [f14175, f8145])).
fof(f8145, plain, ~ (cell72 = cell100), inference(cnf_transformation, [], [f2467])).
fof(f2467, plain, ~ (cell72 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27456)).
fof(f14175, plain, ((cell72 = cell100) | ~ spl177_338), inference(resolution, [], [f13076, f5348])).
fof(f5348, plain, ! [X0, X1] : (~ sP161(X0, X1) | (cell72 = X1)), inference(cnf_transformation, [], [f5156])).
fof(f5156, plain, ! [X0, X1] : (((cell82 = X0) & (cell72 = X1)) | ~ sP161(X0, X1)), inference(nnf_transformation, [], [f5124])).
fof(f13076, plain, (sP161(cell1, cell100) | ~ spl177_338), inference(avatar_component_clause, [], [f13074])).
fof(f14173, plain, ~ spl177_337, inference(avatar_contradiction_clause, [], [f14172])).
fof(f14172, plain, ($false | ~ spl177_337), inference(subsumption_resolution, [], [f14169, f8074])).
fof(f8074, plain, ~ (cell71 = cell100), inference(cnf_transformation, [], [f2396])).
fof(f2396, plain, ~ (cell71 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27385)).
fof(f14169, plain, ((cell71 = cell100) | ~ spl177_337), inference(resolution, [], [f13072, f5350])).
fof(f5350, plain, ! [X0, X1] : (~ sP160(X0, X1) | (cell71 = X1)), inference(cnf_transformation, [], [f5157])).
fof(f5157, plain, ! [X0, X1] : (((cell81 = X0) & (cell71 = X1)) | ~ sP160(X0, X1)), inference(nnf_transformation, [], [f5123])).
fof(f13072, plain, (sP160(cell1, cell100) | ~ spl177_337), inference(avatar_component_clause, [], [f13070])).
fof(f14167, plain, ~ spl177_336, inference(avatar_contradiction_clause, [], [f14166])).
fof(f14166, plain, ($false | ~ spl177_336), inference(subsumption_resolution, [], [f14163, f8215])).
fof(f8215, plain, ~ (cell70 = cell100), inference(cnf_transformation, [], [f2537])).
fof(f2537, plain, ~ (cell70 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27526)).
fof(f14163, plain, ((cell70 = cell100) | ~ spl177_336), inference(resolution, [], [f13068, f5352])).
fof(f5352, plain, ! [X0, X1] : (~ sP159(X0, X1) | (cell70 = X1)), inference(cnf_transformation, [], [f5158])).
fof(f5158, plain, ! [X0, X1] : (((cell80 = X0) & (cell70 = X1)) | ~ sP159(X0, X1)), inference(nnf_transformation, [], [f5122])).
fof(f13068, plain, (sP159(cell1, cell100) | ~ spl177_336), inference(avatar_component_clause, [], [f13066])).
fof(f14161, plain, ~ spl177_335, inference(avatar_contradiction_clause, [], [f14160])).
fof(f14160, plain, ($false | ~ spl177_335), inference(subsumption_resolution, [], [f14157, f8284])).
fof(f8284, plain, ~ (cell69 = cell100), inference(cnf_transformation, [], [f2606])).
fof(f2606, plain, ~ (cell69 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27595)).
fof(f14157, plain, ((cell69 = cell100) | ~ spl177_335), inference(resolution, [], [f13064, f5354])).
fof(f5354, plain, ! [X0, X1] : (~ sP158(X0, X1) | (cell69 = X1)), inference(cnf_transformation, [], [f5159])).
fof(f5159, plain, ! [X0, X1] : (((cell79 = X0) & (cell69 = X1)) | ~ sP158(X0, X1)), inference(nnf_transformation, [], [f5121])).
fof(f13064, plain, (sP158(cell1, cell100) | ~ spl177_335), inference(avatar_component_clause, [], [f13062])).
fof(f14155, plain, ~ spl177_334, inference(avatar_contradiction_clause, [], [f14154])).
fof(f14154, plain, ($false | ~ spl177_334), inference(subsumption_resolution, [], [f14151, f8352])).
fof(f8352, plain, ~ (cell68 = cell100), inference(cnf_transformation, [], [f2674])).
fof(f2674, plain, ~ (cell68 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27663)).
fof(f14151, plain, ((cell68 = cell100) | ~ spl177_334), inference(resolution, [], [f13060, f5356])).
fof(f5356, plain, ! [X0, X1] : (~ sP157(X0, X1) | (cell68 = X1)), inference(cnf_transformation, [], [f5160])).
fof(f5160, plain, ! [X0, X1] : (((cell78 = X0) & (cell68 = X1)) | ~ sP157(X0, X1)), inference(nnf_transformation, [], [f5120])).
fof(f13060, plain, (sP157(cell1, cell100) | ~ spl177_334), inference(avatar_component_clause, [], [f13058])).
fof(f14149, plain, ~ spl177_333, inference(avatar_contradiction_clause, [], [f14148])).
fof(f14148, plain, ($false | ~ spl177_333), inference(subsumption_resolution, [], [f14145, f8419])).
fof(f8419, plain, ~ (cell67 = cell100), inference(cnf_transformation, [], [f2741])).
fof(f2741, plain, ~ (cell67 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27730)).
fof(f14145, plain, ((cell67 = cell100) | ~ spl177_333), inference(resolution, [], [f13056, f5358])).
fof(f5358, plain, ! [X0, X1] : (~ sP156(X0, X1) | (cell67 = X1)), inference(cnf_transformation, [], [f5161])).
fof(f5161, plain, ! [X0, X1] : (((cell77 = X0) & (cell67 = X1)) | ~ sP156(X0, X1)), inference(nnf_transformation, [], [f5119])).
fof(f13056, plain, (sP156(cell1, cell100) | ~ spl177_333), inference(avatar_component_clause, [], [f13054])).
fof(f14143, plain, ~ spl177_332, inference(avatar_contradiction_clause, [], [f14142])).
fof(f14142, plain, ($false | ~ spl177_332), inference(subsumption_resolution, [], [f14139, f8485])).
fof(f8485, plain, ~ (cell66 = cell100), inference(cnf_transformation, [], [f2807])).
fof(f2807, plain, ~ (cell66 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27796)).
fof(f14139, plain, ((cell66 = cell100) | ~ spl177_332), inference(resolution, [], [f13052, f5360])).
fof(f5360, plain, ! [X0, X1] : (~ sP155(X0, X1) | (cell66 = X1)), inference(cnf_transformation, [], [f5162])).
fof(f5162, plain, ! [X0, X1] : (((cell76 = X0) & (cell66 = X1)) | ~ sP155(X0, X1)), inference(nnf_transformation, [], [f5118])).
fof(f13052, plain, (sP155(cell1, cell100) | ~ spl177_332), inference(avatar_component_clause, [], [f13050])).
fof(f14137, plain, ~ spl177_331, inference(avatar_contradiction_clause, [], [f14136])).
fof(f14136, plain, ($false | ~ spl177_331), inference(subsumption_resolution, [], [f14133, f8550])).
fof(f8550, plain, ~ (cell65 = cell100), inference(cnf_transformation, [], [f2872])).
fof(f2872, plain, ~ (cell65 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27861)).
fof(f14133, plain, ((cell65 = cell100) | ~ spl177_331), inference(resolution, [], [f13048, f5362])).
fof(f5362, plain, ! [X0, X1] : (~ sP154(X0, X1) | (cell65 = X1)), inference(cnf_transformation, [], [f5163])).
fof(f5163, plain, ! [X0, X1] : (((cell75 = X0) & (cell65 = X1)) | ~ sP154(X0, X1)), inference(nnf_transformation, [], [f5117])).
fof(f13048, plain, (sP154(cell1, cell100) | ~ spl177_331), inference(avatar_component_clause, [], [f13046])).
fof(f14131, plain, ~ spl177_330, inference(avatar_contradiction_clause, [], [f14130])).
fof(f14130, plain, ($false | ~ spl177_330), inference(subsumption_resolution, [], [f14127, f8614])).
fof(f8614, plain, ~ (cell64 = cell100), inference(cnf_transformation, [], [f2936])).
fof(f2936, plain, ~ (cell64 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27925)).
fof(f14127, plain, ((cell64 = cell100) | ~ spl177_330), inference(resolution, [], [f13044, f5364])).
fof(f5364, plain, ! [X0, X1] : (~ sP153(X0, X1) | (cell64 = X1)), inference(cnf_transformation, [], [f5164])).
fof(f5164, plain, ! [X0, X1] : (((cell74 = X0) & (cell64 = X1)) | ~ sP153(X0, X1)), inference(nnf_transformation, [], [f5116])).
fof(f13044, plain, (sP153(cell1, cell100) | ~ spl177_330), inference(avatar_component_clause, [], [f13042])).
fof(f14125, plain, ~ spl177_329, inference(avatar_contradiction_clause, [], [f14124])).
fof(f14124, plain, ($false | ~ spl177_329), inference(subsumption_resolution, [], [f14121, f8677])).
fof(f8677, plain, ~ (cell63 = cell100), inference(cnf_transformation, [], [f2999])).
fof(f2999, plain, ~ (cell63 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27988)).
fof(f14121, plain, ((cell63 = cell100) | ~ spl177_329), inference(resolution, [], [f13040, f5366])).
fof(f5366, plain, ! [X0, X1] : (~ sP152(X0, X1) | (cell63 = X1)), inference(cnf_transformation, [], [f5165])).
fof(f5165, plain, ! [X0, X1] : (((cell73 = X0) & (cell63 = X1)) | ~ sP152(X0, X1)), inference(nnf_transformation, [], [f5115])).
fof(f13040, plain, (sP152(cell1, cell100) | ~ spl177_329), inference(avatar_component_clause, [], [f13038])).
fof(f14119, plain, ~ spl177_328, inference(avatar_contradiction_clause, [], [f14118])).
fof(f14118, plain, ($false | ~ spl177_328), inference(subsumption_resolution, [], [f14115, f8800])).
fof(f8800, plain, ~ (cell62 = cell100), inference(cnf_transformation, [], [f3122])).
fof(f3122, plain, ~ (cell62 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28111)).
fof(f14115, plain, ((cell62 = cell100) | ~ spl177_328), inference(resolution, [], [f13036, f5368])).
fof(f5368, plain, ! [X0, X1] : (~ sP151(X0, X1) | (cell62 = X1)), inference(cnf_transformation, [], [f5166])).
fof(f5166, plain, ! [X0, X1] : (((cell72 = X0) & (cell62 = X1)) | ~ sP151(X0, X1)), inference(nnf_transformation, [], [f5114])).
fof(f13036, plain, (sP151(cell1, cell100) | ~ spl177_328), inference(avatar_component_clause, [], [f13034])).
fof(f14113, plain, ~ spl177_327, inference(avatar_contradiction_clause, [], [f14112])).
fof(f14112, plain, ($false | ~ spl177_327), inference(subsumption_resolution, [], [f14109, f8739])).
fof(f8739, plain, ~ (cell61 = cell100), inference(cnf_transformation, [], [f3061])).
fof(f3061, plain, ~ (cell61 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28050)).
fof(f14109, plain, ((cell61 = cell100) | ~ spl177_327), inference(resolution, [], [f13032, f5370])).
fof(f5370, plain, ! [X0, X1] : (~ sP150(X0, X1) | (cell61 = X1)), inference(cnf_transformation, [], [f5167])).
fof(f5167, plain, ! [X0, X1] : (((cell71 = X0) & (cell61 = X1)) | ~ sP150(X0, X1)), inference(nnf_transformation, [], [f5113])).
fof(f13032, plain, (sP150(cell1, cell100) | ~ spl177_327), inference(avatar_component_clause, [], [f13030])).
fof(f14107, plain, ~ spl177_326, inference(avatar_contradiction_clause, [], [f14106])).
fof(f14106, plain, ($false | ~ spl177_326), inference(subsumption_resolution, [], [f14103, f8860])).
fof(f8860, plain, ~ (cell60 = cell100), inference(cnf_transformation, [], [f3182])).
fof(f3182, plain, ~ (cell60 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28171)).
fof(f14103, plain, ((cell60 = cell100) | ~ spl177_326), inference(resolution, [], [f13028, f5372])).
fof(f5372, plain, ! [X0, X1] : (~ sP149(X0, X1) | (cell60 = X1)), inference(cnf_transformation, [], [f5168])).
fof(f5168, plain, ! [X0, X1] : (((cell70 = X0) & (cell60 = X1)) | ~ sP149(X0, X1)), inference(nnf_transformation, [], [f5112])).
fof(f13028, plain, (sP149(cell1, cell100) | ~ spl177_326), inference(avatar_component_clause, [], [f13026])).
fof(f14101, plain, ~ spl177_325, inference(avatar_contradiction_clause, [], [f14100])).
fof(f14100, plain, ($false | ~ spl177_325), inference(subsumption_resolution, [], [f14097, f8919])).
fof(f8919, plain, ~ (cell59 = cell100), inference(cnf_transformation, [], [f3241])).
fof(f3241, plain, ~ (cell59 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28230)).
fof(f14097, plain, ((cell59 = cell100) | ~ spl177_325), inference(resolution, [], [f13024, f5374])).
fof(f5374, plain, ! [X0, X1] : (~ sP148(X0, X1) | (cell59 = X1)), inference(cnf_transformation, [], [f5169])).
fof(f5169, plain, ! [X0, X1] : (((cell69 = X0) & (cell59 = X1)) | ~ sP148(X0, X1)), inference(nnf_transformation, [], [f5111])).
fof(f13024, plain, (sP148(cell1, cell100) | ~ spl177_325), inference(avatar_component_clause, [], [f13022])).
fof(f14095, plain, ~ spl177_324, inference(avatar_contradiction_clause, [], [f14094])).
fof(f14094, plain, ($false | ~ spl177_324), inference(subsumption_resolution, [], [f14091, f8977])).
fof(f8977, plain, ~ (cell58 = cell100), inference(cnf_transformation, [], [f3299])).
fof(f3299, plain, ~ (cell58 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28288)).
fof(f14091, plain, ((cell58 = cell100) | ~ spl177_324), inference(resolution, [], [f13020, f5376])).
fof(f5376, plain, ! [X0, X1] : (~ sP147(X0, X1) | (cell58 = X1)), inference(cnf_transformation, [], [f5170])).
fof(f5170, plain, ! [X0, X1] : (((cell68 = X0) & (cell58 = X1)) | ~ sP147(X0, X1)), inference(nnf_transformation, [], [f5110])).
fof(f13020, plain, (sP147(cell1, cell100) | ~ spl177_324), inference(avatar_component_clause, [], [f13018])).
fof(f14089, plain, ~ spl177_323, inference(avatar_contradiction_clause, [], [f14088])).
fof(f14088, plain, ($false | ~ spl177_323), inference(subsumption_resolution, [], [f14085, f9034])).
fof(f9034, plain, ~ (cell57 = cell100), inference(cnf_transformation, [], [f3356])).
fof(f3356, plain, ~ (cell57 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28345)).
fof(f14085, plain, ((cell57 = cell100) | ~ spl177_323), inference(resolution, [], [f13016, f5378])).
fof(f5378, plain, ! [X0, X1] : (~ sP146(X0, X1) | (cell57 = X1)), inference(cnf_transformation, [], [f5171])).
fof(f5171, plain, ! [X0, X1] : (((cell67 = X0) & (cell57 = X1)) | ~ sP146(X0, X1)), inference(nnf_transformation, [], [f5109])).
fof(f13016, plain, (sP146(cell1, cell100) | ~ spl177_323), inference(avatar_component_clause, [], [f13014])).
fof(f14083, plain, ~ spl177_322, inference(avatar_contradiction_clause, [], [f14082])).
fof(f14082, plain, ($false | ~ spl177_322), inference(subsumption_resolution, [], [f14079, f9090])).
fof(f9090, plain, ~ (cell56 = cell100), inference(cnf_transformation, [], [f3412])).
fof(f3412, plain, ~ (cell56 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28401)).
fof(f14079, plain, ((cell56 = cell100) | ~ spl177_322), inference(resolution, [], [f13012, f5380])).
fof(f5380, plain, ! [X0, X1] : (~ sP145(X0, X1) | (cell56 = X1)), inference(cnf_transformation, [], [f5172])).
fof(f5172, plain, ! [X0, X1] : (((cell66 = X0) & (cell56 = X1)) | ~ sP145(X0, X1)), inference(nnf_transformation, [], [f5108])).
fof(f13012, plain, (sP145(cell1, cell100) | ~ spl177_322), inference(avatar_component_clause, [], [f13010])).
fof(f14077, plain, ~ spl177_321, inference(avatar_contradiction_clause, [], [f14076])).
fof(f14076, plain, ($false | ~ spl177_321), inference(subsumption_resolution, [], [f14073, f9145])).
fof(f9145, plain, ~ (cell55 = cell100), inference(cnf_transformation, [], [f3467])).
fof(f3467, plain, ~ (cell55 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28456)).
fof(f14073, plain, ((cell55 = cell100) | ~ spl177_321), inference(resolution, [], [f13008, f5382])).
fof(f5382, plain, ! [X0, X1] : (~ sP144(X0, X1) | (cell55 = X1)), inference(cnf_transformation, [], [f5173])).
fof(f5173, plain, ! [X0, X1] : (((cell65 = X0) & (cell55 = X1)) | ~ sP144(X0, X1)), inference(nnf_transformation, [], [f5107])).
fof(f13008, plain, (sP144(cell1, cell100) | ~ spl177_321), inference(avatar_component_clause, [], [f13006])).
fof(f14071, plain, ~ spl177_320, inference(avatar_contradiction_clause, [], [f14070])).
fof(f14070, plain, ($false | ~ spl177_320), inference(subsumption_resolution, [], [f14067, f9199])).
fof(f9199, plain, ~ (cell54 = cell100), inference(cnf_transformation, [], [f3521])).
fof(f3521, plain, ~ (cell54 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28510)).
fof(f14067, plain, ((cell54 = cell100) | ~ spl177_320), inference(resolution, [], [f13004, f5384])).
fof(f5384, plain, ! [X0, X1] : (~ sP143(X0, X1) | (cell54 = X1)), inference(cnf_transformation, [], [f5174])).
fof(f5174, plain, ! [X0, X1] : (((cell64 = X0) & (cell54 = X1)) | ~ sP143(X0, X1)), inference(nnf_transformation, [], [f5106])).
fof(f13004, plain, (sP143(cell1, cell100) | ~ spl177_320), inference(avatar_component_clause, [], [f13002])).
fof(f14065, plain, ~ spl177_319, inference(avatar_contradiction_clause, [], [f14064])).
fof(f14064, plain, ($false | ~ spl177_319), inference(subsumption_resolution, [], [f14061, f9252])).
fof(f9252, plain, ~ (cell53 = cell100), inference(cnf_transformation, [], [f3574])).
fof(f3574, plain, ~ (cell53 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28563)).
fof(f14061, plain, ((cell53 = cell100) | ~ spl177_319), inference(resolution, [], [f13000, f5386])).
fof(f5386, plain, ! [X0, X1] : (~ sP142(X0, X1) | (cell53 = X1)), inference(cnf_transformation, [], [f5175])).
fof(f5175, plain, ! [X0, X1] : (((cell63 = X0) & (cell53 = X1)) | ~ sP142(X0, X1)), inference(nnf_transformation, [], [f5105])).
fof(f13000, plain, (sP142(cell1, cell100) | ~ spl177_319), inference(avatar_component_clause, [], [f12998])).
fof(f14059, plain, ~ spl177_318, inference(avatar_contradiction_clause, [], [f14058])).
fof(f14058, plain, ($false | ~ spl177_318), inference(subsumption_resolution, [], [f14055, f9355])).
fof(f9355, plain, ~ (cell52 = cell100), inference(cnf_transformation, [], [f3677])).
fof(f3677, plain, ~ (cell52 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28666)).
fof(f14055, plain, ((cell52 = cell100) | ~ spl177_318), inference(resolution, [], [f12996, f5388])).
fof(f5388, plain, ! [X0, X1] : (~ sP141(X0, X1) | (cell52 = X1)), inference(cnf_transformation, [], [f5176])).
fof(f5176, plain, ! [X0, X1] : (((cell62 = X0) & (cell52 = X1)) | ~ sP141(X0, X1)), inference(nnf_transformation, [], [f5104])).
fof(f12996, plain, (sP141(cell1, cell100) | ~ spl177_318), inference(avatar_component_clause, [], [f12994])).
fof(f14053, plain, ~ spl177_317, inference(avatar_contradiction_clause, [], [f14052])).
fof(f14052, plain, ($false | ~ spl177_317), inference(subsumption_resolution, [], [f14049, f9304])).
fof(f9304, plain, ~ (cell51 = cell100), inference(cnf_transformation, [], [f3626])).
fof(f3626, plain, ~ (cell51 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28615)).
fof(f14049, plain, ((cell51 = cell100) | ~ spl177_317), inference(resolution, [], [f12992, f5390])).
fof(f5390, plain, ! [X0, X1] : (~ sP140(X0, X1) | (cell51 = X1)), inference(cnf_transformation, [], [f5177])).
fof(f5177, plain, ! [X0, X1] : (((cell61 = X0) & (cell51 = X1)) | ~ sP140(X0, X1)), inference(nnf_transformation, [], [f5103])).
fof(f12992, plain, (sP140(cell1, cell100) | ~ spl177_317), inference(avatar_component_clause, [], [f12990])).
fof(f14047, plain, ~ spl177_316, inference(avatar_contradiction_clause, [], [f14046])).
fof(f14046, plain, ($false | ~ spl177_316), inference(subsumption_resolution, [], [f14043, f9405])).
fof(f9405, plain, ~ (cell50 = cell100), inference(cnf_transformation, [], [f3727])).
fof(f3727, plain, ~ (cell50 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28716)).
fof(f14043, plain, ((cell50 = cell100) | ~ spl177_316), inference(resolution, [], [f12988, f5392])).
fof(f5392, plain, ! [X0, X1] : (~ sP139(X0, X1) | (cell50 = X1)), inference(cnf_transformation, [], [f5178])).
fof(f5178, plain, ! [X0, X1] : (((cell60 = X0) & (cell50 = X1)) | ~ sP139(X0, X1)), inference(nnf_transformation, [], [f5102])).
fof(f12988, plain, (sP139(cell1, cell100) | ~ spl177_316), inference(avatar_component_clause, [], [f12986])).
fof(f14041, plain, ~ spl177_315, inference(avatar_contradiction_clause, [], [f14040])).
fof(f14040, plain, ($false | ~ spl177_315), inference(subsumption_resolution, [], [f14037, f9454])).
fof(f9454, plain, ~ (cell49 = cell100), inference(cnf_transformation, [], [f3776])).
fof(f3776, plain, ~ (cell49 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28765)).
fof(f14037, plain, ((cell49 = cell100) | ~ spl177_315), inference(resolution, [], [f12984, f5394])).
fof(f5394, plain, ! [X0, X1] : (~ sP138(X0, X1) | (cell49 = X1)), inference(cnf_transformation, [], [f5179])).
fof(f5179, plain, ! [X0, X1] : (((cell59 = X0) & (cell49 = X1)) | ~ sP138(X0, X1)), inference(nnf_transformation, [], [f5101])).
fof(f12984, plain, (sP138(cell1, cell100) | ~ spl177_315), inference(avatar_component_clause, [], [f12982])).
fof(f14035, plain, ~ spl177_314, inference(avatar_contradiction_clause, [], [f14034])).
fof(f14034, plain, ($false | ~ spl177_314), inference(subsumption_resolution, [], [f14031, f9502])).
fof(f9502, plain, ~ (cell48 = cell100), inference(cnf_transformation, [], [f3824])).
fof(f3824, plain, ~ (cell48 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28813)).
fof(f14031, plain, ((cell48 = cell100) | ~ spl177_314), inference(resolution, [], [f12980, f5396])).
fof(f5396, plain, ! [X0, X1] : (~ sP137(X0, X1) | (cell48 = X1)), inference(cnf_transformation, [], [f5180])).
fof(f5180, plain, ! [X0, X1] : (((cell58 = X0) & (cell48 = X1)) | ~ sP137(X0, X1)), inference(nnf_transformation, [], [f5100])).
fof(f12980, plain, (sP137(cell1, cell100) | ~ spl177_314), inference(avatar_component_clause, [], [f12978])).
fof(f14029, plain, ~ spl177_313, inference(avatar_contradiction_clause, [], [f14028])).
fof(f14028, plain, ($false | ~ spl177_313), inference(subsumption_resolution, [], [f14025, f9549])).
fof(f9549, plain, ~ (cell47 = cell100), inference(cnf_transformation, [], [f3871])).
fof(f3871, plain, ~ (cell47 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28860)).
fof(f14025, plain, ((cell47 = cell100) | ~ spl177_313), inference(resolution, [], [f12976, f5398])).
fof(f5398, plain, ! [X0, X1] : (~ sP136(X0, X1) | (cell47 = X1)), inference(cnf_transformation, [], [f5181])).
fof(f5181, plain, ! [X0, X1] : (((cell57 = X0) & (cell47 = X1)) | ~ sP136(X0, X1)), inference(nnf_transformation, [], [f5099])).
fof(f12976, plain, (sP136(cell1, cell100) | ~ spl177_313), inference(avatar_component_clause, [], [f12974])).
fof(f14023, plain, ~ spl177_312, inference(avatar_contradiction_clause, [], [f14022])).
fof(f14022, plain, ($false | ~ spl177_312), inference(subsumption_resolution, [], [f14019, f9595])).
fof(f9595, plain, ~ (cell46 = cell100), inference(cnf_transformation, [], [f3917])).
fof(f3917, plain, ~ (cell46 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28906)).
fof(f14019, plain, ((cell46 = cell100) | ~ spl177_312), inference(resolution, [], [f12972, f5400])).
fof(f5400, plain, ! [X0, X1] : (~ sP135(X0, X1) | (cell46 = X1)), inference(cnf_transformation, [], [f5182])).
fof(f5182, plain, ! [X0, X1] : (((cell56 = X0) & (cell46 = X1)) | ~ sP135(X0, X1)), inference(nnf_transformation, [], [f5098])).
fof(f12972, plain, (sP135(cell1, cell100) | ~ spl177_312), inference(avatar_component_clause, [], [f12970])).
fof(f14017, plain, ~ spl177_311, inference(avatar_contradiction_clause, [], [f14016])).
fof(f14016, plain, ($false | ~ spl177_311), inference(subsumption_resolution, [], [f14013, f9640])).
fof(f9640, plain, ~ (cell45 = cell100), inference(cnf_transformation, [], [f3962])).
fof(f3962, plain, ~ (cell45 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28951)).
fof(f14013, plain, ((cell45 = cell100) | ~ spl177_311), inference(resolution, [], [f12968, f5402])).
fof(f5402, plain, ! [X0, X1] : (~ sP134(X0, X1) | (cell45 = X1)), inference(cnf_transformation, [], [f5183])).
fof(f5183, plain, ! [X0, X1] : (((cell55 = X0) & (cell45 = X1)) | ~ sP134(X0, X1)), inference(nnf_transformation, [], [f5097])).
fof(f12968, plain, (sP134(cell1, cell100) | ~ spl177_311), inference(avatar_component_clause, [], [f12966])).
fof(f14011, plain, ~ spl177_310, inference(avatar_contradiction_clause, [], [f14010])).
fof(f14010, plain, ($false | ~ spl177_310), inference(subsumption_resolution, [], [f14007, f9684])).
fof(f9684, plain, ~ (cell44 = cell100), inference(cnf_transformation, [], [f4006])).
fof(f4006, plain, ~ (cell44 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28995)).
fof(f14007, plain, ((cell44 = cell100) | ~ spl177_310), inference(resolution, [], [f12964, f5404])).
fof(f5404, plain, ! [X0, X1] : (~ sP133(X0, X1) | (cell44 = X1)), inference(cnf_transformation, [], [f5184])).
fof(f5184, plain, ! [X0, X1] : (((cell54 = X0) & (cell44 = X1)) | ~ sP133(X0, X1)), inference(nnf_transformation, [], [f5096])).
fof(f12964, plain, (sP133(cell1, cell100) | ~ spl177_310), inference(avatar_component_clause, [], [f12962])).
fof(f14005, plain, ~ spl177_309, inference(avatar_contradiction_clause, [], [f14004])).
fof(f14004, plain, ($false | ~ spl177_309), inference(subsumption_resolution, [], [f14001, f9727])).
fof(f9727, plain, ~ (cell43 = cell100), inference(cnf_transformation, [], [f4049])).
fof(f4049, plain, ~ (cell43 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29038)).
fof(f14001, plain, ((cell43 = cell100) | ~ spl177_309), inference(resolution, [], [f12960, f5406])).
fof(f5406, plain, ! [X0, X1] : (~ sP132(X0, X1) | (cell43 = X1)), inference(cnf_transformation, [], [f5185])).
fof(f5185, plain, ! [X0, X1] : (((cell53 = X0) & (cell43 = X1)) | ~ sP132(X0, X1)), inference(nnf_transformation, [], [f5095])).
fof(f12960, plain, (sP132(cell1, cell100) | ~ spl177_309), inference(avatar_component_clause, [], [f12958])).
fof(f13999, plain, ~ spl177_308, inference(avatar_contradiction_clause, [], [f13998])).
fof(f13998, plain, ($false | ~ spl177_308), inference(subsumption_resolution, [], [f13995, f9810])).
fof(f9810, plain, ~ (cell42 = cell100), inference(cnf_transformation, [], [f4132])).
fof(f4132, plain, ~ (cell42 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29121)).
fof(f13995, plain, ((cell42 = cell100) | ~ spl177_308), inference(resolution, [], [f12956, f5408])).
fof(f5408, plain, ! [X0, X1] : (~ sP131(X0, X1) | (cell42 = X1)), inference(cnf_transformation, [], [f5186])).
fof(f5186, plain, ! [X0, X1] : (((cell52 = X0) & (cell42 = X1)) | ~ sP131(X0, X1)), inference(nnf_transformation, [], [f5094])).
fof(f12956, plain, (sP131(cell1, cell100) | ~ spl177_308), inference(avatar_component_clause, [], [f12954])).
fof(f13993, plain, ~ spl177_307, inference(avatar_contradiction_clause, [], [f13992])).
fof(f13992, plain, ($false | ~ spl177_307), inference(subsumption_resolution, [], [f13989, f9769])).
fof(f9769, plain, ~ (cell41 = cell100), inference(cnf_transformation, [], [f4091])).
fof(f4091, plain, ~ (cell41 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29080)).
fof(f13989, plain, ((cell41 = cell100) | ~ spl177_307), inference(resolution, [], [f12952, f5410])).
fof(f5410, plain, ! [X0, X1] : (~ sP130(X0, X1) | (cell41 = X1)), inference(cnf_transformation, [], [f5187])).
fof(f5187, plain, ! [X0, X1] : (((cell51 = X0) & (cell41 = X1)) | ~ sP130(X0, X1)), inference(nnf_transformation, [], [f5093])).
fof(f12952, plain, (sP130(cell1, cell100) | ~ spl177_307), inference(avatar_component_clause, [], [f12950])).
fof(f13987, plain, ~ spl177_306, inference(avatar_contradiction_clause, [], [f13986])).
fof(f13986, plain, ($false | ~ spl177_306), inference(subsumption_resolution, [], [f13983, f9850])).
fof(f9850, plain, ~ (cell40 = cell100), inference(cnf_transformation, [], [f4172])).
fof(f4172, plain, ~ (cell40 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29161)).
fof(f13983, plain, ((cell40 = cell100) | ~ spl177_306), inference(resolution, [], [f12948, f5412])).
fof(f5412, plain, ! [X0, X1] : (~ sP129(X0, X1) | (cell40 = X1)), inference(cnf_transformation, [], [f5188])).
fof(f5188, plain, ! [X0, X1] : (((cell50 = X0) & (cell40 = X1)) | ~ sP129(X0, X1)), inference(nnf_transformation, [], [f5092])).
fof(f12948, plain, (sP129(cell1, cell100) | ~ spl177_306), inference(avatar_component_clause, [], [f12946])).
fof(f13981, plain, ~ spl177_305, inference(avatar_contradiction_clause, [], [f13980])).
fof(f13980, plain, ($false | ~ spl177_305), inference(subsumption_resolution, [], [f13977, f9889])).
fof(f9889, plain, ~ (cell39 = cell100), inference(cnf_transformation, [], [f4211])).
fof(f4211, plain, ~ (cell39 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29200)).
fof(f13977, plain, ((cell39 = cell100) | ~ spl177_305), inference(resolution, [], [f12944, f5414])).
fof(f5414, plain, ! [X0, X1] : (~ sP128(X0, X1) | (cell39 = X1)), inference(cnf_transformation, [], [f5189])).
fof(f5189, plain, ! [X0, X1] : (((cell49 = X0) & (cell39 = X1)) | ~ sP128(X0, X1)), inference(nnf_transformation, [], [f5091])).
fof(f12944, plain, (sP128(cell1, cell100) | ~ spl177_305), inference(avatar_component_clause, [], [f12942])).
fof(f13975, plain, ~ spl177_304, inference(avatar_contradiction_clause, [], [f13974])).
fof(f13974, plain, ($false | ~ spl177_304), inference(subsumption_resolution, [], [f13971, f9927])).
fof(f9927, plain, ~ (cell38 = cell100), inference(cnf_transformation, [], [f4249])).
fof(f4249, plain, ~ (cell38 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29238)).
fof(f13971, plain, ((cell38 = cell100) | ~ spl177_304), inference(resolution, [], [f12940, f5416])).
fof(f5416, plain, ! [X0, X1] : (~ sP127(X0, X1) | (cell38 = X1)), inference(cnf_transformation, [], [f5190])).
fof(f5190, plain, ! [X0, X1] : (((cell48 = X0) & (cell38 = X1)) | ~ sP127(X0, X1)), inference(nnf_transformation, [], [f5090])).
fof(f12940, plain, (sP127(cell1, cell100) | ~ spl177_304), inference(avatar_component_clause, [], [f12938])).
fof(f13969, plain, ~ spl177_303, inference(avatar_contradiction_clause, [], [f13968])).
fof(f13968, plain, ($false | ~ spl177_303), inference(subsumption_resolution, [], [f13965, f9964])).
fof(f9964, plain, ~ (cell37 = cell100), inference(cnf_transformation, [], [f4286])).
fof(f4286, plain, ~ (cell37 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29275)).
fof(f13965, plain, ((cell37 = cell100) | ~ spl177_303), inference(resolution, [], [f12936, f5418])).
fof(f5418, plain, ! [X0, X1] : (~ sP126(X0, X1) | (cell37 = X1)), inference(cnf_transformation, [], [f5191])).
fof(f5191, plain, ! [X0, X1] : (((cell47 = X0) & (cell37 = X1)) | ~ sP126(X0, X1)), inference(nnf_transformation, [], [f5089])).
fof(f12936, plain, (sP126(cell1, cell100) | ~ spl177_303), inference(avatar_component_clause, [], [f12934])).
fof(f13963, plain, ~ spl177_302, inference(avatar_contradiction_clause, [], [f13962])).
fof(f13962, plain, ($false | ~ spl177_302), inference(subsumption_resolution, [], [f13959, f10000])).
fof(f10000, plain, ~ (cell36 = cell100), inference(cnf_transformation, [], [f4322])).
fof(f4322, plain, ~ (cell36 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29311)).
fof(f13959, plain, ((cell36 = cell100) | ~ spl177_302), inference(resolution, [], [f12932, f5420])).
fof(f5420, plain, ! [X0, X1] : (~ sP125(X0, X1) | (cell36 = X1)), inference(cnf_transformation, [], [f5192])).
fof(f5192, plain, ! [X0, X1] : (((cell46 = X0) & (cell36 = X1)) | ~ sP125(X0, X1)), inference(nnf_transformation, [], [f5088])).
fof(f12932, plain, (sP125(cell1, cell100) | ~ spl177_302), inference(avatar_component_clause, [], [f12930])).
fof(f13957, plain, ~ spl177_301, inference(avatar_contradiction_clause, [], [f13956])).
fof(f13956, plain, ($false | ~ spl177_301), inference(subsumption_resolution, [], [f13953, f10035])).
fof(f10035, plain, ~ (cell35 = cell100), inference(cnf_transformation, [], [f4357])).
fof(f4357, plain, ~ (cell35 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29346)).
fof(f13953, plain, ((cell35 = cell100) | ~ spl177_301), inference(resolution, [], [f12928, f5422])).
fof(f5422, plain, ! [X0, X1] : (~ sP124(X0, X1) | (cell35 = X1)), inference(cnf_transformation, [], [f5193])).
fof(f5193, plain, ! [X0, X1] : (((cell45 = X0) & (cell35 = X1)) | ~ sP124(X0, X1)), inference(nnf_transformation, [], [f5087])).
fof(f12928, plain, (sP124(cell1, cell100) | ~ spl177_301), inference(avatar_component_clause, [], [f12926])).
fof(f13951, plain, ~ spl177_300, inference(avatar_contradiction_clause, [], [f13950])).
fof(f13950, plain, ($false | ~ spl177_300), inference(subsumption_resolution, [], [f13947, f10069])).
fof(f10069, plain, ~ (cell34 = cell100), inference(cnf_transformation, [], [f4391])).
fof(f4391, plain, ~ (cell34 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29380)).
fof(f13947, plain, ((cell34 = cell100) | ~ spl177_300), inference(resolution, [], [f12924, f5424])).
fof(f5424, plain, ! [X0, X1] : (~ sP123(X0, X1) | (cell34 = X1)), inference(cnf_transformation, [], [f5194])).
fof(f5194, plain, ! [X0, X1] : (((cell44 = X0) & (cell34 = X1)) | ~ sP123(X0, X1)), inference(nnf_transformation, [], [f5086])).
fof(f12924, plain, (sP123(cell1, cell100) | ~ spl177_300), inference(avatar_component_clause, [], [f12922])).
fof(f13945, plain, ~ spl177_299, inference(avatar_contradiction_clause, [], [f13944])).
fof(f13944, plain, ($false | ~ spl177_299), inference(subsumption_resolution, [], [f13941, f10102])).
fof(f10102, plain, ~ (cell33 = cell100), inference(cnf_transformation, [], [f4424])).
fof(f4424, plain, ~ (cell33 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29413)).
fof(f13941, plain, ((cell33 = cell100) | ~ spl177_299), inference(resolution, [], [f12920, f5426])).
fof(f5426, plain, ! [X0, X1] : (~ sP122(X0, X1) | (cell33 = X1)), inference(cnf_transformation, [], [f5195])).
fof(f5195, plain, ! [X0, X1] : (((cell43 = X0) & (cell33 = X1)) | ~ sP122(X0, X1)), inference(nnf_transformation, [], [f5085])).
fof(f12920, plain, (sP122(cell1, cell100) | ~ spl177_299), inference(avatar_component_clause, [], [f12918])).
fof(f13939, plain, ~ spl177_298, inference(avatar_contradiction_clause, [], [f13938])).
fof(f13938, plain, ($false | ~ spl177_298), inference(subsumption_resolution, [], [f13935, f10165])).
fof(f10165, plain, ~ (cell32 = cell100), inference(cnf_transformation, [], [f4487])).
fof(f4487, plain, ~ (cell32 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29476)).
fof(f13935, plain, ((cell32 = cell100) | ~ spl177_298), inference(resolution, [], [f12916, f5428])).
fof(f5428, plain, ! [X0, X1] : (~ sP121(X0, X1) | (cell32 = X1)), inference(cnf_transformation, [], [f5196])).
fof(f5196, plain, ! [X0, X1] : (((cell42 = X0) & (cell32 = X1)) | ~ sP121(X0, X1)), inference(nnf_transformation, [], [f5084])).
fof(f12916, plain, (sP121(cell1, cell100) | ~ spl177_298), inference(avatar_component_clause, [], [f12914])).
fof(f13933, plain, ~ spl177_297, inference(avatar_contradiction_clause, [], [f13932])).
fof(f13932, plain, ($false | ~ spl177_297), inference(subsumption_resolution, [], [f13929, f10134])).
fof(f10134, plain, ~ (cell31 = cell100), inference(cnf_transformation, [], [f4456])).
fof(f4456, plain, ~ (cell31 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29445)).
fof(f13929, plain, ((cell31 = cell100) | ~ spl177_297), inference(resolution, [], [f12912, f5430])).
fof(f5430, plain, ! [X0, X1] : (~ sP120(X0, X1) | (cell31 = X1)), inference(cnf_transformation, [], [f5197])).
fof(f5197, plain, ! [X0, X1] : (((cell41 = X0) & (cell31 = X1)) | ~ sP120(X0, X1)), inference(nnf_transformation, [], [f5083])).
fof(f12912, plain, (sP120(cell1, cell100) | ~ spl177_297), inference(avatar_component_clause, [], [f12910])).
fof(f13927, plain, ~ spl177_296, inference(avatar_contradiction_clause, [], [f13926])).
fof(f13926, plain, ($false | ~ spl177_296), inference(subsumption_resolution, [], [f13923, f10195])).
fof(f10195, plain, ~ (cell30 = cell100), inference(cnf_transformation, [], [f4517])).
fof(f4517, plain, ~ (cell30 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29506)).
fof(f13923, plain, ((cell30 = cell100) | ~ spl177_296), inference(resolution, [], [f12908, f5432])).
fof(f5432, plain, ! [X0, X1] : (~ sP119(X0, X1) | (cell30 = X1)), inference(cnf_transformation, [], [f5198])).
fof(f5198, plain, ! [X0, X1] : (((cell40 = X0) & (cell30 = X1)) | ~ sP119(X0, X1)), inference(nnf_transformation, [], [f5082])).
fof(f12908, plain, (sP119(cell1, cell100) | ~ spl177_296), inference(avatar_component_clause, [], [f12906])).
fof(f13921, plain, ~ spl177_295, inference(avatar_contradiction_clause, [], [f13920])).
fof(f13920, plain, ($false | ~ spl177_295), inference(subsumption_resolution, [], [f13917, f10224])).
fof(f10224, plain, ~ (cell29 = cell100), inference(cnf_transformation, [], [f4546])).
fof(f4546, plain, ~ (cell29 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29535)).
fof(f13917, plain, ((cell29 = cell100) | ~ spl177_295), inference(resolution, [], [f12904, f5434])).
fof(f5434, plain, ! [X0, X1] : (~ sP118(X0, X1) | (cell29 = X1)), inference(cnf_transformation, [], [f5199])).
fof(f5199, plain, ! [X0, X1] : (((cell39 = X0) & (cell29 = X1)) | ~ sP118(X0, X1)), inference(nnf_transformation, [], [f5081])).
fof(f12904, plain, (sP118(cell1, cell100) | ~ spl177_295), inference(avatar_component_clause, [], [f12902])).
fof(f13915, plain, ~ spl177_294, inference(avatar_contradiction_clause, [], [f13914])).
fof(f13914, plain, ($false | ~ spl177_294), inference(subsumption_resolution, [], [f13911, f10252])).
fof(f10252, plain, ~ (cell28 = cell100), inference(cnf_transformation, [], [f4574])).
fof(f4574, plain, ~ (cell28 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29563)).
fof(f13911, plain, ((cell28 = cell100) | ~ spl177_294), inference(resolution, [], [f12900, f5436])).
fof(f5436, plain, ! [X0, X1] : (~ sP117(X0, X1) | (cell28 = X1)), inference(cnf_transformation, [], [f5200])).
fof(f5200, plain, ! [X0, X1] : (((cell38 = X0) & (cell28 = X1)) | ~ sP117(X0, X1)), inference(nnf_transformation, [], [f5080])).
fof(f12900, plain, (sP117(cell1, cell100) | ~ spl177_294), inference(avatar_component_clause, [], [f12898])).
fof(f13909, plain, ~ spl177_293, inference(avatar_contradiction_clause, [], [f13908])).
fof(f13908, plain, ($false | ~ spl177_293), inference(subsumption_resolution, [], [f13905, f10279])).
fof(f10279, plain, ~ (cell27 = cell100), inference(cnf_transformation, [], [f4601])).
fof(f4601, plain, ~ (cell27 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29590)).
fof(f13905, plain, ((cell27 = cell100) | ~ spl177_293), inference(resolution, [], [f12896, f5438])).
fof(f5438, plain, ! [X0, X1] : (~ sP116(X0, X1) | (cell27 = X1)), inference(cnf_transformation, [], [f5201])).
fof(f5201, plain, ! [X0, X1] : (((cell37 = X0) & (cell27 = X1)) | ~ sP116(X0, X1)), inference(nnf_transformation, [], [f5079])).
fof(f12896, plain, (sP116(cell1, cell100) | ~ spl177_293), inference(avatar_component_clause, [], [f12894])).
fof(f13903, plain, ~ spl177_292, inference(avatar_contradiction_clause, [], [f13902])).
fof(f13902, plain, ($false | ~ spl177_292), inference(subsumption_resolution, [], [f13899, f10305])).
fof(f10305, plain, ~ (cell26 = cell100), inference(cnf_transformation, [], [f4627])).
fof(f4627, plain, ~ (cell26 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29616)).
fof(f13899, plain, ((cell26 = cell100) | ~ spl177_292), inference(resolution, [], [f12892, f5440])).
fof(f5440, plain, ! [X0, X1] : (~ sP115(X0, X1) | (cell26 = X1)), inference(cnf_transformation, [], [f5202])).
fof(f5202, plain, ! [X0, X1] : (((cell36 = X0) & (cell26 = X1)) | ~ sP115(X0, X1)), inference(nnf_transformation, [], [f5078])).
fof(f12892, plain, (sP115(cell1, cell100) | ~ spl177_292), inference(avatar_component_clause, [], [f12890])).
fof(f13897, plain, ~ spl177_291, inference(avatar_contradiction_clause, [], [f13896])).
fof(f13896, plain, ($false | ~ spl177_291), inference(subsumption_resolution, [], [f13893, f10330])).
fof(f10330, plain, ~ (cell25 = cell100), inference(cnf_transformation, [], [f4652])).
fof(f4652, plain, ~ (cell25 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29641)).
fof(f13893, plain, ((cell25 = cell100) | ~ spl177_291), inference(resolution, [], [f12888, f5442])).
fof(f5442, plain, ! [X0, X1] : (~ sP114(X0, X1) | (cell25 = X1)), inference(cnf_transformation, [], [f5203])).
fof(f5203, plain, ! [X0, X1] : (((cell35 = X0) & (cell25 = X1)) | ~ sP114(X0, X1)), inference(nnf_transformation, [], [f5077])).
fof(f12888, plain, (sP114(cell1, cell100) | ~ spl177_291), inference(avatar_component_clause, [], [f12886])).
fof(f13891, plain, ~ spl177_290, inference(avatar_contradiction_clause, [], [f13890])).
fof(f13890, plain, ($false | ~ spl177_290), inference(subsumption_resolution, [], [f13887, f10354])).
fof(f10354, plain, ~ (cell24 = cell100), inference(cnf_transformation, [], [f4676])).
fof(f4676, plain, ~ (cell24 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29665)).
fof(f13887, plain, ((cell24 = cell100) | ~ spl177_290), inference(resolution, [], [f12884, f5444])).
fof(f5444, plain, ! [X0, X1] : (~ sP113(X0, X1) | (cell24 = X1)), inference(cnf_transformation, [], [f5204])).
fof(f5204, plain, ! [X0, X1] : (((cell34 = X0) & (cell24 = X1)) | ~ sP113(X0, X1)), inference(nnf_transformation, [], [f5076])).
fof(f12884, plain, (sP113(cell1, cell100) | ~ spl177_290), inference(avatar_component_clause, [], [f12882])).
fof(f13885, plain, ~ spl177_289, inference(avatar_contradiction_clause, [], [f13884])).
fof(f13884, plain, ($false | ~ spl177_289), inference(subsumption_resolution, [], [f13881, f10377])).
fof(f10377, plain, ~ (cell23 = cell100), inference(cnf_transformation, [], [f4699])).
fof(f4699, plain, ~ (cell23 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29688)).
fof(f13881, plain, ((cell23 = cell100) | ~ spl177_289), inference(resolution, [], [f12880, f5446])).
fof(f5446, plain, ! [X0, X1] : (~ sP112(X0, X1) | (cell23 = X1)), inference(cnf_transformation, [], [f5205])).
fof(f5205, plain, ! [X0, X1] : (((cell33 = X0) & (cell23 = X1)) | ~ sP112(X0, X1)), inference(nnf_transformation, [], [f5075])).
fof(f12880, plain, (sP112(cell1, cell100) | ~ spl177_289), inference(avatar_component_clause, [], [f12878])).
fof(f13879, plain, ~ spl177_288, inference(avatar_contradiction_clause, [], [f13878])).
fof(f13878, plain, ($false | ~ spl177_288), inference(subsumption_resolution, [], [f13875, f10420])).
fof(f10420, plain, ~ (cell22 = cell100), inference(cnf_transformation, [], [f4742])).
fof(f4742, plain, ~ (cell22 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29731)).
fof(f13875, plain, ((cell22 = cell100) | ~ spl177_288), inference(resolution, [], [f12876, f5448])).
fof(f5448, plain, ! [X0, X1] : (~ sP111(X0, X1) | (cell22 = X1)), inference(cnf_transformation, [], [f5206])).
fof(f5206, plain, ! [X0, X1] : (((cell32 = X0) & (cell22 = X1)) | ~ sP111(X0, X1)), inference(nnf_transformation, [], [f5074])).
fof(f12876, plain, (sP111(cell1, cell100) | ~ spl177_288), inference(avatar_component_clause, [], [f12874])).
fof(f13873, plain, ~ spl177_287, inference(avatar_contradiction_clause, [], [f13872])).
fof(f13872, plain, ($false | ~ spl177_287), inference(subsumption_resolution, [], [f13869, f10399])).
fof(f10399, plain, ~ (cell21 = cell100), inference(cnf_transformation, [], [f4721])).
fof(f4721, plain, ~ (cell21 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29710)).
fof(f13869, plain, ((cell21 = cell100) | ~ spl177_287), inference(resolution, [], [f12872, f5450])).
fof(f5450, plain, ! [X0, X1] : (~ sP110(X0, X1) | (cell21 = X1)), inference(cnf_transformation, [], [f5207])).
fof(f5207, plain, ! [X0, X1] : (((cell31 = X0) & (cell21 = X1)) | ~ sP110(X0, X1)), inference(nnf_transformation, [], [f5073])).
fof(f12872, plain, (sP110(cell1, cell100) | ~ spl177_287), inference(avatar_component_clause, [], [f12870])).
fof(f13867, plain, ~ spl177_286, inference(avatar_contradiction_clause, [], [f13866])).
fof(f13866, plain, ($false | ~ spl177_286), inference(subsumption_resolution, [], [f13863, f10440])).
fof(f10440, plain, ~ (cell20 = cell100), inference(cnf_transformation, [], [f4762])).
fof(f4762, plain, ~ (cell20 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29751)).
fof(f13863, plain, ((cell20 = cell100) | ~ spl177_286), inference(resolution, [], [f12868, f5452])).
fof(f5452, plain, ! [X0, X1] : (~ sP109(X0, X1) | (cell20 = X1)), inference(cnf_transformation, [], [f5208])).
fof(f5208, plain, ! [X0, X1] : (((cell30 = X0) & (cell20 = X1)) | ~ sP109(X0, X1)), inference(nnf_transformation, [], [f5072])).
fof(f12868, plain, (sP109(cell1, cell100) | ~ spl177_286), inference(avatar_component_clause, [], [f12866])).
fof(f13861, plain, ~ spl177_285, inference(avatar_contradiction_clause, [], [f13860])).
fof(f13860, plain, ($false | ~ spl177_285), inference(subsumption_resolution, [], [f13857, f10459])).
fof(f10459, plain, ~ (cell19 = cell100), inference(cnf_transformation, [], [f4781])).
fof(f4781, plain, ~ (cell19 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29770)).
fof(f13857, plain, ((cell19 = cell100) | ~ spl177_285), inference(resolution, [], [f12864, f5454])).
fof(f5454, plain, ! [X0, X1] : (~ sP108(X0, X1) | (cell19 = X1)), inference(cnf_transformation, [], [f5209])).
fof(f5209, plain, ! [X0, X1] : (((cell29 = X0) & (cell19 = X1)) | ~ sP108(X0, X1)), inference(nnf_transformation, [], [f5071])).
fof(f12864, plain, (sP108(cell1, cell100) | ~ spl177_285), inference(avatar_component_clause, [], [f12862])).
fof(f13855, plain, ~ spl177_284, inference(avatar_contradiction_clause, [], [f13854])).
fof(f13854, plain, ($false | ~ spl177_284), inference(subsumption_resolution, [], [f13851, f10477])).
fof(f10477, plain, ~ (cell18 = cell100), inference(cnf_transformation, [], [f4799])).
fof(f4799, plain, ~ (cell18 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29788)).
fof(f13851, plain, ((cell18 = cell100) | ~ spl177_284), inference(resolution, [], [f12860, f5456])).
fof(f5456, plain, ! [X0, X1] : (~ sP107(X0, X1) | (cell18 = X1)), inference(cnf_transformation, [], [f5210])).
fof(f5210, plain, ! [X0, X1] : (((cell28 = X0) & (cell18 = X1)) | ~ sP107(X0, X1)), inference(nnf_transformation, [], [f5070])).
fof(f12860, plain, (sP107(cell1, cell100) | ~ spl177_284), inference(avatar_component_clause, [], [f12858])).
fof(f13849, plain, ~ spl177_283, inference(avatar_contradiction_clause, [], [f13848])).
fof(f13848, plain, ($false | ~ spl177_283), inference(subsumption_resolution, [], [f13845, f10494])).
fof(f10494, plain, ~ (cell17 = cell100), inference(cnf_transformation, [], [f4816])).
fof(f4816, plain, ~ (cell17 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29805)).
fof(f13845, plain, ((cell17 = cell100) | ~ spl177_283), inference(resolution, [], [f12856, f5458])).
fof(f5458, plain, ! [X0, X1] : (~ sP106(X0, X1) | (cell17 = X1)), inference(cnf_transformation, [], [f5211])).
fof(f5211, plain, ! [X0, X1] : (((cell27 = X0) & (cell17 = X1)) | ~ sP106(X0, X1)), inference(nnf_transformation, [], [f5069])).
fof(f12856, plain, (sP106(cell1, cell100) | ~ spl177_283), inference(avatar_component_clause, [], [f12854])).
fof(f13843, plain, ~ spl177_282, inference(avatar_contradiction_clause, [], [f13842])).
fof(f13842, plain, ($false | ~ spl177_282), inference(subsumption_resolution, [], [f13839, f10510])).
fof(f10510, plain, ~ (cell16 = cell100), inference(cnf_transformation, [], [f4832])).
fof(f4832, plain, ~ (cell16 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29821)).
fof(f13839, plain, ((cell16 = cell100) | ~ spl177_282), inference(resolution, [], [f12852, f5460])).
fof(f5460, plain, ! [X0, X1] : (~ sP105(X0, X1) | (cell16 = X1)), inference(cnf_transformation, [], [f5212])).
fof(f5212, plain, ! [X0, X1] : (((cell26 = X0) & (cell16 = X1)) | ~ sP105(X0, X1)), inference(nnf_transformation, [], [f5068])).
fof(f12852, plain, (sP105(cell1, cell100) | ~ spl177_282), inference(avatar_component_clause, [], [f12850])).
fof(f13837, plain, ~ spl177_281, inference(avatar_contradiction_clause, [], [f13836])).
fof(f13836, plain, ($false | ~ spl177_281), inference(subsumption_resolution, [], [f13833, f10525])).
fof(f10525, plain, ~ (cell15 = cell100), inference(cnf_transformation, [], [f4847])).
fof(f4847, plain, ~ (cell15 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29836)).
fof(f13833, plain, ((cell15 = cell100) | ~ spl177_281), inference(resolution, [], [f12848, f5462])).
fof(f5462, plain, ! [X0, X1] : (~ sP104(X0, X1) | (cell15 = X1)), inference(cnf_transformation, [], [f5213])).
fof(f5213, plain, ! [X0, X1] : (((cell25 = X0) & (cell15 = X1)) | ~ sP104(X0, X1)), inference(nnf_transformation, [], [f5067])).
fof(f12848, plain, (sP104(cell1, cell100) | ~ spl177_281), inference(avatar_component_clause, [], [f12846])).
fof(f13831, plain, ~ spl177_280, inference(avatar_contradiction_clause, [], [f13830])).
fof(f13830, plain, ($false | ~ spl177_280), inference(subsumption_resolution, [], [f13827, f10539])).
fof(f10539, plain, ~ (cell14 = cell100), inference(cnf_transformation, [], [f4861])).
fof(f4861, plain, ~ (cell14 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29850)).
fof(f13827, plain, ((cell14 = cell100) | ~ spl177_280), inference(resolution, [], [f12844, f5464])).
fof(f5464, plain, ! [X0, X1] : (~ sP103(X0, X1) | (cell14 = X1)), inference(cnf_transformation, [], [f5214])).
fof(f5214, plain, ! [X0, X1] : (((cell24 = X0) & (cell14 = X1)) | ~ sP103(X0, X1)), inference(nnf_transformation, [], [f5066])).
fof(f12844, plain, (sP103(cell1, cell100) | ~ spl177_280), inference(avatar_component_clause, [], [f12842])).
fof(f13825, plain, ~ spl177_279, inference(avatar_contradiction_clause, [], [f13824])).
fof(f13824, plain, ($false | ~ spl177_279), inference(subsumption_resolution, [], [f13821, f10552])).
fof(f10552, plain, ~ (cell13 = cell100), inference(cnf_transformation, [], [f4874])).
fof(f4874, plain, ~ (cell13 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29863)).
fof(f13821, plain, ((cell13 = cell100) | ~ spl177_279), inference(resolution, [], [f12840, f5466])).
fof(f5466, plain, ! [X0, X1] : (~ sP102(X0, X1) | (cell13 = X1)), inference(cnf_transformation, [], [f5215])).
fof(f5215, plain, ! [X0, X1] : (((cell23 = X0) & (cell13 = X1)) | ~ sP102(X0, X1)), inference(nnf_transformation, [], [f5065])).
fof(f12840, plain, (sP102(cell1, cell100) | ~ spl177_279), inference(avatar_component_clause, [], [f12838])).
fof(f13819, plain, ~ spl177_278, inference(avatar_contradiction_clause, [], [f13818])).
fof(f13818, plain, ($false | ~ spl177_278), inference(subsumption_resolution, [], [f13815, f10575])).
fof(f10575, plain, ~ (cell12 = cell100), inference(cnf_transformation, [], [f4897])).
fof(f4897, plain, ~ (cell12 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29886)).
fof(f13815, plain, ((cell12 = cell100) | ~ spl177_278), inference(resolution, [], [f12836, f5468])).
fof(f5468, plain, ! [X0, X1] : (~ sP101(X0, X1) | (cell12 = X1)), inference(cnf_transformation, [], [f5216])).
fof(f5216, plain, ! [X0, X1] : (((cell22 = X0) & (cell12 = X1)) | ~ sP101(X0, X1)), inference(nnf_transformation, [], [f5064])).
fof(f12836, plain, (sP101(cell1, cell100) | ~ spl177_278), inference(avatar_component_clause, [], [f12834])).
fof(f13813, plain, ~ spl177_277, inference(avatar_contradiction_clause, [], [f13812])).
fof(f13812, plain, ($false | ~ spl177_277), inference(subsumption_resolution, [], [f13809, f10564])).
fof(f10564, plain, ~ (cell11 = cell100), inference(cnf_transformation, [], [f4886])).
fof(f4886, plain, ~ (cell11 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29875)).
fof(f13809, plain, ((cell11 = cell100) | ~ spl177_277), inference(resolution, [], [f12832, f5470])).
fof(f5470, plain, ! [X0, X1] : (~ sP100(X0, X1) | (cell11 = X1)), inference(cnf_transformation, [], [f5217])).
fof(f5217, plain, ! [X0, X1] : (((cell21 = X0) & (cell11 = X1)) | ~ sP100(X0, X1)), inference(nnf_transformation, [], [f5063])).
fof(f12832, plain, (sP100(cell1, cell100) | ~ spl177_277), inference(avatar_component_clause, [], [f12830])).
fof(f13807, plain, ~ spl177_276, inference(avatar_contradiction_clause, [], [f13806])).
fof(f13806, plain, ($false | ~ spl177_276), inference(subsumption_resolution, [], [f13803, f10585])).
fof(f10585, plain, ~ (cell10 = cell100), inference(cnf_transformation, [], [f4907])).
fof(f4907, plain, ~ (cell10 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29896)).
fof(f13803, plain, ((cell10 = cell100) | ~ spl177_276), inference(resolution, [], [f12828, f5472])).
fof(f5472, plain, ! [X0, X1] : (~ sP99(X0, X1) | (cell10 = X1)), inference(cnf_transformation, [], [f5218])).
fof(f5218, plain, ! [X0, X1] : (((cell20 = X0) & (cell10 = X1)) | ~ sP99(X0, X1)), inference(nnf_transformation, [], [f5062])).
fof(f12828, plain, (sP99(cell1, cell100) | ~ spl177_276), inference(avatar_component_clause, [], [f12826])).
fof(f13801, plain, ~ spl177_275, inference(avatar_contradiction_clause, [], [f13800])).
fof(f13800, plain, ($false | ~ spl177_275), inference(subsumption_resolution, [], [f13797, f10594])).
fof(f10594, plain, ~ (cell9 = cell100), inference(cnf_transformation, [], [f4916])).
fof(f4916, plain, ~ (cell9 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29905)).
fof(f13797, plain, ((cell9 = cell100) | ~ spl177_275), inference(resolution, [], [f12824, f5474])).
fof(f5474, plain, ! [X0, X1] : (~ sP98(X0, X1) | (cell9 = X1)), inference(cnf_transformation, [], [f5219])).
fof(f5219, plain, ! [X0, X1] : (((cell19 = X0) & (cell9 = X1)) | ~ sP98(X0, X1)), inference(nnf_transformation, [], [f5061])).
fof(f12824, plain, (sP98(cell1, cell100) | ~ spl177_275), inference(avatar_component_clause, [], [f12822])).
fof(f13795, plain, ~ spl177_274, inference(avatar_contradiction_clause, [], [f13794])).
fof(f13794, plain, ($false | ~ spl177_274), inference(subsumption_resolution, [], [f13791, f10602])).
fof(f10602, plain, ~ (cell8 = cell100), inference(cnf_transformation, [], [f4924])).
fof(f4924, plain, ~ (cell8 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29913)).
fof(f13791, plain, ((cell8 = cell100) | ~ spl177_274), inference(resolution, [], [f12820, f5476])).
fof(f5476, plain, ! [X0, X1] : (~ sP97(X0, X1) | (cell8 = X1)), inference(cnf_transformation, [], [f5220])).
fof(f5220, plain, ! [X0, X1] : (((cell18 = X0) & (cell8 = X1)) | ~ sP97(X0, X1)), inference(nnf_transformation, [], [f5060])).
fof(f12820, plain, (sP97(cell1, cell100) | ~ spl177_274), inference(avatar_component_clause, [], [f12818])).
fof(f13789, plain, ~ spl177_273, inference(avatar_contradiction_clause, [], [f13788])).
fof(f13788, plain, ($false | ~ spl177_273), inference(subsumption_resolution, [], [f13785, f10609])).
fof(f10609, plain, ~ (cell7 = cell100), inference(cnf_transformation, [], [f4931])).
fof(f4931, plain, ~ (cell7 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29920)).
fof(f13785, plain, ((cell7 = cell100) | ~ spl177_273), inference(resolution, [], [f12816, f5478])).
fof(f5478, plain, ! [X0, X1] : (~ sP96(X0, X1) | (cell7 = X1)), inference(cnf_transformation, [], [f5221])).
fof(f5221, plain, ! [X0, X1] : (((cell17 = X0) & (cell7 = X1)) | ~ sP96(X0, X1)), inference(nnf_transformation, [], [f5059])).
fof(f12816, plain, (sP96(cell1, cell100) | ~ spl177_273), inference(avatar_component_clause, [], [f12814])).
fof(f13783, plain, ~ spl177_272, inference(avatar_contradiction_clause, [], [f13782])).
fof(f13782, plain, ($false | ~ spl177_272), inference(subsumption_resolution, [], [f13779, f10615])).
fof(f10615, plain, ~ (cell6 = cell100), inference(cnf_transformation, [], [f4937])).
fof(f4937, plain, ~ (cell6 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29926)).
fof(f13779, plain, ((cell6 = cell100) | ~ spl177_272), inference(resolution, [], [f12812, f5480])).
fof(f5480, plain, ! [X0, X1] : (~ sP95(X0, X1) | (cell6 = X1)), inference(cnf_transformation, [], [f5222])).
fof(f5222, plain, ! [X0, X1] : (((cell16 = X0) & (cell6 = X1)) | ~ sP95(X0, X1)), inference(nnf_transformation, [], [f5058])).
fof(f12812, plain, (sP95(cell1, cell100) | ~ spl177_272), inference(avatar_component_clause, [], [f12810])).
fof(f13777, plain, ~ spl177_271, inference(avatar_contradiction_clause, [], [f13776])).
fof(f13776, plain, ($false | ~ spl177_271), inference(subsumption_resolution, [], [f13773, f10620])).
fof(f10620, plain, ~ (cell5 = cell100), inference(cnf_transformation, [], [f4942])).
fof(f4942, plain, ~ (cell5 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29931)).
fof(f13773, plain, ((cell5 = cell100) | ~ spl177_271), inference(resolution, [], [f12808, f5482])).
fof(f5482, plain, ! [X0, X1] : (~ sP94(X0, X1) | (cell5 = X1)), inference(cnf_transformation, [], [f5223])).
fof(f5223, plain, ! [X0, X1] : (((cell15 = X0) & (cell5 = X1)) | ~ sP94(X0, X1)), inference(nnf_transformation, [], [f5057])).
fof(f12808, plain, (sP94(cell1, cell100) | ~ spl177_271), inference(avatar_component_clause, [], [f12806])).
fof(f13771, plain, ~ spl177_270, inference(avatar_contradiction_clause, [], [f13770])).
fof(f13770, plain, ($false | ~ spl177_270), inference(subsumption_resolution, [], [f13767, f10624])).
fof(f10624, plain, ~ (cell4 = cell100), inference(cnf_transformation, [], [f4946])).
fof(f4946, plain, ~ (cell4 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29935)).
fof(f13767, plain, ((cell4 = cell100) | ~ spl177_270), inference(resolution, [], [f12804, f5484])).
fof(f5484, plain, ! [X0, X1] : (~ sP93(X0, X1) | (cell4 = X1)), inference(cnf_transformation, [], [f5224])).
fof(f5224, plain, ! [X0, X1] : (((cell14 = X0) & (cell4 = X1)) | ~ sP93(X0, X1)), inference(nnf_transformation, [], [f5056])).
fof(f12804, plain, (sP93(cell1, cell100) | ~ spl177_270), inference(avatar_component_clause, [], [f12802])).
fof(f13765, plain, ~ spl177_269, inference(avatar_contradiction_clause, [], [f13764])).
fof(f13764, plain, ($false | ~ spl177_269), inference(subsumption_resolution, [], [f13761, f10627])).
fof(f10627, plain, ~ (cell3 = cell100), inference(cnf_transformation, [], [f4949])).
fof(f4949, plain, ~ (cell3 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29938)).
fof(f13761, plain, ((cell3 = cell100) | ~ spl177_269), inference(resolution, [], [f12800, f5486])).
fof(f5486, plain, ! [X0, X1] : (~ sP92(X0, X1) | (cell3 = X1)), inference(cnf_transformation, [], [f5225])).
fof(f5225, plain, ! [X0, X1] : (((cell13 = X0) & (cell3 = X1)) | ~ sP92(X0, X1)), inference(nnf_transformation, [], [f5055])).
fof(f12800, plain, (sP92(cell1, cell100) | ~ spl177_269), inference(avatar_component_clause, [], [f12798])).
fof(f13759, plain, ~ spl177_268, inference(avatar_contradiction_clause, [], [f13758])).
fof(f13758, plain, ($false | ~ spl177_268), inference(subsumption_resolution, [], [f13755, f10629])).
fof(f10629, plain, ~ (cell2 = cell100), inference(cnf_transformation, [], [f4951])).
fof(f4951, plain, ~ (cell2 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29940)).
fof(f13755, plain, ((cell2 = cell100) | ~ spl177_268), inference(resolution, [], [f12796, f5488])).
fof(f5488, plain, ! [X0, X1] : (~ sP91(X0, X1) | (cell2 = X1)), inference(cnf_transformation, [], [f5226])).
fof(f5226, plain, ! [X0, X1] : (((cell12 = X0) & (cell2 = X1)) | ~ sP91(X0, X1)), inference(nnf_transformation, [], [f5054])).
fof(f12796, plain, (sP91(cell1, cell100) | ~ spl177_268), inference(avatar_component_clause, [], [f12794])).
fof(f13753, plain, ~ spl177_267, inference(avatar_contradiction_clause, [], [f13752])).
fof(f13752, plain, ($false | ~ spl177_267), inference(subsumption_resolution, [], [f13749, f10630])).
fof(f13749, plain, ((cell1 = cell100) | ~ spl177_267), inference(resolution, [], [f12792, f5490])).
fof(f5490, plain, ! [X0, X1] : (~ sP90(X0, X1) | (cell1 = X1)), inference(cnf_transformation, [], [f5227])).
fof(f5227, plain, ! [X0, X1] : (((cell11 = X0) & (cell1 = X1)) | ~ sP90(X0, X1)), inference(nnf_transformation, [], [f5053])).
fof(f12792, plain, (sP90(cell1, cell100) | ~ spl177_267), inference(avatar_component_clause, [], [f12790])).
fof(f13747, plain, ~ spl177_266, inference(avatar_contradiction_clause, [], [f13746])).
fof(f13746, plain, ($false | ~ spl177_266), inference(subsumption_resolution, [], [f13743, f6070])).
fof(f6070, plain, ~ (cell96 = cell100), inference(cnf_transformation, [], [f392])).
fof(f392, plain, ~ (cell96 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25381)).
fof(f13743, plain, ((cell96 = cell100) | ~ spl177_266), inference(resolution, [], [f12788, f5498])).
fof(f5498, plain, ! [X0, X1] : (~ sP86(X0, X1) | (cell96 = X1)), inference(cnf_transformation, [], [f5231])).
fof(f5231, plain, ! [X0, X1] : (((cell97 = X0) & (cell96 = X1)) | ~ sP86(X0, X1)), inference(nnf_transformation, [], [f5049])).
fof(f12788, plain, (sP86(cell1, cell100) | ~ spl177_266), inference(avatar_component_clause, [], [f12786])).
fof(f13741, plain, ~ spl177_265, inference(avatar_contradiction_clause, [], [f13740])).
fof(f13740, plain, ($false | ~ spl177_265), inference(subsumption_resolution, [], [f13737, f6165])).
fof(f6165, plain, ~ (cell95 = cell100), inference(cnf_transformation, [], [f487])).
fof(f487, plain, ~ (cell95 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25476)).
fof(f13737, plain, ((cell95 = cell100) | ~ spl177_265), inference(resolution, [], [f12784, f5500])).
fof(f5500, plain, ! [X0, X1] : (~ sP85(X0, X1) | (cell95 = X1)), inference(cnf_transformation, [], [f5232])).
fof(f5232, plain, ! [X0, X1] : (((cell96 = X0) & (cell95 = X1)) | ~ sP85(X0, X1)), inference(nnf_transformation, [], [f5048])).
fof(f12784, plain, (sP85(cell1, cell100) | ~ spl177_265), inference(avatar_component_clause, [], [f12782])).
fof(f13735, plain, ~ spl177_264, inference(avatar_contradiction_clause, [], [f13734])).
fof(f13734, plain, ($false | ~ spl177_264), inference(subsumption_resolution, [], [f13731, f6259])).
fof(f6259, plain, ~ (cell94 = cell100), inference(cnf_transformation, [], [f581])).
fof(f581, plain, ~ (cell94 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25570)).
fof(f13731, plain, ((cell94 = cell100) | ~ spl177_264), inference(resolution, [], [f12780, f5502])).
fof(f5502, plain, ! [X0, X1] : (~ sP84(X0, X1) | (cell94 = X1)), inference(cnf_transformation, [], [f5233])).
fof(f5233, plain, ! [X0, X1] : (((cell95 = X0) & (cell94 = X1)) | ~ sP84(X0, X1)), inference(nnf_transformation, [], [f5047])).
fof(f12780, plain, (sP84(cell1, cell100) | ~ spl177_264), inference(avatar_component_clause, [], [f12778])).
fof(f13729, plain, ~ spl177_263, inference(avatar_contradiction_clause, [], [f13728])).
fof(f13728, plain, ($false | ~ spl177_263), inference(subsumption_resolution, [], [f13725, f6352])).
fof(f6352, plain, ~ (cell93 = cell100), inference(cnf_transformation, [], [f674])).
fof(f674, plain, ~ (cell93 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25663)).
fof(f13725, plain, ((cell93 = cell100) | ~ spl177_263), inference(resolution, [], [f12776, f5504])).
fof(f5504, plain, ! [X0, X1] : (~ sP83(X0, X1) | (cell93 = X1)), inference(cnf_transformation, [], [f5234])).
fof(f5234, plain, ! [X0, X1] : (((cell94 = X0) & (cell93 = X1)) | ~ sP83(X0, X1)), inference(nnf_transformation, [], [f5046])).
fof(f12776, plain, (sP83(cell1, cell100) | ~ spl177_263), inference(avatar_component_clause, [], [f12774])).
fof(f13723, plain, ~ spl177_262, inference(avatar_contradiction_clause, [], [f13722])).
fof(f13722, plain, ($false | ~ spl177_262), inference(subsumption_resolution, [], [f13719, f6535])).
fof(f6535, plain, ~ (cell92 = cell100), inference(cnf_transformation, [], [f857])).
fof(f857, plain, ~ (cell92 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25846)).
fof(f13719, plain, ((cell92 = cell100) | ~ spl177_262), inference(resolution, [], [f12772, f5506])).
fof(f5506, plain, ! [X0, X1] : (~ sP82(X0, X1) | (cell92 = X1)), inference(cnf_transformation, [], [f5235])).
fof(f5235, plain, ! [X0, X1] : (((cell93 = X0) & (cell92 = X1)) | ~ sP82(X0, X1)), inference(nnf_transformation, [], [f5045])).
fof(f12772, plain, (sP82(cell1, cell100) | ~ spl177_262), inference(avatar_component_clause, [], [f12770])).
fof(f13717, plain, ~ spl177_261, inference(avatar_contradiction_clause, [], [f13716])).
fof(f13716, plain, ($false | ~ spl177_261), inference(subsumption_resolution, [], [f13713, f6444])).
fof(f6444, plain, ~ (cell91 = cell100), inference(cnf_transformation, [], [f766])).
fof(f766, plain, ~ (cell91 = cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25755)).
fof(f13713, plain, ((cell91 = cell100) | ~ spl177_261), inference(resolution, [], [f12768, f5508])).
fof(f5508, plain, ! [X0, X1] : (~ sP81(X0, X1) | (cell91 = X1)), inference(cnf_transformation, [], [f5236])).
fof(f5236, plain, ! [X0, X1] : (((cell92 = X0) & (cell91 = X1)) | ~ sP81(X0, X1)), inference(nnf_transformation, [], [f5044])).
fof(f12768, plain, (sP81(cell1, cell100) | ~ spl177_261), inference(avatar_component_clause, [], [f12766])).
fof(f13711, plain, ~ spl177_258, inference(avatar_contradiction_clause, [], [f13710])).
fof(f13710, plain, ($false | ~ spl177_258), inference(subsumption_resolution, [], [f13707, f6889])).
fof(f13707, plain, ((cell87 = cell100) | ~ spl177_258), inference(resolution, [], [f12756, f5514])).
fof(f5514, plain, ! [X0, X1] : (~ sP78(X0, X1) | (cell87 = X1)), inference(cnf_transformation, [], [f5239])).
fof(f5239, plain, ! [X0, X1] : (((cell88 = X0) & (cell87 = X1)) | ~ sP78(X0, X1)), inference(nnf_transformation, [], [f5041])).
fof(f12756, plain, (sP78(cell1, cell100) | ~ spl177_258), inference(avatar_component_clause, [], [f12754])).
fof(f13705, plain, ~ spl177_257, inference(avatar_contradiction_clause, [], [f13704])).
fof(f13704, plain, ($false | ~ spl177_257), inference(subsumption_resolution, [], [f13701, f6975])).
fof(f13701, plain, ((cell86 = cell100) | ~ spl177_257), inference(resolution, [], [f12752, f5516])).
fof(f5516, plain, ! [X0, X1] : (~ sP77(X0, X1) | (cell86 = X1)), inference(cnf_transformation, [], [f5240])).
fof(f5240, plain, ! [X0, X1] : (((cell87 = X0) & (cell86 = X1)) | ~ sP77(X0, X1)), inference(nnf_transformation, [], [f5040])).
fof(f12752, plain, (sP77(cell1, cell100) | ~ spl177_257), inference(avatar_component_clause, [], [f12750])).
fof(f13699, plain, ~ spl177_256, inference(avatar_contradiction_clause, [], [f13698])).
fof(f13698, plain, ($false | ~ spl177_256), inference(subsumption_resolution, [], [f13695, f7060])).
fof(f13695, plain, ((cell85 = cell100) | ~ spl177_256), inference(resolution, [], [f12748, f5518])).
fof(f5518, plain, ! [X0, X1] : (~ sP76(X0, X1) | (cell85 = X1)), inference(cnf_transformation, [], [f5241])).
fof(f5241, plain, ! [X0, X1] : (((cell86 = X0) & (cell85 = X1)) | ~ sP76(X0, X1)), inference(nnf_transformation, [], [f5039])).
fof(f12748, plain, (sP76(cell1, cell100) | ~ spl177_256), inference(avatar_component_clause, [], [f12746])).
fof(f13693, plain, ~ spl177_255, inference(avatar_contradiction_clause, [], [f13692])).
fof(f13692, plain, ($false | ~ spl177_255), inference(subsumption_resolution, [], [f13689, f7144])).
fof(f13689, plain, ((cell84 = cell100) | ~ spl177_255), inference(resolution, [], [f12744, f5520])).
fof(f5520, plain, ! [X0, X1] : (~ sP75(X0, X1) | (cell84 = X1)), inference(cnf_transformation, [], [f5242])).
fof(f5242, plain, ! [X0, X1] : (((cell85 = X0) & (cell84 = X1)) | ~ sP75(X0, X1)), inference(nnf_transformation, [], [f5038])).
fof(f12744, plain, (sP75(cell1, cell100) | ~ spl177_255), inference(avatar_component_clause, [], [f12742])).
fof(f13687, plain, ~ spl177_254, inference(avatar_contradiction_clause, [], [f13686])).
fof(f13686, plain, ($false | ~ spl177_254), inference(subsumption_resolution, [], [f13683, f7227])).
fof(f13683, plain, ((cell83 = cell100) | ~ spl177_254), inference(resolution, [], [f12740, f5522])).
fof(f5522, plain, ! [X0, X1] : (~ sP74(X0, X1) | (cell83 = X1)), inference(cnf_transformation, [], [f5243])).
fof(f5243, plain, ! [X0, X1] : (((cell84 = X0) & (cell83 = X1)) | ~ sP74(X0, X1)), inference(nnf_transformation, [], [f5037])).
fof(f12740, plain, (sP74(cell1, cell100) | ~ spl177_254), inference(avatar_component_clause, [], [f12738])).
fof(f13681, plain, ~ spl177_253, inference(avatar_contradiction_clause, [], [f13680])).
fof(f13680, plain, ($false | ~ spl177_253), inference(subsumption_resolution, [], [f13677, f7390])).
fof(f13677, plain, ((cell82 = cell100) | ~ spl177_253), inference(resolution, [], [f12736, f5524])).
fof(f5524, plain, ! [X0, X1] : (~ sP73(X0, X1) | (cell82 = X1)), inference(cnf_transformation, [], [f5244])).
fof(f5244, plain, ! [X0, X1] : (((cell83 = X0) & (cell82 = X1)) | ~ sP73(X0, X1)), inference(nnf_transformation, [], [f5036])).
fof(f12736, plain, (sP73(cell1, cell100) | ~ spl177_253), inference(avatar_component_clause, [], [f12734])).
fof(f13675, plain, ~ spl177_252, inference(avatar_contradiction_clause, [], [f13674])).
fof(f13674, plain, ($false | ~ spl177_252), inference(subsumption_resolution, [], [f13671, f7309])).
fof(f13671, plain, ((cell81 = cell100) | ~ spl177_252), inference(resolution, [], [f12732, f5526])).
fof(f5526, plain, ! [X0, X1] : (~ sP72(X0, X1) | (cell81 = X1)), inference(cnf_transformation, [], [f5245])).
fof(f5245, plain, ! [X0, X1] : (((cell82 = X0) & (cell81 = X1)) | ~ sP72(X0, X1)), inference(nnf_transformation, [], [f5035])).
fof(f12732, plain, (sP72(cell1, cell100) | ~ spl177_252), inference(avatar_component_clause, [], [f12730])).
fof(f13669, plain, ~ spl177_251, inference(avatar_contradiction_clause, [], [f13668])).
fof(f13668, plain, ($false | ~ spl177_251), inference(subsumption_resolution, [], [f13665, f7549])).
fof(f13665, plain, ((cell79 = cell100) | ~ spl177_251), inference(resolution, [], [f12728, f5528])).
fof(f5528, plain, ! [X0, X1] : (~ sP71(X0, X1) | (cell79 = X1)), inference(cnf_transformation, [], [f5246])).
fof(f5246, plain, ! [X0, X1] : (((cell80 = X0) & (cell79 = X1)) | ~ sP71(X0, X1)), inference(nnf_transformation, [], [f5034])).
fof(f12728, plain, (sP71(cell1, cell100) | ~ spl177_251), inference(avatar_component_clause, [], [f12726])).
fof(f13663, plain, ~ spl177_250, inference(avatar_contradiction_clause, [], [f13662])).
fof(f13662, plain, ($false | ~ spl177_250), inference(subsumption_resolution, [], [f13659, f7627])).
fof(f13659, plain, ((cell78 = cell100) | ~ spl177_250), inference(resolution, [], [f12724, f5530])).
fof(f5530, plain, ! [X0, X1] : (~ sP70(X0, X1) | (cell78 = X1)), inference(cnf_transformation, [], [f5247])).
fof(f5247, plain, ! [X0, X1] : (((cell79 = X0) & (cell78 = X1)) | ~ sP70(X0, X1)), inference(nnf_transformation, [], [f5033])).
fof(f12724, plain, (sP70(cell1, cell100) | ~ spl177_250), inference(avatar_component_clause, [], [f12722])).
fof(f13657, plain, ~ spl177_249, inference(avatar_contradiction_clause, [], [f13656])).
fof(f13656, plain, ($false | ~ spl177_249), inference(subsumption_resolution, [], [f13653, f7704])).
fof(f13653, plain, ((cell77 = cell100) | ~ spl177_249), inference(resolution, [], [f12720, f5532])).
fof(f5532, plain, ! [X0, X1] : (~ sP69(X0, X1) | (cell77 = X1)), inference(cnf_transformation, [], [f5248])).
fof(f5248, plain, ! [X0, X1] : (((cell78 = X0) & (cell77 = X1)) | ~ sP69(X0, X1)), inference(nnf_transformation, [], [f5032])).
fof(f12720, plain, (sP69(cell1, cell100) | ~ spl177_249), inference(avatar_component_clause, [], [f12718])).
fof(f13651, plain, ~ spl177_248, inference(avatar_contradiction_clause, [], [f13650])).
fof(f13650, plain, ($false | ~ spl177_248), inference(subsumption_resolution, [], [f13647, f7780])).
fof(f13647, plain, ((cell76 = cell100) | ~ spl177_248), inference(resolution, [], [f12716, f5534])).
fof(f5534, plain, ! [X0, X1] : (~ sP68(X0, X1) | (cell76 = X1)), inference(cnf_transformation, [], [f5249])).
fof(f5249, plain, ! [X0, X1] : (((cell77 = X0) & (cell76 = X1)) | ~ sP68(X0, X1)), inference(nnf_transformation, [], [f5031])).
fof(f12716, plain, (sP68(cell1, cell100) | ~ spl177_248), inference(avatar_component_clause, [], [f12714])).
fof(f13645, plain, ~ spl177_247, inference(avatar_contradiction_clause, [], [f13644])).
fof(f13644, plain, ($false | ~ spl177_247), inference(subsumption_resolution, [], [f13641, f7855])).
fof(f13641, plain, ((cell75 = cell100) | ~ spl177_247), inference(resolution, [], [f12712, f5536])).
fof(f5536, plain, ! [X0, X1] : (~ sP67(X0, X1) | (cell75 = X1)), inference(cnf_transformation, [], [f5250])).
fof(f5250, plain, ! [X0, X1] : (((cell76 = X0) & (cell75 = X1)) | ~ sP67(X0, X1)), inference(nnf_transformation, [], [f5030])).
fof(f12712, plain, (sP67(cell1, cell100) | ~ spl177_247), inference(avatar_component_clause, [], [f12710])).
fof(f13639, plain, ~ spl177_246, inference(avatar_contradiction_clause, [], [f13638])).
fof(f13638, plain, ($false | ~ spl177_246), inference(subsumption_resolution, [], [f13635, f7929])).
fof(f13635, plain, ((cell74 = cell100) | ~ spl177_246), inference(resolution, [], [f12708, f5538])).
fof(f5538, plain, ! [X0, X1] : (~ sP66(X0, X1) | (cell74 = X1)), inference(cnf_transformation, [], [f5251])).
fof(f5251, plain, ! [X0, X1] : (((cell75 = X0) & (cell74 = X1)) | ~ sP66(X0, X1)), inference(nnf_transformation, [], [f5029])).
fof(f12708, plain, (sP66(cell1, cell100) | ~ spl177_246), inference(avatar_component_clause, [], [f12706])).
fof(f13633, plain, ~ spl177_245, inference(avatar_contradiction_clause, [], [f13632])).
fof(f13632, plain, ($false | ~ spl177_245), inference(subsumption_resolution, [], [f13629, f8002])).
fof(f13629, plain, ((cell73 = cell100) | ~ spl177_245), inference(resolution, [], [f12704, f5540])).
fof(f5540, plain, ! [X0, X1] : (~ sP65(X0, X1) | (cell73 = X1)), inference(cnf_transformation, [], [f5252])).
fof(f5252, plain, ! [X0, X1] : (((cell74 = X0) & (cell73 = X1)) | ~ sP65(X0, X1)), inference(nnf_transformation, [], [f5028])).
fof(f12704, plain, (sP65(cell1, cell100) | ~ spl177_245), inference(avatar_component_clause, [], [f12702])).
fof(f13627, plain, ~ spl177_244, inference(avatar_contradiction_clause, [], [f13626])).
fof(f13626, plain, ($false | ~ spl177_244), inference(subsumption_resolution, [], [f13623, f8145])).
fof(f13623, plain, ((cell72 = cell100) | ~ spl177_244), inference(resolution, [], [f12700, f5542])).
fof(f5542, plain, ! [X0, X1] : (~ sP64(X0, X1) | (cell72 = X1)), inference(cnf_transformation, [], [f5253])).
fof(f5253, plain, ! [X0, X1] : (((cell73 = X0) & (cell72 = X1)) | ~ sP64(X0, X1)), inference(nnf_transformation, [], [f5027])).
fof(f12700, plain, (sP64(cell1, cell100) | ~ spl177_244), inference(avatar_component_clause, [], [f12698])).
fof(f13621, plain, ~ spl177_243, inference(avatar_contradiction_clause, [], [f13620])).
fof(f13620, plain, ($false | ~ spl177_243), inference(subsumption_resolution, [], [f13617, f8074])).
fof(f13617, plain, ((cell71 = cell100) | ~ spl177_243), inference(resolution, [], [f12696, f5544])).
fof(f5544, plain, ! [X0, X1] : (~ sP63(X0, X1) | (cell71 = X1)), inference(cnf_transformation, [], [f5254])).
fof(f5254, plain, ! [X0, X1] : (((cell72 = X0) & (cell71 = X1)) | ~ sP63(X0, X1)), inference(nnf_transformation, [], [f5026])).
fof(f12696, plain, (sP63(cell1, cell100) | ~ spl177_243), inference(avatar_component_clause, [], [f12694])).
fof(f13615, plain, ~ spl177_242, inference(avatar_contradiction_clause, [], [f13614])).
fof(f13614, plain, ($false | ~ spl177_242), inference(subsumption_resolution, [], [f13611, f8284])).
fof(f13611, plain, ((cell69 = cell100) | ~ spl177_242), inference(resolution, [], [f12692, f5546])).
fof(f5546, plain, ! [X0, X1] : (~ sP62(X0, X1) | (cell69 = X1)), inference(cnf_transformation, [], [f5255])).
fof(f5255, plain, ! [X0, X1] : (((cell70 = X0) & (cell69 = X1)) | ~ sP62(X0, X1)), inference(nnf_transformation, [], [f5025])).
fof(f12692, plain, (sP62(cell1, cell100) | ~ spl177_242), inference(avatar_component_clause, [], [f12690])).
fof(f13609, plain, ~ spl177_241, inference(avatar_contradiction_clause, [], [f13608])).
fof(f13608, plain, ($false | ~ spl177_241), inference(subsumption_resolution, [], [f13605, f8352])).
fof(f13605, plain, ((cell68 = cell100) | ~ spl177_241), inference(resolution, [], [f12688, f5548])).
fof(f5548, plain, ! [X0, X1] : (~ sP61(X0, X1) | (cell68 = X1)), inference(cnf_transformation, [], [f5256])).
fof(f5256, plain, ! [X0, X1] : (((cell69 = X0) & (cell68 = X1)) | ~ sP61(X0, X1)), inference(nnf_transformation, [], [f5024])).
fof(f12688, plain, (sP61(cell1, cell100) | ~ spl177_241), inference(avatar_component_clause, [], [f12686])).
fof(f13603, plain, ~ spl177_240, inference(avatar_contradiction_clause, [], [f13602])).
fof(f13602, plain, ($false | ~ spl177_240), inference(subsumption_resolution, [], [f13599, f8419])).
fof(f13599, plain, ((cell67 = cell100) | ~ spl177_240), inference(resolution, [], [f12684, f5550])).
fof(f5550, plain, ! [X0, X1] : (~ sP60(X0, X1) | (cell67 = X1)), inference(cnf_transformation, [], [f5257])).
fof(f5257, plain, ! [X0, X1] : (((cell68 = X0) & (cell67 = X1)) | ~ sP60(X0, X1)), inference(nnf_transformation, [], [f5023])).
fof(f12684, plain, (sP60(cell1, cell100) | ~ spl177_240), inference(avatar_component_clause, [], [f12682])).
fof(f13597, plain, ~ spl177_239, inference(avatar_contradiction_clause, [], [f13596])).
fof(f13596, plain, ($false | ~ spl177_239), inference(subsumption_resolution, [], [f13593, f8485])).
fof(f13593, plain, ((cell66 = cell100) | ~ spl177_239), inference(resolution, [], [f12680, f5552])).
fof(f5552, plain, ! [X0, X1] : (~ sP59(X0, X1) | (cell66 = X1)), inference(cnf_transformation, [], [f5258])).
fof(f5258, plain, ! [X0, X1] : (((cell67 = X0) & (cell66 = X1)) | ~ sP59(X0, X1)), inference(nnf_transformation, [], [f5022])).
fof(f12680, plain, (sP59(cell1, cell100) | ~ spl177_239), inference(avatar_component_clause, [], [f12678])).
fof(f13591, plain, ~ spl177_238, inference(avatar_contradiction_clause, [], [f13590])).
fof(f13590, plain, ($false | ~ spl177_238), inference(subsumption_resolution, [], [f13587, f8550])).
fof(f13587, plain, ((cell65 = cell100) | ~ spl177_238), inference(resolution, [], [f12676, f5554])).
fof(f5554, plain, ! [X0, X1] : (~ sP58(X0, X1) | (cell65 = X1)), inference(cnf_transformation, [], [f5259])).
fof(f5259, plain, ! [X0, X1] : (((cell66 = X0) & (cell65 = X1)) | ~ sP58(X0, X1)), inference(nnf_transformation, [], [f5021])).
fof(f12676, plain, (sP58(cell1, cell100) | ~ spl177_238), inference(avatar_component_clause, [], [f12674])).
fof(f13585, plain, ~ spl177_237, inference(avatar_contradiction_clause, [], [f13584])).
fof(f13584, plain, ($false | ~ spl177_237), inference(subsumption_resolution, [], [f13581, f8614])).
fof(f13581, plain, ((cell64 = cell100) | ~ spl177_237), inference(resolution, [], [f12672, f5556])).
fof(f5556, plain, ! [X0, X1] : (~ sP57(X0, X1) | (cell64 = X1)), inference(cnf_transformation, [], [f5260])).
fof(f5260, plain, ! [X0, X1] : (((cell65 = X0) & (cell64 = X1)) | ~ sP57(X0, X1)), inference(nnf_transformation, [], [f5020])).
fof(f12672, plain, (sP57(cell1, cell100) | ~ spl177_237), inference(avatar_component_clause, [], [f12670])).
fof(f13579, plain, ~ spl177_236, inference(avatar_contradiction_clause, [], [f13578])).
fof(f13578, plain, ($false | ~ spl177_236), inference(subsumption_resolution, [], [f13575, f8677])).
fof(f13575, plain, ((cell63 = cell100) | ~ spl177_236), inference(resolution, [], [f12668, f5558])).
fof(f5558, plain, ! [X0, X1] : (~ sP56(X0, X1) | (cell63 = X1)), inference(cnf_transformation, [], [f5261])).
fof(f5261, plain, ! [X0, X1] : (((cell64 = X0) & (cell63 = X1)) | ~ sP56(X0, X1)), inference(nnf_transformation, [], [f5019])).
fof(f12668, plain, (sP56(cell1, cell100) | ~ spl177_236), inference(avatar_component_clause, [], [f12666])).
fof(f13573, plain, ~ spl177_235, inference(avatar_contradiction_clause, [], [f13572])).
fof(f13572, plain, ($false | ~ spl177_235), inference(subsumption_resolution, [], [f13569, f8800])).
fof(f13569, plain, ((cell62 = cell100) | ~ spl177_235), inference(resolution, [], [f12664, f5560])).
fof(f5560, plain, ! [X0, X1] : (~ sP55(X0, X1) | (cell62 = X1)), inference(cnf_transformation, [], [f5262])).
fof(f5262, plain, ! [X0, X1] : (((cell63 = X0) & (cell62 = X1)) | ~ sP55(X0, X1)), inference(nnf_transformation, [], [f5018])).
fof(f12664, plain, (sP55(cell1, cell100) | ~ spl177_235), inference(avatar_component_clause, [], [f12662])).
fof(f13567, plain, ~ spl177_234, inference(avatar_contradiction_clause, [], [f13566])).
fof(f13566, plain, ($false | ~ spl177_234), inference(subsumption_resolution, [], [f13563, f8739])).
fof(f13563, plain, ((cell61 = cell100) | ~ spl177_234), inference(resolution, [], [f12660, f5562])).
fof(f5562, plain, ! [X0, X1] : (~ sP54(X0, X1) | (cell61 = X1)), inference(cnf_transformation, [], [f5263])).
fof(f5263, plain, ! [X0, X1] : (((cell62 = X0) & (cell61 = X1)) | ~ sP54(X0, X1)), inference(nnf_transformation, [], [f5017])).
fof(f12660, plain, (sP54(cell1, cell100) | ~ spl177_234), inference(avatar_component_clause, [], [f12658])).
fof(f13561, plain, ~ spl177_233, inference(avatar_contradiction_clause, [], [f13560])).
fof(f13560, plain, ($false | ~ spl177_233), inference(subsumption_resolution, [], [f13557, f8919])).
fof(f13557, plain, ((cell59 = cell100) | ~ spl177_233), inference(resolution, [], [f12656, f5564])).
fof(f5564, plain, ! [X0, X1] : (~ sP53(X0, X1) | (cell59 = X1)), inference(cnf_transformation, [], [f5264])).
fof(f5264, plain, ! [X0, X1] : (((cell60 = X0) & (cell59 = X1)) | ~ sP53(X0, X1)), inference(nnf_transformation, [], [f5016])).
fof(f12656, plain, (sP53(cell1, cell100) | ~ spl177_233), inference(avatar_component_clause, [], [f12654])).
fof(f13555, plain, ~ spl177_232, inference(avatar_contradiction_clause, [], [f13554])).
fof(f13554, plain, ($false | ~ spl177_232), inference(subsumption_resolution, [], [f13551, f8977])).
fof(f13551, plain, ((cell58 = cell100) | ~ spl177_232), inference(resolution, [], [f12652, f5566])).
fof(f5566, plain, ! [X0, X1] : (~ sP52(X0, X1) | (cell58 = X1)), inference(cnf_transformation, [], [f5265])).
fof(f5265, plain, ! [X0, X1] : (((cell59 = X0) & (cell58 = X1)) | ~ sP52(X0, X1)), inference(nnf_transformation, [], [f5015])).
fof(f12652, plain, (sP52(cell1, cell100) | ~ spl177_232), inference(avatar_component_clause, [], [f12650])).
fof(f13549, plain, ~ spl177_231, inference(avatar_contradiction_clause, [], [f13548])).
fof(f13548, plain, ($false | ~ spl177_231), inference(subsumption_resolution, [], [f13545, f9034])).
fof(f13545, plain, ((cell57 = cell100) | ~ spl177_231), inference(resolution, [], [f12648, f5568])).
fof(f5568, plain, ! [X0, X1] : (~ sP51(X0, X1) | (cell57 = X1)), inference(cnf_transformation, [], [f5266])).
fof(f5266, plain, ! [X0, X1] : (((cell58 = X0) & (cell57 = X1)) | ~ sP51(X0, X1)), inference(nnf_transformation, [], [f5014])).
fof(f12648, plain, (sP51(cell1, cell100) | ~ spl177_231), inference(avatar_component_clause, [], [f12646])).
fof(f13543, plain, ~ spl177_230, inference(avatar_contradiction_clause, [], [f13542])).
fof(f13542, plain, ($false | ~ spl177_230), inference(subsumption_resolution, [], [f13539, f9090])).
fof(f13539, plain, ((cell56 = cell100) | ~ spl177_230), inference(resolution, [], [f12644, f5570])).
fof(f5570, plain, ! [X0, X1] : (~ sP50(X0, X1) | (cell56 = X1)), inference(cnf_transformation, [], [f5267])).
fof(f5267, plain, ! [X0, X1] : (((cell57 = X0) & (cell56 = X1)) | ~ sP50(X0, X1)), inference(nnf_transformation, [], [f5013])).
fof(f12644, plain, (sP50(cell1, cell100) | ~ spl177_230), inference(avatar_component_clause, [], [f12642])).
fof(f13537, plain, ~ spl177_229, inference(avatar_contradiction_clause, [], [f13536])).
fof(f13536, plain, ($false | ~ spl177_229), inference(subsumption_resolution, [], [f13533, f9145])).
fof(f13533, plain, ((cell55 = cell100) | ~ spl177_229), inference(resolution, [], [f12640, f5572])).
fof(f5572, plain, ! [X0, X1] : (~ sP49(X0, X1) | (cell55 = X1)), inference(cnf_transformation, [], [f5268])).
fof(f5268, plain, ! [X0, X1] : (((cell56 = X0) & (cell55 = X1)) | ~ sP49(X0, X1)), inference(nnf_transformation, [], [f5012])).
fof(f12640, plain, (sP49(cell1, cell100) | ~ spl177_229), inference(avatar_component_clause, [], [f12638])).
fof(f13531, plain, ~ spl177_228, inference(avatar_contradiction_clause, [], [f13530])).
fof(f13530, plain, ($false | ~ spl177_228), inference(subsumption_resolution, [], [f13527, f9199])).
fof(f13527, plain, ((cell54 = cell100) | ~ spl177_228), inference(resolution, [], [f12636, f5574])).
fof(f5574, plain, ! [X0, X1] : (~ sP48(X0, X1) | (cell54 = X1)), inference(cnf_transformation, [], [f5269])).
fof(f5269, plain, ! [X0, X1] : (((cell55 = X0) & (cell54 = X1)) | ~ sP48(X0, X1)), inference(nnf_transformation, [], [f5011])).
fof(f12636, plain, (sP48(cell1, cell100) | ~ spl177_228), inference(avatar_component_clause, [], [f12634])).
fof(f13525, plain, ~ spl177_227, inference(avatar_contradiction_clause, [], [f13524])).
fof(f13524, plain, ($false | ~ spl177_227), inference(subsumption_resolution, [], [f13521, f9252])).
fof(f13521, plain, ((cell53 = cell100) | ~ spl177_227), inference(resolution, [], [f12632, f5576])).
fof(f5576, plain, ! [X0, X1] : (~ sP47(X0, X1) | (cell53 = X1)), inference(cnf_transformation, [], [f5270])).
fof(f5270, plain, ! [X0, X1] : (((cell54 = X0) & (cell53 = X1)) | ~ sP47(X0, X1)), inference(nnf_transformation, [], [f5010])).
fof(f12632, plain, (sP47(cell1, cell100) | ~ spl177_227), inference(avatar_component_clause, [], [f12630])).
fof(f13519, plain, ~ spl177_226, inference(avatar_contradiction_clause, [], [f13518])).
fof(f13518, plain, ($false | ~ spl177_226), inference(subsumption_resolution, [], [f13515, f9355])).
fof(f13515, plain, ((cell52 = cell100) | ~ spl177_226), inference(resolution, [], [f12628, f5578])).
fof(f5578, plain, ! [X0, X1] : (~ sP46(X0, X1) | (cell52 = X1)), inference(cnf_transformation, [], [f5271])).
fof(f5271, plain, ! [X0, X1] : (((cell53 = X0) & (cell52 = X1)) | ~ sP46(X0, X1)), inference(nnf_transformation, [], [f5009])).
fof(f12628, plain, (sP46(cell1, cell100) | ~ spl177_226), inference(avatar_component_clause, [], [f12626])).
fof(f13513, plain, ~ spl177_225, inference(avatar_contradiction_clause, [], [f13512])).
fof(f13512, plain, ($false | ~ spl177_225), inference(subsumption_resolution, [], [f13509, f9304])).
fof(f13509, plain, ((cell51 = cell100) | ~ spl177_225), inference(resolution, [], [f12624, f5580])).
fof(f5580, plain, ! [X0, X1] : (~ sP45(X0, X1) | (cell51 = X1)), inference(cnf_transformation, [], [f5272])).
fof(f5272, plain, ! [X0, X1] : (((cell52 = X0) & (cell51 = X1)) | ~ sP45(X0, X1)), inference(nnf_transformation, [], [f5008])).
fof(f12624, plain, (sP45(cell1, cell100) | ~ spl177_225), inference(avatar_component_clause, [], [f12622])).
fof(f13507, plain, ~ spl177_224, inference(avatar_contradiction_clause, [], [f13506])).
fof(f13506, plain, ($false | ~ spl177_224), inference(subsumption_resolution, [], [f13503, f9454])).
fof(f13503, plain, ((cell49 = cell100) | ~ spl177_224), inference(resolution, [], [f12620, f5582])).
fof(f5582, plain, ! [X0, X1] : (~ sP44(X0, X1) | (cell49 = X1)), inference(cnf_transformation, [], [f5273])).
fof(f5273, plain, ! [X0, X1] : (((cell50 = X0) & (cell49 = X1)) | ~ sP44(X0, X1)), inference(nnf_transformation, [], [f5007])).
fof(f12620, plain, (sP44(cell1, cell100) | ~ spl177_224), inference(avatar_component_clause, [], [f12618])).
fof(f13501, plain, ~ spl177_223, inference(avatar_contradiction_clause, [], [f13500])).
fof(f13500, plain, ($false | ~ spl177_223), inference(subsumption_resolution, [], [f13497, f9502])).
fof(f13497, plain, ((cell48 = cell100) | ~ spl177_223), inference(resolution, [], [f12616, f5584])).
fof(f5584, plain, ! [X0, X1] : (~ sP43(X0, X1) | (cell48 = X1)), inference(cnf_transformation, [], [f5274])).
fof(f5274, plain, ! [X0, X1] : (((cell49 = X0) & (cell48 = X1)) | ~ sP43(X0, X1)), inference(nnf_transformation, [], [f5006])).
fof(f12616, plain, (sP43(cell1, cell100) | ~ spl177_223), inference(avatar_component_clause, [], [f12614])).
fof(f13495, plain, ~ spl177_222, inference(avatar_contradiction_clause, [], [f13494])).
fof(f13494, plain, ($false | ~ spl177_222), inference(subsumption_resolution, [], [f13491, f9549])).
fof(f13491, plain, ((cell47 = cell100) | ~ spl177_222), inference(resolution, [], [f12612, f5586])).
fof(f5586, plain, ! [X0, X1] : (~ sP42(X0, X1) | (cell47 = X1)), inference(cnf_transformation, [], [f5275])).
fof(f5275, plain, ! [X0, X1] : (((cell48 = X0) & (cell47 = X1)) | ~ sP42(X0, X1)), inference(nnf_transformation, [], [f5005])).
fof(f12612, plain, (sP42(cell1, cell100) | ~ spl177_222), inference(avatar_component_clause, [], [f12610])).
fof(f13489, plain, ~ spl177_221, inference(avatar_contradiction_clause, [], [f13488])).
fof(f13488, plain, ($false | ~ spl177_221), inference(subsumption_resolution, [], [f13485, f9595])).
fof(f13485, plain, ((cell46 = cell100) | ~ spl177_221), inference(resolution, [], [f12608, f5588])).
fof(f5588, plain, ! [X0, X1] : (~ sP41(X0, X1) | (cell46 = X1)), inference(cnf_transformation, [], [f5276])).
fof(f5276, plain, ! [X0, X1] : (((cell47 = X0) & (cell46 = X1)) | ~ sP41(X0, X1)), inference(nnf_transformation, [], [f5004])).
fof(f12608, plain, (sP41(cell1, cell100) | ~ spl177_221), inference(avatar_component_clause, [], [f12606])).
fof(f13483, plain, ~ spl177_220, inference(avatar_contradiction_clause, [], [f13482])).
fof(f13482, plain, ($false | ~ spl177_220), inference(subsumption_resolution, [], [f13479, f9640])).
fof(f13479, plain, ((cell45 = cell100) | ~ spl177_220), inference(resolution, [], [f12604, f5590])).
fof(f5590, plain, ! [X0, X1] : (~ sP40(X0, X1) | (cell45 = X1)), inference(cnf_transformation, [], [f5277])).
fof(f5277, plain, ! [X0, X1] : (((cell46 = X0) & (cell45 = X1)) | ~ sP40(X0, X1)), inference(nnf_transformation, [], [f5003])).
fof(f12604, plain, (sP40(cell1, cell100) | ~ spl177_220), inference(avatar_component_clause, [], [f12602])).
fof(f13477, plain, ~ spl177_219, inference(avatar_contradiction_clause, [], [f13476])).
fof(f13476, plain, ($false | ~ spl177_219), inference(subsumption_resolution, [], [f13473, f9684])).
fof(f13473, plain, ((cell44 = cell100) | ~ spl177_219), inference(resolution, [], [f12600, f5592])).
fof(f5592, plain, ! [X0, X1] : (~ sP39(X0, X1) | (cell44 = X1)), inference(cnf_transformation, [], [f5278])).
fof(f5278, plain, ! [X0, X1] : (((cell45 = X0) & (cell44 = X1)) | ~ sP39(X0, X1)), inference(nnf_transformation, [], [f5002])).
fof(f12600, plain, (sP39(cell1, cell100) | ~ spl177_219), inference(avatar_component_clause, [], [f12598])).
fof(f13471, plain, ~ spl177_218, inference(avatar_contradiction_clause, [], [f13470])).
fof(f13470, plain, ($false | ~ spl177_218), inference(subsumption_resolution, [], [f13467, f9727])).
fof(f13467, plain, ((cell43 = cell100) | ~ spl177_218), inference(resolution, [], [f12596, f5594])).
fof(f5594, plain, ! [X0, X1] : (~ sP38(X0, X1) | (cell43 = X1)), inference(cnf_transformation, [], [f5279])).
fof(f5279, plain, ! [X0, X1] : (((cell44 = X0) & (cell43 = X1)) | ~ sP38(X0, X1)), inference(nnf_transformation, [], [f5001])).
fof(f12596, plain, (sP38(cell1, cell100) | ~ spl177_218), inference(avatar_component_clause, [], [f12594])).
fof(f13465, plain, ~ spl177_217, inference(avatar_contradiction_clause, [], [f13464])).
fof(f13464, plain, ($false | ~ spl177_217), inference(subsumption_resolution, [], [f13461, f9810])).
fof(f13461, plain, ((cell42 = cell100) | ~ spl177_217), inference(resolution, [], [f12592, f5596])).
fof(f5596, plain, ! [X0, X1] : (~ sP37(X0, X1) | (cell42 = X1)), inference(cnf_transformation, [], [f5280])).
fof(f5280, plain, ! [X0, X1] : (((cell43 = X0) & (cell42 = X1)) | ~ sP37(X0, X1)), inference(nnf_transformation, [], [f5000])).
fof(f12592, plain, (sP37(cell1, cell100) | ~ spl177_217), inference(avatar_component_clause, [], [f12590])).
fof(f13459, plain, ~ spl177_216, inference(avatar_contradiction_clause, [], [f13458])).
fof(f13458, plain, ($false | ~ spl177_216), inference(subsumption_resolution, [], [f13455, f9769])).
fof(f13455, plain, ((cell41 = cell100) | ~ spl177_216), inference(resolution, [], [f12588, f5598])).
fof(f5598, plain, ! [X0, X1] : (~ sP36(X0, X1) | (cell41 = X1)), inference(cnf_transformation, [], [f5281])).
fof(f5281, plain, ! [X0, X1] : (((cell42 = X0) & (cell41 = X1)) | ~ sP36(X0, X1)), inference(nnf_transformation, [], [f4999])).
fof(f12588, plain, (sP36(cell1, cell100) | ~ spl177_216), inference(avatar_component_clause, [], [f12586])).
fof(f13453, plain, ~ spl177_215, inference(avatar_contradiction_clause, [], [f13452])).
fof(f13452, plain, ($false | ~ spl177_215), inference(subsumption_resolution, [], [f13449, f9889])).
fof(f13449, plain, ((cell39 = cell100) | ~ spl177_215), inference(resolution, [], [f12584, f5600])).
fof(f5600, plain, ! [X0, X1] : (~ sP35(X0, X1) | (cell39 = X1)), inference(cnf_transformation, [], [f5282])).
fof(f5282, plain, ! [X0, X1] : (((cell40 = X0) & (cell39 = X1)) | ~ sP35(X0, X1)), inference(nnf_transformation, [], [f4998])).
fof(f12584, plain, (sP35(cell1, cell100) | ~ spl177_215), inference(avatar_component_clause, [], [f12582])).
fof(f13447, plain, ~ spl177_214, inference(avatar_contradiction_clause, [], [f13446])).
fof(f13446, plain, ($false | ~ spl177_214), inference(subsumption_resolution, [], [f13443, f9927])).
fof(f13443, plain, ((cell38 = cell100) | ~ spl177_214), inference(resolution, [], [f12580, f5602])).
fof(f5602, plain, ! [X0, X1] : (~ sP34(X0, X1) | (cell38 = X1)), inference(cnf_transformation, [], [f5283])).
fof(f5283, plain, ! [X0, X1] : (((cell39 = X0) & (cell38 = X1)) | ~ sP34(X0, X1)), inference(nnf_transformation, [], [f4997])).
fof(f12580, plain, (sP34(cell1, cell100) | ~ spl177_214), inference(avatar_component_clause, [], [f12578])).
fof(f13441, plain, ~ spl177_213, inference(avatar_contradiction_clause, [], [f13440])).
fof(f13440, plain, ($false | ~ spl177_213), inference(subsumption_resolution, [], [f13437, f9964])).
fof(f13437, plain, ((cell37 = cell100) | ~ spl177_213), inference(resolution, [], [f12576, f5604])).
fof(f5604, plain, ! [X0, X1] : (~ sP33(X0, X1) | (cell37 = X1)), inference(cnf_transformation, [], [f5284])).
fof(f5284, plain, ! [X0, X1] : (((cell38 = X0) & (cell37 = X1)) | ~ sP33(X0, X1)), inference(nnf_transformation, [], [f4996])).
fof(f12576, plain, (sP33(cell1, cell100) | ~ spl177_213), inference(avatar_component_clause, [], [f12574])).
fof(f13435, plain, ~ spl177_212, inference(avatar_contradiction_clause, [], [f13434])).
fof(f13434, plain, ($false | ~ spl177_212), inference(subsumption_resolution, [], [f13431, f10000])).
fof(f13431, plain, ((cell36 = cell100) | ~ spl177_212), inference(resolution, [], [f12572, f5606])).
fof(f5606, plain, ! [X0, X1] : (~ sP32(X0, X1) | (cell36 = X1)), inference(cnf_transformation, [], [f5285])).
fof(f5285, plain, ! [X0, X1] : (((cell37 = X0) & (cell36 = X1)) | ~ sP32(X0, X1)), inference(nnf_transformation, [], [f4995])).
fof(f12572, plain, (sP32(cell1, cell100) | ~ spl177_212), inference(avatar_component_clause, [], [f12570])).
fof(f13429, plain, ~ spl177_211, inference(avatar_contradiction_clause, [], [f13428])).
fof(f13428, plain, ($false | ~ spl177_211), inference(subsumption_resolution, [], [f13425, f10035])).
fof(f13425, plain, ((cell35 = cell100) | ~ spl177_211), inference(resolution, [], [f12568, f5608])).
fof(f5608, plain, ! [X0, X1] : (~ sP31(X0, X1) | (cell35 = X1)), inference(cnf_transformation, [], [f5286])).
fof(f5286, plain, ! [X0, X1] : (((cell36 = X0) & (cell35 = X1)) | ~ sP31(X0, X1)), inference(nnf_transformation, [], [f4994])).
fof(f12568, plain, (sP31(cell1, cell100) | ~ spl177_211), inference(avatar_component_clause, [], [f12566])).
fof(f13423, plain, ~ spl177_210, inference(avatar_contradiction_clause, [], [f13422])).
fof(f13422, plain, ($false | ~ spl177_210), inference(subsumption_resolution, [], [f13419, f10069])).
fof(f13419, plain, ((cell34 = cell100) | ~ spl177_210), inference(resolution, [], [f12564, f5610])).
fof(f5610, plain, ! [X0, X1] : (~ sP30(X0, X1) | (cell34 = X1)), inference(cnf_transformation, [], [f5287])).
fof(f5287, plain, ! [X0, X1] : (((cell35 = X0) & (cell34 = X1)) | ~ sP30(X0, X1)), inference(nnf_transformation, [], [f4993])).
fof(f12564, plain, (sP30(cell1, cell100) | ~ spl177_210), inference(avatar_component_clause, [], [f12562])).
fof(f13417, plain, ~ spl177_209, inference(avatar_contradiction_clause, [], [f13416])).
fof(f13416, plain, ($false | ~ spl177_209), inference(subsumption_resolution, [], [f13413, f10102])).
fof(f13413, plain, ((cell33 = cell100) | ~ spl177_209), inference(resolution, [], [f12560, f5612])).
fof(f5612, plain, ! [X0, X1] : (~ sP29(X0, X1) | (cell33 = X1)), inference(cnf_transformation, [], [f5288])).
fof(f5288, plain, ! [X0, X1] : (((cell34 = X0) & (cell33 = X1)) | ~ sP29(X0, X1)), inference(nnf_transformation, [], [f4992])).
fof(f12560, plain, (sP29(cell1, cell100) | ~ spl177_209), inference(avatar_component_clause, [], [f12558])).
fof(f13411, plain, ~ spl177_208, inference(avatar_contradiction_clause, [], [f13410])).
fof(f13410, plain, ($false | ~ spl177_208), inference(subsumption_resolution, [], [f13407, f10165])).
fof(f13407, plain, ((cell32 = cell100) | ~ spl177_208), inference(resolution, [], [f12556, f5614])).
fof(f5614, plain, ! [X0, X1] : (~ sP28(X0, X1) | (cell32 = X1)), inference(cnf_transformation, [], [f5289])).
fof(f5289, plain, ! [X0, X1] : (((cell33 = X0) & (cell32 = X1)) | ~ sP28(X0, X1)), inference(nnf_transformation, [], [f4991])).
fof(f12556, plain, (sP28(cell1, cell100) | ~ spl177_208), inference(avatar_component_clause, [], [f12554])).
fof(f13405, plain, ~ spl177_207, inference(avatar_contradiction_clause, [], [f13404])).
fof(f13404, plain, ($false | ~ spl177_207), inference(subsumption_resolution, [], [f13401, f10134])).
fof(f13401, plain, ((cell31 = cell100) | ~ spl177_207), inference(resolution, [], [f12552, f5616])).
fof(f5616, plain, ! [X0, X1] : (~ sP27(X0, X1) | (cell31 = X1)), inference(cnf_transformation, [], [f5290])).
fof(f5290, plain, ! [X0, X1] : (((cell32 = X0) & (cell31 = X1)) | ~ sP27(X0, X1)), inference(nnf_transformation, [], [f4990])).
fof(f12552, plain, (sP27(cell1, cell100) | ~ spl177_207), inference(avatar_component_clause, [], [f12550])).
fof(f13399, plain, ~ spl177_206, inference(avatar_contradiction_clause, [], [f13398])).
fof(f13398, plain, ($false | ~ spl177_206), inference(subsumption_resolution, [], [f13395, f10224])).
fof(f13395, plain, ((cell29 = cell100) | ~ spl177_206), inference(resolution, [], [f12548, f5618])).
fof(f5618, plain, ! [X0, X1] : (~ sP26(X0, X1) | (cell29 = X1)), inference(cnf_transformation, [], [f5291])).
fof(f5291, plain, ! [X0, X1] : (((cell30 = X0) & (cell29 = X1)) | ~ sP26(X0, X1)), inference(nnf_transformation, [], [f4989])).
fof(f12548, plain, (sP26(cell1, cell100) | ~ spl177_206), inference(avatar_component_clause, [], [f12546])).
fof(f13393, plain, ~ spl177_205, inference(avatar_contradiction_clause, [], [f13392])).
fof(f13392, plain, ($false | ~ spl177_205), inference(subsumption_resolution, [], [f13389, f10252])).
fof(f13389, plain, ((cell28 = cell100) | ~ spl177_205), inference(resolution, [], [f12544, f5620])).
fof(f5620, plain, ! [X0, X1] : (~ sP25(X0, X1) | (cell28 = X1)), inference(cnf_transformation, [], [f5292])).
fof(f5292, plain, ! [X0, X1] : (((cell29 = X0) & (cell28 = X1)) | ~ sP25(X0, X1)), inference(nnf_transformation, [], [f4988])).
fof(f12544, plain, (sP25(cell1, cell100) | ~ spl177_205), inference(avatar_component_clause, [], [f12542])).
fof(f13387, plain, ~ spl177_204, inference(avatar_contradiction_clause, [], [f13386])).
fof(f13386, plain, ($false | ~ spl177_204), inference(subsumption_resolution, [], [f13383, f10279])).
fof(f13383, plain, ((cell27 = cell100) | ~ spl177_204), inference(resolution, [], [f12540, f5622])).
fof(f5622, plain, ! [X0, X1] : (~ sP24(X0, X1) | (cell27 = X1)), inference(cnf_transformation, [], [f5293])).
fof(f5293, plain, ! [X0, X1] : (((cell28 = X0) & (cell27 = X1)) | ~ sP24(X0, X1)), inference(nnf_transformation, [], [f4987])).
fof(f12540, plain, (sP24(cell1, cell100) | ~ spl177_204), inference(avatar_component_clause, [], [f12538])).
fof(f13381, plain, ~ spl177_203, inference(avatar_contradiction_clause, [], [f13380])).
fof(f13380, plain, ($false | ~ spl177_203), inference(subsumption_resolution, [], [f13377, f10305])).
fof(f13377, plain, ((cell26 = cell100) | ~ spl177_203), inference(resolution, [], [f12536, f5624])).
fof(f5624, plain, ! [X0, X1] : (~ sP23(X0, X1) | (cell26 = X1)), inference(cnf_transformation, [], [f5294])).
fof(f5294, plain, ! [X0, X1] : (((cell27 = X0) & (cell26 = X1)) | ~ sP23(X0, X1)), inference(nnf_transformation, [], [f4986])).
fof(f12536, plain, (sP23(cell1, cell100) | ~ spl177_203), inference(avatar_component_clause, [], [f12534])).
fof(f13375, plain, ~ spl177_202, inference(avatar_contradiction_clause, [], [f13374])).
fof(f13374, plain, ($false | ~ spl177_202), inference(subsumption_resolution, [], [f13371, f10330])).
fof(f13371, plain, ((cell25 = cell100) | ~ spl177_202), inference(resolution, [], [f12532, f5626])).
fof(f5626, plain, ! [X0, X1] : (~ sP22(X0, X1) | (cell25 = X1)), inference(cnf_transformation, [], [f5295])).
fof(f5295, plain, ! [X0, X1] : (((cell26 = X0) & (cell25 = X1)) | ~ sP22(X0, X1)), inference(nnf_transformation, [], [f4985])).
fof(f12532, plain, (sP22(cell1, cell100) | ~ spl177_202), inference(avatar_component_clause, [], [f12530])).
fof(f13369, plain, ~ spl177_201, inference(avatar_contradiction_clause, [], [f13368])).
fof(f13368, plain, ($false | ~ spl177_201), inference(subsumption_resolution, [], [f13365, f10354])).
fof(f13365, plain, ((cell24 = cell100) | ~ spl177_201), inference(resolution, [], [f12528, f5628])).
fof(f5628, plain, ! [X0, X1] : (~ sP21(X0, X1) | (cell24 = X1)), inference(cnf_transformation, [], [f5296])).
fof(f5296, plain, ! [X0, X1] : (((cell25 = X0) & (cell24 = X1)) | ~ sP21(X0, X1)), inference(nnf_transformation, [], [f4984])).
fof(f12528, plain, (sP21(cell1, cell100) | ~ spl177_201), inference(avatar_component_clause, [], [f12526])).
fof(f13363, plain, ~ spl177_200, inference(avatar_contradiction_clause, [], [f13362])).
fof(f13362, plain, ($false | ~ spl177_200), inference(subsumption_resolution, [], [f13359, f10377])).
fof(f13359, plain, ((cell23 = cell100) | ~ spl177_200), inference(resolution, [], [f12524, f5630])).
fof(f5630, plain, ! [X0, X1] : (~ sP20(X0, X1) | (cell23 = X1)), inference(cnf_transformation, [], [f5297])).
fof(f5297, plain, ! [X0, X1] : (((cell24 = X0) & (cell23 = X1)) | ~ sP20(X0, X1)), inference(nnf_transformation, [], [f4983])).
fof(f12524, plain, (sP20(cell1, cell100) | ~ spl177_200), inference(avatar_component_clause, [], [f12522])).
fof(f13357, plain, ~ spl177_199, inference(avatar_contradiction_clause, [], [f13356])).
fof(f13356, plain, ($false | ~ spl177_199), inference(subsumption_resolution, [], [f13353, f10420])).
fof(f13353, plain, ((cell22 = cell100) | ~ spl177_199), inference(resolution, [], [f12520, f5632])).
fof(f5632, plain, ! [X0, X1] : (~ sP19(X0, X1) | (cell22 = X1)), inference(cnf_transformation, [], [f5298])).
fof(f5298, plain, ! [X0, X1] : (((cell23 = X0) & (cell22 = X1)) | ~ sP19(X0, X1)), inference(nnf_transformation, [], [f4982])).
fof(f12520, plain, (sP19(cell1, cell100) | ~ spl177_199), inference(avatar_component_clause, [], [f12518])).
fof(f13351, plain, ~ spl177_198, inference(avatar_contradiction_clause, [], [f13350])).
fof(f13350, plain, ($false | ~ spl177_198), inference(subsumption_resolution, [], [f13347, f10399])).
fof(f13347, plain, ((cell21 = cell100) | ~ spl177_198), inference(resolution, [], [f12516, f5634])).
fof(f5634, plain, ! [X0, X1] : (~ sP18(X0, X1) | (cell21 = X1)), inference(cnf_transformation, [], [f5299])).
fof(f5299, plain, ! [X0, X1] : (((cell22 = X0) & (cell21 = X1)) | ~ sP18(X0, X1)), inference(nnf_transformation, [], [f4981])).
fof(f12516, plain, (sP18(cell1, cell100) | ~ spl177_198), inference(avatar_component_clause, [], [f12514])).
fof(f13345, plain, ~ spl177_197, inference(avatar_contradiction_clause, [], [f13344])).
fof(f13344, plain, ($false | ~ spl177_197), inference(subsumption_resolution, [], [f13341, f10459])).
fof(f13341, plain, ((cell19 = cell100) | ~ spl177_197), inference(resolution, [], [f12512, f5636])).
fof(f5636, plain, ! [X0, X1] : (~ sP17(X0, X1) | (cell19 = X1)), inference(cnf_transformation, [], [f5300])).
fof(f5300, plain, ! [X0, X1] : (((cell20 = X0) & (cell19 = X1)) | ~ sP17(X0, X1)), inference(nnf_transformation, [], [f4980])).
fof(f12512, plain, (sP17(cell1, cell100) | ~ spl177_197), inference(avatar_component_clause, [], [f12510])).
fof(f13339, plain, ~ spl177_196, inference(avatar_contradiction_clause, [], [f13338])).
fof(f13338, plain, ($false | ~ spl177_196), inference(subsumption_resolution, [], [f13335, f10477])).
fof(f13335, plain, ((cell18 = cell100) | ~ spl177_196), inference(resolution, [], [f12508, f5638])).
fof(f5638, plain, ! [X0, X1] : (~ sP16(X0, X1) | (cell18 = X1)), inference(cnf_transformation, [], [f5301])).
fof(f5301, plain, ! [X0, X1] : (((cell19 = X0) & (cell18 = X1)) | ~ sP16(X0, X1)), inference(nnf_transformation, [], [f4979])).
fof(f12508, plain, (sP16(cell1, cell100) | ~ spl177_196), inference(avatar_component_clause, [], [f12506])).
fof(f13333, plain, ~ spl177_195, inference(avatar_contradiction_clause, [], [f13332])).
fof(f13332, plain, ($false | ~ spl177_195), inference(subsumption_resolution, [], [f13329, f10494])).
fof(f13329, plain, ((cell17 = cell100) | ~ spl177_195), inference(resolution, [], [f12504, f5640])).
fof(f5640, plain, ! [X0, X1] : (~ sP15(X0, X1) | (cell17 = X1)), inference(cnf_transformation, [], [f5302])).
fof(f5302, plain, ! [X0, X1] : (((cell18 = X0) & (cell17 = X1)) | ~ sP15(X0, X1)), inference(nnf_transformation, [], [f4978])).
fof(f12504, plain, (sP15(cell1, cell100) | ~ spl177_195), inference(avatar_component_clause, [], [f12502])).
fof(f13327, plain, ~ spl177_194, inference(avatar_contradiction_clause, [], [f13326])).
fof(f13326, plain, ($false | ~ spl177_194), inference(subsumption_resolution, [], [f13323, f10510])).
fof(f13323, plain, ((cell16 = cell100) | ~ spl177_194), inference(resolution, [], [f12500, f5642])).
fof(f5642, plain, ! [X0, X1] : (~ sP14(X0, X1) | (cell16 = X1)), inference(cnf_transformation, [], [f5303])).
fof(f5303, plain, ! [X0, X1] : (((cell17 = X0) & (cell16 = X1)) | ~ sP14(X0, X1)), inference(nnf_transformation, [], [f4977])).
fof(f12500, plain, (sP14(cell1, cell100) | ~ spl177_194), inference(avatar_component_clause, [], [f12498])).
fof(f13321, plain, ~ spl177_193, inference(avatar_contradiction_clause, [], [f13320])).
fof(f13320, plain, ($false | ~ spl177_193), inference(subsumption_resolution, [], [f13317, f10525])).
fof(f13317, plain, ((cell15 = cell100) | ~ spl177_193), inference(resolution, [], [f12496, f5644])).
fof(f5644, plain, ! [X0, X1] : (~ sP13(X0, X1) | (cell15 = X1)), inference(cnf_transformation, [], [f5304])).
fof(f5304, plain, ! [X0, X1] : (((cell16 = X0) & (cell15 = X1)) | ~ sP13(X0, X1)), inference(nnf_transformation, [], [f4976])).
fof(f12496, plain, (sP13(cell1, cell100) | ~ spl177_193), inference(avatar_component_clause, [], [f12494])).
fof(f13315, plain, ~ spl177_192, inference(avatar_contradiction_clause, [], [f13314])).
fof(f13314, plain, ($false | ~ spl177_192), inference(subsumption_resolution, [], [f13311, f10539])).
fof(f13311, plain, ((cell14 = cell100) | ~ spl177_192), inference(resolution, [], [f12492, f5646])).
fof(f5646, plain, ! [X0, X1] : (~ sP12(X0, X1) | (cell14 = X1)), inference(cnf_transformation, [], [f5305])).
fof(f5305, plain, ! [X0, X1] : (((cell15 = X0) & (cell14 = X1)) | ~ sP12(X0, X1)), inference(nnf_transformation, [], [f4975])).
fof(f12492, plain, (sP12(cell1, cell100) | ~ spl177_192), inference(avatar_component_clause, [], [f12490])).
fof(f13309, plain, ~ spl177_191, inference(avatar_contradiction_clause, [], [f13308])).
fof(f13308, plain, ($false | ~ spl177_191), inference(subsumption_resolution, [], [f13305, f10552])).
fof(f13305, plain, ((cell13 = cell100) | ~ spl177_191), inference(resolution, [], [f12488, f5648])).
fof(f5648, plain, ! [X0, X1] : (~ sP11(X0, X1) | (cell13 = X1)), inference(cnf_transformation, [], [f5306])).
fof(f5306, plain, ! [X0, X1] : (((cell14 = X0) & (cell13 = X1)) | ~ sP11(X0, X1)), inference(nnf_transformation, [], [f4974])).
fof(f12488, plain, (sP11(cell1, cell100) | ~ spl177_191), inference(avatar_component_clause, [], [f12486])).
fof(f13303, plain, ~ spl177_190, inference(avatar_contradiction_clause, [], [f13302])).
fof(f13302, plain, ($false | ~ spl177_190), inference(subsumption_resolution, [], [f13299, f10575])).
fof(f13299, plain, ((cell12 = cell100) | ~ spl177_190), inference(resolution, [], [f12484, f5650])).
fof(f5650, plain, ! [X0, X1] : (~ sP10(X0, X1) | (cell12 = X1)), inference(cnf_transformation, [], [f5307])).
fof(f5307, plain, ! [X0, X1] : (((cell13 = X0) & (cell12 = X1)) | ~ sP10(X0, X1)), inference(nnf_transformation, [], [f4973])).
fof(f12484, plain, (sP10(cell1, cell100) | ~ spl177_190), inference(avatar_component_clause, [], [f12482])).
fof(f13297, plain, ~ spl177_189, inference(avatar_contradiction_clause, [], [f13296])).
fof(f13296, plain, ($false | ~ spl177_189), inference(subsumption_resolution, [], [f13293, f10564])).
fof(f13293, plain, ((cell11 = cell100) | ~ spl177_189), inference(resolution, [], [f12480, f5652])).
fof(f5652, plain, ! [X0, X1] : (~ sP9(X0, X1) | (cell11 = X1)), inference(cnf_transformation, [], [f5308])).
fof(f5308, plain, ! [X0, X1] : (((cell12 = X0) & (cell11 = X1)) | ~ sP9(X0, X1)), inference(nnf_transformation, [], [f4972])).
fof(f12480, plain, (sP9(cell1, cell100) | ~ spl177_189), inference(avatar_component_clause, [], [f12478])).
fof(f13291, plain, ~ spl177_188, inference(avatar_contradiction_clause, [], [f13290])).
fof(f13290, plain, ($false | ~ spl177_188), inference(subsumption_resolution, [], [f13287, f10594])).
fof(f13287, plain, ((cell9 = cell100) | ~ spl177_188), inference(resolution, [], [f12476, f5654])).
fof(f5654, plain, ! [X0, X1] : (~ sP8(X0, X1) | (cell9 = X1)), inference(cnf_transformation, [], [f5309])).
fof(f5309, plain, ! [X0, X1] : (((cell10 = X0) & (cell9 = X1)) | ~ sP8(X0, X1)), inference(nnf_transformation, [], [f4971])).
fof(f12476, plain, (sP8(cell1, cell100) | ~ spl177_188), inference(avatar_component_clause, [], [f12474])).
fof(f13285, plain, ~ spl177_187, inference(avatar_contradiction_clause, [], [f13284])).
fof(f13284, plain, ($false | ~ spl177_187), inference(subsumption_resolution, [], [f13281, f10602])).
fof(f13281, plain, ((cell8 = cell100) | ~ spl177_187), inference(resolution, [], [f12472, f5656])).
fof(f5656, plain, ! [X0, X1] : (~ sP7(X0, X1) | (cell8 = X1)), inference(cnf_transformation, [], [f5310])).
fof(f5310, plain, ! [X0, X1] : (((cell9 = X0) & (cell8 = X1)) | ~ sP7(X0, X1)), inference(nnf_transformation, [], [f4970])).
fof(f12472, plain, (sP7(cell1, cell100) | ~ spl177_187), inference(avatar_component_clause, [], [f12470])).
fof(f13279, plain, ~ spl177_186, inference(avatar_contradiction_clause, [], [f13278])).
fof(f13278, plain, ($false | ~ spl177_186), inference(subsumption_resolution, [], [f13275, f10609])).
fof(f13275, plain, ((cell7 = cell100) | ~ spl177_186), inference(resolution, [], [f12468, f5658])).
fof(f5658, plain, ! [X0, X1] : (~ sP6(X0, X1) | (cell7 = X1)), inference(cnf_transformation, [], [f5311])).
fof(f5311, plain, ! [X0, X1] : (((cell8 = X0) & (cell7 = X1)) | ~ sP6(X0, X1)), inference(nnf_transformation, [], [f4969])).
fof(f12468, plain, (sP6(cell1, cell100) | ~ spl177_186), inference(avatar_component_clause, [], [f12466])).
fof(f13273, plain, ~ spl177_185, inference(avatar_contradiction_clause, [], [f13272])).
fof(f13272, plain, ($false | ~ spl177_185), inference(subsumption_resolution, [], [f13269, f10615])).
fof(f13269, plain, ((cell6 = cell100) | ~ spl177_185), inference(resolution, [], [f12464, f5660])).
fof(f5660, plain, ! [X0, X1] : (~ sP5(X0, X1) | (cell6 = X1)), inference(cnf_transformation, [], [f5312])).
fof(f5312, plain, ! [X0, X1] : (((cell7 = X0) & (cell6 = X1)) | ~ sP5(X0, X1)), inference(nnf_transformation, [], [f4968])).
fof(f12464, plain, (sP5(cell1, cell100) | ~ spl177_185), inference(avatar_component_clause, [], [f12462])).
fof(f13267, plain, ~ spl177_184, inference(avatar_contradiction_clause, [], [f13266])).
fof(f13266, plain, ($false | ~ spl177_184), inference(subsumption_resolution, [], [f13263, f10620])).
fof(f13263, plain, ((cell5 = cell100) | ~ spl177_184), inference(resolution, [], [f12460, f5662])).
fof(f5662, plain, ! [X0, X1] : (~ sP4(X0, X1) | (cell5 = X1)), inference(cnf_transformation, [], [f5313])).
fof(f5313, plain, ! [X0, X1] : (((cell6 = X0) & (cell5 = X1)) | ~ sP4(X0, X1)), inference(nnf_transformation, [], [f4967])).
fof(f12460, plain, (sP4(cell1, cell100) | ~ spl177_184), inference(avatar_component_clause, [], [f12458])).
fof(f13261, plain, ~ spl177_183, inference(avatar_contradiction_clause, [], [f13260])).
fof(f13260, plain, ($false | ~ spl177_183), inference(subsumption_resolution, [], [f13257, f10624])).
fof(f13257, plain, ((cell4 = cell100) | ~ spl177_183), inference(resolution, [], [f12456, f5664])).
fof(f5664, plain, ! [X0, X1] : (~ sP3(X0, X1) | (cell4 = X1)), inference(cnf_transformation, [], [f5314])).
fof(f5314, plain, ! [X0, X1] : (((cell5 = X0) & (cell4 = X1)) | ~ sP3(X0, X1)), inference(nnf_transformation, [], [f4966])).
fof(f12456, plain, (sP3(cell1, cell100) | ~ spl177_183), inference(avatar_component_clause, [], [f12454])).
fof(f13255, plain, ~ spl177_182, inference(avatar_contradiction_clause, [], [f13254])).
fof(f13254, plain, ($false | ~ spl177_182), inference(subsumption_resolution, [], [f13251, f10627])).
fof(f13251, plain, ((cell3 = cell100) | ~ spl177_182), inference(resolution, [], [f12452, f5666])).
fof(f5666, plain, ! [X0, X1] : (~ sP2(X0, X1) | (cell3 = X1)), inference(cnf_transformation, [], [f5315])).
fof(f5315, plain, ! [X0, X1] : (((cell4 = X0) & (cell3 = X1)) | ~ sP2(X0, X1)), inference(nnf_transformation, [], [f4965])).
fof(f12452, plain, (sP2(cell1, cell100) | ~ spl177_182), inference(avatar_component_clause, [], [f12450])).
fof(f13249, plain, ~ spl177_181, inference(avatar_contradiction_clause, [], [f13248])).
fof(f13248, plain, ($false | ~ spl177_181), inference(subsumption_resolution, [], [f13245, f10629])).
fof(f13245, plain, ((cell2 = cell100) | ~ spl177_181), inference(resolution, [], [f12448, f5668])).
fof(f5668, plain, ! [X0, X1] : (~ sP1(X0, X1) | (cell2 = X1)), inference(cnf_transformation, [], [f5316])).
fof(f5316, plain, ! [X0, X1] : (((cell3 = X0) & (cell2 = X1)) | ~ sP1(X0, X1)), inference(nnf_transformation, [], [f4964])).
fof(f12448, plain, (sP1(cell1, cell100) | ~ spl177_181), inference(avatar_component_clause, [], [f12446])).
fof(f13211, plain, ~ spl177_180, inference(avatar_contradiction_clause, [], [f13210])).
fof(f13210, plain, ($false | ~ spl177_180), inference(subsumption_resolution, [], [f13207, f10630])).
fof(f13207, plain, ((cell1 = cell100) | ~ spl177_180), inference(resolution, [], [f12444, f5670])).
fof(f5670, plain, ! [X0, X1] : (~ sP0(X0, X1) | (cell1 = X1)), inference(cnf_transformation, [], [f5317])).
fof(f5317, plain, ! [X0, X1] : (((cell2 = X0) & (cell1 = X1)) | ~ sP0(X0, X1)), inference(nnf_transformation, [], [f4963])).
fof(f12444, plain, (sP0(cell1, cell100) | ~ spl177_180), inference(avatar_component_clause, [], [f12442])).
fof(f13205, plain, ~ spl177_178, inference(avatar_contradiction_clause, [], [f13204])).
fof(f13204, plain, ($false | ~ spl177_178), inference(subsumption_resolution, [], [f13201, f5876])).
fof(f13201, plain, ((cell1 = cell98) | ~ spl177_178), inference(resolution, [], [f11390, f5494])).
fof(f5494, plain, ! [X0, X1] : (~ sP88(X0, X1) | (cell98 = X1)), inference(cnf_transformation, [], [f5229])).
fof(f11390, plain, (sP88(cell100, cell1) | ~ spl177_178), inference(avatar_component_clause, [], [f11388])).
fof(f11388, plain, (spl177_178 <=> sP88(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_178])])).
fof(f12429, plain, ~ spl177_177, inference(avatar_contradiction_clause, [], [f12428])).
fof(f12428, plain, ($false | ~ spl177_177), inference(subsumption_resolution, [], [f12425, f6888])).
fof(f6888, plain, ~ (cell1 = cell87), inference(cnf_transformation, [], [f1210])).
fof(f1210, plain, ~ (cell1 = cell87), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26199)).
fof(f12425, plain, ((cell1 = cell87) | ~ spl177_177), inference(resolution, [], [f11369, f5318])).
fof(f11369, plain, (sP176(cell100, cell1) | ~ spl177_177), inference(avatar_component_clause, [], [f11367])).
fof(f11367, plain, (spl177_177 <=> sP176(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_177])])).
fof(f12423, plain, ~ spl177_176, inference(avatar_contradiction_clause, [], [f12422])).
fof(f12422, plain, ($false | ~ spl177_176), inference(subsumption_resolution, [], [f12419, f6974])).
fof(f6974, plain, ~ (cell1 = cell86), inference(cnf_transformation, [], [f1296])).
fof(f1296, plain, ~ (cell1 = cell86), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26285)).
fof(f12419, plain, ((cell1 = cell86) | ~ spl177_176), inference(resolution, [], [f11365, f5320])).
fof(f11365, plain, (sP175(cell100, cell1) | ~ spl177_176), inference(avatar_component_clause, [], [f11363])).
fof(f11363, plain, (spl177_176 <=> sP175(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_176])])).
fof(f12417, plain, ~ spl177_175, inference(avatar_contradiction_clause, [], [f12416])).
fof(f12416, plain, ($false | ~ spl177_175), inference(subsumption_resolution, [], [f12413, f7059])).
fof(f7059, plain, ~ (cell1 = cell85), inference(cnf_transformation, [], [f1381])).
fof(f1381, plain, ~ (cell1 = cell85), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26370)).
fof(f12413, plain, ((cell1 = cell85) | ~ spl177_175), inference(resolution, [], [f11361, f5322])).
fof(f11361, plain, (sP174(cell100, cell1) | ~ spl177_175), inference(avatar_component_clause, [], [f11359])).
fof(f11359, plain, (spl177_175 <=> sP174(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_175])])).
fof(f12411, plain, ~ spl177_174, inference(avatar_contradiction_clause, [], [f12410])).
fof(f12410, plain, ($false | ~ spl177_174), inference(subsumption_resolution, [], [f12407, f7143])).
fof(f7143, plain, ~ (cell1 = cell84), inference(cnf_transformation, [], [f1465])).
fof(f1465, plain, ~ (cell1 = cell84), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26454)).
fof(f12407, plain, ((cell1 = cell84) | ~ spl177_174), inference(resolution, [], [f11357, f5324])).
fof(f11357, plain, (sP173(cell100, cell1) | ~ spl177_174), inference(avatar_component_clause, [], [f11355])).
fof(f11355, plain, (spl177_174 <=> sP173(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_174])])).
fof(f12405, plain, ~ spl177_173, inference(avatar_contradiction_clause, [], [f12404])).
fof(f12404, plain, ($false | ~ spl177_173), inference(subsumption_resolution, [], [f12401, f7226])).
fof(f7226, plain, ~ (cell1 = cell83), inference(cnf_transformation, [], [f1548])).
fof(f1548, plain, ~ (cell1 = cell83), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26537)).
fof(f12401, plain, ((cell1 = cell83) | ~ spl177_173), inference(resolution, [], [f11353, f5326])).
fof(f11353, plain, (sP172(cell100, cell1) | ~ spl177_173), inference(avatar_component_clause, [], [f11351])).
fof(f11351, plain, (spl177_173 <=> sP172(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_173])])).
fof(f12399, plain, ~ spl177_172, inference(avatar_contradiction_clause, [], [f12398])).
fof(f12398, plain, ($false | ~ spl177_172), inference(subsumption_resolution, [], [f12395, f7389])).
fof(f7389, plain, ~ (cell1 = cell82), inference(cnf_transformation, [], [f1711])).
fof(f1711, plain, ~ (cell1 = cell82), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26700)).
fof(f12395, plain, ((cell1 = cell82) | ~ spl177_172), inference(resolution, [], [f11349, f5328])).
fof(f11349, plain, (sP171(cell100, cell1) | ~ spl177_172), inference(avatar_component_clause, [], [f11347])).
fof(f11347, plain, (spl177_172 <=> sP171(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_172])])).
fof(f12393, plain, ~ spl177_171, inference(avatar_contradiction_clause, [], [f12392])).
fof(f12392, plain, ($false | ~ spl177_171), inference(subsumption_resolution, [], [f12389, f7308])).
fof(f7308, plain, ~ (cell1 = cell81), inference(cnf_transformation, [], [f1630])).
fof(f1630, plain, ~ (cell1 = cell81), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26619)).
fof(f12389, plain, ((cell1 = cell81) | ~ spl177_171), inference(resolution, [], [f11345, f5330])).
fof(f11345, plain, (sP170(cell100, cell1) | ~ spl177_171), inference(avatar_component_clause, [], [f11343])).
fof(f11343, plain, (spl177_171 <=> sP170(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_171])])).
fof(f12387, plain, ~ spl177_170, inference(avatar_contradiction_clause, [], [f12386])).
fof(f12386, plain, ($false | ~ spl177_170), inference(subsumption_resolution, [], [f12383, f7469])).
fof(f7469, plain, ~ (cell1 = cell80), inference(cnf_transformation, [], [f1791])).
fof(f1791, plain, ~ (cell1 = cell80), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26780)).
fof(f12383, plain, ((cell1 = cell80) | ~ spl177_170), inference(resolution, [], [f11341, f5332])).
fof(f11341, plain, (sP169(cell100, cell1) | ~ spl177_170), inference(avatar_component_clause, [], [f11339])).
fof(f11339, plain, (spl177_170 <=> sP169(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_170])])).
fof(f12381, plain, ~ spl177_169, inference(avatar_contradiction_clause, [], [f12380])).
fof(f12380, plain, ($false | ~ spl177_169), inference(subsumption_resolution, [], [f12377, f7548])).
fof(f7548, plain, ~ (cell1 = cell79), inference(cnf_transformation, [], [f1870])).
fof(f1870, plain, ~ (cell1 = cell79), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26859)).
fof(f12377, plain, ((cell1 = cell79) | ~ spl177_169), inference(resolution, [], [f11337, f5334])).
fof(f11337, plain, (sP168(cell100, cell1) | ~ spl177_169), inference(avatar_component_clause, [], [f11335])).
fof(f11335, plain, (spl177_169 <=> sP168(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_169])])).
fof(f12375, plain, ~ spl177_168, inference(avatar_contradiction_clause, [], [f12374])).
fof(f12374, plain, ($false | ~ spl177_168), inference(subsumption_resolution, [], [f12371, f7626])).
fof(f7626, plain, ~ (cell1 = cell78), inference(cnf_transformation, [], [f1948])).
fof(f1948, plain, ~ (cell1 = cell78), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26937)).
fof(f12371, plain, ((cell1 = cell78) | ~ spl177_168), inference(resolution, [], [f11333, f5336])).
fof(f11333, plain, (sP167(cell100, cell1) | ~ spl177_168), inference(avatar_component_clause, [], [f11331])).
fof(f11331, plain, (spl177_168 <=> sP167(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_168])])).
fof(f12369, plain, ~ spl177_167, inference(avatar_contradiction_clause, [], [f12368])).
fof(f12368, plain, ($false | ~ spl177_167), inference(subsumption_resolution, [], [f12365, f7703])).
fof(f7703, plain, ~ (cell1 = cell77), inference(cnf_transformation, [], [f2025])).
fof(f2025, plain, ~ (cell1 = cell77), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27014)).
fof(f12365, plain, ((cell1 = cell77) | ~ spl177_167), inference(resolution, [], [f11329, f5338])).
fof(f11329, plain, (sP166(cell100, cell1) | ~ spl177_167), inference(avatar_component_clause, [], [f11327])).
fof(f11327, plain, (spl177_167 <=> sP166(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_167])])).
fof(f12363, plain, ~ spl177_166, inference(avatar_contradiction_clause, [], [f12362])).
fof(f12362, plain, ($false | ~ spl177_166), inference(subsumption_resolution, [], [f12359, f7779])).
fof(f7779, plain, ~ (cell1 = cell76), inference(cnf_transformation, [], [f2101])).
fof(f2101, plain, ~ (cell1 = cell76), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27090)).
fof(f12359, plain, ((cell1 = cell76) | ~ spl177_166), inference(resolution, [], [f11325, f5340])).
fof(f11325, plain, (sP165(cell100, cell1) | ~ spl177_166), inference(avatar_component_clause, [], [f11323])).
fof(f11323, plain, (spl177_166 <=> sP165(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_166])])).
fof(f12357, plain, ~ spl177_165, inference(avatar_contradiction_clause, [], [f12356])).
fof(f12356, plain, ($false | ~ spl177_165), inference(subsumption_resolution, [], [f12353, f7854])).
fof(f7854, plain, ~ (cell1 = cell75), inference(cnf_transformation, [], [f2176])).
fof(f2176, plain, ~ (cell1 = cell75), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27165)).
fof(f12353, plain, ((cell1 = cell75) | ~ spl177_165), inference(resolution, [], [f11321, f5342])).
fof(f11321, plain, (sP164(cell100, cell1) | ~ spl177_165), inference(avatar_component_clause, [], [f11319])).
fof(f11319, plain, (spl177_165 <=> sP164(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_165])])).
fof(f12351, plain, ~ spl177_164, inference(avatar_contradiction_clause, [], [f12350])).
fof(f12350, plain, ($false | ~ spl177_164), inference(subsumption_resolution, [], [f12347, f7928])).
fof(f7928, plain, ~ (cell1 = cell74), inference(cnf_transformation, [], [f2250])).
fof(f2250, plain, ~ (cell1 = cell74), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27239)).
fof(f12347, plain, ((cell1 = cell74) | ~ spl177_164), inference(resolution, [], [f11317, f5344])).
fof(f11317, plain, (sP163(cell100, cell1) | ~ spl177_164), inference(avatar_component_clause, [], [f11315])).
fof(f11315, plain, (spl177_164 <=> sP163(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_164])])).
fof(f12345, plain, ~ spl177_163, inference(avatar_contradiction_clause, [], [f12344])).
fof(f12344, plain, ($false | ~ spl177_163), inference(subsumption_resolution, [], [f12341, f8001])).
fof(f8001, plain, ~ (cell1 = cell73), inference(cnf_transformation, [], [f2323])).
fof(f2323, plain, ~ (cell1 = cell73), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27312)).
fof(f12341, plain, ((cell1 = cell73) | ~ spl177_163), inference(resolution, [], [f11313, f5346])).
fof(f11313, plain, (sP162(cell100, cell1) | ~ spl177_163), inference(avatar_component_clause, [], [f11311])).
fof(f11311, plain, (spl177_163 <=> sP162(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_163])])).
fof(f12339, plain, ~ spl177_162, inference(avatar_contradiction_clause, [], [f12338])).
fof(f12338, plain, ($false | ~ spl177_162), inference(subsumption_resolution, [], [f12335, f8144])).
fof(f8144, plain, ~ (cell1 = cell72), inference(cnf_transformation, [], [f2466])).
fof(f2466, plain, ~ (cell1 = cell72), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27455)).
fof(f12335, plain, ((cell1 = cell72) | ~ spl177_162), inference(resolution, [], [f11309, f5348])).
fof(f11309, plain, (sP161(cell100, cell1) | ~ spl177_162), inference(avatar_component_clause, [], [f11307])).
fof(f11307, plain, (spl177_162 <=> sP161(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_162])])).
fof(f12333, plain, ~ spl177_161, inference(avatar_contradiction_clause, [], [f12332])).
fof(f12332, plain, ($false | ~ spl177_161), inference(subsumption_resolution, [], [f12329, f8073])).
fof(f8073, plain, ~ (cell1 = cell71), inference(cnf_transformation, [], [f2395])).
fof(f2395, plain, ~ (cell1 = cell71), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27384)).
fof(f12329, plain, ((cell1 = cell71) | ~ spl177_161), inference(resolution, [], [f11305, f5350])).
fof(f11305, plain, (sP160(cell100, cell1) | ~ spl177_161), inference(avatar_component_clause, [], [f11303])).
fof(f11303, plain, (spl177_161 <=> sP160(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_161])])).
fof(f12327, plain, ~ spl177_160, inference(avatar_contradiction_clause, [], [f12326])).
fof(f12326, plain, ($false | ~ spl177_160), inference(subsumption_resolution, [], [f12323, f8214])).
fof(f8214, plain, ~ (cell1 = cell70), inference(cnf_transformation, [], [f2536])).
fof(f2536, plain, ~ (cell1 = cell70), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27525)).
fof(f12323, plain, ((cell1 = cell70) | ~ spl177_160), inference(resolution, [], [f11301, f5352])).
fof(f11301, plain, (sP159(cell100, cell1) | ~ spl177_160), inference(avatar_component_clause, [], [f11299])).
fof(f11299, plain, (spl177_160 <=> sP159(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_160])])).
fof(f12321, plain, ~ spl177_159, inference(avatar_contradiction_clause, [], [f12320])).
fof(f12320, plain, ($false | ~ spl177_159), inference(subsumption_resolution, [], [f12317, f8283])).
fof(f8283, plain, ~ (cell1 = cell69), inference(cnf_transformation, [], [f2605])).
fof(f2605, plain, ~ (cell1 = cell69), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27594)).
fof(f12317, plain, ((cell1 = cell69) | ~ spl177_159), inference(resolution, [], [f11297, f5354])).
fof(f11297, plain, (sP158(cell100, cell1) | ~ spl177_159), inference(avatar_component_clause, [], [f11295])).
fof(f11295, plain, (spl177_159 <=> sP158(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_159])])).
fof(f12315, plain, ~ spl177_158, inference(avatar_contradiction_clause, [], [f12314])).
fof(f12314, plain, ($false | ~ spl177_158), inference(subsumption_resolution, [], [f12311, f8351])).
fof(f8351, plain, ~ (cell1 = cell68), inference(cnf_transformation, [], [f2673])).
fof(f2673, plain, ~ (cell1 = cell68), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27662)).
fof(f12311, plain, ((cell1 = cell68) | ~ spl177_158), inference(resolution, [], [f11293, f5356])).
fof(f11293, plain, (sP157(cell100, cell1) | ~ spl177_158), inference(avatar_component_clause, [], [f11291])).
fof(f11291, plain, (spl177_158 <=> sP157(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_158])])).
fof(f12309, plain, ~ spl177_157, inference(avatar_contradiction_clause, [], [f12308])).
fof(f12308, plain, ($false | ~ spl177_157), inference(subsumption_resolution, [], [f12305, f8418])).
fof(f8418, plain, ~ (cell1 = cell67), inference(cnf_transformation, [], [f2740])).
fof(f2740, plain, ~ (cell1 = cell67), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27729)).
fof(f12305, plain, ((cell1 = cell67) | ~ spl177_157), inference(resolution, [], [f11289, f5358])).
fof(f11289, plain, (sP156(cell100, cell1) | ~ spl177_157), inference(avatar_component_clause, [], [f11287])).
fof(f11287, plain, (spl177_157 <=> sP156(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_157])])).
fof(f12303, plain, ~ spl177_156, inference(avatar_contradiction_clause, [], [f12302])).
fof(f12302, plain, ($false | ~ spl177_156), inference(subsumption_resolution, [], [f12299, f8484])).
fof(f8484, plain, ~ (cell1 = cell66), inference(cnf_transformation, [], [f2806])).
fof(f2806, plain, ~ (cell1 = cell66), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27795)).
fof(f12299, plain, ((cell1 = cell66) | ~ spl177_156), inference(resolution, [], [f11285, f5360])).
fof(f11285, plain, (sP155(cell100, cell1) | ~ spl177_156), inference(avatar_component_clause, [], [f11283])).
fof(f11283, plain, (spl177_156 <=> sP155(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_156])])).
fof(f12297, plain, ~ spl177_155, inference(avatar_contradiction_clause, [], [f12296])).
fof(f12296, plain, ($false | ~ spl177_155), inference(subsumption_resolution, [], [f12293, f8549])).
fof(f8549, plain, ~ (cell1 = cell65), inference(cnf_transformation, [], [f2871])).
fof(f2871, plain, ~ (cell1 = cell65), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27860)).
fof(f12293, plain, ((cell1 = cell65) | ~ spl177_155), inference(resolution, [], [f11281, f5362])).
fof(f11281, plain, (sP154(cell100, cell1) | ~ spl177_155), inference(avatar_component_clause, [], [f11279])).
fof(f11279, plain, (spl177_155 <=> sP154(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_155])])).
fof(f12291, plain, ~ spl177_154, inference(avatar_contradiction_clause, [], [f12290])).
fof(f12290, plain, ($false | ~ spl177_154), inference(subsumption_resolution, [], [f12287, f8613])).
fof(f8613, plain, ~ (cell1 = cell64), inference(cnf_transformation, [], [f2935])).
fof(f2935, plain, ~ (cell1 = cell64), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27924)).
fof(f12287, plain, ((cell1 = cell64) | ~ spl177_154), inference(resolution, [], [f11277, f5364])).
fof(f11277, plain, (sP153(cell100, cell1) | ~ spl177_154), inference(avatar_component_clause, [], [f11275])).
fof(f11275, plain, (spl177_154 <=> sP153(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_154])])).
fof(f12285, plain, ~ spl177_153, inference(avatar_contradiction_clause, [], [f12284])).
fof(f12284, plain, ($false | ~ spl177_153), inference(subsumption_resolution, [], [f12281, f8676])).
fof(f8676, plain, ~ (cell1 = cell63), inference(cnf_transformation, [], [f2998])).
fof(f2998, plain, ~ (cell1 = cell63), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof27987)).
fof(f12281, plain, ((cell1 = cell63) | ~ spl177_153), inference(resolution, [], [f11273, f5366])).
fof(f11273, plain, (sP152(cell100, cell1) | ~ spl177_153), inference(avatar_component_clause, [], [f11271])).
fof(f11271, plain, (spl177_153 <=> sP152(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_153])])).
fof(f12279, plain, ~ spl177_152, inference(avatar_contradiction_clause, [], [f12278])).
fof(f12278, plain, ($false | ~ spl177_152), inference(subsumption_resolution, [], [f12275, f8799])).
fof(f8799, plain, ~ (cell1 = cell62), inference(cnf_transformation, [], [f3121])).
fof(f3121, plain, ~ (cell1 = cell62), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28110)).
fof(f12275, plain, ((cell1 = cell62) | ~ spl177_152), inference(resolution, [], [f11269, f5368])).
fof(f11269, plain, (sP151(cell100, cell1) | ~ spl177_152), inference(avatar_component_clause, [], [f11267])).
fof(f11267, plain, (spl177_152 <=> sP151(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_152])])).
fof(f12273, plain, ~ spl177_151, inference(avatar_contradiction_clause, [], [f12272])).
fof(f12272, plain, ($false | ~ spl177_151), inference(subsumption_resolution, [], [f12269, f8738])).
fof(f8738, plain, ~ (cell1 = cell61), inference(cnf_transformation, [], [f3060])).
fof(f3060, plain, ~ (cell1 = cell61), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28049)).
fof(f12269, plain, ((cell1 = cell61) | ~ spl177_151), inference(resolution, [], [f11265, f5370])).
fof(f11265, plain, (sP150(cell100, cell1) | ~ spl177_151), inference(avatar_component_clause, [], [f11263])).
fof(f11263, plain, (spl177_151 <=> sP150(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_151])])).
fof(f12267, plain, ~ spl177_150, inference(avatar_contradiction_clause, [], [f12266])).
fof(f12266, plain, ($false | ~ spl177_150), inference(subsumption_resolution, [], [f12263, f8859])).
fof(f8859, plain, ~ (cell1 = cell60), inference(cnf_transformation, [], [f3181])).
fof(f3181, plain, ~ (cell1 = cell60), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28170)).
fof(f12263, plain, ((cell1 = cell60) | ~ spl177_150), inference(resolution, [], [f11261, f5372])).
fof(f11261, plain, (sP149(cell100, cell1) | ~ spl177_150), inference(avatar_component_clause, [], [f11259])).
fof(f11259, plain, (spl177_150 <=> sP149(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_150])])).
fof(f12261, plain, ~ spl177_149, inference(avatar_contradiction_clause, [], [f12260])).
fof(f12260, plain, ($false | ~ spl177_149), inference(subsumption_resolution, [], [f12257, f8918])).
fof(f8918, plain, ~ (cell1 = cell59), inference(cnf_transformation, [], [f3240])).
fof(f3240, plain, ~ (cell1 = cell59), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28229)).
fof(f12257, plain, ((cell1 = cell59) | ~ spl177_149), inference(resolution, [], [f11257, f5374])).
fof(f11257, plain, (sP148(cell100, cell1) | ~ spl177_149), inference(avatar_component_clause, [], [f11255])).
fof(f11255, plain, (spl177_149 <=> sP148(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_149])])).
fof(f12255, plain, ~ spl177_148, inference(avatar_contradiction_clause, [], [f12254])).
fof(f12254, plain, ($false | ~ spl177_148), inference(subsumption_resolution, [], [f12251, f8976])).
fof(f8976, plain, ~ (cell1 = cell58), inference(cnf_transformation, [], [f3298])).
fof(f3298, plain, ~ (cell1 = cell58), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28287)).
fof(f12251, plain, ((cell1 = cell58) | ~ spl177_148), inference(resolution, [], [f11253, f5376])).
fof(f11253, plain, (sP147(cell100, cell1) | ~ spl177_148), inference(avatar_component_clause, [], [f11251])).
fof(f11251, plain, (spl177_148 <=> sP147(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_148])])).
fof(f12249, plain, ~ spl177_147, inference(avatar_contradiction_clause, [], [f12248])).
fof(f12248, plain, ($false | ~ spl177_147), inference(subsumption_resolution, [], [f12245, f9033])).
fof(f9033, plain, ~ (cell1 = cell57), inference(cnf_transformation, [], [f3355])).
fof(f3355, plain, ~ (cell1 = cell57), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28344)).
fof(f12245, plain, ((cell1 = cell57) | ~ spl177_147), inference(resolution, [], [f11249, f5378])).
fof(f11249, plain, (sP146(cell100, cell1) | ~ spl177_147), inference(avatar_component_clause, [], [f11247])).
fof(f11247, plain, (spl177_147 <=> sP146(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_147])])).
fof(f12243, plain, ~ spl177_146, inference(avatar_contradiction_clause, [], [f12242])).
fof(f12242, plain, ($false | ~ spl177_146), inference(subsumption_resolution, [], [f12239, f9089])).
fof(f9089, plain, ~ (cell1 = cell56), inference(cnf_transformation, [], [f3411])).
fof(f3411, plain, ~ (cell1 = cell56), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28400)).
fof(f12239, plain, ((cell1 = cell56) | ~ spl177_146), inference(resolution, [], [f11245, f5380])).
fof(f11245, plain, (sP145(cell100, cell1) | ~ spl177_146), inference(avatar_component_clause, [], [f11243])).
fof(f11243, plain, (spl177_146 <=> sP145(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_146])])).
fof(f12237, plain, ~ spl177_145, inference(avatar_contradiction_clause, [], [f12236])).
fof(f12236, plain, ($false | ~ spl177_145), inference(subsumption_resolution, [], [f12233, f9144])).
fof(f9144, plain, ~ (cell1 = cell55), inference(cnf_transformation, [], [f3466])).
fof(f3466, plain, ~ (cell1 = cell55), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28455)).
fof(f12233, plain, ((cell1 = cell55) | ~ spl177_145), inference(resolution, [], [f11241, f5382])).
fof(f11241, plain, (sP144(cell100, cell1) | ~ spl177_145), inference(avatar_component_clause, [], [f11239])).
fof(f11239, plain, (spl177_145 <=> sP144(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_145])])).
fof(f12231, plain, ~ spl177_144, inference(avatar_contradiction_clause, [], [f12230])).
fof(f12230, plain, ($false | ~ spl177_144), inference(subsumption_resolution, [], [f12227, f9198])).
fof(f9198, plain, ~ (cell1 = cell54), inference(cnf_transformation, [], [f3520])).
fof(f3520, plain, ~ (cell1 = cell54), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28509)).
fof(f12227, plain, ((cell1 = cell54) | ~ spl177_144), inference(resolution, [], [f11237, f5384])).
fof(f11237, plain, (sP143(cell100, cell1) | ~ spl177_144), inference(avatar_component_clause, [], [f11235])).
fof(f11235, plain, (spl177_144 <=> sP143(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_144])])).
fof(f12225, plain, ~ spl177_143, inference(avatar_contradiction_clause, [], [f12224])).
fof(f12224, plain, ($false | ~ spl177_143), inference(subsumption_resolution, [], [f12221, f9251])).
fof(f9251, plain, ~ (cell1 = cell53), inference(cnf_transformation, [], [f3573])).
fof(f3573, plain, ~ (cell1 = cell53), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28562)).
fof(f12221, plain, ((cell1 = cell53) | ~ spl177_143), inference(resolution, [], [f11233, f5386])).
fof(f11233, plain, (sP142(cell100, cell1) | ~ spl177_143), inference(avatar_component_clause, [], [f11231])).
fof(f11231, plain, (spl177_143 <=> sP142(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_143])])).
fof(f12219, plain, ~ spl177_142, inference(avatar_contradiction_clause, [], [f12218])).
fof(f12218, plain, ($false | ~ spl177_142), inference(subsumption_resolution, [], [f12215, f9354])).
fof(f9354, plain, ~ (cell1 = cell52), inference(cnf_transformation, [], [f3676])).
fof(f3676, plain, ~ (cell1 = cell52), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28665)).
fof(f12215, plain, ((cell1 = cell52) | ~ spl177_142), inference(resolution, [], [f11229, f5388])).
fof(f11229, plain, (sP141(cell100, cell1) | ~ spl177_142), inference(avatar_component_clause, [], [f11227])).
fof(f11227, plain, (spl177_142 <=> sP141(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_142])])).
fof(f12213, plain, ~ spl177_141, inference(avatar_contradiction_clause, [], [f12212])).
fof(f12212, plain, ($false | ~ spl177_141), inference(subsumption_resolution, [], [f12209, f9303])).
fof(f9303, plain, ~ (cell1 = cell51), inference(cnf_transformation, [], [f3625])).
fof(f3625, plain, ~ (cell1 = cell51), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28614)).
fof(f12209, plain, ((cell1 = cell51) | ~ spl177_141), inference(resolution, [], [f11225, f5390])).
fof(f11225, plain, (sP140(cell100, cell1) | ~ spl177_141), inference(avatar_component_clause, [], [f11223])).
fof(f11223, plain, (spl177_141 <=> sP140(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_141])])).
fof(f12207, plain, ~ spl177_140, inference(avatar_contradiction_clause, [], [f12206])).
fof(f12206, plain, ($false | ~ spl177_140), inference(subsumption_resolution, [], [f12203, f9404])).
fof(f9404, plain, ~ (cell1 = cell50), inference(cnf_transformation, [], [f3726])).
fof(f3726, plain, ~ (cell1 = cell50), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28715)).
fof(f12203, plain, ((cell1 = cell50) | ~ spl177_140), inference(resolution, [], [f11221, f5392])).
fof(f11221, plain, (sP139(cell100, cell1) | ~ spl177_140), inference(avatar_component_clause, [], [f11219])).
fof(f11219, plain, (spl177_140 <=> sP139(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_140])])).
fof(f12201, plain, ~ spl177_139, inference(avatar_contradiction_clause, [], [f12200])).
fof(f12200, plain, ($false | ~ spl177_139), inference(subsumption_resolution, [], [f12197, f9453])).
fof(f9453, plain, ~ (cell1 = cell49), inference(cnf_transformation, [], [f3775])).
fof(f3775, plain, ~ (cell1 = cell49), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28764)).
fof(f12197, plain, ((cell1 = cell49) | ~ spl177_139), inference(resolution, [], [f11217, f5394])).
fof(f11217, plain, (sP138(cell100, cell1) | ~ spl177_139), inference(avatar_component_clause, [], [f11215])).
fof(f11215, plain, (spl177_139 <=> sP138(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_139])])).
fof(f12195, plain, ~ spl177_138, inference(avatar_contradiction_clause, [], [f12194])).
fof(f12194, plain, ($false | ~ spl177_138), inference(subsumption_resolution, [], [f12191, f9501])).
fof(f9501, plain, ~ (cell1 = cell48), inference(cnf_transformation, [], [f3823])).
fof(f3823, plain, ~ (cell1 = cell48), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28812)).
fof(f12191, plain, ((cell1 = cell48) | ~ spl177_138), inference(resolution, [], [f11213, f5396])).
fof(f11213, plain, (sP137(cell100, cell1) | ~ spl177_138), inference(avatar_component_clause, [], [f11211])).
fof(f11211, plain, (spl177_138 <=> sP137(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_138])])).
fof(f12189, plain, ~ spl177_137, inference(avatar_contradiction_clause, [], [f12188])).
fof(f12188, plain, ($false | ~ spl177_137), inference(subsumption_resolution, [], [f12185, f9548])).
fof(f9548, plain, ~ (cell1 = cell47), inference(cnf_transformation, [], [f3870])).
fof(f3870, plain, ~ (cell1 = cell47), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28859)).
fof(f12185, plain, ((cell1 = cell47) | ~ spl177_137), inference(resolution, [], [f11209, f5398])).
fof(f11209, plain, (sP136(cell100, cell1) | ~ spl177_137), inference(avatar_component_clause, [], [f11207])).
fof(f11207, plain, (spl177_137 <=> sP136(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_137])])).
fof(f12183, plain, ~ spl177_136, inference(avatar_contradiction_clause, [], [f12182])).
fof(f12182, plain, ($false | ~ spl177_136), inference(subsumption_resolution, [], [f12179, f9594])).
fof(f9594, plain, ~ (cell1 = cell46), inference(cnf_transformation, [], [f3916])).
fof(f3916, plain, ~ (cell1 = cell46), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28905)).
fof(f12179, plain, ((cell1 = cell46) | ~ spl177_136), inference(resolution, [], [f11205, f5400])).
fof(f11205, plain, (sP135(cell100, cell1) | ~ spl177_136), inference(avatar_component_clause, [], [f11203])).
fof(f11203, plain, (spl177_136 <=> sP135(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_136])])).
fof(f12177, plain, ~ spl177_135, inference(avatar_contradiction_clause, [], [f12176])).
fof(f12176, plain, ($false | ~ spl177_135), inference(subsumption_resolution, [], [f12173, f9639])).
fof(f9639, plain, ~ (cell1 = cell45), inference(cnf_transformation, [], [f3961])).
fof(f3961, plain, ~ (cell1 = cell45), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28950)).
fof(f12173, plain, ((cell1 = cell45) | ~ spl177_135), inference(resolution, [], [f11201, f5402])).
fof(f11201, plain, (sP134(cell100, cell1) | ~ spl177_135), inference(avatar_component_clause, [], [f11199])).
fof(f11199, plain, (spl177_135 <=> sP134(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_135])])).
fof(f12171, plain, ~ spl177_134, inference(avatar_contradiction_clause, [], [f12170])).
fof(f12170, plain, ($false | ~ spl177_134), inference(subsumption_resolution, [], [f12167, f9683])).
fof(f9683, plain, ~ (cell1 = cell44), inference(cnf_transformation, [], [f4005])).
fof(f4005, plain, ~ (cell1 = cell44), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof28994)).
fof(f12167, plain, ((cell1 = cell44) | ~ spl177_134), inference(resolution, [], [f11197, f5404])).
fof(f11197, plain, (sP133(cell100, cell1) | ~ spl177_134), inference(avatar_component_clause, [], [f11195])).
fof(f11195, plain, (spl177_134 <=> sP133(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_134])])).
fof(f12165, plain, ~ spl177_133, inference(avatar_contradiction_clause, [], [f12164])).
fof(f12164, plain, ($false | ~ spl177_133), inference(subsumption_resolution, [], [f12161, f9726])).
fof(f9726, plain, ~ (cell1 = cell43), inference(cnf_transformation, [], [f4048])).
fof(f4048, plain, ~ (cell1 = cell43), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29037)).
fof(f12161, plain, ((cell1 = cell43) | ~ spl177_133), inference(resolution, [], [f11193, f5406])).
fof(f11193, plain, (sP132(cell100, cell1) | ~ spl177_133), inference(avatar_component_clause, [], [f11191])).
fof(f11191, plain, (spl177_133 <=> sP132(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_133])])).
fof(f12159, plain, ~ spl177_132, inference(avatar_contradiction_clause, [], [f12158])).
fof(f12158, plain, ($false | ~ spl177_132), inference(subsumption_resolution, [], [f12155, f9809])).
fof(f9809, plain, ~ (cell1 = cell42), inference(cnf_transformation, [], [f4131])).
fof(f4131, plain, ~ (cell1 = cell42), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29120)).
fof(f12155, plain, ((cell1 = cell42) | ~ spl177_132), inference(resolution, [], [f11189, f5408])).
fof(f11189, plain, (sP131(cell100, cell1) | ~ spl177_132), inference(avatar_component_clause, [], [f11187])).
fof(f11187, plain, (spl177_132 <=> sP131(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_132])])).
fof(f12153, plain, ~ spl177_131, inference(avatar_contradiction_clause, [], [f12152])).
fof(f12152, plain, ($false | ~ spl177_131), inference(subsumption_resolution, [], [f12149, f9768])).
fof(f9768, plain, ~ (cell1 = cell41), inference(cnf_transformation, [], [f4090])).
fof(f4090, plain, ~ (cell1 = cell41), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29079)).
fof(f12149, plain, ((cell1 = cell41) | ~ spl177_131), inference(resolution, [], [f11185, f5410])).
fof(f11185, plain, (sP130(cell100, cell1) | ~ spl177_131), inference(avatar_component_clause, [], [f11183])).
fof(f11183, plain, (spl177_131 <=> sP130(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_131])])).
fof(f12147, plain, ~ spl177_130, inference(avatar_contradiction_clause, [], [f12146])).
fof(f12146, plain, ($false | ~ spl177_130), inference(subsumption_resolution, [], [f12143, f9849])).
fof(f9849, plain, ~ (cell1 = cell40), inference(cnf_transformation, [], [f4171])).
fof(f4171, plain, ~ (cell1 = cell40), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29160)).
fof(f12143, plain, ((cell1 = cell40) | ~ spl177_130), inference(resolution, [], [f11181, f5412])).
fof(f11181, plain, (sP129(cell100, cell1) | ~ spl177_130), inference(avatar_component_clause, [], [f11179])).
fof(f11179, plain, (spl177_130 <=> sP129(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_130])])).
fof(f12141, plain, ~ spl177_129, inference(avatar_contradiction_clause, [], [f12140])).
fof(f12140, plain, ($false | ~ spl177_129), inference(subsumption_resolution, [], [f12137, f9888])).
fof(f9888, plain, ~ (cell1 = cell39), inference(cnf_transformation, [], [f4210])).
fof(f4210, plain, ~ (cell1 = cell39), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29199)).
fof(f12137, plain, ((cell1 = cell39) | ~ spl177_129), inference(resolution, [], [f11177, f5414])).
fof(f11177, plain, (sP128(cell100, cell1) | ~ spl177_129), inference(avatar_component_clause, [], [f11175])).
fof(f11175, plain, (spl177_129 <=> sP128(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_129])])).
fof(f12135, plain, ~ spl177_128, inference(avatar_contradiction_clause, [], [f12134])).
fof(f12134, plain, ($false | ~ spl177_128), inference(subsumption_resolution, [], [f12131, f9926])).
fof(f9926, plain, ~ (cell1 = cell38), inference(cnf_transformation, [], [f4248])).
fof(f4248, plain, ~ (cell1 = cell38), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29237)).
fof(f12131, plain, ((cell1 = cell38) | ~ spl177_128), inference(resolution, [], [f11173, f5416])).
fof(f11173, plain, (sP127(cell100, cell1) | ~ spl177_128), inference(avatar_component_clause, [], [f11171])).
fof(f11171, plain, (spl177_128 <=> sP127(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_128])])).
fof(f12129, plain, ~ spl177_127, inference(avatar_contradiction_clause, [], [f12128])).
fof(f12128, plain, ($false | ~ spl177_127), inference(subsumption_resolution, [], [f12125, f9963])).
fof(f9963, plain, ~ (cell1 = cell37), inference(cnf_transformation, [], [f4285])).
fof(f4285, plain, ~ (cell1 = cell37), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29274)).
fof(f12125, plain, ((cell1 = cell37) | ~ spl177_127), inference(resolution, [], [f11169, f5418])).
fof(f11169, plain, (sP126(cell100, cell1) | ~ spl177_127), inference(avatar_component_clause, [], [f11167])).
fof(f11167, plain, (spl177_127 <=> sP126(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_127])])).
fof(f12123, plain, ~ spl177_126, inference(avatar_contradiction_clause, [], [f12122])).
fof(f12122, plain, ($false | ~ spl177_126), inference(subsumption_resolution, [], [f12119, f9999])).
fof(f9999, plain, ~ (cell1 = cell36), inference(cnf_transformation, [], [f4321])).
fof(f4321, plain, ~ (cell1 = cell36), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29310)).
fof(f12119, plain, ((cell1 = cell36) | ~ spl177_126), inference(resolution, [], [f11165, f5420])).
fof(f11165, plain, (sP125(cell100, cell1) | ~ spl177_126), inference(avatar_component_clause, [], [f11163])).
fof(f11163, plain, (spl177_126 <=> sP125(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_126])])).
fof(f12117, plain, ~ spl177_125, inference(avatar_contradiction_clause, [], [f12116])).
fof(f12116, plain, ($false | ~ spl177_125), inference(subsumption_resolution, [], [f12113, f10034])).
fof(f10034, plain, ~ (cell1 = cell35), inference(cnf_transformation, [], [f4356])).
fof(f4356, plain, ~ (cell1 = cell35), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29345)).
fof(f12113, plain, ((cell1 = cell35) | ~ spl177_125), inference(resolution, [], [f11161, f5422])).
fof(f11161, plain, (sP124(cell100, cell1) | ~ spl177_125), inference(avatar_component_clause, [], [f11159])).
fof(f11159, plain, (spl177_125 <=> sP124(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_125])])).
fof(f12111, plain, ~ spl177_124, inference(avatar_contradiction_clause, [], [f12110])).
fof(f12110, plain, ($false | ~ spl177_124), inference(subsumption_resolution, [], [f12107, f10068])).
fof(f10068, plain, ~ (cell1 = cell34), inference(cnf_transformation, [], [f4390])).
fof(f4390, plain, ~ (cell1 = cell34), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29379)).
fof(f12107, plain, ((cell1 = cell34) | ~ spl177_124), inference(resolution, [], [f11157, f5424])).
fof(f11157, plain, (sP123(cell100, cell1) | ~ spl177_124), inference(avatar_component_clause, [], [f11155])).
fof(f11155, plain, (spl177_124 <=> sP123(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_124])])).
fof(f12105, plain, ~ spl177_123, inference(avatar_contradiction_clause, [], [f12104])).
fof(f12104, plain, ($false | ~ spl177_123), inference(subsumption_resolution, [], [f12101, f10101])).
fof(f10101, plain, ~ (cell1 = cell33), inference(cnf_transformation, [], [f4423])).
fof(f4423, plain, ~ (cell1 = cell33), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29412)).
fof(f12101, plain, ((cell1 = cell33) | ~ spl177_123), inference(resolution, [], [f11153, f5426])).
fof(f11153, plain, (sP122(cell100, cell1) | ~ spl177_123), inference(avatar_component_clause, [], [f11151])).
fof(f11151, plain, (spl177_123 <=> sP122(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_123])])).
fof(f12099, plain, ~ spl177_122, inference(avatar_contradiction_clause, [], [f12098])).
fof(f12098, plain, ($false | ~ spl177_122), inference(subsumption_resolution, [], [f12095, f10164])).
fof(f10164, plain, ~ (cell1 = cell32), inference(cnf_transformation, [], [f4486])).
fof(f4486, plain, ~ (cell1 = cell32), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29475)).
fof(f12095, plain, ((cell1 = cell32) | ~ spl177_122), inference(resolution, [], [f11149, f5428])).
fof(f11149, plain, (sP121(cell100, cell1) | ~ spl177_122), inference(avatar_component_clause, [], [f11147])).
fof(f11147, plain, (spl177_122 <=> sP121(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_122])])).
fof(f12093, plain, ~ spl177_121, inference(avatar_contradiction_clause, [], [f12092])).
fof(f12092, plain, ($false | ~ spl177_121), inference(subsumption_resolution, [], [f12089, f10133])).
fof(f10133, plain, ~ (cell1 = cell31), inference(cnf_transformation, [], [f4455])).
fof(f4455, plain, ~ (cell1 = cell31), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29444)).
fof(f12089, plain, ((cell1 = cell31) | ~ spl177_121), inference(resolution, [], [f11145, f5430])).
fof(f11145, plain, (sP120(cell100, cell1) | ~ spl177_121), inference(avatar_component_clause, [], [f11143])).
fof(f11143, plain, (spl177_121 <=> sP120(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_121])])).
fof(f12087, plain, ~ spl177_120, inference(avatar_contradiction_clause, [], [f12086])).
fof(f12086, plain, ($false | ~ spl177_120), inference(subsumption_resolution, [], [f12083, f10194])).
fof(f10194, plain, ~ (cell1 = cell30), inference(cnf_transformation, [], [f4516])).
fof(f4516, plain, ~ (cell1 = cell30), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29505)).
fof(f12083, plain, ((cell1 = cell30) | ~ spl177_120), inference(resolution, [], [f11141, f5432])).
fof(f11141, plain, (sP119(cell100, cell1) | ~ spl177_120), inference(avatar_component_clause, [], [f11139])).
fof(f11139, plain, (spl177_120 <=> sP119(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_120])])).
fof(f12081, plain, ~ spl177_119, inference(avatar_contradiction_clause, [], [f12080])).
fof(f12080, plain, ($false | ~ spl177_119), inference(subsumption_resolution, [], [f12077, f10223])).
fof(f10223, plain, ~ (cell1 = cell29), inference(cnf_transformation, [], [f4545])).
fof(f4545, plain, ~ (cell1 = cell29), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29534)).
fof(f12077, plain, ((cell1 = cell29) | ~ spl177_119), inference(resolution, [], [f11137, f5434])).
fof(f11137, plain, (sP118(cell100, cell1) | ~ spl177_119), inference(avatar_component_clause, [], [f11135])).
fof(f11135, plain, (spl177_119 <=> sP118(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_119])])).
fof(f12075, plain, ~ spl177_118, inference(avatar_contradiction_clause, [], [f12074])).
fof(f12074, plain, ($false | ~ spl177_118), inference(subsumption_resolution, [], [f12071, f10251])).
fof(f10251, plain, ~ (cell1 = cell28), inference(cnf_transformation, [], [f4573])).
fof(f4573, plain, ~ (cell1 = cell28), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29562)).
fof(f12071, plain, ((cell1 = cell28) | ~ spl177_118), inference(resolution, [], [f11133, f5436])).
fof(f11133, plain, (sP117(cell100, cell1) | ~ spl177_118), inference(avatar_component_clause, [], [f11131])).
fof(f11131, plain, (spl177_118 <=> sP117(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_118])])).
fof(f12069, plain, ~ spl177_117, inference(avatar_contradiction_clause, [], [f12068])).
fof(f12068, plain, ($false | ~ spl177_117), inference(subsumption_resolution, [], [f12065, f10278])).
fof(f10278, plain, ~ (cell1 = cell27), inference(cnf_transformation, [], [f4600])).
fof(f4600, plain, ~ (cell1 = cell27), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29589)).
fof(f12065, plain, ((cell1 = cell27) | ~ spl177_117), inference(resolution, [], [f11129, f5438])).
fof(f11129, plain, (sP116(cell100, cell1) | ~ spl177_117), inference(avatar_component_clause, [], [f11127])).
fof(f11127, plain, (spl177_117 <=> sP116(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_117])])).
fof(f12063, plain, ~ spl177_116, inference(avatar_contradiction_clause, [], [f12062])).
fof(f12062, plain, ($false | ~ spl177_116), inference(subsumption_resolution, [], [f12059, f10304])).
fof(f10304, plain, ~ (cell1 = cell26), inference(cnf_transformation, [], [f4626])).
fof(f4626, plain, ~ (cell1 = cell26), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29615)).
fof(f12059, plain, ((cell1 = cell26) | ~ spl177_116), inference(resolution, [], [f11125, f5440])).
fof(f11125, plain, (sP115(cell100, cell1) | ~ spl177_116), inference(avatar_component_clause, [], [f11123])).
fof(f11123, plain, (spl177_116 <=> sP115(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_116])])).
fof(f12057, plain, ~ spl177_115, inference(avatar_contradiction_clause, [], [f12056])).
fof(f12056, plain, ($false | ~ spl177_115), inference(subsumption_resolution, [], [f12053, f10329])).
fof(f10329, plain, ~ (cell1 = cell25), inference(cnf_transformation, [], [f4651])).
fof(f4651, plain, ~ (cell1 = cell25), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29640)).
fof(f12053, plain, ((cell1 = cell25) | ~ spl177_115), inference(resolution, [], [f11121, f5442])).
fof(f11121, plain, (sP114(cell100, cell1) | ~ spl177_115), inference(avatar_component_clause, [], [f11119])).
fof(f11119, plain, (spl177_115 <=> sP114(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_115])])).
fof(f12051, plain, ~ spl177_114, inference(avatar_contradiction_clause, [], [f12050])).
fof(f12050, plain, ($false | ~ spl177_114), inference(subsumption_resolution, [], [f12047, f10353])).
fof(f10353, plain, ~ (cell1 = cell24), inference(cnf_transformation, [], [f4675])).
fof(f4675, plain, ~ (cell1 = cell24), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29664)).
fof(f12047, plain, ((cell1 = cell24) | ~ spl177_114), inference(resolution, [], [f11117, f5444])).
fof(f11117, plain, (sP113(cell100, cell1) | ~ spl177_114), inference(avatar_component_clause, [], [f11115])).
fof(f11115, plain, (spl177_114 <=> sP113(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_114])])).
fof(f12045, plain, ~ spl177_113, inference(avatar_contradiction_clause, [], [f12044])).
fof(f12044, plain, ($false | ~ spl177_113), inference(subsumption_resolution, [], [f12041, f10376])).
fof(f10376, plain, ~ (cell1 = cell23), inference(cnf_transformation, [], [f4698])).
fof(f4698, plain, ~ (cell1 = cell23), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29687)).
fof(f12041, plain, ((cell1 = cell23) | ~ spl177_113), inference(resolution, [], [f11113, f5446])).
fof(f11113, plain, (sP112(cell100, cell1) | ~ spl177_113), inference(avatar_component_clause, [], [f11111])).
fof(f11111, plain, (spl177_113 <=> sP112(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_113])])).
fof(f12039, plain, ~ spl177_112, inference(avatar_contradiction_clause, [], [f12038])).
fof(f12038, plain, ($false | ~ spl177_112), inference(subsumption_resolution, [], [f12035, f10419])).
fof(f10419, plain, ~ (cell1 = cell22), inference(cnf_transformation, [], [f4741])).
fof(f4741, plain, ~ (cell1 = cell22), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29730)).
fof(f12035, plain, ((cell1 = cell22) | ~ spl177_112), inference(resolution, [], [f11109, f5448])).
fof(f11109, plain, (sP111(cell100, cell1) | ~ spl177_112), inference(avatar_component_clause, [], [f11107])).
fof(f11107, plain, (spl177_112 <=> sP111(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_112])])).
fof(f12033, plain, ~ spl177_111, inference(avatar_contradiction_clause, [], [f12032])).
fof(f12032, plain, ($false | ~ spl177_111), inference(subsumption_resolution, [], [f12029, f10398])).
fof(f10398, plain, ~ (cell1 = cell21), inference(cnf_transformation, [], [f4720])).
fof(f4720, plain, ~ (cell1 = cell21), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29709)).
fof(f12029, plain, ((cell1 = cell21) | ~ spl177_111), inference(resolution, [], [f11105, f5450])).
fof(f11105, plain, (sP110(cell100, cell1) | ~ spl177_111), inference(avatar_component_clause, [], [f11103])).
fof(f11103, plain, (spl177_111 <=> sP110(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_111])])).
fof(f12027, plain, ~ spl177_110, inference(avatar_contradiction_clause, [], [f12026])).
fof(f12026, plain, ($false | ~ spl177_110), inference(subsumption_resolution, [], [f12023, f10439])).
fof(f10439, plain, ~ (cell1 = cell20), inference(cnf_transformation, [], [f4761])).
fof(f4761, plain, ~ (cell1 = cell20), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29750)).
fof(f12023, plain, ((cell1 = cell20) | ~ spl177_110), inference(resolution, [], [f11101, f5452])).
fof(f11101, plain, (sP109(cell100, cell1) | ~ spl177_110), inference(avatar_component_clause, [], [f11099])).
fof(f11099, plain, (spl177_110 <=> sP109(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_110])])).
fof(f12021, plain, ~ spl177_109, inference(avatar_contradiction_clause, [], [f12020])).
fof(f12020, plain, ($false | ~ spl177_109), inference(subsumption_resolution, [], [f12017, f10458])).
fof(f10458, plain, ~ (cell1 = cell19), inference(cnf_transformation, [], [f4780])).
fof(f4780, plain, ~ (cell1 = cell19), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29769)).
fof(f12017, plain, ((cell1 = cell19) | ~ spl177_109), inference(resolution, [], [f11097, f5454])).
fof(f11097, plain, (sP108(cell100, cell1) | ~ spl177_109), inference(avatar_component_clause, [], [f11095])).
fof(f11095, plain, (spl177_109 <=> sP108(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_109])])).
fof(f12015, plain, ~ spl177_108, inference(avatar_contradiction_clause, [], [f12014])).
fof(f12014, plain, ($false | ~ spl177_108), inference(subsumption_resolution, [], [f12011, f10476])).
fof(f10476, plain, ~ (cell1 = cell18), inference(cnf_transformation, [], [f4798])).
fof(f4798, plain, ~ (cell1 = cell18), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29787)).
fof(f12011, plain, ((cell1 = cell18) | ~ spl177_108), inference(resolution, [], [f11093, f5456])).
fof(f11093, plain, (sP107(cell100, cell1) | ~ spl177_108), inference(avatar_component_clause, [], [f11091])).
fof(f11091, plain, (spl177_108 <=> sP107(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_108])])).
fof(f12009, plain, ~ spl177_107, inference(avatar_contradiction_clause, [], [f12008])).
fof(f12008, plain, ($false | ~ spl177_107), inference(subsumption_resolution, [], [f12005, f10493])).
fof(f10493, plain, ~ (cell1 = cell17), inference(cnf_transformation, [], [f4815])).
fof(f4815, plain, ~ (cell1 = cell17), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29804)).
fof(f12005, plain, ((cell1 = cell17) | ~ spl177_107), inference(resolution, [], [f11089, f5458])).
fof(f11089, plain, (sP106(cell100, cell1) | ~ spl177_107), inference(avatar_component_clause, [], [f11087])).
fof(f11087, plain, (spl177_107 <=> sP106(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_107])])).
fof(f12003, plain, ~ spl177_106, inference(avatar_contradiction_clause, [], [f12002])).
fof(f12002, plain, ($false | ~ spl177_106), inference(subsumption_resolution, [], [f11999, f10509])).
fof(f10509, plain, ~ (cell1 = cell16), inference(cnf_transformation, [], [f4831])).
fof(f4831, plain, ~ (cell1 = cell16), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29820)).
fof(f11999, plain, ((cell1 = cell16) | ~ spl177_106), inference(resolution, [], [f11085, f5460])).
fof(f11085, plain, (sP105(cell100, cell1) | ~ spl177_106), inference(avatar_component_clause, [], [f11083])).
fof(f11083, plain, (spl177_106 <=> sP105(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_106])])).
fof(f11997, plain, ~ spl177_105, inference(avatar_contradiction_clause, [], [f11996])).
fof(f11996, plain, ($false | ~ spl177_105), inference(subsumption_resolution, [], [f11993, f10524])).
fof(f10524, plain, ~ (cell1 = cell15), inference(cnf_transformation, [], [f4846])).
fof(f4846, plain, ~ (cell1 = cell15), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29835)).
fof(f11993, plain, ((cell1 = cell15) | ~ spl177_105), inference(resolution, [], [f11081, f5462])).
fof(f11081, plain, (sP104(cell100, cell1) | ~ spl177_105), inference(avatar_component_clause, [], [f11079])).
fof(f11079, plain, (spl177_105 <=> sP104(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_105])])).
fof(f11991, plain, ~ spl177_104, inference(avatar_contradiction_clause, [], [f11990])).
fof(f11990, plain, ($false | ~ spl177_104), inference(subsumption_resolution, [], [f11987, f10538])).
fof(f10538, plain, ~ (cell1 = cell14), inference(cnf_transformation, [], [f4860])).
fof(f4860, plain, ~ (cell1 = cell14), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29849)).
fof(f11987, plain, ((cell1 = cell14) | ~ spl177_104), inference(resolution, [], [f11077, f5464])).
fof(f11077, plain, (sP103(cell100, cell1) | ~ spl177_104), inference(avatar_component_clause, [], [f11075])).
fof(f11075, plain, (spl177_104 <=> sP103(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_104])])).
fof(f11985, plain, ~ spl177_103, inference(avatar_contradiction_clause, [], [f11984])).
fof(f11984, plain, ($false | ~ spl177_103), inference(subsumption_resolution, [], [f11981, f10551])).
fof(f10551, plain, ~ (cell1 = cell13), inference(cnf_transformation, [], [f4873])).
fof(f4873, plain, ~ (cell1 = cell13), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29862)).
fof(f11981, plain, ((cell1 = cell13) | ~ spl177_103), inference(resolution, [], [f11073, f5466])).
fof(f11073, plain, (sP102(cell100, cell1) | ~ spl177_103), inference(avatar_component_clause, [], [f11071])).
fof(f11071, plain, (spl177_103 <=> sP102(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_103])])).
fof(f11979, plain, ~ spl177_102, inference(avatar_contradiction_clause, [], [f11978])).
fof(f11978, plain, ($false | ~ spl177_102), inference(subsumption_resolution, [], [f11975, f10574])).
fof(f10574, plain, ~ (cell1 = cell12), inference(cnf_transformation, [], [f4896])).
fof(f4896, plain, ~ (cell1 = cell12), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29885)).
fof(f11975, plain, ((cell1 = cell12) | ~ spl177_102), inference(resolution, [], [f11069, f5468])).
fof(f11069, plain, (sP101(cell100, cell1) | ~ spl177_102), inference(avatar_component_clause, [], [f11067])).
fof(f11067, plain, (spl177_102 <=> sP101(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_102])])).
fof(f11973, plain, ~ spl177_101, inference(avatar_contradiction_clause, [], [f11972])).
fof(f11972, plain, ($false | ~ spl177_101), inference(subsumption_resolution, [], [f11969, f10563])).
fof(f10563, plain, ~ (cell1 = cell11), inference(cnf_transformation, [], [f4885])).
fof(f4885, plain, ~ (cell1 = cell11), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29874)).
fof(f11969, plain, ((cell1 = cell11) | ~ spl177_101), inference(resolution, [], [f11065, f5470])).
fof(f11065, plain, (sP100(cell100, cell1) | ~ spl177_101), inference(avatar_component_clause, [], [f11063])).
fof(f11063, plain, (spl177_101 <=> sP100(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_101])])).
fof(f11967, plain, ~ spl177_100, inference(avatar_contradiction_clause, [], [f11966])).
fof(f11966, plain, ($false | ~ spl177_100), inference(subsumption_resolution, [], [f11963, f10584])).
fof(f10584, plain, ~ (cell1 = cell10), inference(cnf_transformation, [], [f4906])).
fof(f4906, plain, ~ (cell1 = cell10), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29895)).
fof(f11963, plain, ((cell1 = cell10) | ~ spl177_100), inference(resolution, [], [f11061, f5472])).
fof(f11061, plain, (sP99(cell100, cell1) | ~ spl177_100), inference(avatar_component_clause, [], [f11059])).
fof(f11059, plain, (spl177_100 <=> sP99(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_100])])).
fof(f11961, plain, ~ spl177_99, inference(avatar_contradiction_clause, [], [f11960])).
fof(f11960, plain, ($false | ~ spl177_99), inference(subsumption_resolution, [], [f11957, f10593])).
fof(f10593, plain, ~ (cell1 = cell9), inference(cnf_transformation, [], [f4915])).
fof(f4915, plain, ~ (cell1 = cell9), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29904)).
fof(f11957, plain, ((cell1 = cell9) | ~ spl177_99), inference(resolution, [], [f11057, f5474])).
fof(f11057, plain, (sP98(cell100, cell1) | ~ spl177_99), inference(avatar_component_clause, [], [f11055])).
fof(f11055, plain, (spl177_99 <=> sP98(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_99])])).
fof(f11955, plain, ~ spl177_98, inference(avatar_contradiction_clause, [], [f11954])).
fof(f11954, plain, ($false | ~ spl177_98), inference(subsumption_resolution, [], [f11951, f10601])).
fof(f10601, plain, ~ (cell1 = cell8), inference(cnf_transformation, [], [f4923])).
fof(f4923, plain, ~ (cell1 = cell8), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29912)).
fof(f11951, plain, ((cell1 = cell8) | ~ spl177_98), inference(resolution, [], [f11053, f5476])).
fof(f11053, plain, (sP97(cell100, cell1) | ~ spl177_98), inference(avatar_component_clause, [], [f11051])).
fof(f11051, plain, (spl177_98 <=> sP97(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_98])])).
fof(f11949, plain, ~ spl177_97, inference(avatar_contradiction_clause, [], [f11948])).
fof(f11948, plain, ($false | ~ spl177_97), inference(subsumption_resolution, [], [f11945, f10608])).
fof(f10608, plain, ~ (cell1 = cell7), inference(cnf_transformation, [], [f4930])).
fof(f4930, plain, ~ (cell1 = cell7), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29919)).
fof(f11945, plain, ((cell1 = cell7) | ~ spl177_97), inference(resolution, [], [f11049, f5478])).
fof(f11049, plain, (sP96(cell100, cell1) | ~ spl177_97), inference(avatar_component_clause, [], [f11047])).
fof(f11047, plain, (spl177_97 <=> sP96(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_97])])).
fof(f11943, plain, ~ spl177_96, inference(avatar_contradiction_clause, [], [f11942])).
fof(f11942, plain, ($false | ~ spl177_96), inference(subsumption_resolution, [], [f11939, f10614])).
fof(f10614, plain, ~ (cell1 = cell6), inference(cnf_transformation, [], [f4936])).
fof(f4936, plain, ~ (cell1 = cell6), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29925)).
fof(f11939, plain, ((cell1 = cell6) | ~ spl177_96), inference(resolution, [], [f11045, f5480])).
fof(f11045, plain, (sP95(cell100, cell1) | ~ spl177_96), inference(avatar_component_clause, [], [f11043])).
fof(f11043, plain, (spl177_96 <=> sP95(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_96])])).
fof(f11937, plain, ~ spl177_95, inference(avatar_contradiction_clause, [], [f11936])).
fof(f11936, plain, ($false | ~ spl177_95), inference(subsumption_resolution, [], [f11933, f10619])).
fof(f10619, plain, ~ (cell1 = cell5), inference(cnf_transformation, [], [f4941])).
fof(f4941, plain, ~ (cell1 = cell5), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29930)).
fof(f11933, plain, ((cell1 = cell5) | ~ spl177_95), inference(resolution, [], [f11041, f5482])).
fof(f11041, plain, (sP94(cell100, cell1) | ~ spl177_95), inference(avatar_component_clause, [], [f11039])).
fof(f11039, plain, (spl177_95 <=> sP94(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_95])])).
fof(f11931, plain, ~ spl177_94, inference(avatar_contradiction_clause, [], [f11930])).
fof(f11930, plain, ($false | ~ spl177_94), inference(subsumption_resolution, [], [f11927, f10623])).
fof(f10623, plain, ~ (cell1 = cell4), inference(cnf_transformation, [], [f4945])).
fof(f4945, plain, ~ (cell1 = cell4), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29934)).
fof(f11927, plain, ((cell1 = cell4) | ~ spl177_94), inference(resolution, [], [f11037, f5484])).
fof(f11037, plain, (sP93(cell100, cell1) | ~ spl177_94), inference(avatar_component_clause, [], [f11035])).
fof(f11035, plain, (spl177_94 <=> sP93(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_94])])).
fof(f11925, plain, ~ spl177_93, inference(avatar_contradiction_clause, [], [f11924])).
fof(f11924, plain, ($false | ~ spl177_93), inference(subsumption_resolution, [], [f11921, f10626])).
fof(f10626, plain, ~ (cell1 = cell3), inference(cnf_transformation, [], [f4948])).
fof(f4948, plain, ~ (cell1 = cell3), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29937)).
fof(f11921, plain, ((cell1 = cell3) | ~ spl177_93), inference(resolution, [], [f11033, f5486])).
fof(f11033, plain, (sP92(cell100, cell1) | ~ spl177_93), inference(avatar_component_clause, [], [f11031])).
fof(f11031, plain, (spl177_93 <=> sP92(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_93])])).
fof(f11919, plain, ~ spl177_92, inference(avatar_contradiction_clause, [], [f11918])).
fof(f11918, plain, ($false | ~ spl177_92), inference(subsumption_resolution, [], [f11915, f10628])).
fof(f10628, plain, ~ (cell1 = cell2), inference(cnf_transformation, [], [f4950])).
fof(f4950, plain, ~ (cell1 = cell2), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29939)).
fof(f11915, plain, ((cell1 = cell2) | ~ spl177_92), inference(resolution, [], [f11029, f5488])).
fof(f11029, plain, (sP91(cell100, cell1) | ~ spl177_92), inference(avatar_component_clause, [], [f11027])).
fof(f11027, plain, (spl177_92 <=> sP91(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_92])])).
fof(f11913, plain, ~ spl177_91, inference(avatar_contradiction_clause, [], [f11912])).
fof(f11912, plain, ($false | ~ spl177_91), inference(subsumption_resolution, [], [f11910, f10564])).
fof(f11910, plain, ((cell11 = cell100) | ~ spl177_91), inference(resolution, [], [f11025, f5491])).
fof(f5491, plain, ! [X0, X1] : (~ sP90(X0, X1) | (cell11 = X0)), inference(cnf_transformation, [], [f5227])).
fof(f11025, plain, (sP90(cell100, cell1) | ~ spl177_91), inference(avatar_component_clause, [], [f11023])).
fof(f11023, plain, (spl177_91 <=> sP90(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_91])])).
fof(f11909, plain, ~ spl177_90, inference(avatar_contradiction_clause, [], [f11908])).
fof(f11908, plain, ($false | ~ spl177_90), inference(subsumption_resolution, [], [f11907, f5778])).
fof(f11907, plain, ((cell1 = cell99) | ~ spl177_90), inference(resolution, [], [f11021, f5492])).
fof(f5492, plain, ! [X0, X1] : (~ sP89(X0, X1) | (cell99 = X1)), inference(cnf_transformation, [], [f5228])).
fof(f11021, plain, (sP89(cell100, cell1) | ~ spl177_90), inference(avatar_component_clause, [], [f11019])).
fof(f11019, plain, (spl177_90 <=> sP89(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_90])])).
fof(f11905, plain, ~ spl177_89, inference(avatar_contradiction_clause, [], [f11904])).
fof(f11904, plain, ($false | ~ spl177_89), inference(subsumption_resolution, [], [f11901, f6069])).
fof(f6069, plain, ~ (cell1 = cell96), inference(cnf_transformation, [], [f391])).
fof(f391, plain, ~ (cell1 = cell96), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25380)).
fof(f11901, plain, ((cell1 = cell96) | ~ spl177_89), inference(resolution, [], [f11017, f5498])).
fof(f11017, plain, (sP86(cell100, cell1) | ~ spl177_89), inference(avatar_component_clause, [], [f11015])).
fof(f11015, plain, (spl177_89 <=> sP86(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_89])])).
fof(f11899, plain, ~ spl177_88, inference(avatar_contradiction_clause, [], [f11898])).
fof(f11898, plain, ($false | ~ spl177_88), inference(subsumption_resolution, [], [f11895, f6164])).
fof(f6164, plain, ~ (cell1 = cell95), inference(cnf_transformation, [], [f486])).
fof(f486, plain, ~ (cell1 = cell95), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25475)).
fof(f11895, plain, ((cell1 = cell95) | ~ spl177_88), inference(resolution, [], [f11013, f5500])).
fof(f11013, plain, (sP85(cell100, cell1) | ~ spl177_88), inference(avatar_component_clause, [], [f11011])).
fof(f11011, plain, (spl177_88 <=> sP85(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_88])])).
fof(f11893, plain, ~ spl177_87, inference(avatar_contradiction_clause, [], [f11892])).
fof(f11892, plain, ($false | ~ spl177_87), inference(subsumption_resolution, [], [f11889, f6258])).
fof(f6258, plain, ~ (cell1 = cell94), inference(cnf_transformation, [], [f580])).
fof(f580, plain, ~ (cell1 = cell94), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25569)).
fof(f11889, plain, ((cell1 = cell94) | ~ spl177_87), inference(resolution, [], [f11009, f5502])).
fof(f11009, plain, (sP84(cell100, cell1) | ~ spl177_87), inference(avatar_component_clause, [], [f11007])).
fof(f11007, plain, (spl177_87 <=> sP84(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_87])])).
fof(f11887, plain, ~ spl177_86, inference(avatar_contradiction_clause, [], [f11886])).
fof(f11886, plain, ($false | ~ spl177_86), inference(subsumption_resolution, [], [f11883, f6351])).
fof(f6351, plain, ~ (cell1 = cell93), inference(cnf_transformation, [], [f673])).
fof(f673, plain, ~ (cell1 = cell93), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25662)).
fof(f11883, plain, ((cell1 = cell93) | ~ spl177_86), inference(resolution, [], [f11005, f5504])).
fof(f11005, plain, (sP83(cell100, cell1) | ~ spl177_86), inference(avatar_component_clause, [], [f11003])).
fof(f11003, plain, (spl177_86 <=> sP83(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_86])])).
fof(f11881, plain, ~ spl177_85, inference(avatar_contradiction_clause, [], [f11880])).
fof(f11880, plain, ($false | ~ spl177_85), inference(subsumption_resolution, [], [f11877, f6534])).
fof(f6534, plain, ~ (cell1 = cell92), inference(cnf_transformation, [], [f856])).
fof(f856, plain, ~ (cell1 = cell92), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25845)).
fof(f11877, plain, ((cell1 = cell92) | ~ spl177_85), inference(resolution, [], [f11001, f5506])).
fof(f11001, plain, (sP82(cell100, cell1) | ~ spl177_85), inference(avatar_component_clause, [], [f10999])).
fof(f10999, plain, (spl177_85 <=> sP82(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_85])])).
fof(f11875, plain, ~ spl177_84, inference(avatar_contradiction_clause, [], [f11874])).
fof(f11874, plain, ($false | ~ spl177_84), inference(subsumption_resolution, [], [f11871, f6443])).
fof(f6443, plain, ~ (cell1 = cell91), inference(cnf_transformation, [], [f765])).
fof(f765, plain, ~ (cell1 = cell91), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25754)).
fof(f11871, plain, ((cell1 = cell91) | ~ spl177_84), inference(resolution, [], [f10997, f5508])).
fof(f10997, plain, (sP81(cell100, cell1) | ~ spl177_84), inference(avatar_component_clause, [], [f10995])).
fof(f10995, plain, (spl177_84 <=> sP81(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_84])])).
fof(f11869, plain, ~ spl177_81, inference(avatar_contradiction_clause, [], [f11868])).
fof(f11868, plain, ($false | ~ spl177_81), inference(subsumption_resolution, [], [f11865, f6888])).
fof(f11865, plain, ((cell1 = cell87) | ~ spl177_81), inference(resolution, [], [f10985, f5514])).
fof(f10985, plain, (sP78(cell100, cell1) | ~ spl177_81), inference(avatar_component_clause, [], [f10983])).
fof(f10983, plain, (spl177_81 <=> sP78(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_81])])).
fof(f11863, plain, ~ spl177_80, inference(avatar_contradiction_clause, [], [f11862])).
fof(f11862, plain, ($false | ~ spl177_80), inference(subsumption_resolution, [], [f11859, f6974])).
fof(f11859, plain, ((cell1 = cell86) | ~ spl177_80), inference(resolution, [], [f10981, f5516])).
fof(f10981, plain, (sP77(cell100, cell1) | ~ spl177_80), inference(avatar_component_clause, [], [f10979])).
fof(f10979, plain, (spl177_80 <=> sP77(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_80])])).
fof(f11857, plain, ~ spl177_79, inference(avatar_contradiction_clause, [], [f11856])).
fof(f11856, plain, ($false | ~ spl177_79), inference(subsumption_resolution, [], [f11853, f7059])).
fof(f11853, plain, ((cell1 = cell85) | ~ spl177_79), inference(resolution, [], [f10977, f5518])).
fof(f10977, plain, (sP76(cell100, cell1) | ~ spl177_79), inference(avatar_component_clause, [], [f10975])).
fof(f10975, plain, (spl177_79 <=> sP76(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_79])])).
fof(f11851, plain, ~ spl177_78, inference(avatar_contradiction_clause, [], [f11850])).
fof(f11850, plain, ($false | ~ spl177_78), inference(subsumption_resolution, [], [f11847, f7143])).
fof(f11847, plain, ((cell1 = cell84) | ~ spl177_78), inference(resolution, [], [f10973, f5520])).
fof(f10973, plain, (sP75(cell100, cell1) | ~ spl177_78), inference(avatar_component_clause, [], [f10971])).
fof(f10971, plain, (spl177_78 <=> sP75(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_78])])).
fof(f11845, plain, ~ spl177_77, inference(avatar_contradiction_clause, [], [f11844])).
fof(f11844, plain, ($false | ~ spl177_77), inference(subsumption_resolution, [], [f11841, f7226])).
fof(f11841, plain, ((cell1 = cell83) | ~ spl177_77), inference(resolution, [], [f10969, f5522])).
fof(f10969, plain, (sP74(cell100, cell1) | ~ spl177_77), inference(avatar_component_clause, [], [f10967])).
fof(f10967, plain, (spl177_77 <=> sP74(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_77])])).
fof(f11839, plain, ~ spl177_76, inference(avatar_contradiction_clause, [], [f11838])).
fof(f11838, plain, ($false | ~ spl177_76), inference(subsumption_resolution, [], [f11835, f7389])).
fof(f11835, plain, ((cell1 = cell82) | ~ spl177_76), inference(resolution, [], [f10965, f5524])).
fof(f10965, plain, (sP73(cell100, cell1) | ~ spl177_76), inference(avatar_component_clause, [], [f10963])).
fof(f10963, plain, (spl177_76 <=> sP73(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_76])])).
fof(f11833, plain, ~ spl177_75, inference(avatar_contradiction_clause, [], [f11832])).
fof(f11832, plain, ($false | ~ spl177_75), inference(subsumption_resolution, [], [f11829, f7308])).
fof(f11829, plain, ((cell1 = cell81) | ~ spl177_75), inference(resolution, [], [f10961, f5526])).
fof(f10961, plain, (sP72(cell100, cell1) | ~ spl177_75), inference(avatar_component_clause, [], [f10959])).
fof(f10959, plain, (spl177_75 <=> sP72(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_75])])).
fof(f11827, plain, ~ spl177_74, inference(avatar_contradiction_clause, [], [f11826])).
fof(f11826, plain, ($false | ~ spl177_74), inference(subsumption_resolution, [], [f11823, f7548])).
fof(f11823, plain, ((cell1 = cell79) | ~ spl177_74), inference(resolution, [], [f10957, f5528])).
fof(f10957, plain, (sP71(cell100, cell1) | ~ spl177_74), inference(avatar_component_clause, [], [f10955])).
fof(f10955, plain, (spl177_74 <=> sP71(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_74])])).
fof(f11821, plain, ~ spl177_73, inference(avatar_contradiction_clause, [], [f11820])).
fof(f11820, plain, ($false | ~ spl177_73), inference(subsumption_resolution, [], [f11817, f7626])).
fof(f11817, plain, ((cell1 = cell78) | ~ spl177_73), inference(resolution, [], [f10953, f5530])).
fof(f10953, plain, (sP70(cell100, cell1) | ~ spl177_73), inference(avatar_component_clause, [], [f10951])).
fof(f10951, plain, (spl177_73 <=> sP70(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_73])])).
fof(f11815, plain, ~ spl177_72, inference(avatar_contradiction_clause, [], [f11814])).
fof(f11814, plain, ($false | ~ spl177_72), inference(subsumption_resolution, [], [f11811, f7703])).
fof(f11811, plain, ((cell1 = cell77) | ~ spl177_72), inference(resolution, [], [f10949, f5532])).
fof(f10949, plain, (sP69(cell100, cell1) | ~ spl177_72), inference(avatar_component_clause, [], [f10947])).
fof(f10947, plain, (spl177_72 <=> sP69(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_72])])).
fof(f11809, plain, ~ spl177_71, inference(avatar_contradiction_clause, [], [f11808])).
fof(f11808, plain, ($false | ~ spl177_71), inference(subsumption_resolution, [], [f11805, f7779])).
fof(f11805, plain, ((cell1 = cell76) | ~ spl177_71), inference(resolution, [], [f10945, f5534])).
fof(f10945, plain, (sP68(cell100, cell1) | ~ spl177_71), inference(avatar_component_clause, [], [f10943])).
fof(f10943, plain, (spl177_71 <=> sP68(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_71])])).
fof(f11803, plain, ~ spl177_70, inference(avatar_contradiction_clause, [], [f11802])).
fof(f11802, plain, ($false | ~ spl177_70), inference(subsumption_resolution, [], [f11799, f7854])).
fof(f11799, plain, ((cell1 = cell75) | ~ spl177_70), inference(resolution, [], [f10941, f5536])).
fof(f10941, plain, (sP67(cell100, cell1) | ~ spl177_70), inference(avatar_component_clause, [], [f10939])).
fof(f10939, plain, (spl177_70 <=> sP67(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_70])])).
fof(f11797, plain, ~ spl177_69, inference(avatar_contradiction_clause, [], [f11796])).
fof(f11796, plain, ($false | ~ spl177_69), inference(subsumption_resolution, [], [f11793, f7928])).
fof(f11793, plain, ((cell1 = cell74) | ~ spl177_69), inference(resolution, [], [f10937, f5538])).
fof(f10937, plain, (sP66(cell100, cell1) | ~ spl177_69), inference(avatar_component_clause, [], [f10935])).
fof(f10935, plain, (spl177_69 <=> sP66(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_69])])).
fof(f11791, plain, ~ spl177_68, inference(avatar_contradiction_clause, [], [f11790])).
fof(f11790, plain, ($false | ~ spl177_68), inference(subsumption_resolution, [], [f11787, f8001])).
fof(f11787, plain, ((cell1 = cell73) | ~ spl177_68), inference(resolution, [], [f10933, f5540])).
fof(f10933, plain, (sP65(cell100, cell1) | ~ spl177_68), inference(avatar_component_clause, [], [f10931])).
fof(f10931, plain, (spl177_68 <=> sP65(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_68])])).
fof(f11785, plain, ~ spl177_67, inference(avatar_contradiction_clause, [], [f11784])).
fof(f11784, plain, ($false | ~ spl177_67), inference(subsumption_resolution, [], [f11781, f8144])).
fof(f11781, plain, ((cell1 = cell72) | ~ spl177_67), inference(resolution, [], [f10929, f5542])).
fof(f10929, plain, (sP64(cell100, cell1) | ~ spl177_67), inference(avatar_component_clause, [], [f10927])).
fof(f10927, plain, (spl177_67 <=> sP64(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_67])])).
fof(f11779, plain, ~ spl177_66, inference(avatar_contradiction_clause, [], [f11778])).
fof(f11778, plain, ($false | ~ spl177_66), inference(subsumption_resolution, [], [f11775, f8073])).
fof(f11775, plain, ((cell1 = cell71) | ~ spl177_66), inference(resolution, [], [f10925, f5544])).
fof(f10925, plain, (sP63(cell100, cell1) | ~ spl177_66), inference(avatar_component_clause, [], [f10923])).
fof(f10923, plain, (spl177_66 <=> sP63(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_66])])).
fof(f11773, plain, ~ spl177_65, inference(avatar_contradiction_clause, [], [f11772])).
fof(f11772, plain, ($false | ~ spl177_65), inference(subsumption_resolution, [], [f11769, f8283])).
fof(f11769, plain, ((cell1 = cell69) | ~ spl177_65), inference(resolution, [], [f10921, f5546])).
fof(f10921, plain, (sP62(cell100, cell1) | ~ spl177_65), inference(avatar_component_clause, [], [f10919])).
fof(f10919, plain, (spl177_65 <=> sP62(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_65])])).
fof(f11767, plain, ~ spl177_64, inference(avatar_contradiction_clause, [], [f11766])).
fof(f11766, plain, ($false | ~ spl177_64), inference(subsumption_resolution, [], [f11763, f8351])).
fof(f11763, plain, ((cell1 = cell68) | ~ spl177_64), inference(resolution, [], [f10917, f5548])).
fof(f10917, plain, (sP61(cell100, cell1) | ~ spl177_64), inference(avatar_component_clause, [], [f10915])).
fof(f10915, plain, (spl177_64 <=> sP61(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_64])])).
fof(f11761, plain, ~ spl177_63, inference(avatar_contradiction_clause, [], [f11760])).
fof(f11760, plain, ($false | ~ spl177_63), inference(subsumption_resolution, [], [f11757, f8418])).
fof(f11757, plain, ((cell1 = cell67) | ~ spl177_63), inference(resolution, [], [f10913, f5550])).
fof(f10913, plain, (sP60(cell100, cell1) | ~ spl177_63), inference(avatar_component_clause, [], [f10911])).
fof(f10911, plain, (spl177_63 <=> sP60(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_63])])).
fof(f11755, plain, ~ spl177_62, inference(avatar_contradiction_clause, [], [f11754])).
fof(f11754, plain, ($false | ~ spl177_62), inference(subsumption_resolution, [], [f11751, f8484])).
fof(f11751, plain, ((cell1 = cell66) | ~ spl177_62), inference(resolution, [], [f10909, f5552])).
fof(f10909, plain, (sP59(cell100, cell1) | ~ spl177_62), inference(avatar_component_clause, [], [f10907])).
fof(f10907, plain, (spl177_62 <=> sP59(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_62])])).
fof(f11749, plain, ~ spl177_61, inference(avatar_contradiction_clause, [], [f11748])).
fof(f11748, plain, ($false | ~ spl177_61), inference(subsumption_resolution, [], [f11745, f8549])).
fof(f11745, plain, ((cell1 = cell65) | ~ spl177_61), inference(resolution, [], [f10905, f5554])).
fof(f10905, plain, (sP58(cell100, cell1) | ~ spl177_61), inference(avatar_component_clause, [], [f10903])).
fof(f10903, plain, (spl177_61 <=> sP58(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_61])])).
fof(f11743, plain, ~ spl177_60, inference(avatar_contradiction_clause, [], [f11742])).
fof(f11742, plain, ($false | ~ spl177_60), inference(subsumption_resolution, [], [f11739, f8613])).
fof(f11739, plain, ((cell1 = cell64) | ~ spl177_60), inference(resolution, [], [f10901, f5556])).
fof(f10901, plain, (sP57(cell100, cell1) | ~ spl177_60), inference(avatar_component_clause, [], [f10899])).
fof(f10899, plain, (spl177_60 <=> sP57(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_60])])).
fof(f11737, plain, ~ spl177_59, inference(avatar_contradiction_clause, [], [f11736])).
fof(f11736, plain, ($false | ~ spl177_59), inference(subsumption_resolution, [], [f11733, f8676])).
fof(f11733, plain, ((cell1 = cell63) | ~ spl177_59), inference(resolution, [], [f10897, f5558])).
fof(f10897, plain, (sP56(cell100, cell1) | ~ spl177_59), inference(avatar_component_clause, [], [f10895])).
fof(f10895, plain, (spl177_59 <=> sP56(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_59])])).
fof(f11731, plain, ~ spl177_58, inference(avatar_contradiction_clause, [], [f11730])).
fof(f11730, plain, ($false | ~ spl177_58), inference(subsumption_resolution, [], [f11727, f8799])).
fof(f11727, plain, ((cell1 = cell62) | ~ spl177_58), inference(resolution, [], [f10893, f5560])).
fof(f10893, plain, (sP55(cell100, cell1) | ~ spl177_58), inference(avatar_component_clause, [], [f10891])).
fof(f10891, plain, (spl177_58 <=> sP55(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_58])])).
fof(f11725, plain, ~ spl177_57, inference(avatar_contradiction_clause, [], [f11724])).
fof(f11724, plain, ($false | ~ spl177_57), inference(subsumption_resolution, [], [f11721, f8738])).
fof(f11721, plain, ((cell1 = cell61) | ~ spl177_57), inference(resolution, [], [f10889, f5562])).
fof(f10889, plain, (sP54(cell100, cell1) | ~ spl177_57), inference(avatar_component_clause, [], [f10887])).
fof(f10887, plain, (spl177_57 <=> sP54(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_57])])).
fof(f11719, plain, ~ spl177_56, inference(avatar_contradiction_clause, [], [f11718])).
fof(f11718, plain, ($false | ~ spl177_56), inference(subsumption_resolution, [], [f11715, f8918])).
fof(f11715, plain, ((cell1 = cell59) | ~ spl177_56), inference(resolution, [], [f10885, f5564])).
fof(f10885, plain, (sP53(cell100, cell1) | ~ spl177_56), inference(avatar_component_clause, [], [f10883])).
fof(f10883, plain, (spl177_56 <=> sP53(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_56])])).
fof(f11713, plain, ~ spl177_55, inference(avatar_contradiction_clause, [], [f11712])).
fof(f11712, plain, ($false | ~ spl177_55), inference(subsumption_resolution, [], [f11709, f8976])).
fof(f11709, plain, ((cell1 = cell58) | ~ spl177_55), inference(resolution, [], [f10881, f5566])).
fof(f10881, plain, (sP52(cell100, cell1) | ~ spl177_55), inference(avatar_component_clause, [], [f10879])).
fof(f10879, plain, (spl177_55 <=> sP52(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_55])])).
fof(f11707, plain, ~ spl177_54, inference(avatar_contradiction_clause, [], [f11706])).
fof(f11706, plain, ($false | ~ spl177_54), inference(subsumption_resolution, [], [f11703, f9033])).
fof(f11703, plain, ((cell1 = cell57) | ~ spl177_54), inference(resolution, [], [f10877, f5568])).
fof(f10877, plain, (sP51(cell100, cell1) | ~ spl177_54), inference(avatar_component_clause, [], [f10875])).
fof(f10875, plain, (spl177_54 <=> sP51(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_54])])).
fof(f11701, plain, ~ spl177_53, inference(avatar_contradiction_clause, [], [f11700])).
fof(f11700, plain, ($false | ~ spl177_53), inference(subsumption_resolution, [], [f11697, f9089])).
fof(f11697, plain, ((cell1 = cell56) | ~ spl177_53), inference(resolution, [], [f10873, f5570])).
fof(f10873, plain, (sP50(cell100, cell1) | ~ spl177_53), inference(avatar_component_clause, [], [f10871])).
fof(f10871, plain, (spl177_53 <=> sP50(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_53])])).
fof(f11695, plain, ~ spl177_52, inference(avatar_contradiction_clause, [], [f11694])).
fof(f11694, plain, ($false | ~ spl177_52), inference(subsumption_resolution, [], [f11691, f9144])).
fof(f11691, plain, ((cell1 = cell55) | ~ spl177_52), inference(resolution, [], [f10869, f5572])).
fof(f10869, plain, (sP49(cell100, cell1) | ~ spl177_52), inference(avatar_component_clause, [], [f10867])).
fof(f10867, plain, (spl177_52 <=> sP49(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_52])])).
fof(f11689, plain, ~ spl177_51, inference(avatar_contradiction_clause, [], [f11688])).
fof(f11688, plain, ($false | ~ spl177_51), inference(subsumption_resolution, [], [f11685, f9198])).
fof(f11685, plain, ((cell1 = cell54) | ~ spl177_51), inference(resolution, [], [f10865, f5574])).
fof(f10865, plain, (sP48(cell100, cell1) | ~ spl177_51), inference(avatar_component_clause, [], [f10863])).
fof(f10863, plain, (spl177_51 <=> sP48(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_51])])).
fof(f11683, plain, ~ spl177_50, inference(avatar_contradiction_clause, [], [f11682])).
fof(f11682, plain, ($false | ~ spl177_50), inference(subsumption_resolution, [], [f11679, f9251])).
fof(f11679, plain, ((cell1 = cell53) | ~ spl177_50), inference(resolution, [], [f10861, f5576])).
fof(f10861, plain, (sP47(cell100, cell1) | ~ spl177_50), inference(avatar_component_clause, [], [f10859])).
fof(f10859, plain, (spl177_50 <=> sP47(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_50])])).
fof(f11677, plain, ~ spl177_49, inference(avatar_contradiction_clause, [], [f11676])).
fof(f11676, plain, ($false | ~ spl177_49), inference(subsumption_resolution, [], [f11673, f9354])).
fof(f11673, plain, ((cell1 = cell52) | ~ spl177_49), inference(resolution, [], [f10857, f5578])).
fof(f10857, plain, (sP46(cell100, cell1) | ~ spl177_49), inference(avatar_component_clause, [], [f10855])).
fof(f10855, plain, (spl177_49 <=> sP46(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_49])])).
fof(f11671, plain, ~ spl177_48, inference(avatar_contradiction_clause, [], [f11670])).
fof(f11670, plain, ($false | ~ spl177_48), inference(subsumption_resolution, [], [f11667, f9303])).
fof(f11667, plain, ((cell1 = cell51) | ~ spl177_48), inference(resolution, [], [f10853, f5580])).
fof(f10853, plain, (sP45(cell100, cell1) | ~ spl177_48), inference(avatar_component_clause, [], [f10851])).
fof(f10851, plain, (spl177_48 <=> sP45(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_48])])).
fof(f11665, plain, ~ spl177_47, inference(avatar_contradiction_clause, [], [f11664])).
fof(f11664, plain, ($false | ~ spl177_47), inference(subsumption_resolution, [], [f11661, f9453])).
fof(f11661, plain, ((cell1 = cell49) | ~ spl177_47), inference(resolution, [], [f10849, f5582])).
fof(f10849, plain, (sP44(cell100, cell1) | ~ spl177_47), inference(avatar_component_clause, [], [f10847])).
fof(f10847, plain, (spl177_47 <=> sP44(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_47])])).
fof(f11659, plain, ~ spl177_46, inference(avatar_contradiction_clause, [], [f11658])).
fof(f11658, plain, ($false | ~ spl177_46), inference(subsumption_resolution, [], [f11655, f9501])).
fof(f11655, plain, ((cell1 = cell48) | ~ spl177_46), inference(resolution, [], [f10845, f5584])).
fof(f10845, plain, (sP43(cell100, cell1) | ~ spl177_46), inference(avatar_component_clause, [], [f10843])).
fof(f10843, plain, (spl177_46 <=> sP43(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_46])])).
fof(f11653, plain, ~ spl177_45, inference(avatar_contradiction_clause, [], [f11652])).
fof(f11652, plain, ($false | ~ spl177_45), inference(subsumption_resolution, [], [f11649, f9548])).
fof(f11649, plain, ((cell1 = cell47) | ~ spl177_45), inference(resolution, [], [f10841, f5586])).
fof(f10841, plain, (sP42(cell100, cell1) | ~ spl177_45), inference(avatar_component_clause, [], [f10839])).
fof(f10839, plain, (spl177_45 <=> sP42(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_45])])).
fof(f11647, plain, ~ spl177_44, inference(avatar_contradiction_clause, [], [f11646])).
fof(f11646, plain, ($false | ~ spl177_44), inference(subsumption_resolution, [], [f11643, f9594])).
fof(f11643, plain, ((cell1 = cell46) | ~ spl177_44), inference(resolution, [], [f10837, f5588])).
fof(f10837, plain, (sP41(cell100, cell1) | ~ spl177_44), inference(avatar_component_clause, [], [f10835])).
fof(f10835, plain, (spl177_44 <=> sP41(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_44])])).
fof(f11641, plain, ~ spl177_43, inference(avatar_contradiction_clause, [], [f11640])).
fof(f11640, plain, ($false | ~ spl177_43), inference(subsumption_resolution, [], [f11637, f9639])).
fof(f11637, plain, ((cell1 = cell45) | ~ spl177_43), inference(resolution, [], [f10833, f5590])).
fof(f10833, plain, (sP40(cell100, cell1) | ~ spl177_43), inference(avatar_component_clause, [], [f10831])).
fof(f10831, plain, (spl177_43 <=> sP40(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_43])])).
fof(f11635, plain, ~ spl177_42, inference(avatar_contradiction_clause, [], [f11634])).
fof(f11634, plain, ($false | ~ spl177_42), inference(subsumption_resolution, [], [f11631, f9683])).
fof(f11631, plain, ((cell1 = cell44) | ~ spl177_42), inference(resolution, [], [f10829, f5592])).
fof(f10829, plain, (sP39(cell100, cell1) | ~ spl177_42), inference(avatar_component_clause, [], [f10827])).
fof(f10827, plain, (spl177_42 <=> sP39(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_42])])).
fof(f11629, plain, ~ spl177_41, inference(avatar_contradiction_clause, [], [f11628])).
fof(f11628, plain, ($false | ~ spl177_41), inference(subsumption_resolution, [], [f11625, f9726])).
fof(f11625, plain, ((cell1 = cell43) | ~ spl177_41), inference(resolution, [], [f10825, f5594])).
fof(f10825, plain, (sP38(cell100, cell1) | ~ spl177_41), inference(avatar_component_clause, [], [f10823])).
fof(f10823, plain, (spl177_41 <=> sP38(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_41])])).
fof(f11623, plain, ~ spl177_40, inference(avatar_contradiction_clause, [], [f11622])).
fof(f11622, plain, ($false | ~ spl177_40), inference(subsumption_resolution, [], [f11619, f9809])).
fof(f11619, plain, ((cell1 = cell42) | ~ spl177_40), inference(resolution, [], [f10821, f5596])).
fof(f10821, plain, (sP37(cell100, cell1) | ~ spl177_40), inference(avatar_component_clause, [], [f10819])).
fof(f10819, plain, (spl177_40 <=> sP37(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_40])])).
fof(f11617, plain, ~ spl177_39, inference(avatar_contradiction_clause, [], [f11616])).
fof(f11616, plain, ($false | ~ spl177_39), inference(subsumption_resolution, [], [f11613, f9768])).
fof(f11613, plain, ((cell1 = cell41) | ~ spl177_39), inference(resolution, [], [f10817, f5598])).
fof(f10817, plain, (sP36(cell100, cell1) | ~ spl177_39), inference(avatar_component_clause, [], [f10815])).
fof(f10815, plain, (spl177_39 <=> sP36(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_39])])).
fof(f11611, plain, ~ spl177_38, inference(avatar_contradiction_clause, [], [f11610])).
fof(f11610, plain, ($false | ~ spl177_38), inference(subsumption_resolution, [], [f11607, f9888])).
fof(f11607, plain, ((cell1 = cell39) | ~ spl177_38), inference(resolution, [], [f10813, f5600])).
fof(f10813, plain, (sP35(cell100, cell1) | ~ spl177_38), inference(avatar_component_clause, [], [f10811])).
fof(f10811, plain, (spl177_38 <=> sP35(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_38])])).
fof(f11605, plain, ~ spl177_37, inference(avatar_contradiction_clause, [], [f11604])).
fof(f11604, plain, ($false | ~ spl177_37), inference(subsumption_resolution, [], [f11601, f9926])).
fof(f11601, plain, ((cell1 = cell38) | ~ spl177_37), inference(resolution, [], [f10809, f5602])).
fof(f10809, plain, (sP34(cell100, cell1) | ~ spl177_37), inference(avatar_component_clause, [], [f10807])).
fof(f10807, plain, (spl177_37 <=> sP34(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_37])])).
fof(f11599, plain, ~ spl177_36, inference(avatar_contradiction_clause, [], [f11598])).
fof(f11598, plain, ($false | ~ spl177_36), inference(subsumption_resolution, [], [f11595, f9963])).
fof(f11595, plain, ((cell1 = cell37) | ~ spl177_36), inference(resolution, [], [f10805, f5604])).
fof(f10805, plain, (sP33(cell100, cell1) | ~ spl177_36), inference(avatar_component_clause, [], [f10803])).
fof(f10803, plain, (spl177_36 <=> sP33(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_36])])).
fof(f11593, plain, ~ spl177_35, inference(avatar_contradiction_clause, [], [f11592])).
fof(f11592, plain, ($false | ~ spl177_35), inference(subsumption_resolution, [], [f11589, f9999])).
fof(f11589, plain, ((cell1 = cell36) | ~ spl177_35), inference(resolution, [], [f10801, f5606])).
fof(f10801, plain, (sP32(cell100, cell1) | ~ spl177_35), inference(avatar_component_clause, [], [f10799])).
fof(f10799, plain, (spl177_35 <=> sP32(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_35])])).
fof(f11587, plain, ~ spl177_34, inference(avatar_contradiction_clause, [], [f11586])).
fof(f11586, plain, ($false | ~ spl177_34), inference(subsumption_resolution, [], [f11583, f10034])).
fof(f11583, plain, ((cell1 = cell35) | ~ spl177_34), inference(resolution, [], [f10797, f5608])).
fof(f10797, plain, (sP31(cell100, cell1) | ~ spl177_34), inference(avatar_component_clause, [], [f10795])).
fof(f10795, plain, (spl177_34 <=> sP31(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_34])])).
fof(f11581, plain, ~ spl177_33, inference(avatar_contradiction_clause, [], [f11580])).
fof(f11580, plain, ($false | ~ spl177_33), inference(subsumption_resolution, [], [f11577, f10068])).
fof(f11577, plain, ((cell1 = cell34) | ~ spl177_33), inference(resolution, [], [f10793, f5610])).
fof(f10793, plain, (sP30(cell100, cell1) | ~ spl177_33), inference(avatar_component_clause, [], [f10791])).
fof(f10791, plain, (spl177_33 <=> sP30(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_33])])).
fof(f11575, plain, ~ spl177_32, inference(avatar_contradiction_clause, [], [f11574])).
fof(f11574, plain, ($false | ~ spl177_32), inference(subsumption_resolution, [], [f11571, f10101])).
fof(f11571, plain, ((cell1 = cell33) | ~ spl177_32), inference(resolution, [], [f10789, f5612])).
fof(f10789, plain, (sP29(cell100, cell1) | ~ spl177_32), inference(avatar_component_clause, [], [f10787])).
fof(f10787, plain, (spl177_32 <=> sP29(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_32])])).
fof(f11569, plain, ~ spl177_31, inference(avatar_contradiction_clause, [], [f11568])).
fof(f11568, plain, ($false | ~ spl177_31), inference(subsumption_resolution, [], [f11565, f10164])).
fof(f11565, plain, ((cell1 = cell32) | ~ spl177_31), inference(resolution, [], [f10785, f5614])).
fof(f10785, plain, (sP28(cell100, cell1) | ~ spl177_31), inference(avatar_component_clause, [], [f10783])).
fof(f10783, plain, (spl177_31 <=> sP28(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_31])])).
fof(f11563, plain, ~ spl177_30, inference(avatar_contradiction_clause, [], [f11562])).
fof(f11562, plain, ($false | ~ spl177_30), inference(subsumption_resolution, [], [f11559, f10133])).
fof(f11559, plain, ((cell1 = cell31) | ~ spl177_30), inference(resolution, [], [f10781, f5616])).
fof(f10781, plain, (sP27(cell100, cell1) | ~ spl177_30), inference(avatar_component_clause, [], [f10779])).
fof(f10779, plain, (spl177_30 <=> sP27(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_30])])).
fof(f11557, plain, ~ spl177_29, inference(avatar_contradiction_clause, [], [f11556])).
fof(f11556, plain, ($false | ~ spl177_29), inference(subsumption_resolution, [], [f11553, f10223])).
fof(f11553, plain, ((cell1 = cell29) | ~ spl177_29), inference(resolution, [], [f10777, f5618])).
fof(f10777, plain, (sP26(cell100, cell1) | ~ spl177_29), inference(avatar_component_clause, [], [f10775])).
fof(f10775, plain, (spl177_29 <=> sP26(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_29])])).
fof(f11551, plain, ~ spl177_28, inference(avatar_contradiction_clause, [], [f11550])).
fof(f11550, plain, ($false | ~ spl177_28), inference(subsumption_resolution, [], [f11547, f10251])).
fof(f11547, plain, ((cell1 = cell28) | ~ spl177_28), inference(resolution, [], [f10773, f5620])).
fof(f10773, plain, (sP25(cell100, cell1) | ~ spl177_28), inference(avatar_component_clause, [], [f10771])).
fof(f10771, plain, (spl177_28 <=> sP25(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_28])])).
fof(f11545, plain, ~ spl177_27, inference(avatar_contradiction_clause, [], [f11544])).
fof(f11544, plain, ($false | ~ spl177_27), inference(subsumption_resolution, [], [f11541, f10278])).
fof(f11541, plain, ((cell1 = cell27) | ~ spl177_27), inference(resolution, [], [f10769, f5622])).
fof(f10769, plain, (sP24(cell100, cell1) | ~ spl177_27), inference(avatar_component_clause, [], [f10767])).
fof(f10767, plain, (spl177_27 <=> sP24(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_27])])).
fof(f11539, plain, ~ spl177_26, inference(avatar_contradiction_clause, [], [f11538])).
fof(f11538, plain, ($false | ~ spl177_26), inference(subsumption_resolution, [], [f11535, f10304])).
fof(f11535, plain, ((cell1 = cell26) | ~ spl177_26), inference(resolution, [], [f10765, f5624])).
fof(f10765, plain, (sP23(cell100, cell1) | ~ spl177_26), inference(avatar_component_clause, [], [f10763])).
fof(f10763, plain, (spl177_26 <=> sP23(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_26])])).
fof(f11533, plain, ~ spl177_25, inference(avatar_contradiction_clause, [], [f11532])).
fof(f11532, plain, ($false | ~ spl177_25), inference(subsumption_resolution, [], [f11529, f10329])).
fof(f11529, plain, ((cell1 = cell25) | ~ spl177_25), inference(resolution, [], [f10761, f5626])).
fof(f10761, plain, (sP22(cell100, cell1) | ~ spl177_25), inference(avatar_component_clause, [], [f10759])).
fof(f10759, plain, (spl177_25 <=> sP22(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_25])])).
fof(f11527, plain, ~ spl177_24, inference(avatar_contradiction_clause, [], [f11526])).
fof(f11526, plain, ($false | ~ spl177_24), inference(subsumption_resolution, [], [f11523, f10353])).
fof(f11523, plain, ((cell1 = cell24) | ~ spl177_24), inference(resolution, [], [f10757, f5628])).
fof(f10757, plain, (sP21(cell100, cell1) | ~ spl177_24), inference(avatar_component_clause, [], [f10755])).
fof(f10755, plain, (spl177_24 <=> sP21(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_24])])).
fof(f11521, plain, ~ spl177_23, inference(avatar_contradiction_clause, [], [f11520])).
fof(f11520, plain, ($false | ~ spl177_23), inference(subsumption_resolution, [], [f11517, f10376])).
fof(f11517, plain, ((cell1 = cell23) | ~ spl177_23), inference(resolution, [], [f10753, f5630])).
fof(f10753, plain, (sP20(cell100, cell1) | ~ spl177_23), inference(avatar_component_clause, [], [f10751])).
fof(f10751, plain, (spl177_23 <=> sP20(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_23])])).
fof(f11515, plain, ~ spl177_22, inference(avatar_contradiction_clause, [], [f11514])).
fof(f11514, plain, ($false | ~ spl177_22), inference(subsumption_resolution, [], [f11511, f10419])).
fof(f11511, plain, ((cell1 = cell22) | ~ spl177_22), inference(resolution, [], [f10749, f5632])).
fof(f10749, plain, (sP19(cell100, cell1) | ~ spl177_22), inference(avatar_component_clause, [], [f10747])).
fof(f10747, plain, (spl177_22 <=> sP19(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_22])])).
fof(f11509, plain, ~ spl177_21, inference(avatar_contradiction_clause, [], [f11508])).
fof(f11508, plain, ($false | ~ spl177_21), inference(subsumption_resolution, [], [f11505, f10398])).
fof(f11505, plain, ((cell1 = cell21) | ~ spl177_21), inference(resolution, [], [f10745, f5634])).
fof(f10745, plain, (sP18(cell100, cell1) | ~ spl177_21), inference(avatar_component_clause, [], [f10743])).
fof(f10743, plain, (spl177_21 <=> sP18(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_21])])).
fof(f11503, plain, ~ spl177_20, inference(avatar_contradiction_clause, [], [f11502])).
fof(f11502, plain, ($false | ~ spl177_20), inference(subsumption_resolution, [], [f11499, f10458])).
fof(f11499, plain, ((cell1 = cell19) | ~ spl177_20), inference(resolution, [], [f10741, f5636])).
fof(f10741, plain, (sP17(cell100, cell1) | ~ spl177_20), inference(avatar_component_clause, [], [f10739])).
fof(f10739, plain, (spl177_20 <=> sP17(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_20])])).
fof(f11497, plain, ~ spl177_19, inference(avatar_contradiction_clause, [], [f11496])).
fof(f11496, plain, ($false | ~ spl177_19), inference(subsumption_resolution, [], [f11493, f10476])).
fof(f11493, plain, ((cell1 = cell18) | ~ spl177_19), inference(resolution, [], [f10737, f5638])).
fof(f10737, plain, (sP16(cell100, cell1) | ~ spl177_19), inference(avatar_component_clause, [], [f10735])).
fof(f10735, plain, (spl177_19 <=> sP16(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_19])])).
fof(f11491, plain, ~ spl177_18, inference(avatar_contradiction_clause, [], [f11490])).
fof(f11490, plain, ($false | ~ spl177_18), inference(subsumption_resolution, [], [f11487, f10493])).
fof(f11487, plain, ((cell1 = cell17) | ~ spl177_18), inference(resolution, [], [f10733, f5640])).
fof(f10733, plain, (sP15(cell100, cell1) | ~ spl177_18), inference(avatar_component_clause, [], [f10731])).
fof(f10731, plain, (spl177_18 <=> sP15(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_18])])).
fof(f11485, plain, ~ spl177_17, inference(avatar_contradiction_clause, [], [f11484])).
fof(f11484, plain, ($false | ~ spl177_17), inference(subsumption_resolution, [], [f11481, f10509])).
fof(f11481, plain, ((cell1 = cell16) | ~ spl177_17), inference(resolution, [], [f10729, f5642])).
fof(f10729, plain, (sP14(cell100, cell1) | ~ spl177_17), inference(avatar_component_clause, [], [f10727])).
fof(f10727, plain, (spl177_17 <=> sP14(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_17])])).
fof(f11479, plain, ~ spl177_16, inference(avatar_contradiction_clause, [], [f11478])).
fof(f11478, plain, ($false | ~ spl177_16), inference(subsumption_resolution, [], [f11475, f10524])).
fof(f11475, plain, ((cell1 = cell15) | ~ spl177_16), inference(resolution, [], [f10725, f5644])).
fof(f10725, plain, (sP13(cell100, cell1) | ~ spl177_16), inference(avatar_component_clause, [], [f10723])).
fof(f10723, plain, (spl177_16 <=> sP13(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_16])])).
fof(f11473, plain, ~ spl177_15, inference(avatar_contradiction_clause, [], [f11472])).
fof(f11472, plain, ($false | ~ spl177_15), inference(subsumption_resolution, [], [f11469, f10538])).
fof(f11469, plain, ((cell1 = cell14) | ~ spl177_15), inference(resolution, [], [f10721, f5646])).
fof(f10721, plain, (sP12(cell100, cell1) | ~ spl177_15), inference(avatar_component_clause, [], [f10719])).
fof(f10719, plain, (spl177_15 <=> sP12(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_15])])).
fof(f11467, plain, ~ spl177_14, inference(avatar_contradiction_clause, [], [f11466])).
fof(f11466, plain, ($false | ~ spl177_14), inference(subsumption_resolution, [], [f11463, f10551])).
fof(f11463, plain, ((cell1 = cell13) | ~ spl177_14), inference(resolution, [], [f10717, f5648])).
fof(f10717, plain, (sP11(cell100, cell1) | ~ spl177_14), inference(avatar_component_clause, [], [f10715])).
fof(f10715, plain, (spl177_14 <=> sP11(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_14])])).
fof(f11461, plain, ~ spl177_13, inference(avatar_contradiction_clause, [], [f11460])).
fof(f11460, plain, ($false | ~ spl177_13), inference(subsumption_resolution, [], [f11457, f10574])).
fof(f11457, plain, ((cell1 = cell12) | ~ spl177_13), inference(resolution, [], [f10713, f5650])).
fof(f10713, plain, (sP10(cell100, cell1) | ~ spl177_13), inference(avatar_component_clause, [], [f10711])).
fof(f10711, plain, (spl177_13 <=> sP10(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_13])])).
fof(f11455, plain, ~ spl177_12, inference(avatar_contradiction_clause, [], [f11454])).
fof(f11454, plain, ($false | ~ spl177_12), inference(subsumption_resolution, [], [f11451, f10563])).
fof(f11451, plain, ((cell1 = cell11) | ~ spl177_12), inference(resolution, [], [f10709, f5652])).
fof(f10709, plain, (sP9(cell100, cell1) | ~ spl177_12), inference(avatar_component_clause, [], [f10707])).
fof(f10707, plain, (spl177_12 <=> sP9(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_12])])).
fof(f11449, plain, ~ spl177_11, inference(avatar_contradiction_clause, [], [f11448])).
fof(f11448, plain, ($false | ~ spl177_11), inference(subsumption_resolution, [], [f11445, f10593])).
fof(f11445, plain, ((cell1 = cell9) | ~ spl177_11), inference(resolution, [], [f10705, f5654])).
fof(f10705, plain, (sP8(cell100, cell1) | ~ spl177_11), inference(avatar_component_clause, [], [f10703])).
fof(f10703, plain, (spl177_11 <=> sP8(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_11])])).
fof(f11443, plain, ~ spl177_10, inference(avatar_contradiction_clause, [], [f11442])).
fof(f11442, plain, ($false | ~ spl177_10), inference(subsumption_resolution, [], [f11439, f10601])).
fof(f11439, plain, ((cell1 = cell8) | ~ spl177_10), inference(resolution, [], [f10701, f5656])).
fof(f10701, plain, (sP7(cell100, cell1) | ~ spl177_10), inference(avatar_component_clause, [], [f10699])).
fof(f10699, plain, (spl177_10 <=> sP7(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_10])])).
fof(f11437, plain, ~ spl177_9, inference(avatar_contradiction_clause, [], [f11436])).
fof(f11436, plain, ($false | ~ spl177_9), inference(subsumption_resolution, [], [f11433, f10608])).
fof(f11433, plain, ((cell1 = cell7) | ~ spl177_9), inference(resolution, [], [f10697, f5658])).
fof(f10697, plain, (sP6(cell100, cell1) | ~ spl177_9), inference(avatar_component_clause, [], [f10695])).
fof(f10695, plain, (spl177_9 <=> sP6(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_9])])).
fof(f11431, plain, ~ spl177_8, inference(avatar_contradiction_clause, [], [f11430])).
fof(f11430, plain, ($false | ~ spl177_8), inference(subsumption_resolution, [], [f11427, f10614])).
fof(f11427, plain, ((cell1 = cell6) | ~ spl177_8), inference(resolution, [], [f10693, f5660])).
fof(f10693, plain, (sP5(cell100, cell1) | ~ spl177_8), inference(avatar_component_clause, [], [f10691])).
fof(f10691, plain, (spl177_8 <=> sP5(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_8])])).
fof(f11425, plain, ~ spl177_7, inference(avatar_contradiction_clause, [], [f11424])).
fof(f11424, plain, ($false | ~ spl177_7), inference(subsumption_resolution, [], [f11421, f10619])).
fof(f11421, plain, ((cell1 = cell5) | ~ spl177_7), inference(resolution, [], [f10689, f5662])).
fof(f10689, plain, (sP4(cell100, cell1) | ~ spl177_7), inference(avatar_component_clause, [], [f10687])).
fof(f10687, plain, (spl177_7 <=> sP4(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_7])])).
fof(f11419, plain, ~ spl177_6, inference(avatar_contradiction_clause, [], [f11418])).
fof(f11418, plain, ($false | ~ spl177_6), inference(subsumption_resolution, [], [f11415, f10623])).
fof(f11415, plain, ((cell1 = cell4) | ~ spl177_6), inference(resolution, [], [f10685, f5664])).
fof(f10685, plain, (sP3(cell100, cell1) | ~ spl177_6), inference(avatar_component_clause, [], [f10683])).
fof(f10683, plain, (spl177_6 <=> sP3(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_6])])).
fof(f11413, plain, (spl177_5 | spl177_6 | spl177_7 | spl177_8 | spl177_9 | spl177_10 | spl177_11 | spl177_12 | spl177_13 | spl177_14 | spl177_15 | spl177_16 | spl177_17 | spl177_18 | spl177_19 | spl177_20 | spl177_21 | spl177_22 | spl177_23 | spl177_24 | spl177_25 | spl177_26 | spl177_27 | spl177_28 | spl177_29 | spl177_30 | spl177_31 | spl177_32 | spl177_33 | spl177_34 | spl177_35 | spl177_36 | spl177_37 | spl177_38 | spl177_39 | spl177_40 | spl177_41 | spl177_42 | spl177_43 | spl177_44 | spl177_45 | spl177_46 | spl177_47 | spl177_48 | spl177_49 | spl177_50 | spl177_51 | spl177_52 | spl177_53 | spl177_54 | spl177_55 | spl177_56 | spl177_57 | spl177_58 | spl177_59 | spl177_60 | spl177_61 | spl177_62 | spl177_63 | spl177_64 | spl177_65 | spl177_66 | spl177_67 | spl177_68 | spl177_69 | spl177_70 | spl177_71 | spl177_72 | spl177_73 | spl177_74 | spl177_75 | spl177_76 | spl177_77 | spl177_78 | spl177_79 | spl177_80 | spl177_81 | spl177_84 | spl177_85 | spl177_86 | spl177_87 | spl177_88 | spl177_89 | spl177_179 | spl177_178 | spl177_90 | spl177_91 | spl177_92 | spl177_93 | spl177_94 | spl177_95 | spl177_96 | spl177_97 | spl177_98 | spl177_99 | spl177_100 | spl177_101 | spl177_102 | spl177_103 | spl177_104 | spl177_105 | spl177_106 | spl177_107 | spl177_108 | spl177_109 | spl177_110 | spl177_111 | spl177_112 | spl177_113 | spl177_114 | spl177_115 | spl177_116 | spl177_117 | spl177_118 | spl177_119 | spl177_120 | spl177_121 | spl177_122 | spl177_123 | spl177_124 | spl177_125 | spl177_126 | spl177_127 | spl177_128 | spl177_129 | spl177_130 | spl177_131 | spl177_132 | spl177_133 | spl177_134 | spl177_135 | spl177_136 | spl177_137 | spl177_138 | spl177_139 | spl177_140 | spl177_141 | spl177_142 | spl177_143 | spl177_144 | spl177_145 | spl177_146 | spl177_147 | spl177_148 | spl177_149 | spl177_150 | spl177_151 | spl177_152 | spl177_153 | spl177_154 | spl177_155 | spl177_156 | spl177_157 | spl177_158 | spl177_159 | spl177_160 | spl177_161 | spl177_162 | spl177_163 | spl177_164 | spl177_165 | spl177_166 | spl177_167 | spl177_168 | spl177_169 | spl177_170 | spl177_171 | spl177_172 | spl177_173 | spl177_174 | spl177_175 | spl177_176 | spl177_177 | ~ spl177_2 | spl177_3 | spl177_4), inference(avatar_split_clause, [], [f11412, f10675, f10671, f10658, f11367, f11363, f11359, f11355, f11351, f11347, f11343, f11339, f11335, f11331, f11327, f11323, f11319, f11315, f11311, f11307, f11303, f11299, f11295, f11291, f11287, f11283, f11279, f11275, f11271, f11267, f11263, f11259, f11255, f11251, f11247, f11243, f11239, f11235, f11231, f11227, f11223, f11219, f11215, f11211, f11207, f11203, f11199, f11195, f11191, f11187, f11183, f11179, f11175, f11171, f11167, f11163, f11159, f11155, f11151, f11147, f11143, f11139, f11135, f11131, f11127, f11123, f11119, f11115, f11111, f11107, f11103, f11099, f11095, f11091, f11087, f11083, f11079, f11075, f11071, f11067, f11063, f11059, f11055, f11051, f11047, f11043, f11039, f11035, f11031, f11027, f11023, f11019, f11388, f11397, f11015, f11011, f11007, f11003, f10999, f10995, f10983, f10979, f10975, f10971, f10967, f10963, f10959, f10955, f10951, f10947, f10943, f10939, f10935, f10931, f10927, f10923, f10919, f10915, f10911, f10907, f10903, f10899, f10895, f10891, f10887, f10883, f10879, f10875, f10871, f10867, f10863, f10859, f10855, f10851, f10847, f10843, f10839, f10835, f10831, f10827, f10823, f10819, f10815, f10811, f10807, f10803, f10799, f10795, f10791, f10787, f10783, f10779, f10775, f10771, f10767, f10763, f10759, f10755, f10751, f10747, f10743, f10739, f10735, f10731, f10727, f10723, f10719, f10715, f10711, f10707, f10703, f10699, f10695, f10691, f10687, f10683, f10679])).
fof(f10679, plain, (spl177_5 <=> sP2(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_5])])).
fof(f10658, plain, (spl177_2 <=> edge(cell1, cell100)), introduced(avatar_definition, [new_symbols(naming, [spl177_2])])).
fof(f10671, plain, (spl177_3 <=> sP0(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_3])])).
fof(f10675, plain, (spl177_4 <=> sP1(cell100, cell1)), introduced(avatar_definition, [new_symbols(naming, [spl177_4])])).
fof(f11412, plain, (sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | (~ spl177_2 | spl177_3 | spl177_4)), inference(subsumption_resolution, [], [f11411, f6624])).
fof(f6624, plain, ~ (cell1 = cell90), inference(cnf_transformation, [], [f946])).
fof(f946, plain, ~ (cell1 = cell90), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof25935)).
fof(f11411, plain, (sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | (cell1 = cell90) | (~ spl177_2 | spl177_3 | spl177_4)), inference(subsumption_resolution, [], [f11410, f10672])).
fof(f10672, plain, (~ sP0(cell100, cell1) | spl177_3), inference(avatar_component_clause, [], [f10671])).
fof(f11410, plain, (sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | sP0(cell100, cell1) | (cell1 = cell90) | (~ spl177_2 | spl177_4)), inference(subsumption_resolution, [], [f11409, f10676])).
fof(f10676, plain, (~ sP1(cell100, cell1) | spl177_4), inference(avatar_component_clause, [], [f10675])).
fof(f11409, plain, (sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | sP1(cell100, cell1) | sP0(cell100, cell1) | (cell1 = cell90) | ~ spl177_2), inference(subsumption_resolution, [], [f11408, f6801])).
fof(f6801, plain, ~ (cell1 = cell88), inference(cnf_transformation, [], [f1123])).
fof(f1123, plain, ~ (cell1 = cell88), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26112)).
fof(f11408, plain, ((cell1 = cell88) | sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | sP1(cell100, cell1) | sP0(cell100, cell1) | (cell1 = cell90) | ~ spl177_2), inference(subsumption_resolution, [], [f11401, f6713])).
fof(f6713, plain, ~ (cell1 = cell89), inference(cnf_transformation, [], [f1035])).
fof(f1035, plain, ~ (cell1 = cell89), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof26024)).
fof(f11401, plain, ((cell1 = cell89) | (cell1 = cell88) | sP176(cell100, cell1) | sP175(cell100, cell1) | sP174(cell100, cell1) | sP173(cell100, cell1) | sP172(cell100, cell1) | sP171(cell100, cell1) | sP170(cell100, cell1) | sP169(cell100, cell1) | sP168(cell100, cell1) | sP167(cell100, cell1) | sP166(cell100, cell1) | sP165(cell100, cell1) | sP164(cell100, cell1) | sP163(cell100, cell1) | sP162(cell100, cell1) | sP161(cell100, cell1) | sP160(cell100, cell1) | sP159(cell100, cell1) | sP158(cell100, cell1) | sP157(cell100, cell1) | sP156(cell100, cell1) | sP155(cell100, cell1) | sP154(cell100, cell1) | sP153(cell100, cell1) | sP152(cell100, cell1) | sP151(cell100, cell1) | sP150(cell100, cell1) | sP149(cell100, cell1) | sP148(cell100, cell1) | sP147(cell100, cell1) | sP146(cell100, cell1) | sP145(cell100, cell1) | sP144(cell100, cell1) | sP143(cell100, cell1) | sP142(cell100, cell1) | sP141(cell100, cell1) | sP140(cell100, cell1) | sP139(cell100, cell1) | sP138(cell100, cell1) | sP137(cell100, cell1) | sP136(cell100, cell1) | sP135(cell100, cell1) | sP134(cell100, cell1) | sP133(cell100, cell1) | sP132(cell100, cell1) | sP131(cell100, cell1) | sP130(cell100, cell1) | sP129(cell100, cell1) | sP128(cell100, cell1) | sP127(cell100, cell1) | sP126(cell100, cell1) | sP125(cell100, cell1) | sP124(cell100, cell1) | sP123(cell100, cell1) | sP122(cell100, cell1) | sP121(cell100, cell1) | sP120(cell100, cell1) | sP119(cell100, cell1) | sP118(cell100, cell1) | sP117(cell100, cell1) | sP116(cell100, cell1) | sP115(cell100, cell1) | sP114(cell100, cell1) | sP113(cell100, cell1) | sP112(cell100, cell1) | sP111(cell100, cell1) | sP110(cell100, cell1) | sP109(cell100, cell1) | sP108(cell100, cell1) | sP107(cell100, cell1) | sP106(cell100, cell1) | sP105(cell100, cell1) | sP104(cell100, cell1) | sP103(cell100, cell1) | sP102(cell100, cell1) | sP101(cell100, cell1) | sP100(cell100, cell1) | sP99(cell100, cell1) | sP98(cell100, cell1) | sP97(cell100, cell1) | sP96(cell100, cell1) | sP95(cell100, cell1) | sP94(cell100, cell1) | sP93(cell100, cell1) | sP92(cell100, cell1) | sP91(cell100, cell1) | sP90(cell100, cell1) | sP89(cell100, cell1) | sP88(cell100, cell1) | sP87(cell100, cell1) | sP86(cell100, cell1) | sP85(cell100, cell1) | sP84(cell100, cell1) | sP83(cell100, cell1) | sP82(cell100, cell1) | sP81(cell100, cell1) | sP78(cell100, cell1) | sP77(cell100, cell1) | sP76(cell100, cell1) | sP75(cell100, cell1) | sP74(cell100, cell1) | sP73(cell100, cell1) | sP72(cell100, cell1) | sP71(cell100, cell1) | sP70(cell100, cell1) | sP69(cell100, cell1) | sP68(cell100, cell1) | sP67(cell100, cell1) | sP66(cell100, cell1) | sP65(cell100, cell1) | sP64(cell100, cell1) | sP63(cell100, cell1) | sP62(cell100, cell1) | sP61(cell100, cell1) | sP60(cell100, cell1) | sP59(cell100, cell1) | sP58(cell100, cell1) | sP57(cell100, cell1) | sP56(cell100, cell1) | sP55(cell100, cell1) | sP54(cell100, cell1) | sP53(cell100, cell1) | sP52(cell100, cell1) | sP51(cell100, cell1) | sP50(cell100, cell1) | sP49(cell100, cell1) | sP48(cell100, cell1) | sP47(cell100, cell1) | sP46(cell100, cell1) | sP45(cell100, cell1) | sP44(cell100, cell1) | sP43(cell100, cell1) | sP42(cell100, cell1) | sP41(cell100, cell1) | sP40(cell100, cell1) | sP39(cell100, cell1) | sP38(cell100, cell1) | sP37(cell100, cell1) | sP36(cell100, cell1) | sP35(cell100, cell1) | sP34(cell100, cell1) | sP33(cell100, cell1) | sP32(cell100, cell1) | sP31(cell100, cell1) | sP30(cell100, cell1) | sP29(cell100, cell1) | sP28(cell100, cell1) | sP27(cell100, cell1) | sP26(cell100, cell1) | sP25(cell100, cell1) | sP24(cell100, cell1) | sP23(cell100, cell1) | sP22(cell100, cell1) | sP21(cell100, cell1) | sP20(cell100, cell1) | sP19(cell100, cell1) | sP18(cell100, cell1) | sP17(cell100, cell1) | sP16(cell100, cell1) | sP15(cell100, cell1) | sP14(cell100, cell1) | sP13(cell100, cell1) | sP12(cell100, cell1) | sP11(cell100, cell1) | sP10(cell100, cell1) | sP9(cell100, cell1) | sP8(cell100, cell1) | sP7(cell100, cell1) | sP6(cell100, cell1) | sP5(cell100, cell1) | sP4(cell100, cell1) | sP3(cell100, cell1) | sP2(cell100, cell1) | sP1(cell100, cell1) | sP0(cell100, cell1) | (cell1 = cell90) | ~ spl177_2), inference(resolution, [], [f10651, f10660])).
fof(f10660, plain, (edge(cell1, cell100) | ~ spl177_2), inference(avatar_component_clause, [], [f10658])).
fof(f10651, plain, ! [X0, X1] : (~ edge(X1, X0) | (cell89 = X1) | (cell88 = X1) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP87(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | (cell90 = X1)), inference(subsumption_resolution, [], [f10650, f5510])).
fof(f10650, plain, ! [X0, X1] : ((cell90 = X1) | (cell89 = X1) | (cell88 = X1) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP87(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(subsumption_resolution, [], [f5672, f5512])).
fof(f5672, plain, ! [X0, X1] : ((cell90 = X1) | (cell89 = X1) | (cell88 = X1) | sP176(X0, X1) | sP175(X0, X1) | sP174(X0, X1) | sP173(X0, X1) | sP172(X0, X1) | sP171(X0, X1) | sP170(X0, X1) | sP169(X0, X1) | sP168(X0, X1) | sP167(X0, X1) | sP166(X0, X1) | sP165(X0, X1) | sP164(X0, X1) | sP163(X0, X1) | sP162(X0, X1) | sP161(X0, X1) | sP160(X0, X1) | sP159(X0, X1) | sP158(X0, X1) | sP157(X0, X1) | sP156(X0, X1) | sP155(X0, X1) | sP154(X0, X1) | sP153(X0, X1) | sP152(X0, X1) | sP151(X0, X1) | sP150(X0, X1) | sP149(X0, X1) | sP148(X0, X1) | sP147(X0, X1) | sP146(X0, X1) | sP145(X0, X1) | sP144(X0, X1) | sP143(X0, X1) | sP142(X0, X1) | sP141(X0, X1) | sP140(X0, X1) | sP139(X0, X1) | sP138(X0, X1) | sP137(X0, X1) | sP136(X0, X1) | sP135(X0, X1) | sP134(X0, X1) | sP133(X0, X1) | sP132(X0, X1) | sP131(X0, X1) | sP130(X0, X1) | sP129(X0, X1) | sP128(X0, X1) | sP127(X0, X1) | sP126(X0, X1) | sP125(X0, X1) | sP124(X0, X1) | sP123(X0, X1) | sP122(X0, X1) | sP121(X0, X1) | sP120(X0, X1) | sP119(X0, X1) | sP118(X0, X1) | sP117(X0, X1) | sP116(X0, X1) | sP115(X0, X1) | sP114(X0, X1) | sP113(X0, X1) | sP112(X0, X1) | sP111(X0, X1) | sP110(X0, X1) | sP109(X0, X1) | sP108(X0, X1) | sP107(X0, X1) | sP106(X0, X1) | sP105(X0, X1) | sP104(X0, X1) | sP103(X0, X1) | sP102(X0, X1) | sP101(X0, X1) | sP100(X0, X1) | sP99(X0, X1) | sP98(X0, X1) | sP97(X0, X1) | sP96(X0, X1) | sP95(X0, X1) | sP94(X0, X1) | sP93(X0, X1) | sP92(X0, X1) | sP91(X0, X1) | sP90(X0, X1) | sP89(X0, X1) | sP88(X0, X1) | sP87(X0, X1) | sP86(X0, X1) | sP85(X0, X1) | sP84(X0, X1) | sP83(X0, X1) | sP82(X0, X1) | sP81(X0, X1) | sP80(X0, X1) | sP79(X0, X1) | sP78(X0, X1) | sP77(X0, X1) | sP76(X0, X1) | sP75(X0, X1) | sP74(X0, X1) | sP73(X0, X1) | sP72(X0, X1) | sP71(X0, X1) | sP70(X0, X1) | sP69(X0, X1) | sP68(X0, X1) | sP67(X0, X1) | sP66(X0, X1) | sP65(X0, X1) | sP64(X0, X1) | sP63(X0, X1) | sP62(X0, X1) | sP61(X0, X1) | sP60(X0, X1) | sP59(X0, X1) | sP58(X0, X1) | sP57(X0, X1) | sP56(X0, X1) | sP55(X0, X1) | sP54(X0, X1) | sP53(X0, X1) | sP52(X0, X1) | sP51(X0, X1) | sP50(X0, X1) | sP49(X0, X1) | sP48(X0, X1) | sP47(X0, X1) | sP46(X0, X1) | sP45(X0, X1) | sP44(X0, X1) | sP43(X0, X1) | sP42(X0, X1) | sP41(X0, X1) | sP40(X0, X1) | sP39(X0, X1) | sP38(X0, X1) | sP37(X0, X1) | sP36(X0, X1) | sP35(X0, X1) | sP34(X0, X1) | sP33(X0, X1) | sP32(X0, X1) | sP31(X0, X1) | sP30(X0, X1) | sP29(X0, X1) | sP28(X0, X1) | sP27(X0, X1) | sP26(X0, X1) | sP25(X0, X1) | sP24(X0, X1) | sP23(X0, X1) | sP22(X0, X1) | sP21(X0, X1) | sP20(X0, X1) | sP19(X0, X1) | sP18(X0, X1) | sP17(X0, X1) | sP16(X0, X1) | sP15(X0, X1) | sP14(X0, X1) | sP13(X0, X1) | sP12(X0, X1) | sP11(X0, X1) | sP10(X0, X1) | sP9(X0, X1) | sP8(X0, X1) | sP7(X0, X1) | sP6(X0, X1) | sP5(X0, X1) | sP4(X0, X1) | sP3(X0, X1) | sP2(X0, X1) | sP1(X0, X1) | sP0(X0, X1) | ~ edge(X1, X0)), inference(cnf_transformation, [], [f5140])).
fof(f11407, plain, ~ spl177_5, inference(avatar_contradiction_clause, [], [f11406])).
fof(f11406, plain, ($false | ~ spl177_5), inference(subsumption_resolution, [], [f11403, f10626])).
fof(f11403, plain, ((cell1 = cell3) | ~ spl177_5), inference(resolution, [], [f10681, f5666])).
fof(f10681, plain, (sP2(cell100, cell1) | ~ spl177_5), inference(avatar_component_clause, [], [f10679])).
fof(f11382, plain, ~ spl177_4, inference(avatar_contradiction_clause, [], [f11381])).
fof(f11381, plain, ($false | ~ spl177_4), inference(subsumption_resolution, [], [f11378, f10628])).
fof(f11378, plain, ((cell1 = cell2) | ~ spl177_4), inference(resolution, [], [f10677, f5668])).
fof(f10677, plain, (sP1(cell100, cell1) | ~ spl177_4), inference(avatar_component_clause, [], [f10675])).
fof(f11374, plain, ~ spl177_3, inference(avatar_contradiction_clause, [], [f11373])).
fof(f11373, plain, ($false | ~ spl177_3), inference(subsumption_resolution, [], [f11371, f10629])).
fof(f11371, plain, ((cell2 = cell100) | ~ spl177_3), inference(resolution, [], [f10673, f5671])).
fof(f5671, plain, ! [X0, X1] : (~ sP0(X0, X1) | (cell2 = X0)), inference(cnf_transformation, [], [f5317])).
fof(f10673, plain, (sP0(cell100, cell1) | ~ spl177_3), inference(avatar_component_clause, [], [f10671])).
fof(f10661, plain, (spl177_1 | spl177_2), inference(avatar_split_clause, [], [f10652, f10658, f10654])).
fof(f10652, plain, (edge(cell1, cell100) | edge(cell100, cell1)), inference(resolution, [], [f5680, f10631])).
fof(f10631, plain, adj(cell1, cell100), inference(cnf_transformation, [], [f4956])).
fof(f4956, plain, adj(cell1, cell100), inference(flattening, [], [f4954])).
fof(f4954, plain, ~ ~ adj(cell1, cell100), inference(negated_conjecture, [], [f4953])).
fof(f4953, plain, ~ ~ adj(cell1, cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof29942)).
fof(f5680, plain, ! [X0, X1] : (~ adj(X1, X0) | edge(X1, X0) | edge(X0, X1)), inference(cnf_transformation, [], [f4962])).
fof(f4962, plain, ! [X0, X1] : (edge(X0, X1) | edge(X1, X0) | ~ adj(X1, X0)), inference(flattening, [], [f4961])).
fof(f4961, plain, ! [X0, X1] : ((edge(X0, X1) | edge(X1, X0)) | ~ adj(X1, X0)), inference(ennf_transformation, [], [f4957])).
fof(f4957, plain, ! [X0, X1] : (adj(X1, X0) => (edge(X0, X1) | edge(X1, X0))), inference(unused_predicate_definition_removal, [], [f4955])).
fof(f4955, plain, ! [X0, X1] : (adj(X1, X0) <=> (edge(X0, X1) | edge(X1, X0))), inference(rectify, [], [f2])).
fof(f2, plain, ! [X2, X3] : (adj(X3, X2) <=> (edge(X2, X3) | edge(X3, X2))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ074+1.p', tlhfof24991)).