fof(f12732, plain, $false, inference(subsumption_resolution, [], [f12731, f245])).
fof(f245, plain, ~ westof(cell1, cell100), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ~ westof(cell1, cell100), inference(flattening, [], [f8])).
fof(f8, plain, ~ westof(cell1, cell100), inference(negated_conjecture, [], [f7])).
fof(f7, plain, ~ westof(cell1, cell100), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34981)).
fof(f12731, plain, westof(cell1, cell100), inference(resolution, [], [f12515, f12085])).
fof(f12085, plain, duewest(cell10, cell100), inference(resolution, [], [f11442, f695])).
fof(f695, plain, west(cell90, cell100), inference(resolution, [], [f129, f247])).
fof(f247, plain, ~ sP0(cell100, cell90), inference(equality_resolution, [], [f246])).
fof(f246, plain, ! [X1] : (~ (cell90 = X1) | ~ sP0(cell100, X1)), inference(equality_resolution, [], [f128])).
fof(f128, plain, ! [X0, X1] : (~ (cell100 = X0) | ~ (cell90 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f35, plain, ! [X0, X1] : (((~ (cell100 = X0) | ~ (cell90 = X1)) & (~ (cell99 = X0) | ~ (cell89 = X1)) & (~ (cell98 = X0) | ~ (cell88 = X1)) & (~ (cell97 = X0) | ~ (cell87 = X1)) & (~ (cell96 = X0) | ~ (cell86 = X1)) & (~ (cell95 = X0) | ~ (cell85 = X1)) & (~ (cell94 = X0) | ~ (cell84 = X1)) & (~ (cell93 = X0) | ~ (cell83 = X1)) & (~ (cell92 = X0) | ~ (cell82 = X1)) & (~ (cell91 = X0) | ~ (cell81 = X1)) & (~ (cell90 = X0) | ~ (cell80 = X1)) & (~ (cell89 = X0) | ~ (cell79 = X1)) & (~ (cell88 = X0) | ~ (cell78 = X1)) & (~ (cell87 = X0) | ~ (cell77 = X1)) & (~ (cell86 = X0) | ~ (cell76 = X1)) & (~ (cell85 = X0) | ~ (cell75 = X1)) & (~ (cell84 = X0) | ~ (cell74 = X1)) & (~ (cell83 = X0) | ~ (cell73 = X1)) & (~ (cell82 = X0) | ~ (cell72 = X1)) & (~ (cell81 = X0) | ~ (cell71 = X1)) & (~ (cell80 = X0) | ~ (cell70 = X1)) & (~ (cell79 = X0) | ~ (cell69 = X1)) & (~ (cell78 = X0) | ~ (cell68 = X1)) & (~ (cell77 = X0) | ~ (cell67 = X1)) & (~ (cell76 = X0) | ~ (cell66 = X1)) & (~ (cell75 = X0) | ~ (cell65 = X1)) & (~ (cell74 = X0) | ~ (cell64 = X1)) & (~ (cell73 = X0) | ~ (cell63 = X1)) & (~ (cell72 = X0) | ~ (cell62 = X1)) & (~ (cell71 = X0) | ~ (cell61 = X1)) & (~ (cell70 = X0) | ~ (cell60 = X1)) & (~ (cell69 = X0) | ~ (cell59 = X1)) & (~ (cell68 = X0) | ~ (cell58 = X1)) & (~ (cell67 = X0) | ~ (cell57 = X1)) & (~ (cell66 = X0) | ~ (cell56 = X1)) & (~ (cell65 = X0) | ~ (cell55 = X1)) & (~ (cell64 = X0) | ~ (cell54 = X1)) & (~ (cell63 = X0) | ~ (cell53 = X1)) & (~ (cell62 = X0) | ~ (cell52 = X1)) & (~ (cell61 = X0) | ~ (cell51 = X1)) & (~ (cell60 = X0) | ~ (cell50 = X1)) & (~ (cell59 = X0) | ~ (cell49 = X1)) & (~ (cell58 = X0) | ~ (cell48 = X1)) & (~ (cell57 = X0) | ~ (cell47 = X1)) & (~ (cell56 = X0) | ~ (cell46 = X1)) & (~ (cell55 = X0) | ~ (cell45 = X1)) & (~ (cell54 = X0) | ~ (cell44 = X1)) & (~ (cell53 = X0) | ~ (cell43 = X1)) & (~ (cell52 = X0) | ~ (cell42 = X1)) & (~ (cell51 = X0) | ~ (cell41 = X1)) & (~ (cell50 = X0) | ~ (cell40 = X1)) & (~ (cell49 = X0) | ~ (cell39 = X1)) & (~ (cell48 = X0) | ~ (cell38 = X1)) & (~ (cell47 = X0) | ~ (cell37 = X1)) & (~ (cell46 = X0) | ~ (cell36 = X1)) & (~ (cell45 = X0) | ~ (cell35 = X1)) & (~ (cell44 = X0) | ~ (cell34 = X1)) & (~ (cell43 = X0) | ~ (cell33 = X1)) & (~ (cell42 = X0) | ~ (cell32 = X1)) & (~ (cell41 = X0) | ~ (cell31 = X1)) & (~ (cell40 = X0) | ~ (cell30 = X1)) & (~ (cell39 = X0) | ~ (cell29 = X1)) & (~ (cell38 = X0) | ~ (cell28 = X1)) & (~ (cell37 = X0) | ~ (cell27 = X1)) & (~ (cell36 = X0) | ~ (cell26 = X1)) & (~ (cell35 = X0) | ~ (cell25 = X1)) & (~ (cell34 = X0) | ~ (cell24 = X1)) & (~ (cell33 = X0) | ~ (cell23 = X1)) & (~ (cell32 = X0) | ~ (cell22 = X1)) & (~ (cell31 = X0) | ~ (cell21 = X1)) & (~ (cell30 = X0) | ~ (cell20 = X1)) & (~ (cell29 = X0) | ~ (cell19 = X1)) & (~ (cell28 = X0) | ~ (cell18 = X1)) & (~ (cell27 = X0) | ~ (cell17 = X1)) & (~ (cell26 = X0) | ~ (cell16 = X1)) & (~ (cell25 = X0) | ~ (cell15 = X1)) & (~ (cell24 = X0) | ~ (cell14 = X1)) & (~ (cell23 = X0) | ~ (cell13 = X1)) & (~ (cell22 = X0) | ~ (cell12 = X1)) & (~ (cell21 = X0) | ~ (cell11 = X1)) & (~ (cell20 = X0) | ~ (cell10 = X1)) & (~ (cell19 = X0) | ~ (cell9 = X1)) & (~ (cell18 = X0) | ~ (cell8 = X1)) & (~ (cell17 = X0) | ~ (cell7 = X1)) & (~ (cell16 = X0) | ~ (cell6 = X1)) & (~ (cell15 = X0) | ~ (cell5 = X1)) & (~ (cell14 = X0) | ~ (cell4 = X1)) & (~ (cell13 = X0) | ~ (cell3 = X1)) & (~ (cell12 = X0) | ~ (cell2 = X1)) & (~ (cell11 = X0) | ~ (cell1 = X1))) | ~ sP0(X0, X1)), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ! [X0, X1] : (((~ (cell100 = X0) | ~ (cell90 = X1)) & (~ (cell99 = X0) | ~ (cell89 = X1)) & (~ (cell98 = X0) | ~ (cell88 = X1)) & (~ (cell97 = X0) | ~ (cell87 = X1)) & (~ (cell96 = X0) | ~ (cell86 = X1)) & (~ (cell95 = X0) | ~ (cell85 = X1)) & (~ (cell94 = X0) | ~ (cell84 = X1)) & (~ (cell93 = X0) | ~ (cell83 = X1)) & (~ (cell92 = X0) | ~ (cell82 = X1)) & (~ (cell91 = X0) | ~ (cell81 = X1)) & (~ (cell90 = X0) | ~ (cell80 = X1)) & (~ (cell89 = X0) | ~ (cell79 = X1)) & (~ (cell88 = X0) | ~ (cell78 = X1)) & (~ (cell87 = X0) | ~ (cell77 = X1)) & (~ (cell86 = X0) | ~ (cell76 = X1)) & (~ (cell85 = X0) | ~ (cell75 = X1)) & (~ (cell84 = X0) | ~ (cell74 = X1)) & (~ (cell83 = X0) | ~ (cell73 = X1)) & (~ (cell82 = X0) | ~ (cell72 = X1)) & (~ (cell81 = X0) | ~ (cell71 = X1)) & (~ (cell80 = X0) | ~ (cell70 = X1)) & (~ (cell79 = X0) | ~ (cell69 = X1)) & (~ (cell78 = X0) | ~ (cell68 = X1)) & (~ (cell77 = X0) | ~ (cell67 = X1)) & (~ (cell76 = X0) | ~ (cell66 = X1)) & (~ (cell75 = X0) | ~ (cell65 = X1)) & (~ (cell74 = X0) | ~ (cell64 = X1)) & (~ (cell73 = X0) | ~ (cell63 = X1)) & (~ (cell72 = X0) | ~ (cell62 = X1)) & (~ (cell71 = X0) | ~ (cell61 = X1)) & (~ (cell70 = X0) | ~ (cell60 = X1)) & (~ (cell69 = X0) | ~ (cell59 = X1)) & (~ (cell68 = X0) | ~ (cell58 = X1)) & (~ (cell67 = X0) | ~ (cell57 = X1)) & (~ (cell66 = X0) | ~ (cell56 = X1)) & (~ (cell65 = X0) | ~ (cell55 = X1)) & (~ (cell64 = X0) | ~ (cell54 = X1)) & (~ (cell63 = X0) | ~ (cell53 = X1)) & (~ (cell62 = X0) | ~ (cell52 = X1)) & (~ (cell61 = X0) | ~ (cell51 = X1)) & (~ (cell60 = X0) | ~ (cell50 = X1)) & (~ (cell59 = X0) | ~ (cell49 = X1)) & (~ (cell58 = X0) | ~ (cell48 = X1)) & (~ (cell57 = X0) | ~ (cell47 = X1)) & (~ (cell56 = X0) | ~ (cell46 = X1)) & (~ (cell55 = X0) | ~ (cell45 = X1)) & (~ (cell54 = X0) | ~ (cell44 = X1)) & (~ (cell53 = X0) | ~ (cell43 = X1)) & (~ (cell52 = X0) | ~ (cell42 = X1)) & (~ (cell51 = X0) | ~ (cell41 = X1)) & (~ (cell50 = X0) | ~ (cell40 = X1)) & (~ (cell49 = X0) | ~ (cell39 = X1)) & (~ (cell48 = X0) | ~ (cell38 = X1)) & (~ (cell47 = X0) | ~ (cell37 = X1)) & (~ (cell46 = X0) | ~ (cell36 = X1)) & (~ (cell45 = X0) | ~ (cell35 = X1)) & (~ (cell44 = X0) | ~ (cell34 = X1)) & (~ (cell43 = X0) | ~ (cell33 = X1)) & (~ (cell42 = X0) | ~ (cell32 = X1)) & (~ (cell41 = X0) | ~ (cell31 = X1)) & (~ (cell40 = X0) | ~ (cell30 = X1)) & (~ (cell39 = X0) | ~ (cell29 = X1)) & (~ (cell38 = X0) | ~ (cell28 = X1)) & (~ (cell37 = X0) | ~ (cell27 = X1)) & (~ (cell36 = X0) | ~ (cell26 = X1)) & (~ (cell35 = X0) | ~ (cell25 = X1)) & (~ (cell34 = X0) | ~ (cell24 = X1)) & (~ (cell33 = X0) | ~ (cell23 = X1)) & (~ (cell32 = X0) | ~ (cell22 = X1)) & (~ (cell31 = X0) | ~ (cell21 = X1)) & (~ (cell30 = X0) | ~ (cell20 = X1)) & (~ (cell29 = X0) | ~ (cell19 = X1)) & (~ (cell28 = X0) | ~ (cell18 = X1)) & (~ (cell27 = X0) | ~ (cell17 = X1)) & (~ (cell26 = X0) | ~ (cell16 = X1)) & (~ (cell25 = X0) | ~ (cell15 = X1)) & (~ (cell24 = X0) | ~ (cell14 = X1)) & (~ (cell23 = X0) | ~ (cell13 = X1)) & (~ (cell22 = X0) | ~ (cell12 = X1)) & (~ (cell21 = X0) | ~ (cell11 = X1)) & (~ (cell20 = X0) | ~ (cell10 = X1)) & (~ (cell19 = X0) | ~ (cell9 = X1)) & (~ (cell18 = X0) | ~ (cell8 = X1)) & (~ (cell17 = X0) | ~ (cell7 = X1)) & (~ (cell16 = X0) | ~ (cell6 = X1)) & (~ (cell15 = X0) | ~ (cell5 = X1)) & (~ (cell14 = X0) | ~ (cell4 = X1)) & (~ (cell13 = X0) | ~ (cell3 = X1)) & (~ (cell12 = X0) | ~ (cell2 = X1)) & (~ (cell11 = X0) | ~ (cell1 = X1))) | ~ sP0(X0, X1)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X0, X1] : (sP0(X0, X1) <=> ((~ (cell100 = X0) | ~ (cell90 = X1)) & (~ (cell99 = X0) | ~ (cell89 = X1)) & (~ (cell98 = X0) | ~ (cell88 = X1)) & (~ (cell97 = X0) | ~ (cell87 = X1)) & (~ (cell96 = X0) | ~ (cell86 = X1)) & (~ (cell95 = X0) | ~ (cell85 = X1)) & (~ (cell94 = X0) | ~ (cell84 = X1)) & (~ (cell93 = X0) | ~ (cell83 = X1)) & (~ (cell92 = X0) | ~ (cell82 = X1)) & (~ (cell91 = X0) | ~ (cell81 = X1)) & (~ (cell90 = X0) | ~ (cell80 = X1)) & (~ (cell89 = X0) | ~ (cell79 = X1)) & (~ (cell88 = X0) | ~ (cell78 = X1)) & (~ (cell87 = X0) | ~ (cell77 = X1)) & (~ (cell86 = X0) | ~ (cell76 = X1)) & (~ (cell85 = X0) | ~ (cell75 = X1)) & (~ (cell84 = X0) | ~ (cell74 = X1)) & (~ (cell83 = X0) | ~ (cell73 = X1)) & (~ (cell82 = X0) | ~ (cell72 = X1)) & (~ (cell81 = X0) | ~ (cell71 = X1)) & (~ (cell80 = X0) | ~ (cell70 = X1)) & (~ (cell79 = X0) | ~ (cell69 = X1)) & (~ (cell78 = X0) | ~ (cell68 = X1)) & (~ (cell77 = X0) | ~ (cell67 = X1)) & (~ (cell76 = X0) | ~ (cell66 = X1)) & (~ (cell75 = X0) | ~ (cell65 = X1)) & (~ (cell74 = X0) | ~ (cell64 = X1)) & (~ (cell73 = X0) | ~ (cell63 = X1)) & (~ (cell72 = X0) | ~ (cell62 = X1)) & (~ (cell71 = X0) | ~ (cell61 = X1)) & (~ (cell70 = X0) | ~ (cell60 = X1)) & (~ (cell69 = X0) | ~ (cell59 = X1)) & (~ (cell68 = X0) | ~ (cell58 = X1)) & (~ (cell67 = X0) | ~ (cell57 = X1)) & (~ (cell66 = X0) | ~ (cell56 = X1)) & (~ (cell65 = X0) | ~ (cell55 = X1)) & (~ (cell64 = X0) | ~ (cell54 = X1)) & (~ (cell63 = X0) | ~ (cell53 = X1)) & (~ (cell62 = X0) | ~ (cell52 = X1)) & (~ (cell61 = X0) | ~ (cell51 = X1)) & (~ (cell60 = X0) | ~ (cell50 = X1)) & (~ (cell59 = X0) | ~ (cell49 = X1)) & (~ (cell58 = X0) | ~ (cell48 = X1)) & (~ (cell57 = X0) | ~ (cell47 = X1)) & (~ (cell56 = X0) | ~ (cell46 = X1)) & (~ (cell55 = X0) | ~ (cell45 = X1)) & (~ (cell54 = X0) | ~ (cell44 = X1)) & (~ (cell53 = X0) | ~ (cell43 = X1)) & (~ (cell52 = X0) | ~ (cell42 = X1)) & (~ (cell51 = X0) | ~ (cell41 = X1)) & (~ (cell50 = X0) | ~ (cell40 = X1)) & (~ (cell49 = X0) | ~ (cell39 = X1)) & (~ (cell48 = X0) | ~ (cell38 = X1)) & (~ (cell47 = X0) | ~ (cell37 = X1)) & (~ (cell46 = X0) | ~ (cell36 = X1)) & (~ (cell45 = X0) | ~ (cell35 = X1)) & (~ (cell44 = X0) | ~ (cell34 = X1)) & (~ (cell43 = X0) | ~ (cell33 = X1)) & (~ (cell42 = X0) | ~ (cell32 = X1)) & (~ (cell41 = X0) | ~ (cell31 = X1)) & (~ (cell40 = X0) | ~ (cell30 = X1)) & (~ (cell39 = X0) | ~ (cell29 = X1)) & (~ (cell38 = X0) | ~ (cell28 = X1)) & (~ (cell37 = X0) | ~ (cell27 = X1)) & (~ (cell36 = X0) | ~ (cell26 = X1)) & (~ (cell35 = X0) | ~ (cell25 = X1)) & (~ (cell34 = X0) | ~ (cell24 = X1)) & (~ (cell33 = X0) | ~ (cell23 = X1)) & (~ (cell32 = X0) | ~ (cell22 = X1)) & (~ (cell31 = X0) | ~ (cell21 = X1)) & (~ (cell30 = X0) | ~ (cell20 = X1)) & (~ (cell29 = X0) | ~ (cell19 = X1)) & (~ (cell28 = X0) | ~ (cell18 = X1)) & (~ (cell27 = X0) | ~ (cell17 = X1)) & (~ (cell26 = X0) | ~ (cell16 = X1)) & (~ (cell25 = X0) | ~ (cell15 = X1)) & (~ (cell24 = X0) | ~ (cell14 = X1)) & (~ (cell23 = X0) | ~ (cell13 = X1)) & (~ (cell22 = X0) | ~ (cell12 = X1)) & (~ (cell21 = X0) | ~ (cell11 = X1)) & (~ (cell20 = X0) | ~ (cell10 = X1)) & (~ (cell19 = X0) | ~ (cell9 = X1)) & (~ (cell18 = X0) | ~ (cell8 = X1)) & (~ (cell17 = X0) | ~ (cell7 = X1)) & (~ (cell16 = X0) | ~ (cell6 = X1)) & (~ (cell15 = X0) | ~ (cell5 = X1)) & (~ (cell14 = X0) | ~ (cell4 = X1)) & (~ (cell13 = X0) | ~ (cell3 = X1)) & (~ (cell12 = X0) | ~ (cell2 = X1)) & (~ (cell11 = X0) | ~ (cell1 = X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f129, plain, ! [X0, X1] : (sP0(X0, X1) | west(X1, X0)), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ! [X0, X1] : (west(X1, X0) | sP0(X0, X1)), inference(definition_folding, [], [f21, e27])).
fof(f21, plain, ! [X0, X1] : (west(X1, X0) | ((~ (cell100 = X0) | ~ (cell90 = X1)) & (~ (cell99 = X0) | ~ (cell89 = X1)) & (~ (cell98 = X0) | ~ (cell88 = X1)) & (~ (cell97 = X0) | ~ (cell87 = X1)) & (~ (cell96 = X0) | ~ (cell86 = X1)) & (~ (cell95 = X0) | ~ (cell85 = X1)) & (~ (cell94 = X0) | ~ (cell84 = X1)) & (~ (cell93 = X0) | ~ (cell83 = X1)) & (~ (cell92 = X0) | ~ (cell82 = X1)) & (~ (cell91 = X0) | ~ (cell81 = X1)) & (~ (cell90 = X0) | ~ (cell80 = X1)) & (~ (cell89 = X0) | ~ (cell79 = X1)) & (~ (cell88 = X0) | ~ (cell78 = X1)) & (~ (cell87 = X0) | ~ (cell77 = X1)) & (~ (cell86 = X0) | ~ (cell76 = X1)) & (~ (cell85 = X0) | ~ (cell75 = X1)) & (~ (cell84 = X0) | ~ (cell74 = X1)) & (~ (cell83 = X0) | ~ (cell73 = X1)) & (~ (cell82 = X0) | ~ (cell72 = X1)) & (~ (cell81 = X0) | ~ (cell71 = X1)) & (~ (cell80 = X0) | ~ (cell70 = X1)) & (~ (cell79 = X0) | ~ (cell69 = X1)) & (~ (cell78 = X0) | ~ (cell68 = X1)) & (~ (cell77 = X0) | ~ (cell67 = X1)) & (~ (cell76 = X0) | ~ (cell66 = X1)) & (~ (cell75 = X0) | ~ (cell65 = X1)) & (~ (cell74 = X0) | ~ (cell64 = X1)) & (~ (cell73 = X0) | ~ (cell63 = X1)) & (~ (cell72 = X0) | ~ (cell62 = X1)) & (~ (cell71 = X0) | ~ (cell61 = X1)) & (~ (cell70 = X0) | ~ (cell60 = X1)) & (~ (cell69 = X0) | ~ (cell59 = X1)) & (~ (cell68 = X0) | ~ (cell58 = X1)) & (~ (cell67 = X0) | ~ (cell57 = X1)) & (~ (cell66 = X0) | ~ (cell56 = X1)) & (~ (cell65 = X0) | ~ (cell55 = X1)) & (~ (cell64 = X0) | ~ (cell54 = X1)) & (~ (cell63 = X0) | ~ (cell53 = X1)) & (~ (cell62 = X0) | ~ (cell52 = X1)) & (~ (cell61 = X0) | ~ (cell51 = X1)) & (~ (cell60 = X0) | ~ (cell50 = X1)) & (~ (cell59 = X0) | ~ (cell49 = X1)) & (~ (cell58 = X0) | ~ (cell48 = X1)) & (~ (cell57 = X0) | ~ (cell47 = X1)) & (~ (cell56 = X0) | ~ (cell46 = X1)) & (~ (cell55 = X0) | ~ (cell45 = X1)) & (~ (cell54 = X0) | ~ (cell44 = X1)) & (~ (cell53 = X0) | ~ (cell43 = X1)) & (~ (cell52 = X0) | ~ (cell42 = X1)) & (~ (cell51 = X0) | ~ (cell41 = X1)) & (~ (cell50 = X0) | ~ (cell40 = X1)) & (~ (cell49 = X0) | ~ (cell39 = X1)) & (~ (cell48 = X0) | ~ (cell38 = X1)) & (~ (cell47 = X0) | ~ (cell37 = X1)) & (~ (cell46 = X0) | ~ (cell36 = X1)) & (~ (cell45 = X0) | ~ (cell35 = X1)) & (~ (cell44 = X0) | ~ (cell34 = X1)) & (~ (cell43 = X0) | ~ (cell33 = X1)) & (~ (cell42 = X0) | ~ (cell32 = X1)) & (~ (cell41 = X0) | ~ (cell31 = X1)) & (~ (cell40 = X0) | ~ (cell30 = X1)) & (~ (cell39 = X0) | ~ (cell29 = X1)) & (~ (cell38 = X0) | ~ (cell28 = X1)) & (~ (cell37 = X0) | ~ (cell27 = X1)) & (~ (cell36 = X0) | ~ (cell26 = X1)) & (~ (cell35 = X0) | ~ (cell25 = X1)) & (~ (cell34 = X0) | ~ (cell24 = X1)) & (~ (cell33 = X0) | ~ (cell23 = X1)) & (~ (cell32 = X0) | ~ (cell22 = X1)) & (~ (cell31 = X0) | ~ (cell21 = X1)) & (~ (cell30 = X0) | ~ (cell20 = X1)) & (~ (cell29 = X0) | ~ (cell19 = X1)) & (~ (cell28 = X0) | ~ (cell18 = X1)) & (~ (cell27 = X0) | ~ (cell17 = X1)) & (~ (cell26 = X0) | ~ (cell16 = X1)) & (~ (cell25 = X0) | ~ (cell15 = X1)) & (~ (cell24 = X0) | ~ (cell14 = X1)) & (~ (cell23 = X0) | ~ (cell13 = X1)) & (~ (cell22 = X0) | ~ (cell12 = X1)) & (~ (cell21 = X0) | ~ (cell11 = X1)) & (~ (cell20 = X0) | ~ (cell10 = X1)) & (~ (cell19 = X0) | ~ (cell9 = X1)) & (~ (cell18 = X0) | ~ (cell8 = X1)) & (~ (cell17 = X0) | ~ (cell7 = X1)) & (~ (cell16 = X0) | ~ (cell6 = X1)) & (~ (cell15 = X0) | ~ (cell5 = X1)) & (~ (cell14 = X0) | ~ (cell4 = X1)) & (~ (cell13 = X0) | ~ (cell3 = X1)) & (~ (cell12 = X0) | ~ (cell2 = X1)) & (~ (cell11 = X0) | ~ (cell1 = X1)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : ((((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1))) => west(X1, X0)), inference(unused_predicate_definition_removal, [], [f1])).
fof(f1, plain, ! [X0, X1] : (west(X1, X0) <=> (((cell100 = X0) & (cell90 = X1)) | ((cell99 = X0) & (cell89 = X1)) | ((cell98 = X0) & (cell88 = X1)) | ((cell97 = X0) & (cell87 = X1)) | ((cell96 = X0) & (cell86 = X1)) | ((cell95 = X0) & (cell85 = X1)) | ((cell94 = X0) & (cell84 = X1)) | ((cell93 = X0) & (cell83 = X1)) | ((cell92 = X0) & (cell82 = X1)) | ((cell91 = X0) & (cell81 = X1)) | ((cell90 = X0) & (cell80 = X1)) | ((cell89 = X0) & (cell79 = X1)) | ((cell88 = X0) & (cell78 = X1)) | ((cell87 = X0) & (cell77 = X1)) | ((cell86 = X0) & (cell76 = X1)) | ((cell85 = X0) & (cell75 = X1)) | ((cell84 = X0) & (cell74 = X1)) | ((cell83 = X0) & (cell73 = X1)) | ((cell82 = X0) & (cell72 = X1)) | ((cell81 = X0) & (cell71 = X1)) | ((cell80 = X0) & (cell70 = X1)) | ((cell79 = X0) & (cell69 = X1)) | ((cell78 = X0) & (cell68 = X1)) | ((cell77 = X0) & (cell67 = X1)) | ((cell76 = X0) & (cell66 = X1)) | ((cell75 = X0) & (cell65 = X1)) | ((cell74 = X0) & (cell64 = X1)) | ((cell73 = X0) & (cell63 = X1)) | ((cell72 = X0) & (cell62 = X1)) | ((cell71 = X0) & (cell61 = X1)) | ((cell70 = X0) & (cell60 = X1)) | ((cell69 = X0) & (cell59 = X1)) | ((cell68 = X0) & (cell58 = X1)) | ((cell67 = X0) & (cell57 = X1)) | ((cell66 = X0) & (cell56 = X1)) | ((cell65 = X0) & (cell55 = X1)) | ((cell64 = X0) & (cell54 = X1)) | ((cell63 = X0) & (cell53 = X1)) | ((cell62 = X0) & (cell52 = X1)) | ((cell61 = X0) & (cell51 = X1)) | ((cell60 = X0) & (cell50 = X1)) | ((cell59 = X0) & (cell49 = X1)) | ((cell58 = X0) & (cell48 = X1)) | ((cell57 = X0) & (cell47 = X1)) | ((cell56 = X0) & (cell46 = X1)) | ((cell55 = X0) & (cell45 = X1)) | ((cell54 = X0) & (cell44 = X1)) | ((cell53 = X0) & (cell43 = X1)) | ((cell52 = X0) & (cell42 = X1)) | ((cell51 = X0) & (cell41 = X1)) | ((cell50 = X0) & (cell40 = X1)) | ((cell49 = X0) & (cell39 = X1)) | ((cell48 = X0) & (cell38 = X1)) | ((cell47 = X0) & (cell37 = X1)) | ((cell46 = X0) & (cell36 = X1)) | ((cell45 = X0) & (cell35 = X1)) | ((cell44 = X0) & (cell34 = X1)) | ((cell43 = X0) & (cell33 = X1)) | ((cell42 = X0) & (cell32 = X1)) | ((cell41 = X0) & (cell31 = X1)) | ((cell40 = X0) & (cell30 = X1)) | ((cell39 = X0) & (cell29 = X1)) | ((cell38 = X0) & (cell28 = X1)) | ((cell37 = X0) & (cell27 = X1)) | ((cell36 = X0) & (cell26 = X1)) | ((cell35 = X0) & (cell25 = X1)) | ((cell34 = X0) & (cell24 = X1)) | ((cell33 = X0) & (cell23 = X1)) | ((cell32 = X0) & (cell22 = X1)) | ((cell31 = X0) & (cell21 = X1)) | ((cell30 = X0) & (cell20 = X1)) | ((cell29 = X0) & (cell19 = X1)) | ((cell28 = X0) & (cell18 = X1)) | ((cell27 = X0) & (cell17 = X1)) | ((cell26 = X0) & (cell16 = X1)) | ((cell25 = X0) & (cell15 = X1)) | ((cell24 = X0) & (cell14 = X1)) | ((cell23 = X0) & (cell13 = X1)) | ((cell22 = X0) & (cell12 = X1)) | ((cell21 = X0) & (cell11 = X1)) | ((cell20 = X0) & (cell10 = X1)) | ((cell19 = X0) & (cell9 = X1)) | ((cell18 = X0) & (cell8 = X1)) | ((cell17 = X0) & (cell7 = X1)) | ((cell16 = X0) & (cell6 = X1)) | ((cell15 = X0) & (cell5 = X1)) | ((cell14 = X0) & (cell4 = X1)) | ((cell13 = X0) & (cell3 = X1)) | ((cell12 = X0) & (cell2 = X1)) | ((cell11 = X0) & (cell1 = X1)))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34975)).
fof(f11442, plain, ! [X0] : (~ west(cell90, X0) | duewest(cell10, X0)), inference(resolution, [], [f10652, f685])).
fof(f685, plain, west(cell80, cell90), inference(resolution, [], [f129, f267])).
fof(f267, plain, ~ sP0(cell90, cell80), inference(equality_resolution, [], [f266])).
fof(f266, plain, ! [X1] : (~ (cell80 = X1) | ~ sP0(cell90, X1)), inference(equality_resolution, [], [f118])).
fof(f118, plain, ! [X0, X1] : (~ (cell90 = X0) | ~ (cell80 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f10652, plain, ! [X0, X1] : (~ west(cell80, X0) | ~ west(X0, X1) | duewest(cell10, X1)), inference(resolution, [], [f9904, f675])).
fof(f675, plain, west(cell70, cell80), inference(resolution, [], [f129, f287])).
fof(f287, plain, ~ sP0(cell80, cell70), inference(equality_resolution, [], [f286])).
fof(f286, plain, ! [X1] : (~ (cell70 = X1) | ~ sP0(cell80, X1)), inference(equality_resolution, [], [f108])).
fof(f108, plain, ! [X0, X1] : (~ (cell80 = X0) | ~ (cell70 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f9904, plain, ! [X2, X0, X1] : (~ west(cell70, X0) | ~ west(X0, X1) | ~ west(X1, X2) | duewest(cell10, X2)), inference(resolution, [], [f8897, f665])).
fof(f665, plain, west(cell60, cell70), inference(resolution, [], [f129, f307])).
fof(f307, plain, ~ sP0(cell70, cell60), inference(equality_resolution, [], [f306])).
fof(f306, plain, ! [X1] : (~ (cell60 = X1) | ~ sP0(cell70, X1)), inference(equality_resolution, [], [f98])).
fof(f98, plain, ! [X0, X1] : (~ (cell70 = X0) | ~ (cell60 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f8897, plain, ! [X2, X0, X3, X1] : (~ west(cell60, X0) | ~ west(X0, X1) | ~ west(X1, X2) | ~ west(X2, X3) | duewest(cell10, X3)), inference(resolution, [], [f7822, f655])).
fof(f655, plain, west(cell50, cell60), inference(resolution, [], [f129, f327])).
fof(f327, plain, ~ sP0(cell60, cell50), inference(equality_resolution, [], [f326])).
fof(f326, plain, ! [X1] : (~ (cell50 = X1) | ~ sP0(cell60, X1)), inference(equality_resolution, [], [f88])).
fof(f88, plain, ! [X0, X1] : (~ (cell60 = X0) | ~ (cell50 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f7822, plain, ! [X4, X2, X0, X3, X1] : (~ west(cell50, X0) | ~ west(X0, X1) | ~ west(X1, X2) | ~ west(X2, X3) | ~ west(X3, X4) | duewest(cell10, X4)), inference(resolution, [], [f6625, f645])).
fof(f645, plain, west(cell40, cell50), inference(resolution, [], [f129, f347])).
fof(f347, plain, ~ sP0(cell50, cell40), inference(equality_resolution, [], [f346])).
fof(f346, plain, ! [X1] : (~ (cell40 = X1) | ~ sP0(cell50, X1)), inference(equality_resolution, [], [f78])).
fof(f78, plain, ! [X0, X1] : (~ (cell50 = X0) | ~ (cell40 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f6625, plain, ! [X4, X2, X0, X5, X3, X1] : (~ west(cell40, X0) | ~ west(X0, X1) | ~ west(X1, X2) | ~ west(X2, X3) | ~ west(X3, X4) | ~ west(X4, X5) | duewest(cell10, X5)), inference(resolution, [], [f4977, f635])).
fof(f635, plain, west(cell30, cell40), inference(resolution, [], [f129, f367])).
fof(f367, plain, ~ sP0(cell40, cell30), inference(equality_resolution, [], [f366])).
fof(f366, plain, ! [X1] : (~ (cell30 = X1) | ~ sP0(cell40, X1)), inference(equality_resolution, [], [f68])).
fof(f68, plain, ! [X0, X1] : (~ (cell40 = X0) | ~ (cell30 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f4977, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ west(cell30, X3) | ~ west(X2, X0) | ~ west(X3, X2) | ~ west(X0, X1) | ~ west(X1, X4) | ~ west(X4, X5) | ~ west(X5, X6) | duewest(cell10, X6)), inference(resolution, [], [f2278, f625])).
fof(f625, plain, west(cell20, cell30), inference(resolution, [], [f129, f387])).
fof(f387, plain, ~ sP0(cell30, cell20), inference(equality_resolution, [], [f386])).
fof(f386, plain, ! [X1] : (~ (cell20 = X1) | ~ sP0(cell30, X1)), inference(equality_resolution, [], [f58])).
fof(f58, plain, ! [X0, X1] : (~ (cell30 = X0) | ~ (cell20 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f2278, plain, ! [X146, X144, X151, X149, X147, X145, X150, X148] : (~ west(cell20, X149) | ~ west(X146, X144) | ~ west(X147, X146) | ~ west(X148, X147) | ~ west(X149, X148) | ~ west(X144, X145) | ~ west(X145, X150) | ~ west(X150, X151) | duewest(cell10, X151)), inference(resolution, [], [f802, f615])).
fof(f615, plain, west(cell10, cell20), inference(resolution, [], [f129, f407])).
fof(f407, plain, ~ sP0(cell20, cell10), inference(equality_resolution, [], [f406])).
fof(f406, plain, ! [X1] : (~ (cell10 = X1) | ~ sP0(cell20, X1)), inference(equality_resolution, [], [f48])).
fof(f48, plain, ! [X0, X1] : (~ (cell20 = X0) | ~ (cell10 = X1) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f35])).
fof(f802, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ west(X8, X7) | ~ west(X2, X0) | ~ west(X3, X2) | ~ west(X4, X3) | ~ west(X5, X4) | ~ west(X6, X5) | ~ west(X7, X6) | ~ west(X0, X1) | ~ west(X1, X9) | duewest(X8, X9)), inference(resolution, [], [f229, f230])).
fof(f230, plain, ! [X0, X1] : (sP2(X0, X1) | duewest(X1, X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0, X1] : (duewest(X1, X0) | sP2(X0, X1)), inference(definition_folding, [], [f23, e31])).
fof(f31, plain, ! [X0, X1] : ((! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ west(X9, X0) | ~ west(X8, X9) | ~ west(X7, X8) | ~ west(X6, X7) | ~ west(X5, X6) | ~ west(X4, X5) | ~ west(X3, X4) | ~ west(X2, X3) | ~ west(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ west(X16, X0) | ~ west(X15, X16) | ~ west(X14, X15) | ~ west(X13, X14) | ~ west(X12, X13) | ~ west(X11, X12) | ~ west(X10, X11) | ~ west(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ west(X22, X0) | ~ west(X21, X22) | ~ west(X20, X21) | ~ west(X19, X20) | ~ west(X18, X19) | ~ west(X17, X18) | ~ west(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ west(X27, X0) | ~ west(X26, X27) | ~ west(X25, X26) | ~ west(X24, X25) | ~ west(X23, X24) | ~ west(X1, X23)) & ! [X28, X29, X30, X31] : (~ west(X31, X0) | ~ west(X30, X31) | ~ west(X29, X30) | ~ west(X28, X29) | ~ west(X1, X28)) & ! [X32, X33, X34] : (~ west(X34, X0) | ~ west(X33, X34) | ~ west(X32, X33) | ~ west(X1, X32)) & ! [X35, X36] : (~ west(X36, X0) | ~ west(X35, X36) | ~ west(X1, X35)) & ! [X37] : (~ west(X37, X0) | ~ west(X1, X37)) & ~ west(X1, X0)) | ~ sP2(X0, X1)), inference(usedef, [], [e31])).
fof(e31, plain, ! [X0, X1] : (sP2(X0, X1) <=> (! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ west(X9, X0) | ~ west(X8, X9) | ~ west(X7, X8) | ~ west(X6, X7) | ~ west(X5, X6) | ~ west(X4, X5) | ~ west(X3, X4) | ~ west(X2, X3) | ~ west(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ west(X16, X0) | ~ west(X15, X16) | ~ west(X14, X15) | ~ west(X13, X14) | ~ west(X12, X13) | ~ west(X11, X12) | ~ west(X10, X11) | ~ west(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ west(X22, X0) | ~ west(X21, X22) | ~ west(X20, X21) | ~ west(X19, X20) | ~ west(X18, X19) | ~ west(X17, X18) | ~ west(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ west(X27, X0) | ~ west(X26, X27) | ~ west(X25, X26) | ~ west(X24, X25) | ~ west(X23, X24) | ~ west(X1, X23)) & ! [X28, X29, X30, X31] : (~ west(X31, X0) | ~ west(X30, X31) | ~ west(X29, X30) | ~ west(X28, X29) | ~ west(X1, X28)) & ! [X32, X33, X34] : (~ west(X34, X0) | ~ west(X33, X34) | ~ west(X32, X33) | ~ west(X1, X32)) & ! [X35, X36] : (~ west(X36, X0) | ~ west(X35, X36) | ~ west(X1, X35)) & ! [X37] : (~ west(X37, X0) | ~ west(X1, X37)) & ~ west(X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f23, plain, ! [X0, X1] : (duewest(X1, X0) | (! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ west(X9, X0) | ~ west(X8, X9) | ~ west(X7, X8) | ~ west(X6, X7) | ~ west(X5, X6) | ~ west(X4, X5) | ~ west(X3, X4) | ~ west(X2, X3) | ~ west(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ west(X16, X0) | ~ west(X15, X16) | ~ west(X14, X15) | ~ west(X13, X14) | ~ west(X12, X13) | ~ west(X11, X12) | ~ west(X10, X11) | ~ west(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ west(X22, X0) | ~ west(X21, X22) | ~ west(X20, X21) | ~ west(X19, X20) | ~ west(X18, X19) | ~ west(X17, X18) | ~ west(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ west(X27, X0) | ~ west(X26, X27) | ~ west(X25, X26) | ~ west(X24, X25) | ~ west(X23, X24) | ~ west(X1, X23)) & ! [X28, X29, X30, X31] : (~ west(X31, X0) | ~ west(X30, X31) | ~ west(X29, X30) | ~ west(X28, X29) | ~ west(X1, X28)) & ! [X32, X33, X34] : (~ west(X34, X0) | ~ west(X33, X34) | ~ west(X32, X33) | ~ west(X1, X32)) & ! [X35, X36] : (~ west(X36, X0) | ~ west(X35, X36) | ~ west(X1, X35)) & ! [X37] : (~ west(X37, X0) | ~ west(X1, X37)) & ~ west(X1, X0))), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0, X1] : ((? [X2, X3, X4, X5, X6, X7, X8, X9] : (west(X9, X0) & west(X8, X9) & west(X7, X8) & west(X6, X7) & west(X5, X6) & west(X4, X5) & west(X3, X4) & west(X2, X3) & west(X1, X2)) | ? [X10, X11, X12, X13, X14, X15, X16] : (west(X16, X0) & west(X15, X16) & west(X14, X15) & west(X13, X14) & west(X12, X13) & west(X11, X12) & west(X10, X11) & west(X1, X10)) | ? [X17, X18, X19, X20, X21, X22] : (west(X22, X0) & west(X21, X22) & west(X20, X21) & west(X19, X20) & west(X18, X19) & west(X17, X18) & west(X1, X17)) | ? [X23, X24, X25, X26, X27] : (west(X27, X0) & west(X26, X27) & west(X25, X26) & west(X24, X25) & west(X23, X24) & west(X1, X23)) | ? [X28, X29, X30, X31] : (west(X31, X0) & west(X30, X31) & west(X29, X30) & west(X28, X29) & west(X1, X28)) | ? [X32, X33, X34] : (west(X34, X0) & west(X33, X34) & west(X32, X33) & west(X1, X32)) | ? [X35, X36] : (west(X36, X0) & west(X35, X36) & west(X1, X35)) | ? [X37] : (west(X37, X0) & west(X1, X37)) | west(X1, X0)) => duewest(X1, X0)), inference(unused_predicate_definition_removal, [], [f10])).
fof(f10, plain, ! [X0, X1] : (duewest(X1, X0) <=> (? [X2, X3, X4, X5, X6, X7, X8, X9] : (west(X9, X0) & west(X8, X9) & west(X7, X8) & west(X6, X7) & west(X5, X6) & west(X4, X5) & west(X3, X4) & west(X2, X3) & west(X1, X2)) | ? [X10, X11, X12, X13, X14, X15, X16] : (west(X16, X0) & west(X15, X16) & west(X14, X15) & west(X13, X14) & west(X12, X13) & west(X11, X12) & west(X10, X11) & west(X1, X10)) | ? [X17, X18, X19, X20, X21, X22] : (west(X22, X0) & west(X21, X22) & west(X20, X21) & west(X19, X20) & west(X18, X19) & west(X17, X18) & west(X1, X17)) | ? [X23, X24, X25, X26, X27] : (west(X27, X0) & west(X26, X27) & west(X25, X26) & west(X24, X25) & west(X23, X24) & west(X1, X23)) | ? [X28, X29, X30, X31] : (west(X31, X0) & west(X30, X31) & west(X29, X30) & west(X28, X29) & west(X1, X28)) | ? [X32, X33, X34] : (west(X34, X0) & west(X33, X34) & west(X32, X33) & west(X1, X32)) | ? [X35, X36] : (west(X36, X0) & west(X35, X36) & west(X1, X35)) | ? [X37] : (west(X37, X0) & west(X1, X37)) | west(X1, X0))), inference(rectify, [], [f3])).
fof(f3, plain, ! [X4, X5] : (duewest(X5, X4) <=> (? [X34, X35, X36, X37, X38, X39, X40, X41] : (west(X41, X4) & west(X40, X41) & west(X39, X40) & west(X38, X39) & west(X37, X38) & west(X36, X37) & west(X35, X36) & west(X34, X35) & west(X5, X34)) | ? [X27, X28, X29, X30, X31, X32, X33] : (west(X33, X4) & west(X32, X33) & west(X31, X32) & west(X30, X31) & west(X29, X30) & west(X28, X29) & west(X27, X28) & west(X5, X27)) | ? [X21, X22, X23, X24, X25, X26] : (west(X26, X4) & west(X25, X26) & west(X24, X25) & west(X23, X24) & west(X22, X23) & west(X21, X22) & west(X5, X21)) | ? [X16, X17, X18, X19, X20] : (west(X20, X4) & west(X19, X20) & west(X18, X19) & west(X17, X18) & west(X16, X17) & west(X5, X16)) | ? [X12, X13, X14, X15] : (west(X15, X4) & west(X14, X15) & west(X13, X14) & west(X12, X13) & west(X5, X12)) | ? [X9, X10, X11] : (west(X11, X4) & west(X10, X11) & west(X9, X10) & west(X5, X9)) | ? [X7, X8] : (west(X8, X4) & west(X7, X8) & west(X5, X7)) | ? [X6] : (west(X6, X4) & west(X5, X6)) | west(X5, X4))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34977)).
fof(f229, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ sP2(X0, X1) | ~ west(X8, X9) | ~ west(X7, X8) | ~ west(X6, X7) | ~ west(X5, X6) | ~ west(X4, X5) | ~ west(X3, X4) | ~ west(X2, X3) | ~ west(X1, X2) | ~ west(X9, X0)), inference(cnf_transformation, [], [f37])).
fof(f37, plain, ! [X0, X1] : ((! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ west(X9, X0) | ~ west(X8, X9) | ~ west(X7, X8) | ~ west(X6, X7) | ~ west(X5, X6) | ~ west(X4, X5) | ~ west(X3, X4) | ~ west(X2, X3) | ~ west(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ west(X16, X0) | ~ west(X15, X16) | ~ west(X14, X15) | ~ west(X13, X14) | ~ west(X12, X13) | ~ west(X11, X12) | ~ west(X10, X11) | ~ west(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ west(X22, X0) | ~ west(X21, X22) | ~ west(X20, X21) | ~ west(X19, X20) | ~ west(X18, X19) | ~ west(X17, X18) | ~ west(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ west(X27, X0) | ~ west(X26, X27) | ~ west(X25, X26) | ~ west(X24, X25) | ~ west(X23, X24) | ~ west(X1, X23)) & ! [X28, X29, X30, X31] : (~ west(X31, X0) | ~ west(X30, X31) | ~ west(X29, X30) | ~ west(X28, X29) | ~ west(X1, X28)) & ! [X32, X33, X34] : (~ west(X34, X0) | ~ west(X33, X34) | ~ west(X32, X33) | ~ west(X1, X32)) & ! [X35, X36] : (~ west(X36, X0) | ~ west(X35, X36) | ~ west(X1, X35)) & ! [X37] : (~ west(X37, X0) | ~ west(X1, X37)) & ~ west(X1, X0)) | ~ sP2(X0, X1)), inference(nnf_transformation, [], [f31])).
fof(f12515, plain, ! [X0] : (~ duewest(cell10, X0) | westof(cell1, X0)), inference(resolution, [], [f12225, f244])).
fof(f244, plain, ! [X2, X0, X1] : (~ vert(X1, X2) | ~ duewest(X2, X0) | westof(X1, X0)), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ! [X0, X1] : (westof(X1, X0) | (! [X2] : (~ duewest(X2, X0) | ~ vert(X1, X2)) & ~ duewest(X1, X0))), inference(ennf_transformation, [], [f15])).
fof(f15, plain, ! [X0, X1] : ((? [X2] : (duewest(X2, X0) & vert(X1, X2)) | duewest(X1, X0)) => westof(X1, X0)), inference(unused_predicate_definition_removal, [], [f13])).
fof(f13, plain, ! [X0, X1] : (westof(X1, X0) <=> (? [X2] : (duewest(X2, X0) & vert(X1, X2)) | duewest(X1, X0))), inference(rectify, [], [f6])).
fof(f6, plain, ! [X80, X81] : (westof(X81, X80) <=> (? [X82] : (duewest(X82, X80) & vert(X81, X82)) | duewest(X81, X80))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34980)).
fof(f12225, plain, vert(cell1, cell10), inference(resolution, [], [f12086, f241])).
fof(f241, plain, ! [X0, X1] : (~ duenorth(X1, X0) | vert(X1, X0)), inference(cnf_transformation, [], [f25])).
fof(f25, plain, ! [X0, X1] : (vert(X1, X0) | (~ duenorth(X0, X1) & ~ duenorth(X1, X0))), inference(ennf_transformation, [], [f18])).
fof(f18, plain, ! [X0, X1] : ((duenorth(X0, X1) | duenorth(X1, X0)) => vert(X1, X0)), inference(unused_predicate_definition_removal, [], [f12])).
fof(f12, plain, ! [X0, X1] : (vert(X1, X0) <=> (duenorth(X0, X1) | duenorth(X1, X0))), inference(rectify, [], [f5])).
fof(f5, plain, ! [X80, X81] : (vert(X81, X80) <=> (duenorth(X80, X81) | duenorth(X81, X80))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34979)).
fof(f12086, plain, duenorth(cell1, cell10), inference(resolution, [], [f11444, f712])).
fof(f712, plain, north(cell9, cell10), inference(resolution, [], [f220, f589])).
fof(f589, plain, ~ sP1(cell10, cell9), inference(equality_resolution, [], [f588])).
fof(f588, plain, ! [X1] : (~ (cell9 = X1) | ~ sP1(cell10, X1)), inference(equality_resolution, [], [f138])).
fof(f138, plain, ! [X0, X1] : (~ (cell10 = X0) | ~ (cell9 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ! [X0, X1] : (((~ (cell100 = X0) | ~ (cell99 = X1)) & (~ (cell99 = X0) | ~ (cell98 = X1)) & (~ (cell98 = X0) | ~ (cell97 = X1)) & (~ (cell97 = X0) | ~ (cell96 = X1)) & (~ (cell96 = X0) | ~ (cell95 = X1)) & (~ (cell95 = X0) | ~ (cell94 = X1)) & (~ (cell94 = X0) | ~ (cell93 = X1)) & (~ (cell93 = X0) | ~ (cell92 = X1)) & (~ (cell92 = X0) | ~ (cell91 = X1)) & (~ (cell90 = X0) | ~ (cell89 = X1)) & (~ (cell89 = X0) | ~ (cell88 = X1)) & (~ (cell88 = X0) | ~ (cell87 = X1)) & (~ (cell87 = X0) | ~ (cell86 = X1)) & (~ (cell86 = X0) | ~ (cell85 = X1)) & (~ (cell85 = X0) | ~ (cell84 = X1)) & (~ (cell84 = X0) | ~ (cell83 = X1)) & (~ (cell83 = X0) | ~ (cell82 = X1)) & (~ (cell82 = X0) | ~ (cell81 = X1)) & (~ (cell80 = X0) | ~ (cell79 = X1)) & (~ (cell79 = X0) | ~ (cell78 = X1)) & (~ (cell78 = X0) | ~ (cell77 = X1)) & (~ (cell77 = X0) | ~ (cell76 = X1)) & (~ (cell76 = X0) | ~ (cell75 = X1)) & (~ (cell75 = X0) | ~ (cell74 = X1)) & (~ (cell74 = X0) | ~ (cell73 = X1)) & (~ (cell73 = X0) | ~ (cell72 = X1)) & (~ (cell72 = X0) | ~ (cell71 = X1)) & (~ (cell70 = X0) | ~ (cell69 = X1)) & (~ (cell69 = X0) | ~ (cell68 = X1)) & (~ (cell68 = X0) | ~ (cell67 = X1)) & (~ (cell67 = X0) | ~ (cell66 = X1)) & (~ (cell66 = X0) | ~ (cell65 = X1)) & (~ (cell65 = X0) | ~ (cell64 = X1)) & (~ (cell64 = X0) | ~ (cell63 = X1)) & (~ (cell63 = X0) | ~ (cell62 = X1)) & (~ (cell62 = X0) | ~ (cell61 = X1)) & (~ (cell60 = X0) | ~ (cell59 = X1)) & (~ (cell59 = X0) | ~ (cell58 = X1)) & (~ (cell58 = X0) | ~ (cell57 = X1)) & (~ (cell57 = X0) | ~ (cell56 = X1)) & (~ (cell56 = X0) | ~ (cell55 = X1)) & (~ (cell55 = X0) | ~ (cell54 = X1)) & (~ (cell54 = X0) | ~ (cell53 = X1)) & (~ (cell53 = X0) | ~ (cell52 = X1)) & (~ (cell52 = X0) | ~ (cell51 = X1)) & (~ (cell50 = X0) | ~ (cell49 = X1)) & (~ (cell49 = X0) | ~ (cell48 = X1)) & (~ (cell48 = X0) | ~ (cell47 = X1)) & (~ (cell47 = X0) | ~ (cell46 = X1)) & (~ (cell46 = X0) | ~ (cell45 = X1)) & (~ (cell45 = X0) | ~ (cell44 = X1)) & (~ (cell44 = X0) | ~ (cell43 = X1)) & (~ (cell43 = X0) | ~ (cell42 = X1)) & (~ (cell42 = X0) | ~ (cell41 = X1)) & (~ (cell40 = X0) | ~ (cell39 = X1)) & (~ (cell39 = X0) | ~ (cell38 = X1)) & (~ (cell38 = X0) | ~ (cell37 = X1)) & (~ (cell37 = X0) | ~ (cell36 = X1)) & (~ (cell36 = X0) | ~ (cell35 = X1)) & (~ (cell35 = X0) | ~ (cell34 = X1)) & (~ (cell34 = X0) | ~ (cell33 = X1)) & (~ (cell33 = X0) | ~ (cell32 = X1)) & (~ (cell32 = X0) | ~ (cell31 = X1)) & (~ (cell30 = X0) | ~ (cell29 = X1)) & (~ (cell29 = X0) | ~ (cell28 = X1)) & (~ (cell28 = X0) | ~ (cell27 = X1)) & (~ (cell27 = X0) | ~ (cell26 = X1)) & (~ (cell26 = X0) | ~ (cell25 = X1)) & (~ (cell25 = X0) | ~ (cell24 = X1)) & (~ (cell24 = X0) | ~ (cell23 = X1)) & (~ (cell23 = X0) | ~ (cell22 = X1)) & (~ (cell22 = X0) | ~ (cell21 = X1)) & (~ (cell20 = X0) | ~ (cell19 = X1)) & (~ (cell19 = X0) | ~ (cell18 = X1)) & (~ (cell18 = X0) | ~ (cell17 = X1)) & (~ (cell17 = X0) | ~ (cell16 = X1)) & (~ (cell16 = X0) | ~ (cell15 = X1)) & (~ (cell15 = X0) | ~ (cell14 = X1)) & (~ (cell14 = X0) | ~ (cell13 = X1)) & (~ (cell13 = X0) | ~ (cell12 = X1)) & (~ (cell12 = X0) | ~ (cell11 = X1)) & (~ (cell10 = X0) | ~ (cell9 = X1)) & (~ (cell9 = X0) | ~ (cell8 = X1)) & (~ (cell8 = X0) | ~ (cell7 = X1)) & (~ (cell7 = X0) | ~ (cell6 = X1)) & (~ (cell6 = X0) | ~ (cell5 = X1)) & (~ (cell5 = X0) | ~ (cell4 = X1)) & (~ (cell4 = X0) | ~ (cell3 = X1)) & (~ (cell3 = X0) | ~ (cell2 = X1)) & (~ (cell2 = X0) | ~ (cell1 = X1))) | ~ sP1(X0, X1)), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1] : (((~ (cell100 = X0) | ~ (cell99 = X1)) & (~ (cell99 = X0) | ~ (cell98 = X1)) & (~ (cell98 = X0) | ~ (cell97 = X1)) & (~ (cell97 = X0) | ~ (cell96 = X1)) & (~ (cell96 = X0) | ~ (cell95 = X1)) & (~ (cell95 = X0) | ~ (cell94 = X1)) & (~ (cell94 = X0) | ~ (cell93 = X1)) & (~ (cell93 = X0) | ~ (cell92 = X1)) & (~ (cell92 = X0) | ~ (cell91 = X1)) & (~ (cell90 = X0) | ~ (cell89 = X1)) & (~ (cell89 = X0) | ~ (cell88 = X1)) & (~ (cell88 = X0) | ~ (cell87 = X1)) & (~ (cell87 = X0) | ~ (cell86 = X1)) & (~ (cell86 = X0) | ~ (cell85 = X1)) & (~ (cell85 = X0) | ~ (cell84 = X1)) & (~ (cell84 = X0) | ~ (cell83 = X1)) & (~ (cell83 = X0) | ~ (cell82 = X1)) & (~ (cell82 = X0) | ~ (cell81 = X1)) & (~ (cell80 = X0) | ~ (cell79 = X1)) & (~ (cell79 = X0) | ~ (cell78 = X1)) & (~ (cell78 = X0) | ~ (cell77 = X1)) & (~ (cell77 = X0) | ~ (cell76 = X1)) & (~ (cell76 = X0) | ~ (cell75 = X1)) & (~ (cell75 = X0) | ~ (cell74 = X1)) & (~ (cell74 = X0) | ~ (cell73 = X1)) & (~ (cell73 = X0) | ~ (cell72 = X1)) & (~ (cell72 = X0) | ~ (cell71 = X1)) & (~ (cell70 = X0) | ~ (cell69 = X1)) & (~ (cell69 = X0) | ~ (cell68 = X1)) & (~ (cell68 = X0) | ~ (cell67 = X1)) & (~ (cell67 = X0) | ~ (cell66 = X1)) & (~ (cell66 = X0) | ~ (cell65 = X1)) & (~ (cell65 = X0) | ~ (cell64 = X1)) & (~ (cell64 = X0) | ~ (cell63 = X1)) & (~ (cell63 = X0) | ~ (cell62 = X1)) & (~ (cell62 = X0) | ~ (cell61 = X1)) & (~ (cell60 = X0) | ~ (cell59 = X1)) & (~ (cell59 = X0) | ~ (cell58 = X1)) & (~ (cell58 = X0) | ~ (cell57 = X1)) & (~ (cell57 = X0) | ~ (cell56 = X1)) & (~ (cell56 = X0) | ~ (cell55 = X1)) & (~ (cell55 = X0) | ~ (cell54 = X1)) & (~ (cell54 = X0) | ~ (cell53 = X1)) & (~ (cell53 = X0) | ~ (cell52 = X1)) & (~ (cell52 = X0) | ~ (cell51 = X1)) & (~ (cell50 = X0) | ~ (cell49 = X1)) & (~ (cell49 = X0) | ~ (cell48 = X1)) & (~ (cell48 = X0) | ~ (cell47 = X1)) & (~ (cell47 = X0) | ~ (cell46 = X1)) & (~ (cell46 = X0) | ~ (cell45 = X1)) & (~ (cell45 = X0) | ~ (cell44 = X1)) & (~ (cell44 = X0) | ~ (cell43 = X1)) & (~ (cell43 = X0) | ~ (cell42 = X1)) & (~ (cell42 = X0) | ~ (cell41 = X1)) & (~ (cell40 = X0) | ~ (cell39 = X1)) & (~ (cell39 = X0) | ~ (cell38 = X1)) & (~ (cell38 = X0) | ~ (cell37 = X1)) & (~ (cell37 = X0) | ~ (cell36 = X1)) & (~ (cell36 = X0) | ~ (cell35 = X1)) & (~ (cell35 = X0) | ~ (cell34 = X1)) & (~ (cell34 = X0) | ~ (cell33 = X1)) & (~ (cell33 = X0) | ~ (cell32 = X1)) & (~ (cell32 = X0) | ~ (cell31 = X1)) & (~ (cell30 = X0) | ~ (cell29 = X1)) & (~ (cell29 = X0) | ~ (cell28 = X1)) & (~ (cell28 = X0) | ~ (cell27 = X1)) & (~ (cell27 = X0) | ~ (cell26 = X1)) & (~ (cell26 = X0) | ~ (cell25 = X1)) & (~ (cell25 = X0) | ~ (cell24 = X1)) & (~ (cell24 = X0) | ~ (cell23 = X1)) & (~ (cell23 = X0) | ~ (cell22 = X1)) & (~ (cell22 = X0) | ~ (cell21 = X1)) & (~ (cell20 = X0) | ~ (cell19 = X1)) & (~ (cell19 = X0) | ~ (cell18 = X1)) & (~ (cell18 = X0) | ~ (cell17 = X1)) & (~ (cell17 = X0) | ~ (cell16 = X1)) & (~ (cell16 = X0) | ~ (cell15 = X1)) & (~ (cell15 = X0) | ~ (cell14 = X1)) & (~ (cell14 = X0) | ~ (cell13 = X1)) & (~ (cell13 = X0) | ~ (cell12 = X1)) & (~ (cell12 = X0) | ~ (cell11 = X1)) & (~ (cell10 = X0) | ~ (cell9 = X1)) & (~ (cell9 = X0) | ~ (cell8 = X1)) & (~ (cell8 = X0) | ~ (cell7 = X1)) & (~ (cell7 = X0) | ~ (cell6 = X1)) & (~ (cell6 = X0) | ~ (cell5 = X1)) & (~ (cell5 = X0) | ~ (cell4 = X1)) & (~ (cell4 = X0) | ~ (cell3 = X1)) & (~ (cell3 = X0) | ~ (cell2 = X1)) & (~ (cell2 = X0) | ~ (cell1 = X1))) | ~ sP1(X0, X1)), inference(usedef, [], [e29])).
fof(e29, plain, ! [X0, X1] : (sP1(X0, X1) <=> ((~ (cell100 = X0) | ~ (cell99 = X1)) & (~ (cell99 = X0) | ~ (cell98 = X1)) & (~ (cell98 = X0) | ~ (cell97 = X1)) & (~ (cell97 = X0) | ~ (cell96 = X1)) & (~ (cell96 = X0) | ~ (cell95 = X1)) & (~ (cell95 = X0) | ~ (cell94 = X1)) & (~ (cell94 = X0) | ~ (cell93 = X1)) & (~ (cell93 = X0) | ~ (cell92 = X1)) & (~ (cell92 = X0) | ~ (cell91 = X1)) & (~ (cell90 = X0) | ~ (cell89 = X1)) & (~ (cell89 = X0) | ~ (cell88 = X1)) & (~ (cell88 = X0) | ~ (cell87 = X1)) & (~ (cell87 = X0) | ~ (cell86 = X1)) & (~ (cell86 = X0) | ~ (cell85 = X1)) & (~ (cell85 = X0) | ~ (cell84 = X1)) & (~ (cell84 = X0) | ~ (cell83 = X1)) & (~ (cell83 = X0) | ~ (cell82 = X1)) & (~ (cell82 = X0) | ~ (cell81 = X1)) & (~ (cell80 = X0) | ~ (cell79 = X1)) & (~ (cell79 = X0) | ~ (cell78 = X1)) & (~ (cell78 = X0) | ~ (cell77 = X1)) & (~ (cell77 = X0) | ~ (cell76 = X1)) & (~ (cell76 = X0) | ~ (cell75 = X1)) & (~ (cell75 = X0) | ~ (cell74 = X1)) & (~ (cell74 = X0) | ~ (cell73 = X1)) & (~ (cell73 = X0) | ~ (cell72 = X1)) & (~ (cell72 = X0) | ~ (cell71 = X1)) & (~ (cell70 = X0) | ~ (cell69 = X1)) & (~ (cell69 = X0) | ~ (cell68 = X1)) & (~ (cell68 = X0) | ~ (cell67 = X1)) & (~ (cell67 = X0) | ~ (cell66 = X1)) & (~ (cell66 = X0) | ~ (cell65 = X1)) & (~ (cell65 = X0) | ~ (cell64 = X1)) & (~ (cell64 = X0) | ~ (cell63 = X1)) & (~ (cell63 = X0) | ~ (cell62 = X1)) & (~ (cell62 = X0) | ~ (cell61 = X1)) & (~ (cell60 = X0) | ~ (cell59 = X1)) & (~ (cell59 = X0) | ~ (cell58 = X1)) & (~ (cell58 = X0) | ~ (cell57 = X1)) & (~ (cell57 = X0) | ~ (cell56 = X1)) & (~ (cell56 = X0) | ~ (cell55 = X1)) & (~ (cell55 = X0) | ~ (cell54 = X1)) & (~ (cell54 = X0) | ~ (cell53 = X1)) & (~ (cell53 = X0) | ~ (cell52 = X1)) & (~ (cell52 = X0) | ~ (cell51 = X1)) & (~ (cell50 = X0) | ~ (cell49 = X1)) & (~ (cell49 = X0) | ~ (cell48 = X1)) & (~ (cell48 = X0) | ~ (cell47 = X1)) & (~ (cell47 = X0) | ~ (cell46 = X1)) & (~ (cell46 = X0) | ~ (cell45 = X1)) & (~ (cell45 = X0) | ~ (cell44 = X1)) & (~ (cell44 = X0) | ~ (cell43 = X1)) & (~ (cell43 = X0) | ~ (cell42 = X1)) & (~ (cell42 = X0) | ~ (cell41 = X1)) & (~ (cell40 = X0) | ~ (cell39 = X1)) & (~ (cell39 = X0) | ~ (cell38 = X1)) & (~ (cell38 = X0) | ~ (cell37 = X1)) & (~ (cell37 = X0) | ~ (cell36 = X1)) & (~ (cell36 = X0) | ~ (cell35 = X1)) & (~ (cell35 = X0) | ~ (cell34 = X1)) & (~ (cell34 = X0) | ~ (cell33 = X1)) & (~ (cell33 = X0) | ~ (cell32 = X1)) & (~ (cell32 = X0) | ~ (cell31 = X1)) & (~ (cell30 = X0) | ~ (cell29 = X1)) & (~ (cell29 = X0) | ~ (cell28 = X1)) & (~ (cell28 = X0) | ~ (cell27 = X1)) & (~ (cell27 = X0) | ~ (cell26 = X1)) & (~ (cell26 = X0) | ~ (cell25 = X1)) & (~ (cell25 = X0) | ~ (cell24 = X1)) & (~ (cell24 = X0) | ~ (cell23 = X1)) & (~ (cell23 = X0) | ~ (cell22 = X1)) & (~ (cell22 = X0) | ~ (cell21 = X1)) & (~ (cell20 = X0) | ~ (cell19 = X1)) & (~ (cell19 = X0) | ~ (cell18 = X1)) & (~ (cell18 = X0) | ~ (cell17 = X1)) & (~ (cell17 = X0) | ~ (cell16 = X1)) & (~ (cell16 = X0) | ~ (cell15 = X1)) & (~ (cell15 = X0) | ~ (cell14 = X1)) & (~ (cell14 = X0) | ~ (cell13 = X1)) & (~ (cell13 = X0) | ~ (cell12 = X1)) & (~ (cell12 = X0) | ~ (cell11 = X1)) & (~ (cell10 = X0) | ~ (cell9 = X1)) & (~ (cell9 = X0) | ~ (cell8 = X1)) & (~ (cell8 = X0) | ~ (cell7 = X1)) & (~ (cell7 = X0) | ~ (cell6 = X1)) & (~ (cell6 = X0) | ~ (cell5 = X1)) & (~ (cell5 = X0) | ~ (cell4 = X1)) & (~ (cell4 = X0) | ~ (cell3 = X1)) & (~ (cell3 = X0) | ~ (cell2 = X1)) & (~ (cell2 = X0) | ~ (cell1 = X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f220, plain, ! [X0, X1] : (sP1(X0, X1) | north(X1, X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0, X1] : (north(X1, X0) | sP1(X0, X1)), inference(definition_folding, [], [f22, e29])).
fof(f22, plain, ! [X0, X1] : (north(X1, X0) | ((~ (cell100 = X0) | ~ (cell99 = X1)) & (~ (cell99 = X0) | ~ (cell98 = X1)) & (~ (cell98 = X0) | ~ (cell97 = X1)) & (~ (cell97 = X0) | ~ (cell96 = X1)) & (~ (cell96 = X0) | ~ (cell95 = X1)) & (~ (cell95 = X0) | ~ (cell94 = X1)) & (~ (cell94 = X0) | ~ (cell93 = X1)) & (~ (cell93 = X0) | ~ (cell92 = X1)) & (~ (cell92 = X0) | ~ (cell91 = X1)) & (~ (cell90 = X0) | ~ (cell89 = X1)) & (~ (cell89 = X0) | ~ (cell88 = X1)) & (~ (cell88 = X0) | ~ (cell87 = X1)) & (~ (cell87 = X0) | ~ (cell86 = X1)) & (~ (cell86 = X0) | ~ (cell85 = X1)) & (~ (cell85 = X0) | ~ (cell84 = X1)) & (~ (cell84 = X0) | ~ (cell83 = X1)) & (~ (cell83 = X0) | ~ (cell82 = X1)) & (~ (cell82 = X0) | ~ (cell81 = X1)) & (~ (cell80 = X0) | ~ (cell79 = X1)) & (~ (cell79 = X0) | ~ (cell78 = X1)) & (~ (cell78 = X0) | ~ (cell77 = X1)) & (~ (cell77 = X0) | ~ (cell76 = X1)) & (~ (cell76 = X0) | ~ (cell75 = X1)) & (~ (cell75 = X0) | ~ (cell74 = X1)) & (~ (cell74 = X0) | ~ (cell73 = X1)) & (~ (cell73 = X0) | ~ (cell72 = X1)) & (~ (cell72 = X0) | ~ (cell71 = X1)) & (~ (cell70 = X0) | ~ (cell69 = X1)) & (~ (cell69 = X0) | ~ (cell68 = X1)) & (~ (cell68 = X0) | ~ (cell67 = X1)) & (~ (cell67 = X0) | ~ (cell66 = X1)) & (~ (cell66 = X0) | ~ (cell65 = X1)) & (~ (cell65 = X0) | ~ (cell64 = X1)) & (~ (cell64 = X0) | ~ (cell63 = X1)) & (~ (cell63 = X0) | ~ (cell62 = X1)) & (~ (cell62 = X0) | ~ (cell61 = X1)) & (~ (cell60 = X0) | ~ (cell59 = X1)) & (~ (cell59 = X0) | ~ (cell58 = X1)) & (~ (cell58 = X0) | ~ (cell57 = X1)) & (~ (cell57 = X0) | ~ (cell56 = X1)) & (~ (cell56 = X0) | ~ (cell55 = X1)) & (~ (cell55 = X0) | ~ (cell54 = X1)) & (~ (cell54 = X0) | ~ (cell53 = X1)) & (~ (cell53 = X0) | ~ (cell52 = X1)) & (~ (cell52 = X0) | ~ (cell51 = X1)) & (~ (cell50 = X0) | ~ (cell49 = X1)) & (~ (cell49 = X0) | ~ (cell48 = X1)) & (~ (cell48 = X0) | ~ (cell47 = X1)) & (~ (cell47 = X0) | ~ (cell46 = X1)) & (~ (cell46 = X0) | ~ (cell45 = X1)) & (~ (cell45 = X0) | ~ (cell44 = X1)) & (~ (cell44 = X0) | ~ (cell43 = X1)) & (~ (cell43 = X0) | ~ (cell42 = X1)) & (~ (cell42 = X0) | ~ (cell41 = X1)) & (~ (cell40 = X0) | ~ (cell39 = X1)) & (~ (cell39 = X0) | ~ (cell38 = X1)) & (~ (cell38 = X0) | ~ (cell37 = X1)) & (~ (cell37 = X0) | ~ (cell36 = X1)) & (~ (cell36 = X0) | ~ (cell35 = X1)) & (~ (cell35 = X0) | ~ (cell34 = X1)) & (~ (cell34 = X0) | ~ (cell33 = X1)) & (~ (cell33 = X0) | ~ (cell32 = X1)) & (~ (cell32 = X0) | ~ (cell31 = X1)) & (~ (cell30 = X0) | ~ (cell29 = X1)) & (~ (cell29 = X0) | ~ (cell28 = X1)) & (~ (cell28 = X0) | ~ (cell27 = X1)) & (~ (cell27 = X0) | ~ (cell26 = X1)) & (~ (cell26 = X0) | ~ (cell25 = X1)) & (~ (cell25 = X0) | ~ (cell24 = X1)) & (~ (cell24 = X0) | ~ (cell23 = X1)) & (~ (cell23 = X0) | ~ (cell22 = X1)) & (~ (cell22 = X0) | ~ (cell21 = X1)) & (~ (cell20 = X0) | ~ (cell19 = X1)) & (~ (cell19 = X0) | ~ (cell18 = X1)) & (~ (cell18 = X0) | ~ (cell17 = X1)) & (~ (cell17 = X0) | ~ (cell16 = X1)) & (~ (cell16 = X0) | ~ (cell15 = X1)) & (~ (cell15 = X0) | ~ (cell14 = X1)) & (~ (cell14 = X0) | ~ (cell13 = X1)) & (~ (cell13 = X0) | ~ (cell12 = X1)) & (~ (cell12 = X0) | ~ (cell11 = X1)) & (~ (cell10 = X0) | ~ (cell9 = X1)) & (~ (cell9 = X0) | ~ (cell8 = X1)) & (~ (cell8 = X0) | ~ (cell7 = X1)) & (~ (cell7 = X0) | ~ (cell6 = X1)) & (~ (cell6 = X0) | ~ (cell5 = X1)) & (~ (cell5 = X0) | ~ (cell4 = X1)) & (~ (cell4 = X0) | ~ (cell3 = X1)) & (~ (cell3 = X0) | ~ (cell2 = X1)) & (~ (cell2 = X0) | ~ (cell1 = X1)))), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X1] : ((((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1))) => north(X1, X0)), inference(unused_predicate_definition_removal, [], [f9])).
fof(f9, plain, ! [X0, X1] : (north(X1, X0) <=> (((cell100 = X0) & (cell99 = X1)) | ((cell99 = X0) & (cell98 = X1)) | ((cell98 = X0) & (cell97 = X1)) | ((cell97 = X0) & (cell96 = X1)) | ((cell96 = X0) & (cell95 = X1)) | ((cell95 = X0) & (cell94 = X1)) | ((cell94 = X0) & (cell93 = X1)) | ((cell93 = X0) & (cell92 = X1)) | ((cell92 = X0) & (cell91 = X1)) | ((cell90 = X0) & (cell89 = X1)) | ((cell89 = X0) & (cell88 = X1)) | ((cell88 = X0) & (cell87 = X1)) | ((cell87 = X0) & (cell86 = X1)) | ((cell86 = X0) & (cell85 = X1)) | ((cell85 = X0) & (cell84 = X1)) | ((cell84 = X0) & (cell83 = X1)) | ((cell83 = X0) & (cell82 = X1)) | ((cell82 = X0) & (cell81 = X1)) | ((cell80 = X0) & (cell79 = X1)) | ((cell79 = X0) & (cell78 = X1)) | ((cell78 = X0) & (cell77 = X1)) | ((cell77 = X0) & (cell76 = X1)) | ((cell76 = X0) & (cell75 = X1)) | ((cell75 = X0) & (cell74 = X1)) | ((cell74 = X0) & (cell73 = X1)) | ((cell73 = X0) & (cell72 = X1)) | ((cell72 = X0) & (cell71 = X1)) | ((cell70 = X0) & (cell69 = X1)) | ((cell69 = X0) & (cell68 = X1)) | ((cell68 = X0) & (cell67 = X1)) | ((cell67 = X0) & (cell66 = X1)) | ((cell66 = X0) & (cell65 = X1)) | ((cell65 = X0) & (cell64 = X1)) | ((cell64 = X0) & (cell63 = X1)) | ((cell63 = X0) & (cell62 = X1)) | ((cell62 = X0) & (cell61 = X1)) | ((cell60 = X0) & (cell59 = X1)) | ((cell59 = X0) & (cell58 = X1)) | ((cell58 = X0) & (cell57 = X1)) | ((cell57 = X0) & (cell56 = X1)) | ((cell56 = X0) & (cell55 = X1)) | ((cell55 = X0) & (cell54 = X1)) | ((cell54 = X0) & (cell53 = X1)) | ((cell53 = X0) & (cell52 = X1)) | ((cell52 = X0) & (cell51 = X1)) | ((cell50 = X0) & (cell49 = X1)) | ((cell49 = X0) & (cell48 = X1)) | ((cell48 = X0) & (cell47 = X1)) | ((cell47 = X0) & (cell46 = X1)) | ((cell46 = X0) & (cell45 = X1)) | ((cell45 = X0) & (cell44 = X1)) | ((cell44 = X0) & (cell43 = X1)) | ((cell43 = X0) & (cell42 = X1)) | ((cell42 = X0) & (cell41 = X1)) | ((cell40 = X0) & (cell39 = X1)) | ((cell39 = X0) & (cell38 = X1)) | ((cell38 = X0) & (cell37 = X1)) | ((cell37 = X0) & (cell36 = X1)) | ((cell36 = X0) & (cell35 = X1)) | ((cell35 = X0) & (cell34 = X1)) | ((cell34 = X0) & (cell33 = X1)) | ((cell33 = X0) & (cell32 = X1)) | ((cell32 = X0) & (cell31 = X1)) | ((cell30 = X0) & (cell29 = X1)) | ((cell29 = X0) & (cell28 = X1)) | ((cell28 = X0) & (cell27 = X1)) | ((cell27 = X0) & (cell26 = X1)) | ((cell26 = X0) & (cell25 = X1)) | ((cell25 = X0) & (cell24 = X1)) | ((cell24 = X0) & (cell23 = X1)) | ((cell23 = X0) & (cell22 = X1)) | ((cell22 = X0) & (cell21 = X1)) | ((cell20 = X0) & (cell19 = X1)) | ((cell19 = X0) & (cell18 = X1)) | ((cell18 = X0) & (cell17 = X1)) | ((cell17 = X0) & (cell16 = X1)) | ((cell16 = X0) & (cell15 = X1)) | ((cell15 = X0) & (cell14 = X1)) | ((cell14 = X0) & (cell13 = X1)) | ((cell13 = X0) & (cell12 = X1)) | ((cell12 = X0) & (cell11 = X1)) | ((cell10 = X0) & (cell9 = X1)) | ((cell9 = X0) & (cell8 = X1)) | ((cell8 = X0) & (cell7 = X1)) | ((cell7 = X0) & (cell6 = X1)) | ((cell6 = X0) & (cell5 = X1)) | ((cell5 = X0) & (cell4 = X1)) | ((cell4 = X0) & (cell3 = X1)) | ((cell3 = X0) & (cell2 = X1)) | ((cell2 = X0) & (cell1 = X1)))), inference(rectify, [], [f2])).
fof(f2, plain, ! [X2, X3] : (north(X3, X2) <=> (((cell100 = X2) & (cell99 = X3)) | ((cell99 = X2) & (cell98 = X3)) | ((cell98 = X2) & (cell97 = X3)) | ((cell97 = X2) & (cell96 = X3)) | ((cell96 = X2) & (cell95 = X3)) | ((cell95 = X2) & (cell94 = X3)) | ((cell94 = X2) & (cell93 = X3)) | ((cell93 = X2) & (cell92 = X3)) | ((cell92 = X2) & (cell91 = X3)) | ((cell90 = X2) & (cell89 = X3)) | ((cell89 = X2) & (cell88 = X3)) | ((cell88 = X2) & (cell87 = X3)) | ((cell87 = X2) & (cell86 = X3)) | ((cell86 = X2) & (cell85 = X3)) | ((cell85 = X2) & (cell84 = X3)) | ((cell84 = X2) & (cell83 = X3)) | ((cell83 = X2) & (cell82 = X3)) | ((cell82 = X2) & (cell81 = X3)) | ((cell80 = X2) & (cell79 = X3)) | ((cell79 = X2) & (cell78 = X3)) | ((cell78 = X2) & (cell77 = X3)) | ((cell77 = X2) & (cell76 = X3)) | ((cell76 = X2) & (cell75 = X3)) | ((cell75 = X2) & (cell74 = X3)) | ((cell74 = X2) & (cell73 = X3)) | ((cell73 = X2) & (cell72 = X3)) | ((cell72 = X2) & (cell71 = X3)) | ((cell70 = X2) & (cell69 = X3)) | ((cell69 = X2) & (cell68 = X3)) | ((cell68 = X2) & (cell67 = X3)) | ((cell67 = X2) & (cell66 = X3)) | ((cell66 = X2) & (cell65 = X3)) | ((cell65 = X2) & (cell64 = X3)) | ((cell64 = X2) & (cell63 = X3)) | ((cell63 = X2) & (cell62 = X3)) | ((cell62 = X2) & (cell61 = X3)) | ((cell60 = X2) & (cell59 = X3)) | ((cell59 = X2) & (cell58 = X3)) | ((cell58 = X2) & (cell57 = X3)) | ((cell57 = X2) & (cell56 = X3)) | ((cell56 = X2) & (cell55 = X3)) | ((cell55 = X2) & (cell54 = X3)) | ((cell54 = X2) & (cell53 = X3)) | ((cell53 = X2) & (cell52 = X3)) | ((cell52 = X2) & (cell51 = X3)) | ((cell50 = X2) & (cell49 = X3)) | ((cell49 = X2) & (cell48 = X3)) | ((cell48 = X2) & (cell47 = X3)) | ((cell47 = X2) & (cell46 = X3)) | ((cell46 = X2) & (cell45 = X3)) | ((cell45 = X2) & (cell44 = X3)) | ((cell44 = X2) & (cell43 = X3)) | ((cell43 = X2) & (cell42 = X3)) | ((cell42 = X2) & (cell41 = X3)) | ((cell40 = X2) & (cell39 = X3)) | ((cell39 = X2) & (cell38 = X3)) | ((cell38 = X2) & (cell37 = X3)) | ((cell37 = X2) & (cell36 = X3)) | ((cell36 = X2) & (cell35 = X3)) | ((cell35 = X2) & (cell34 = X3)) | ((cell34 = X2) & (cell33 = X3)) | ((cell33 = X2) & (cell32 = X3)) | ((cell32 = X2) & (cell31 = X3)) | ((cell30 = X2) & (cell29 = X3)) | ((cell29 = X2) & (cell28 = X3)) | ((cell28 = X2) & (cell27 = X3)) | ((cell27 = X2) & (cell26 = X3)) | ((cell26 = X2) & (cell25 = X3)) | ((cell25 = X2) & (cell24 = X3)) | ((cell24 = X2) & (cell23 = X3)) | ((cell23 = X2) & (cell22 = X3)) | ((cell22 = X2) & (cell21 = X3)) | ((cell20 = X2) & (cell19 = X3)) | ((cell19 = X2) & (cell18 = X3)) | ((cell18 = X2) & (cell17 = X3)) | ((cell17 = X2) & (cell16 = X3)) | ((cell16 = X2) & (cell15 = X3)) | ((cell15 = X2) & (cell14 = X3)) | ((cell14 = X2) & (cell13 = X3)) | ((cell13 = X2) & (cell12 = X3)) | ((cell12 = X2) & (cell11 = X3)) | ((cell10 = X2) & (cell9 = X3)) | ((cell9 = X2) & (cell8 = X3)) | ((cell8 = X2) & (cell7 = X3)) | ((cell7 = X2) & (cell6 = X3)) | ((cell6 = X2) & (cell5 = X3)) | ((cell5 = X2) & (cell4 = X3)) | ((cell4 = X2) & (cell3 = X3)) | ((cell3 = X2) & (cell2 = X3)) | ((cell2 = X2) & (cell1 = X3)))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34976)).
fof(f11444, plain, ! [X0] : (~ north(cell9, X0) | duenorth(cell1, X0)), inference(resolution, [], [f10664, f710])).
fof(f710, plain, north(cell8, cell9), inference(resolution, [], [f220, f591])).
fof(f591, plain, ~ sP1(cell9, cell8), inference(equality_resolution, [], [f590])).
fof(f590, plain, ! [X1] : (~ (cell8 = X1) | ~ sP1(cell9, X1)), inference(equality_resolution, [], [f137])).
fof(f137, plain, ! [X0, X1] : (~ (cell9 = X0) | ~ (cell8 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f10664, plain, ! [X0, X1] : (~ north(cell8, X0) | ~ north(X0, X1) | duenorth(cell1, X1)), inference(resolution, [], [f9926, f708])).
fof(f708, plain, north(cell7, cell8), inference(resolution, [], [f220, f593])).
fof(f593, plain, ~ sP1(cell8, cell7), inference(equality_resolution, [], [f592])).
fof(f592, plain, ! [X1] : (~ (cell7 = X1) | ~ sP1(cell8, X1)), inference(equality_resolution, [], [f136])).
fof(f136, plain, ! [X0, X1] : (~ (cell8 = X0) | ~ (cell7 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f9926, plain, ! [X2, X0, X1] : (~ north(cell7, X0) | ~ north(X0, X1) | ~ north(X1, X2) | duenorth(cell1, X2)), inference(resolution, [], [f8964, f706])).
fof(f706, plain, north(cell6, cell7), inference(resolution, [], [f220, f595])).
fof(f595, plain, ~ sP1(cell7, cell6), inference(equality_resolution, [], [f594])).
fof(f594, plain, ! [X1] : (~ (cell6 = X1) | ~ sP1(cell7, X1)), inference(equality_resolution, [], [f135])).
fof(f135, plain, ! [X0, X1] : (~ (cell7 = X0) | ~ (cell6 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f8964, plain, ! [X2, X0, X3, X1] : (~ north(cell6, X0) | ~ north(X0, X1) | ~ north(X1, X2) | ~ north(X2, X3) | duenorth(cell1, X3)), inference(resolution, [], [f7906, f704])).
fof(f704, plain, north(cell5, cell6), inference(resolution, [], [f220, f597])).
fof(f597, plain, ~ sP1(cell6, cell5), inference(equality_resolution, [], [f596])).
fof(f596, plain, ! [X1] : (~ (cell5 = X1) | ~ sP1(cell6, X1)), inference(equality_resolution, [], [f134])).
fof(f134, plain, ! [X0, X1] : (~ (cell6 = X0) | ~ (cell5 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f7906, plain, ! [X4, X2, X0, X3, X1] : (~ north(cell5, X0) | ~ north(X0, X1) | ~ north(X1, X2) | ~ north(X2, X3) | ~ north(X3, X4) | duenorth(cell1, X4)), inference(resolution, [], [f6750, f702])).
fof(f702, plain, north(cell4, cell5), inference(resolution, [], [f220, f599])).
fof(f599, plain, ~ sP1(cell5, cell4), inference(equality_resolution, [], [f598])).
fof(f598, plain, ! [X1] : (~ (cell4 = X1) | ~ sP1(cell5, X1)), inference(equality_resolution, [], [f133])).
fof(f133, plain, ! [X0, X1] : (~ (cell5 = X0) | ~ (cell4 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f6750, plain, ! [X4, X2, X0, X5, X3, X1] : (~ north(cell4, X0) | ~ north(X0, X1) | ~ north(X1, X2) | ~ north(X2, X3) | ~ north(X3, X4) | ~ north(X4, X5) | duenorth(cell1, X5)), inference(resolution, [], [f5093, f700])).
fof(f700, plain, north(cell3, cell4), inference(resolution, [], [f220, f601])).
fof(f601, plain, ~ sP1(cell4, cell3), inference(equality_resolution, [], [f600])).
fof(f600, plain, ! [X1] : (~ (cell3 = X1) | ~ sP1(cell4, X1)), inference(equality_resolution, [], [f132])).
fof(f132, plain, ! [X0, X1] : (~ (cell4 = X0) | ~ (cell3 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f5093, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ north(cell3, X3) | ~ north(X2, X0) | ~ north(X3, X2) | ~ north(X0, X1) | ~ north(X1, X4) | ~ north(X4, X5) | ~ north(X5, X6) | duenorth(cell1, X6)), inference(resolution, [], [f2351, f698])).
fof(f698, plain, north(cell2, cell3), inference(resolution, [], [f220, f603])).
fof(f603, plain, ~ sP1(cell3, cell2), inference(equality_resolution, [], [f602])).
fof(f602, plain, ! [X1] : (~ (cell2 = X1) | ~ sP1(cell3, X1)), inference(equality_resolution, [], [f131])).
fof(f131, plain, ! [X0, X1] : (~ (cell3 = X0) | ~ (cell2 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f2351, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (~ north(cell2, X5) | ~ north(X2, X0) | ~ north(X3, X2) | ~ north(X4, X3) | ~ north(X5, X4) | ~ north(X0, X1) | ~ north(X1, X6) | ~ north(X6, X7) | duenorth(cell1, X7)), inference(resolution, [], [f803, f696])).
fof(f696, plain, north(cell1, cell2), inference(resolution, [], [f220, f605])).
fof(f605, plain, ~ sP1(cell2, cell1), inference(equality_resolution, [], [f604])).
fof(f604, plain, ! [X1] : (~ (cell1 = X1) | ~ sP1(cell2, X1)), inference(equality_resolution, [], [f130])).
fof(f130, plain, ! [X0, X1] : (~ (cell2 = X0) | ~ (cell1 = X1) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f36])).
fof(f803, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ north(X8, X7) | ~ north(X2, X0) | ~ north(X3, X2) | ~ north(X4, X3) | ~ north(X5, X4) | ~ north(X6, X5) | ~ north(X7, X6) | ~ north(X0, X1) | ~ north(X1, X9) | duenorth(X8, X9)), inference(resolution, [], [f239, f240])).
fof(f240, plain, ! [X0, X1] : (sP3(X0, X1) | duenorth(X1, X0)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ! [X0, X1] : (duenorth(X1, X0) | sP3(X0, X1)), inference(definition_folding, [], [f24, e33])).
fof(f33, plain, ! [X0, X1] : ((! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ north(X9, X0) | ~ north(X8, X9) | ~ north(X7, X8) | ~ north(X6, X7) | ~ north(X5, X6) | ~ north(X4, X5) | ~ north(X3, X4) | ~ north(X2, X3) | ~ north(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ north(X16, X0) | ~ north(X15, X16) | ~ north(X14, X15) | ~ north(X13, X14) | ~ north(X12, X13) | ~ north(X11, X12) | ~ north(X10, X11) | ~ north(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ north(X22, X0) | ~ north(X21, X22) | ~ north(X20, X21) | ~ north(X19, X20) | ~ north(X18, X19) | ~ north(X17, X18) | ~ north(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ north(X27, X0) | ~ north(X26, X27) | ~ north(X25, X26) | ~ north(X24, X25) | ~ north(X23, X24) | ~ north(X1, X23)) & ! [X28, X29, X30, X31] : (~ north(X31, X0) | ~ north(X30, X31) | ~ north(X29, X30) | ~ north(X28, X29) | ~ north(X1, X28)) & ! [X32, X33, X34] : (~ north(X34, X0) | ~ north(X33, X34) | ~ north(X32, X33) | ~ north(X1, X32)) & ! [X35, X36] : (~ north(X36, X0) | ~ north(X35, X36) | ~ north(X1, X35)) & ! [X37] : (~ north(X37, X0) | ~ north(X1, X37)) & ~ north(X1, X0)) | ~ sP3(X0, X1)), inference(usedef, [], [e33])).
fof(e33, plain, ! [X0, X1] : (sP3(X0, X1) <=> (! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ north(X9, X0) | ~ north(X8, X9) | ~ north(X7, X8) | ~ north(X6, X7) | ~ north(X5, X6) | ~ north(X4, X5) | ~ north(X3, X4) | ~ north(X2, X3) | ~ north(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ north(X16, X0) | ~ north(X15, X16) | ~ north(X14, X15) | ~ north(X13, X14) | ~ north(X12, X13) | ~ north(X11, X12) | ~ north(X10, X11) | ~ north(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ north(X22, X0) | ~ north(X21, X22) | ~ north(X20, X21) | ~ north(X19, X20) | ~ north(X18, X19) | ~ north(X17, X18) | ~ north(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ north(X27, X0) | ~ north(X26, X27) | ~ north(X25, X26) | ~ north(X24, X25) | ~ north(X23, X24) | ~ north(X1, X23)) & ! [X28, X29, X30, X31] : (~ north(X31, X0) | ~ north(X30, X31) | ~ north(X29, X30) | ~ north(X28, X29) | ~ north(X1, X28)) & ! [X32, X33, X34] : (~ north(X34, X0) | ~ north(X33, X34) | ~ north(X32, X33) | ~ north(X1, X32)) & ! [X35, X36] : (~ north(X36, X0) | ~ north(X35, X36) | ~ north(X1, X35)) & ! [X37] : (~ north(X37, X0) | ~ north(X1, X37)) & ~ north(X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f24, plain, ! [X0, X1] : (duenorth(X1, X0) | (! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ north(X9, X0) | ~ north(X8, X9) | ~ north(X7, X8) | ~ north(X6, X7) | ~ north(X5, X6) | ~ north(X4, X5) | ~ north(X3, X4) | ~ north(X2, X3) | ~ north(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ north(X16, X0) | ~ north(X15, X16) | ~ north(X14, X15) | ~ north(X13, X14) | ~ north(X12, X13) | ~ north(X11, X12) | ~ north(X10, X11) | ~ north(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ north(X22, X0) | ~ north(X21, X22) | ~ north(X20, X21) | ~ north(X19, X20) | ~ north(X18, X19) | ~ north(X17, X18) | ~ north(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ north(X27, X0) | ~ north(X26, X27) | ~ north(X25, X26) | ~ north(X24, X25) | ~ north(X23, X24) | ~ north(X1, X23)) & ! [X28, X29, X30, X31] : (~ north(X31, X0) | ~ north(X30, X31) | ~ north(X29, X30) | ~ north(X28, X29) | ~ north(X1, X28)) & ! [X32, X33, X34] : (~ north(X34, X0) | ~ north(X33, X34) | ~ north(X32, X33) | ~ north(X1, X32)) & ! [X35, X36] : (~ north(X36, X0) | ~ north(X35, X36) | ~ north(X1, X35)) & ! [X37] : (~ north(X37, X0) | ~ north(X1, X37)) & ~ north(X1, X0))), inference(ennf_transformation, [], [f19])).
fof(f19, plain, ! [X0, X1] : ((? [X2, X3, X4, X5, X6, X7, X8, X9] : (north(X9, X0) & north(X8, X9) & north(X7, X8) & north(X6, X7) & north(X5, X6) & north(X4, X5) & north(X3, X4) & north(X2, X3) & north(X1, X2)) | ? [X10, X11, X12, X13, X14, X15, X16] : (north(X16, X0) & north(X15, X16) & north(X14, X15) & north(X13, X14) & north(X12, X13) & north(X11, X12) & north(X10, X11) & north(X1, X10)) | ? [X17, X18, X19, X20, X21, X22] : (north(X22, X0) & north(X21, X22) & north(X20, X21) & north(X19, X20) & north(X18, X19) & north(X17, X18) & north(X1, X17)) | ? [X23, X24, X25, X26, X27] : (north(X27, X0) & north(X26, X27) & north(X25, X26) & north(X24, X25) & north(X23, X24) & north(X1, X23)) | ? [X28, X29, X30, X31] : (north(X31, X0) & north(X30, X31) & north(X29, X30) & north(X28, X29) & north(X1, X28)) | ? [X32, X33, X34] : (north(X34, X0) & north(X33, X34) & north(X32, X33) & north(X1, X32)) | ? [X35, X36] : (north(X36, X0) & north(X35, X36) & north(X1, X35)) | ? [X37] : (north(X37, X0) & north(X1, X37)) | north(X1, X0)) => duenorth(X1, X0)), inference(unused_predicate_definition_removal, [], [f11])).
fof(f11, plain, ! [X0, X1] : (duenorth(X1, X0) <=> (? [X2, X3, X4, X5, X6, X7, X8, X9] : (north(X9, X0) & north(X8, X9) & north(X7, X8) & north(X6, X7) & north(X5, X6) & north(X4, X5) & north(X3, X4) & north(X2, X3) & north(X1, X2)) | ? [X10, X11, X12, X13, X14, X15, X16] : (north(X16, X0) & north(X15, X16) & north(X14, X15) & north(X13, X14) & north(X12, X13) & north(X11, X12) & north(X10, X11) & north(X1, X10)) | ? [X17, X18, X19, X20, X21, X22] : (north(X22, X0) & north(X21, X22) & north(X20, X21) & north(X19, X20) & north(X18, X19) & north(X17, X18) & north(X1, X17)) | ? [X23, X24, X25, X26, X27] : (north(X27, X0) & north(X26, X27) & north(X25, X26) & north(X24, X25) & north(X23, X24) & north(X1, X23)) | ? [X28, X29, X30, X31] : (north(X31, X0) & north(X30, X31) & north(X29, X30) & north(X28, X29) & north(X1, X28)) | ? [X32, X33, X34] : (north(X34, X0) & north(X33, X34) & north(X32, X33) & north(X1, X32)) | ? [X35, X36] : (north(X36, X0) & north(X35, X36) & north(X1, X35)) | ? [X37] : (north(X37, X0) & north(X1, X37)) | north(X1, X0))), inference(rectify, [], [f4])).
fof(f4, plain, ! [X42, X43] : (duenorth(X43, X42) <=> (? [X72, X73, X74, X75, X76, X77, X78, X79] : (north(X79, X42) & north(X78, X79) & north(X77, X78) & north(X76, X77) & north(X75, X76) & north(X74, X75) & north(X73, X74) & north(X72, X73) & north(X43, X72)) | ? [X65, X66, X67, X68, X69, X70, X71] : (north(X71, X42) & north(X70, X71) & north(X69, X70) & north(X68, X69) & north(X67, X68) & north(X66, X67) & north(X65, X66) & north(X43, X65)) | ? [X59, X60, X61, X62, X63, X64] : (north(X64, X42) & north(X63, X64) & north(X62, X63) & north(X61, X62) & north(X60, X61) & north(X59, X60) & north(X43, X59)) | ? [X54, X55, X56, X57, X58] : (north(X58, X42) & north(X57, X58) & north(X56, X57) & north(X55, X56) & north(X54, X55) & north(X43, X54)) | ? [X50, X51, X52, X53] : (north(X53, X42) & north(X52, X53) & north(X51, X52) & north(X50, X51) & north(X43, X50)) | ? [X47, X48, X49] : (north(X49, X42) & north(X48, X49) & north(X47, X48) & north(X43, X47)) | ? [X45, X46] : (north(X46, X42) & north(X45, X46) & north(X43, X45)) | ? [X44] : (north(X44, X42) & north(X43, X44)) | north(X43, X42))), file('/home/ubuntu/library/tptp/Problems/PUZ/PUZ077+1.p', tlhfof34978)).
fof(f239, plain, ! [X6, X4, X2, X0, X8, X7, X5, X3, X1, X9] : (~ sP3(X0, X1) | ~ north(X8, X9) | ~ north(X7, X8) | ~ north(X6, X7) | ~ north(X5, X6) | ~ north(X4, X5) | ~ north(X3, X4) | ~ north(X2, X3) | ~ north(X1, X2) | ~ north(X9, X0)), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ! [X0, X1] : ((! [X2, X3, X4, X5, X6, X7, X8, X9] : (~ north(X9, X0) | ~ north(X8, X9) | ~ north(X7, X8) | ~ north(X6, X7) | ~ north(X5, X6) | ~ north(X4, X5) | ~ north(X3, X4) | ~ north(X2, X3) | ~ north(X1, X2)) & ! [X10, X11, X12, X13, X14, X15, X16] : (~ north(X16, X0) | ~ north(X15, X16) | ~ north(X14, X15) | ~ north(X13, X14) | ~ north(X12, X13) | ~ north(X11, X12) | ~ north(X10, X11) | ~ north(X1, X10)) & ! [X17, X18, X19, X20, X21, X22] : (~ north(X22, X0) | ~ north(X21, X22) | ~ north(X20, X21) | ~ north(X19, X20) | ~ north(X18, X19) | ~ north(X17, X18) | ~ north(X1, X17)) & ! [X23, X24, X25, X26, X27] : (~ north(X27, X0) | ~ north(X26, X27) | ~ north(X25, X26) | ~ north(X24, X25) | ~ north(X23, X24) | ~ north(X1, X23)) & ! [X28, X29, X30, X31] : (~ north(X31, X0) | ~ north(X30, X31) | ~ north(X29, X30) | ~ north(X28, X29) | ~ north(X1, X28)) & ! [X32, X33, X34] : (~ north(X34, X0) | ~ north(X33, X34) | ~ north(X32, X33) | ~ north(X1, X32)) & ! [X35, X36] : (~ north(X36, X0) | ~ north(X35, X36) | ~ north(X1, X35)) & ! [X37] : (~ north(X37, X0) | ~ north(X1, X37)) & ~ north(X1, X0)) | ~ sP3(X0, X1)), inference(nnf_transformation, [], [f33])).