fof(f9606, plain, $false, inference(avatar_sat_refutation, [], [f677, f715, f722, f816, f1601, f1604, f1607, f2429, f2435, f2455, f9595])).
fof(f9595, plain, (~ spl18_19 | spl18_23 | ~ spl18_24 | ~ spl18_25), inference(avatar_contradiction_clause, [], [f9594])).
fof(f9594, plain, ($false | (~ spl18_19 | spl18_23 | ~ spl18_24 | ~ spl18_25)), inference(subsumption_resolution, [], [f9593, f1478])).
fof(f1478, plain, (aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ spl18_25), inference(avatar_component_clause, [], [f1476])).
fof(f1476, plain, (spl18_25 <=> aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl18_25])])).
fof(f9593, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | (~ spl18_19 | spl18_23 | ~ spl18_24)), inference(subsumption_resolution, [], [f9590, f1469])).
fof(f1469, plain, (aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ spl18_24), inference(avatar_component_clause, [], [f1468])).
fof(f1468, plain, (spl18_24 <=> aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl18_24])])).
fof(f9590, plain, (~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | (~ spl18_19 | spl18_23)), inference(resolution, [], [f1466, f718])).
fof(f718, plain, (! [X8, X7] : (aElementOf0(sdtpldt0(X8, X7), sdtpldt1(xI, xJ)) | ~ aElementOf0(X8, sdtpldt1(xI, xJ)) | ~ aElementOf0(X7, sdtpldt1(xI, xJ))) | ~ spl18_19), inference(avatar_component_clause, [], [f717])).
fof(f717, plain, (spl18_19 <=> ! [X7, X8] : (~ aElementOf0(X7, sdtpldt1(xI, xJ)) | ~ aElementOf0(X8, sdtpldt1(xI, xJ)) | aElementOf0(sdtpldt0(X8, X7), sdtpldt1(xI, xJ)))), introduced(avatar_definition, [new_symbols(naming, [spl18_19])])).
fof(f1466, plain, (~ aElementOf0(sdtpldt0(sK11(sdtpldt1(xI, xJ)), sK13(sdtpldt1(xI, xJ))), sdtpldt1(xI, xJ)) | spl18_23), inference(avatar_component_clause, [], [f1464])).
fof(f1464, plain, (spl18_23 <=> aElementOf0(sdtpldt0(sK11(sdtpldt1(xI, xJ)), sK13(sdtpldt1(xI, xJ))), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl18_23])])).
fof(f2455, plain, ~ spl18_20, inference(avatar_contradiction_clause, [], [f2436])).
fof(f2436, plain, ($false | ~ spl18_20), inference(resolution, [], [f721, f98])).
fof(f98, plain, aElement0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aElement0(sz00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+1.p', mSortsC)).
fof(f721, plain, (! [X6] : ~ aElement0(X6) | ~ spl18_20), inference(avatar_component_clause, [], [f720])).
fof(f720, plain, (spl18_20 <=> ! [X6] : ~ aElement0(X6)), introduced(avatar_definition, [new_symbols(naming, [spl18_20])])).
fof(f2435, plain, (~ spl18_7 | spl18_10 | spl18_25), inference(avatar_contradiction_clause, [], [f2434])).
fof(f2434, plain, ($false | (~ spl18_7 | spl18_10 | spl18_25)), inference(subsumption_resolution, [], [f2433, f624])).
fof(f624, plain, (aSet0(sdtpldt1(xI, xJ)) | ~ spl18_7), inference(avatar_component_clause, [], [f623])).
fof(f623, plain, (spl18_7 <=> aSet0(sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl18_7])])).
fof(f2433, plain, (~ aSet0(sdtpldt1(xI, xJ)) | (spl18_10 | spl18_25)), inference(subsumption_resolution, [], [f2432, f169])).
fof(f169, plain, ~ aIdeal0(sdtpldt1(xI, xJ)), inference(cnf_transformation, [], [f97])).
fof(f97, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : ((sP2(X0, X2, X1, sK17(X0, X1, X2), sK16(X0, X1, X2)) & (sdtpldt0(sK16(X0, X1, X2), sK17(X0, X1, X2)) = X0) & aElementOf0(sK17(X0, X1, X2), xJ) & aElementOf0(sK16(X0, X1, X2), xI)) | ~ aElement0(X2) | ~ aElementOf0(X1, sdtpldt1(xI, xJ)) | ~ aElementOf0(X0, sdtpldt1(xI, xJ)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK16, sK17])], [f68, f96])).
fof(f96, plain, ! [X2, X1, X0] : (? [X3, X4] : (sP2(X0, X2, X1, X4, X3) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) => (sP2(X0, X2, X1, sK17(X0, X1, X2), sK16(X0, X1, X2)) & (sdtpldt0(sK16(X0, X1, X2), sK17(X0, X1, X2)) = X0) & aElementOf0(sK17(X0, X1, X2), xJ) & aElementOf0(sK16(X0, X1, X2), xI))), introduced(choice_axiom, [])).
fof(f68, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (? [X3, X4] : (sP2(X0, X2, X1, X4, X3) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) | ~ aElement0(X2) | ~ aElementOf0(X1, sdtpldt1(xI, xJ)) | ~ aElementOf0(X0, sdtpldt1(xI, xJ)))), inference(definition_folding, [], [f63, e67])).
fof(f67, plain, ! [X0, X2, X1, X4, X3] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) | ~ sP2(X0, X2, X1, X4, X3)), inference(usedef, [], [e67])).
fof(e67, plain, ! [X0, X2, X1, X4, X3] : (sP2(X0, X2, X1, X4, X3) <=> ? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f63, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) | ~ aElement0(X2) | ~ aElementOf0(X1, sdtpldt1(xI, xJ)) | ~ aElementOf0(X0, sdtpldt1(xI, xJ)))), inference(flattening, [], [f62])).
fof(f62, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) | (~ aElement0(X2) | ~ aElementOf0(X1, sdtpldt1(xI, xJ)) | ~ aElementOf0(X0, sdtpldt1(xI, xJ))))), inference(ennf_transformation, [], [f27])).
fof(f27, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & aElementOf0(X1, sdtpldt1(xI, xJ)) & aElementOf0(X0, sdtpldt1(xI, xJ))) => ? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) => aIdeal0(sdtpldt1(xI, xJ))), inference(negated_conjecture, [], [f26])).
fof(f26, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & aElementOf0(X1, sdtpldt1(xI, xJ)) & aElementOf0(X0, sdtpldt1(xI, xJ))) => ? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) => aIdeal0(sdtpldt1(xI, xJ))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+1.p', m__)).
fof(f2432, plain, (aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | (spl18_10 | spl18_25)), inference(subsumption_resolution, [], [f2431, f638])).
fof(f638, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | spl18_10), inference(avatar_component_clause, [], [f637])).
fof(f637, plain, (spl18_10 <=> aElement0(sK12(sdtpldt1(xI, xJ)))), introduced(avatar_definition, [new_symbols(naming, [spl18_10])])).
fof(f2431, plain, (aElement0(sK12(sdtpldt1(xI, xJ))) | aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | spl18_25), inference(resolution, [], [f1477, f149])).
fof(f149, plain, ! [X0] : (aElementOf0(sK13(X0), X0) | aElement0(sK12(X0)) | aIdeal0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f91, plain, ! [X0] : ((aIdeal0(X0) | (((~ aElementOf0(sdtasdt0(sK12(X0), sK11(X0)), X0) & aElement0(sK12(X0))) | (~ aElementOf0(sdtpldt0(sK11(X0), sK13(X0)), X0) & aElementOf0(sK13(X0), X0))) & aElementOf0(sK11(X0), X0)) | ~ aSet0(X0)) & ((! [X4] : ((! [X5] : (aElementOf0(sdtasdt0(X5, X4), X0) | ~ aElement0(X5)) & ! [X6] : (aElementOf0(sdtpldt0(X4, X6), X0) | ~ aElementOf0(X6, X0))) | ~ aElementOf0(X4, X0)) & aSet0(X0)) | ~ aIdeal0(X0))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12, sK13])], [f87, f90, f89, f88])).
fof(f88, plain, ! [X0] : (? [X1] : ((? [X2] : (~ aElementOf0(sdtasdt0(X2, X1), X0) & aElement0(X2)) | ? [X3] : (~ aElementOf0(sdtpldt0(X1, X3), X0) & aElementOf0(X3, X0))) & aElementOf0(X1, X0)) => ((? [X2] : (~ aElementOf0(sdtasdt0(X2, sK11(X0)), X0) & aElement0(X2)) | ? [X3] : (~ aElementOf0(sdtpldt0(sK11(X0), X3), X0) & aElementOf0(X3, X0))) & aElementOf0(sK11(X0), X0))), introduced(choice_axiom, [])).
fof(f89, plain, ! [X0] : (? [X2] : (~ aElementOf0(sdtasdt0(X2, sK11(X0)), X0) & aElement0(X2)) => (~ aElementOf0(sdtasdt0(sK12(X0), sK11(X0)), X0) & aElement0(sK12(X0)))), introduced(choice_axiom, [])).
fof(f90, plain, ! [X0] : (? [X3] : (~ aElementOf0(sdtpldt0(sK11(X0), X3), X0) & aElementOf0(X3, X0)) => (~ aElementOf0(sdtpldt0(sK11(X0), sK13(X0)), X0) & aElementOf0(sK13(X0), X0))), introduced(choice_axiom, [])).
fof(f87, plain, ! [X0] : ((aIdeal0(X0) | ? [X1] : ((? [X2] : (~ aElementOf0(sdtasdt0(X2, X1), X0) & aElement0(X2)) | ? [X3] : (~ aElementOf0(sdtpldt0(X1, X3), X0) & aElementOf0(X3, X0))) & aElementOf0(X1, X0)) | ~ aSet0(X0)) & ((! [X4] : ((! [X5] : (aElementOf0(sdtasdt0(X5, X4), X0) | ~ aElement0(X5)) & ! [X6] : (aElementOf0(sdtpldt0(X4, X6), X0) | ~ aElementOf0(X6, X0))) | ~ aElementOf0(X4, X0)) & aSet0(X0)) | ~ aIdeal0(X0))), inference(rectify, [], [f86])).
fof(f86, plain, ! [X0] : ((aIdeal0(X0) | ? [X1] : ((? [X2] : (~ aElementOf0(sdtasdt0(X2, X1), X0) & aElement0(X2)) | ? [X3] : (~ aElementOf0(sdtpldt0(X1, X3), X0) & aElementOf0(X3, X0))) & aElementOf0(X1, X0)) | ~ aSet0(X0)) & ((! [X1] : ((! [X2] : (aElementOf0(sdtasdt0(X2, X1), X0) | ~ aElement0(X2)) & ! [X3] : (aElementOf0(sdtpldt0(X1, X3), X0) | ~ aElementOf0(X3, X0))) | ~ aElementOf0(X1, X0)) & aSet0(X0)) | ~ aIdeal0(X0))), inference(flattening, [], [f85])).
fof(f85, plain, ! [X0] : ((aIdeal0(X0) | (? [X1] : ((? [X2] : (~ aElementOf0(sdtasdt0(X2, X1), X0) & aElement0(X2)) | ? [X3] : (~ aElementOf0(sdtpldt0(X1, X3), X0) & aElementOf0(X3, X0))) & aElementOf0(X1, X0)) | ~ aSet0(X0))) & ((! [X1] : ((! [X2] : (aElementOf0(sdtasdt0(X2, X1), X0) | ~ aElement0(X2)) & ! [X3] : (aElementOf0(sdtpldt0(X1, X3), X0) | ~ aElementOf0(X3, X0))) | ~ aElementOf0(X1, X0)) & aSet0(X0)) | ~ aIdeal0(X0))), inference(nnf_transformation, [], [f61])).
fof(f61, plain, ! [X0] : (aIdeal0(X0) <=> (! [X1] : ((! [X2] : (aElementOf0(sdtasdt0(X2, X1), X0) | ~ aElement0(X2)) & ! [X3] : (aElementOf0(sdtpldt0(X1, X3), X0) | ~ aElementOf0(X3, X0))) | ~ aElementOf0(X1, X0)) & aSet0(X0))), inference(ennf_transformation, [], [f31])).
fof(f31, plain, ! [X0] : (aIdeal0(X0) <=> (! [X1] : (aElementOf0(X1, X0) => (! [X2] : (aElement0(X2) => aElementOf0(sdtasdt0(X2, X1), X0)) & ! [X3] : (aElementOf0(X3, X0) => aElementOf0(sdtpldt0(X1, X3), X0)))) & aSet0(X0))), inference(rectify, [], [f24])).
fof(f24, plain, ! [X0] : (aIdeal0(X0) <=> (! [X1] : (aElementOf0(X1, X0) => (! [X2] : (aElement0(X2) => aElementOf0(sdtasdt0(X2, X1), X0)) & ! [X2] : (aElementOf0(X2, X0) => aElementOf0(sdtpldt0(X1, X2), X0)))) & aSet0(X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+1.p', mDefIdeal)).
fof(f1477, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | spl18_25), inference(avatar_component_clause, [], [f1476])).
fof(f2429, plain, (~ spl18_25 | ~ spl18_7 | spl18_10 | ~ spl18_19 | ~ spl18_24), inference(avatar_split_clause, [], [f2428, f1468, f717, f637, f623, f1476])).
fof(f2428, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | (~ spl18_7 | spl18_10 | ~ spl18_19 | ~ spl18_24)), inference(subsumption_resolution, [], [f2427, f624])).
fof(f2427, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | (spl18_10 | ~ spl18_19 | ~ spl18_24)), inference(subsumption_resolution, [], [f2426, f169])).
fof(f2426, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | (spl18_10 | ~ spl18_19 | ~ spl18_24)), inference(subsumption_resolution, [], [f2425, f638])).
fof(f2425, plain, (~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aElement0(sK12(sdtpldt1(xI, xJ))) | aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | (~ spl18_19 | ~ spl18_24)), inference(subsumption_resolution, [], [f2405, f1469])).
fof(f2405, plain, (~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aElement0(sK12(sdtpldt1(xI, xJ))) | aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_19), inference(resolution, [], [f718, f150])).
fof(f150, plain, ! [X0] : (~ aElementOf0(sdtpldt0(sK11(X0), sK13(X0)), X0) | aElement0(sK12(X0)) | aIdeal0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f1607, plain, (~ spl18_23 | ~ spl18_24 | ~ spl18_10 | ~ spl18_7 | ~ spl18_18), inference(avatar_split_clause, [], [f1606, f713, f623, f637, f1468, f1464])).
fof(f713, plain, (spl18_18 <=> ! [X3, X5] : (~ aElement0(X3) | aElementOf0(sdtasdt0(X3, X5), sdtpldt1(xI, xJ)) | ~ aElementOf0(X5, sdtpldt1(xI, xJ)))), introduced(avatar_definition, [new_symbols(naming, [spl18_18])])).
fof(f1606, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aElementOf0(sdtpldt0(sK11(sdtpldt1(xI, xJ)), sK13(sdtpldt1(xI, xJ))), sdtpldt1(xI, xJ)) | (~ spl18_7 | ~ spl18_18)), inference(subsumption_resolution, [], [f1605, f624])).
fof(f1605, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aElementOf0(sdtpldt0(sK11(sdtpldt1(xI, xJ)), sK13(sdtpldt1(xI, xJ))), sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_18), inference(subsumption_resolution, [], [f1426, f169])).
fof(f1426, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aIdeal0(sdtpldt1(xI, xJ)) | ~ aElementOf0(sdtpldt0(sK11(sdtpldt1(xI, xJ)), sK13(sdtpldt1(xI, xJ))), sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_18), inference(resolution, [], [f714, f152])).
fof(f152, plain, ! [X0] : (~ aElementOf0(sdtasdt0(sK12(X0), sK11(X0)), X0) | aIdeal0(X0) | ~ aElementOf0(sdtpldt0(sK11(X0), sK13(X0)), X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f714, plain, (! [X5, X3] : (aElementOf0(sdtasdt0(X3, X5), sdtpldt1(xI, xJ)) | ~ aElement0(X3) | ~ aElementOf0(X5, sdtpldt1(xI, xJ))) | ~ spl18_18), inference(avatar_component_clause, [], [f713])).
fof(f1604, plain, (spl18_25 | ~ spl18_24 | ~ spl18_10 | ~ spl18_7 | ~ spl18_18), inference(avatar_split_clause, [], [f1603, f713, f623, f637, f1468, f1476])).
fof(f1603, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | (~ spl18_7 | ~ spl18_18)), inference(subsumption_resolution, [], [f1602, f624])).
fof(f1602, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_18), inference(subsumption_resolution, [], [f1427, f169])).
fof(f1427, plain, (~ aElement0(sK12(sdtpldt1(xI, xJ))) | ~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | aIdeal0(sdtpldt1(xI, xJ)) | aElementOf0(sK13(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_18), inference(resolution, [], [f714, f151])).
fof(f151, plain, ! [X0] : (~ aElementOf0(sdtasdt0(sK12(X0), sK11(X0)), X0) | aIdeal0(X0) | aElementOf0(sK13(X0), X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f1601, plain, (~ spl18_7 | spl18_24), inference(avatar_contradiction_clause, [], [f1600])).
fof(f1600, plain, ($false | (~ spl18_7 | spl18_24)), inference(subsumption_resolution, [], [f1599, f624])).
fof(f1599, plain, (~ aSet0(sdtpldt1(xI, xJ)) | spl18_24), inference(subsumption_resolution, [], [f1598, f169])).
fof(f1598, plain, (aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | spl18_24), inference(resolution, [], [f1470, f148])).
fof(f148, plain, ! [X0] : (aElementOf0(sK11(X0), X0) | aIdeal0(X0) | ~ aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f1470, plain, (~ aElementOf0(sK11(sdtpldt1(xI, xJ)), sdtpldt1(xI, xJ)) | spl18_24), inference(avatar_component_clause, [], [f1468])).
fof(f816, plain, (~ spl18_7 | ~ spl18_17), inference(avatar_contradiction_clause, [], [f815])).
fof(f815, plain, ($false | (~ spl18_7 | ~ spl18_17)), inference(subsumption_resolution, [], [f814, f624])).
fof(f814, plain, (~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_17), inference(subsumption_resolution, [], [f800, f169])).
fof(f800, plain, (aIdeal0(sdtpldt1(xI, xJ)) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl18_17), inference(resolution, [], [f711, f148])).
fof(f711, plain, (! [X4] : ~ aElementOf0(X4, sdtpldt1(xI, xJ)) | ~ spl18_17), inference(avatar_component_clause, [], [f710])).
fof(f710, plain, (spl18_17 <=> ! [X4] : ~ aElementOf0(X4, sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl18_17])])).
fof(f722, plain, (spl18_19 | spl18_20), inference(avatar_split_clause, [], [f706, f720, f717])).
fof(f706, plain, ! [X6, X8, X7] : (~ aElement0(X6) | ~ aElementOf0(X7, sdtpldt1(xI, xJ)) | ~ aElementOf0(X8, sdtpldt1(xI, xJ)) | aElementOf0(sdtpldt0(X8, X7), sdtpldt1(xI, xJ))), inference(resolution, [], [f168, f163])).
fof(f163, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3, X4) | aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ! [X0, X1, X2, X3, X4] : ((aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, sK14(X0, X1, X2, X3, X4)), sdtpldt0(X3, sK15(X0, X1, X2, X3, X4)))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, sK15(X0, X1, X2, X3, X4)), xJ) & aElementOf0(sdtpldt0(X4, sK14(X0, X1, X2, X3, X4)), xI) & (sdtpldt0(sK14(X0, X1, X2, X3, X4), sK15(X0, X1, X2, X3, X4)) = X2) & aElementOf0(sK15(X0, X1, X2, X3, X4), xJ) & aElementOf0(sK14(X0, X1, X2, X3, X4), xI)) | ~ sP2(X0, X1, X2, X3, X4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK14, sK15])], [f93, f94])).
fof(f94, plain, ! [X4, X3, X2, X1, X0] : (? [X5, X6] : (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, X5), sdtpldt0(X3, X6))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, X6), xJ) & aElementOf0(sdtpldt0(X4, X5), xI) & (sdtpldt0(X5, X6) = X2) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) => (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, sK14(X0, X1, X2, X3, X4)), sdtpldt0(X3, sK15(X0, X1, X2, X3, X4)))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, sK15(X0, X1, X2, X3, X4)), xJ) & aElementOf0(sdtpldt0(X4, sK14(X0, X1, X2, X3, X4)), xI) & (sdtpldt0(sK14(X0, X1, X2, X3, X4), sK15(X0, X1, X2, X3, X4)) = X2) & aElementOf0(sK15(X0, X1, X2, X3, X4), xJ) & aElementOf0(sK14(X0, X1, X2, X3, X4), xI))), introduced(choice_axiom, [])).
fof(f93, plain, ! [X0, X1, X2, X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, X5), sdtpldt0(X3, X6))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, X6), xJ) & aElementOf0(sdtpldt0(X4, X5), xI) & (sdtpldt0(X5, X6) = X2) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) | ~ sP2(X0, X1, X2, X3, X4)), inference(rectify, [], [f92])).
fof(f92, plain, ! [X0, X2, X1, X4, X3] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) | ~ sP2(X0, X2, X1, X4, X3)), inference(nnf_transformation, [], [f67])).
fof(f168, plain, ! [X2, X0, X1] : (sP2(X0, X2, X1, sK17(X0, X1, X2), sK16(X0, X1, X2)) | ~ aElement0(X2) | ~ aElementOf0(X1, sdtpldt1(xI, xJ)) | ~ aElementOf0(X0, sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f97])).
fof(f715, plain, (spl18_17 | spl18_18), inference(avatar_split_clause, [], [f705, f713, f710])).
fof(f705, plain, ! [X4, X5, X3] : (~ aElement0(X3) | ~ aElementOf0(X4, sdtpldt1(xI, xJ)) | ~ aElementOf0(X5, sdtpldt1(xI, xJ)) | aElementOf0(sdtasdt0(X3, X5), sdtpldt1(xI, xJ))), inference(resolution, [], [f168, f164])).
fof(f164, plain, ! [X4, X2, X0, X3, X1] : (~ sP2(X0, X1, X2, X3, X4) | aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f95])).
fof(f677, plain, spl18_7, inference(avatar_contradiction_clause, [], [f676])).
fof(f676, plain, ($false | spl18_7), inference(subsumption_resolution, [], [f675, f176])).
fof(f176, plain, aSet0(xI), inference(resolution, [], [f145, f153])).
fof(f153, plain, aIdeal0(xI), inference(cnf_transformation, [], [f25])).
fof(f25, plain, (aIdeal0(xJ) & aIdeal0(xI)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+1.p', m__870)).
fof(f145, plain, ! [X0] : (~ aIdeal0(X0) | aSet0(X0)), inference(cnf_transformation, [], [f91])).
fof(f675, plain, (~ aSet0(xI) | spl18_7), inference(subsumption_resolution, [], [f674, f177])).
fof(f177, plain, aSet0(xJ), inference(resolution, [], [f145, f154])).
fof(f154, plain, aIdeal0(xJ), inference(cnf_transformation, [], [f25])).
fof(f674, plain, (~ aSet0(xJ) | ~ aSet0(xI) | spl18_7), inference(resolution, [], [f656, f137])).
fof(f137, plain, ! [X0, X1] : (sP1(X0, X1) | ~ aSet0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f66])).
fof(f66, plain, ! [X0, X1] : (sP1(X0, X1) | ~ aSet0(X1) | ~ aSet0(X0)), inference(definition_folding, [], [f58, e65, e64])).
fof(f64, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0))) & aSet0(X2))), inference(usedef, [], [e64])).
fof(e64, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0))) & aSet0(X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f65, plain, ! [X0, X1] : (! [X2] : ((sdtpldt1(X0, X1) = X2) <=> sP0(X1, X0, X2)) | ~ sP1(X0, X1)), inference(usedef, [], [e65])).
fof(e65, plain, ! [X0, X1] : (sP1(X0, X1) <=> ! [X2] : ((sdtpldt1(X0, X1) = X2) <=> sP0(X1, X0, X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f58, plain, ! [X0, X1] : (! [X2] : ((sdtpldt1(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0))) & aSet0(X2))) | ~ aSet0(X1) | ~ aSet0(X0)), inference(flattening, [], [f57])).
fof(f57, plain, ! [X0, X1] : (! [X2] : ((sdtpldt1(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0))) & aSet0(X2))) | (~ aSet0(X1) | ~ aSet0(X0))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ! [X0, X1] : ((aSet0(X1) & aSet0(X0)) => ! [X2] : ((sdtpldt1(X0, X1) = X2) <=> (! [X3] : (aElementOf0(X3, X2) <=> ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0))) & aSet0(X2)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+1.p', mDefSSum)).
fof(f656, plain, (~ sP1(xI, xJ) | spl18_7), inference(resolution, [], [f625, f254])).
fof(f254, plain, ! [X0, X1] : (aSet0(sdtpldt1(X0, X1)) | ~ sP1(X0, X1)), inference(resolution, [], [f170, f128])).
fof(f128, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | aSet0(X2)), inference(cnf_transformation, [], [f79])).
fof(f79, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ((! [X4, X5] : (~ (sdtpldt0(X4, X5) = sK5(X0, X1, X2)) | ~ aElementOf0(X5, X0) | ~ aElementOf0(X4, X1)) | ~ aElementOf0(sK5(X0, X1, X2), X2)) & (((sK5(X0, X1, X2) = sdtpldt0(sK6(X0, X1, X2), sK7(X0, X1, X2))) & aElementOf0(sK7(X0, X1, X2), X0) & aElementOf0(sK6(X0, X1, X2), X1)) | aElementOf0(sK5(X0, X1, X2), X2))) | ~ aSet0(X2)) & ((! [X8] : ((aElementOf0(X8, X2) | ! [X9, X10] : (~ (sdtpldt0(X9, X10) = X8) | ~ aElementOf0(X10, X0) | ~ aElementOf0(X9, X1))) & (((sdtpldt0(sK8(X0, X1, X8), sK9(X0, X1, X8)) = X8) & aElementOf0(sK9(X0, X1, X8), X0) & aElementOf0(sK8(X0, X1, X8), X1)) | ~ aElementOf0(X8, X2))) & aSet0(X2)) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6, sK7, sK8, sK9])], [f75, f78, f77, f76])).
fof(f76, plain, ! [X2, X1, X0] : (? [X3] : ((! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X0) | ~ aElementOf0(X4, X1)) | ~ aElementOf0(X3, X2)) & (? [X6, X7] : ((sdtpldt0(X6, X7) = X3) & aElementOf0(X7, X0) & aElementOf0(X6, X1)) | aElementOf0(X3, X2))) => ((! [X5, X4] : (~ (sdtpldt0(X4, X5) = sK5(X0, X1, X2)) | ~ aElementOf0(X5, X0) | ~ aElementOf0(X4, X1)) | ~ aElementOf0(sK5(X0, X1, X2), X2)) & (? [X7, X6] : ((sdtpldt0(X6, X7) = sK5(X0, X1, X2)) & aElementOf0(X7, X0) & aElementOf0(X6, X1)) | aElementOf0(sK5(X0, X1, X2), X2)))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X2, X1, X0] : (? [X7, X6] : ((sdtpldt0(X6, X7) = sK5(X0, X1, X2)) & aElementOf0(X7, X0) & aElementOf0(X6, X1)) => ((sK5(X0, X1, X2) = sdtpldt0(sK6(X0, X1, X2), sK7(X0, X1, X2))) & aElementOf0(sK7(X0, X1, X2), X0) & aElementOf0(sK6(X0, X1, X2), X1))), introduced(choice_axiom, [])).
fof(f78, plain, ! [X8, X1, X0] : (? [X11, X12] : ((sdtpldt0(X11, X12) = X8) & aElementOf0(X12, X0) & aElementOf0(X11, X1)) => ((sdtpldt0(sK8(X0, X1, X8), sK9(X0, X1, X8)) = X8) & aElementOf0(sK9(X0, X1, X8), X0) & aElementOf0(sK8(X0, X1, X8), X1))), introduced(choice_axiom, [])).
fof(f75, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ? [X3] : ((! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X0) | ~ aElementOf0(X4, X1)) | ~ aElementOf0(X3, X2)) & (? [X6, X7] : ((sdtpldt0(X6, X7) = X3) & aElementOf0(X7, X0) & aElementOf0(X6, X1)) | aElementOf0(X3, X2))) | ~ aSet0(X2)) & ((! [X8] : ((aElementOf0(X8, X2) | ! [X9, X10] : (~ (sdtpldt0(X9, X10) = X8) | ~ aElementOf0(X10, X0) | ~ aElementOf0(X9, X1))) & (? [X11, X12] : ((sdtpldt0(X11, X12) = X8) & aElementOf0(X12, X0) & aElementOf0(X11, X1)) | ~ aElementOf0(X8, X2))) & aSet0(X2)) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f74])).
fof(f74, plain, ! [X1, X0, X2] : ((sP0(X1, X0, X2) | ? [X3] : ((! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X1) | ~ aElementOf0(X4, X0)) | ~ aElementOf0(X3, X2)) & (? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0)) | aElementOf0(X3, X2))) | ~ aSet0(X2)) & ((! [X3] : ((aElementOf0(X3, X2) | ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X1) | ~ aElementOf0(X4, X0))) & (? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0)) | ~ aElementOf0(X3, X2))) & aSet0(X2)) | ~ sP0(X1, X0, X2))), inference(flattening, [], [f73])).
fof(f73, plain, ! [X1, X0, X2] : ((sP0(X1, X0, X2) | (? [X3] : ((! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X1) | ~ aElementOf0(X4, X0)) | ~ aElementOf0(X3, X2)) & (? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0)) | aElementOf0(X3, X2))) | ~ aSet0(X2))) & ((! [X3] : ((aElementOf0(X3, X2) | ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, X1) | ~ aElementOf0(X4, X0))) & (? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, X1) & aElementOf0(X4, X0)) | ~ aElementOf0(X3, X2))) & aSet0(X2)) | ~ sP0(X1, X0, X2))), inference(nnf_transformation, [], [f64])).
fof(f170, plain, ! [X0, X1] : (sP0(X1, X0, sdtpldt1(X0, X1)) | ~ sP1(X0, X1)), inference(equality_resolution, [], [f126])).
fof(f126, plain, ! [X2, X0, X1] : (sP0(X1, X0, X2) | ~ (sdtpldt1(X0, X1) = X2) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ! [X0, X1] : (! [X2] : (((sdtpldt1(X0, X1) = X2) | ~ sP0(X1, X0, X2)) & (sP0(X1, X0, X2) | ~ (sdtpldt1(X0, X1) = X2))) | ~ sP1(X0, X1)), inference(nnf_transformation, [], [f65])).
fof(f625, plain, (~ aSet0(sdtpldt1(xI, xJ)) | spl18_7), inference(avatar_component_clause, [], [f623])).