fof(f66847, plain, $false, inference(avatar_sat_refutation, [], [f240, f245, f250, f251, f1542, f4015, f16550, f66846])).
fof(f66846, plain, (spl30_2 | ~ spl30_4), inference(avatar_contradiction_clause, [], [f66845])).
fof(f66845, plain, ($false | (spl30_2 | ~ spl30_4)), inference(subsumption_resolution, [], [f66844, f249])).
fof(f249, plain, (aElement0(sK26) | ~ spl30_4), inference(avatar_component_clause, [], [f247])).
fof(f247, plain, (spl30_4 <=> aElement0(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl30_4])])).
fof(f66844, plain, (~ aElement0(sK26) | (spl30_2 | ~ spl30_4)), inference(subsumption_resolution, [], [f66839, f16646])).
fof(f16646, plain, (~ aElementOf0(sdtasdt0(sK25, sK26), sdtpldt1(xI, xJ)) | (spl30_2 | ~ spl30_4)), inference(backward_demodulation, [], [f239, f16618])).
fof(f16618, plain, ((sdtasdt0(sK26, sK25) = sdtasdt0(sK25, sK26)) | ~ spl30_4), inference(resolution, [], [f249, f526])).
fof(f526, plain, ! [X10] : (~ aElement0(X10) | (sdtasdt0(X10, sK25) = sdtasdt0(sK25, X10))), inference(resolution, [], [f132, f274])).
fof(f274, plain, aElement0(sK25), inference(subsumption_resolution, [], [f272, f210])).
fof(f210, plain, aSet0(sdtpldt1(xI, xJ)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & (((~ aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ)) & aElement0(sK26)) | (~ aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ)) & aElementOf0(sK27, sdtpldt1(xI, xJ)))) & aElementOf0(sK25, sdtpldt1(xI, xJ))) & ! [X3] : ((aElementOf0(X3, sdtpldt1(xI, xJ)) | ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, xJ) | ~ aElementOf0(X4, xI))) & (((sdtpldt0(sK28(X3), sK29(X3)) = X3) & aElementOf0(sK29(X3), xJ) & aElementOf0(sK28(X3), xI)) | ~ aElementOf0(X3, sdtpldt1(xI, xJ)))) & aSet0(sdtpldt1(xI, xJ)) & ! [X8, X9, X10] : (sP5(X9, X10, X8) | ~ aElement0(X10) | (~ aElementOf0(X9, sdtpldt1(xI, xJ)) & ! [X11, X12] : (~ (sdtpldt0(X11, X12) = X9) | ~ aElementOf0(X12, xJ) | ~ aElementOf0(X11, xI))) | (~ aElementOf0(X8, sdtpldt1(xI, xJ)) & ! [X13, X14] : (~ (sdtpldt0(X13, X14) = X8) | ~ aElementOf0(X14, xJ) | ~ aElementOf0(X13, xI))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25, sK26, sK27, sK28, sK29])], [f115, f119, f118, f117, f116])).
fof(f116, plain, (? [X0] : ((? [X1] : (~ aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & aElementOf0(X2, sdtpldt1(xI, xJ)))) & aElementOf0(X0, sdtpldt1(xI, xJ))) => ((? [X1] : (~ aElementOf0(sdtasdt0(X1, sK25), sdtpldt1(xI, xJ)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(sK25, X2), sdtpldt1(xI, xJ)) & aElementOf0(X2, sdtpldt1(xI, xJ)))) & aElementOf0(sK25, sdtpldt1(xI, xJ)))), introduced(choice_axiom, [])).
fof(f117, plain, (? [X1] : (~ aElementOf0(sdtasdt0(X1, sK25), sdtpldt1(xI, xJ)) & aElement0(X1)) => (~ aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ)) & aElement0(sK26))), introduced(choice_axiom, [])).
fof(f118, plain, (? [X2] : (~ aElementOf0(sdtpldt0(sK25, X2), sdtpldt1(xI, xJ)) & aElementOf0(X2, sdtpldt1(xI, xJ))) => (~ aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ)) & aElementOf0(sK27, sdtpldt1(xI, xJ)))), introduced(choice_axiom, [])).
fof(f119, plain, ! [X3] : (? [X6, X7] : ((sdtpldt0(X6, X7) = X3) & aElementOf0(X7, xJ) & aElementOf0(X6, xI)) => ((sdtpldt0(sK28(X3), sK29(X3)) = X3) & aElementOf0(sK29(X3), xJ) & aElementOf0(sK28(X3), xI))), introduced(choice_axiom, [])).
fof(f115, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ? [X0] : ((? [X1] : (~ aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & aElementOf0(X2, sdtpldt1(xI, xJ)))) & aElementOf0(X0, sdtpldt1(xI, xJ))) & ! [X3] : ((aElementOf0(X3, sdtpldt1(xI, xJ)) | ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, xJ) | ~ aElementOf0(X4, xI))) & (? [X6, X7] : ((sdtpldt0(X6, X7) = X3) & aElementOf0(X7, xJ) & aElementOf0(X6, xI)) | ~ aElementOf0(X3, sdtpldt1(xI, xJ)))) & aSet0(sdtpldt1(xI, xJ)) & ! [X8, X9, X10] : (sP5(X9, X10, X8) | ~ aElement0(X10) | (~ aElementOf0(X9, sdtpldt1(xI, xJ)) & ! [X11, X12] : (~ (sdtpldt0(X11, X12) = X9) | ~ aElementOf0(X12, xJ) | ~ aElementOf0(X11, xI))) | (~ aElementOf0(X8, sdtpldt1(xI, xJ)) & ! [X13, X14] : (~ (sdtpldt0(X13, X14) = X8) | ~ aElementOf0(X14, xJ) | ~ aElementOf0(X13, xI))))), inference(rectify, [], [f114])).
fof(f114, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ? [X18] : ((? [X19] : (~ aElementOf0(sdtasdt0(X19, X18), sdtpldt1(xI, xJ)) & aElement0(X19)) | ? [X20] : (~ aElementOf0(sdtpldt0(X18, X20), sdtpldt1(xI, xJ)) & aElementOf0(X20, sdtpldt1(xI, xJ)))) & aElementOf0(X18, sdtpldt1(xI, xJ))) & ! [X15] : ((aElementOf0(X15, sdtpldt1(xI, xJ)) | ! [X16, X17] : (~ (sdtpldt0(X16, X17) = X15) | ~ aElementOf0(X17, xJ) | ~ aElementOf0(X16, xI))) & (? [X16, X17] : ((sdtpldt0(X16, X17) = X15) & aElementOf0(X17, xJ) & aElementOf0(X16, xI)) | ~ aElementOf0(X15, sdtpldt1(xI, xJ)))) & aSet0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (sP5(X1, X2, X0) | ~ aElement0(X2) | (~ aElementOf0(X1, sdtpldt1(xI, xJ)) & ! [X3, X4] : (~ (sdtpldt0(X3, X4) = X1) | ~ aElementOf0(X4, xJ) | ~ aElementOf0(X3, xI))) | (~ aElementOf0(X0, sdtpldt1(xI, xJ)) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X0) | ~ aElementOf0(X6, xJ) | ~ aElementOf0(X5, xI))))), inference(nnf_transformation, [], [f74])).
fof(f74, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ? [X18] : ((? [X19] : (~ aElementOf0(sdtasdt0(X19, X18), sdtpldt1(xI, xJ)) & aElement0(X19)) | ? [X20] : (~ aElementOf0(sdtpldt0(X18, X20), sdtpldt1(xI, xJ)) & aElementOf0(X20, sdtpldt1(xI, xJ)))) & aElementOf0(X18, sdtpldt1(xI, xJ))) & ! [X15] : (aElementOf0(X15, sdtpldt1(xI, xJ)) <=> ? [X16, X17] : ((sdtpldt0(X16, X17) = X15) & aElementOf0(X17, xJ) & aElementOf0(X16, xI))) & aSet0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (sP5(X1, X2, X0) | ~ aElement0(X2) | (~ aElementOf0(X1, sdtpldt1(xI, xJ)) & ! [X3, X4] : (~ (sdtpldt0(X3, X4) = X1) | ~ aElementOf0(X4, xJ) | ~ aElementOf0(X3, xI))) | (~ aElementOf0(X0, sdtpldt1(xI, xJ)) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X0) | ~ aElementOf0(X6, xJ) | ~ aElementOf0(X5, xI))))), inference(definition_folding, [], [f66, e73, e72, e71, e70])).
fof(f70, plain, ! [X1, X0] : (? [X13, X14] : ((sdtpldt0(X0, X1) = sdtpldt0(X13, X14)) & aElementOf0(X14, xJ) & aElementOf0(X13, xI)) | ~ sP2(X1, X0)), inference(usedef, [], [e70])).
fof(e70, plain, ! [X1, X0] : (sP2(X1, X0) <=> ? [X13, X14] : ((sdtpldt0(X0, X1) = sdtpldt0(X13, X14)) & aElementOf0(X14, xJ) & aElementOf0(X13, xI))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f71, plain, ! [X0, X2] : (? [X11, X12] : ((sdtasdt0(X2, X0) = sdtpldt0(X11, X12)) & aElementOf0(X12, xJ) & aElementOf0(X11, xI)) | ~ sP3(X0, X2)), inference(usedef, [], [e71])).
fof(e71, plain, ! [X0, X2] : (sP3(X0, X2) <=> ? [X11, X12] : ((sdtasdt0(X2, X0) = sdtpldt0(X11, X12)) & aElementOf0(X12, xJ) & aElementOf0(X11, xI))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f72, plain, ! [X0, X2, X1, X8, X7] : (? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & sP2(X1, X0) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI)) | ~ sP4(X0, X2, X1, X8, X7)), inference(usedef, [], [e72])).
fof(e72, plain, ! [X0, X2, X1, X8, X7] : (sP4(X0, X2, X1, X8, X7) <=> ? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & sP2(X1, X0) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f73, plain, ! [X1, X2, X0] : (? [X7, X8] : (sP4(X0, X2, X1, X8, X7) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) | ~ sP5(X1, X2, X0)), inference(usedef, [], [e73])).
fof(e73, plain, ! [X1, X2, X0] : (sP5(X1, X2, X0) <=> ? [X7, X8] : (sP4(X0, X2, X1, X8, X7) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f66, plain, (~ aIdeal0(sdtpldt1(xI, xJ)) & ? [X18] : ((? [X19] : (~ aElementOf0(sdtasdt0(X19, X18), sdtpldt1(xI, xJ)) & aElement0(X19)) | ? [X20] : (~ aElementOf0(sdtpldt0(X18, X20), sdtpldt1(xI, xJ)) & aElementOf0(X20, sdtpldt1(xI, xJ)))) & aElementOf0(X18, sdtpldt1(xI, xJ))) & ! [X15] : (aElementOf0(X15, sdtpldt1(xI, xJ)) <=> ? [X16, X17] : ((sdtpldt0(X16, X17) = X15) & aElementOf0(X17, xJ) & aElementOf0(X16, xI))) & aSet0(sdtpldt1(xI, xJ)) & ! [X0, X1, X2] : (? [X7, X8] : (? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & ? [X11, X12] : ((sdtasdt0(X2, X0) = sdtpldt0(X11, X12)) & aElementOf0(X12, xJ) & aElementOf0(X11, xI)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & ? [X13, X14] : ((sdtpldt0(X0, X1) = sdtpldt0(X13, X14)) & aElementOf0(X14, xJ) & aElementOf0(X13, xI)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI)) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) | ~ aElement0(X2) | (~ aElementOf0(X1, sdtpldt1(xI, xJ)) & ! [X3, X4] : (~ (sdtpldt0(X3, X4) = X1) | ~ aElementOf0(X4, xJ) | ~ aElementOf0(X3, xI))) | (~ aElementOf0(X0, sdtpldt1(xI, xJ)) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X0) | ~ aElementOf0(X6, xJ) | ~ aElementOf0(X5, xI))))), inference(flattening, [], [f65])).
fof(f65, plain, (((~ aIdeal0(sdtpldt1(xI, xJ)) & ? [X18] : ((? [X19] : (~ aElementOf0(sdtasdt0(X19, X18), sdtpldt1(xI, xJ)) & aElement0(X19)) | ? [X20] : (~ aElementOf0(sdtpldt0(X18, X20), sdtpldt1(xI, xJ)) & aElementOf0(X20, sdtpldt1(xI, xJ)))) & aElementOf0(X18, sdtpldt1(xI, xJ)))) & (! [X15] : (aElementOf0(X15, sdtpldt1(xI, xJ)) <=> ? [X16, X17] : ((sdtpldt0(X16, X17) = X15) & aElementOf0(X17, xJ) & aElementOf0(X16, xI))) & aSet0(sdtpldt1(xI, xJ)))) & ! [X0, X1, X2] : (? [X7, X8] : (? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & ? [X11, X12] : ((sdtasdt0(X2, X0) = sdtpldt0(X11, X12)) & aElementOf0(X12, xJ) & aElementOf0(X11, xI)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & ? [X13, X14] : ((sdtpldt0(X0, X1) = sdtpldt0(X13, X14)) & aElementOf0(X14, xJ) & aElementOf0(X13, xI)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI)) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) | (~ aElement0(X2) | (~ aElementOf0(X1, sdtpldt1(xI, xJ)) & ! [X3, X4] : (~ (sdtpldt0(X3, X4) = X1) | ~ aElementOf0(X4, xJ) | ~ aElementOf0(X3, xI))) | (~ aElementOf0(X0, sdtpldt1(xI, xJ)) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X0) | ~ aElementOf0(X6, xJ) | ~ aElementOf0(X5, xI)))))), inference(ennf_transformation, [], [f33])).
fof(f33, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, sdtpldt1(xI, xJ)) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X1) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) & (aElementOf0(X0, sdtpldt1(xI, xJ)) | ? [X5, X6] : ((sdtpldt0(X5, X6) = X0) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)))) => ? [X7, X8] : (? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & ? [X11, X12] : ((sdtasdt0(X2, X0) = sdtpldt0(X11, X12)) & aElementOf0(X12, xJ) & aElementOf0(X11, xI)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & ? [X13, X14] : ((sdtpldt0(X0, X1) = sdtpldt0(X13, X14)) & aElementOf0(X14, xJ) & aElementOf0(X13, xI)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI)) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI))) => ((! [X15] : (aElementOf0(X15, sdtpldt1(xI, xJ)) <=> ? [X16, X17] : ((sdtpldt0(X16, X17) = X15) & aElementOf0(X17, xJ) & aElementOf0(X16, xI))) & aSet0(sdtpldt1(xI, xJ))) => (aIdeal0(sdtpldt1(xI, xJ)) | ! [X18] : (aElementOf0(X18, sdtpldt1(xI, xJ)) => (! [X19] : (aElement0(X19) => aElementOf0(sdtasdt0(X19, X18), sdtpldt1(xI, xJ))) & ! [X20] : (aElementOf0(X20, sdtpldt1(xI, xJ)) => aElementOf0(sdtpldt0(X18, X20), sdtpldt1(xI, xJ)))))))), inference(rectify, [], [f27])).
fof(f27, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, sdtpldt1(xI, xJ)) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X1) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) & (aElementOf0(X0, sdtpldt1(xI, xJ)) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)))) => ? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & ? [X7, X8] : ((sdtasdt0(X2, X0) = sdtpldt0(X7, X8)) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & ? [X7, X8] : ((sdtpldt0(X0, X1) = sdtpldt0(X7, X8)) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) => ((! [X0] : (aElementOf0(X0, sdtpldt1(xI, xJ)) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, xJ) & aElementOf0(X1, xI))) & aSet0(sdtpldt1(xI, xJ))) => (aIdeal0(sdtpldt1(xI, xJ)) | ! [X0] : (aElementOf0(X0, sdtpldt1(xI, xJ)) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ))) & ! [X1] : (aElementOf0(X1, sdtpldt1(xI, xJ)) => aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)))))))), inference(negated_conjecture, [], [f26])).
fof(f26, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, sdtpldt1(xI, xJ)) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X1) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) & (aElementOf0(X0, sdtpldt1(xI, xJ)) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)))) => ? [X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & ? [X7, X8] : ((sdtasdt0(X2, X0) = sdtpldt0(X7, X8)) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & ? [X7, X8] : ((sdtpldt0(X0, X1) = sdtpldt0(X7, X8)) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X3, X5), sdtpldt0(X4, X6))) & aElementOf0(sdtasdt0(X2, X4), xJ) & aElementOf0(sdtasdt0(X2, X3), xI) & aElementOf0(sdtpldt0(X4, X6), xJ) & aElementOf0(sdtpldt0(X3, X5), xI) & (sdtpldt0(X5, X6) = X1) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) & (sdtpldt0(X3, X4) = X0) & aElementOf0(X4, xJ) & aElementOf0(X3, xI))) => ((! [X0] : (aElementOf0(X0, sdtpldt1(xI, xJ)) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, xJ) & aElementOf0(X1, xI))) & aSet0(sdtpldt1(xI, xJ))) => (aIdeal0(sdtpldt1(xI, xJ)) | ! [X0] : (aElementOf0(X0, sdtpldt1(xI, xJ)) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ))) & ! [X1] : (aElementOf0(X1, sdtpldt1(xI, xJ)) => aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)))))))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+2.p', m__)).
fof(f272, plain, (aElement0(sK25) | ~ aSet0(sdtpldt1(xI, xJ))), inference(resolution, [], [f144, f215])).
fof(f215, plain, aElementOf0(sK25, sdtpldt1(xI, xJ)), inference(cnf_transformation, [], [f120])).
fof(f144, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f56])).
fof(f56, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+2.p', mEOfElem)).
fof(f132, plain, ! [X0, X1] : (~ aElement0(X1) | (sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aElement0(X0)), inference(cnf_transformation, [], [f46])).
fof(f46, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f45])).
fof(f45, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | (~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => (sdtasdt0(X0, X1) = sdtasdt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+2.p', mMulComm)).
fof(f239, plain, (~ aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ)) | spl30_2), inference(avatar_component_clause, [], [f237])).
fof(f237, plain, (spl30_2 <=> aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl30_2])])).
fof(f66839, plain, (aElementOf0(sdtasdt0(sK25, sK26), sdtpldt1(xI, xJ)) | ~ aElement0(sK26) | ~ spl30_4), inference(superposition, [], [f3123, f16618])).
fof(f3123, plain, ! [X13] : (aElementOf0(sdtasdt0(X13, sK25), sdtpldt1(xI, xJ)) | ~ aElement0(X13)), inference(resolution, [], [f814, f2460])).
fof(f2460, plain, ! [X4] : (sP5(sK25, X4, sK25) | ~ aElement0(X4)), inference(resolution, [], [f800, f215])).
fof(f800, plain, ! [X0, X1] : (~ aElementOf0(X1, sdtpldt1(xI, xJ)) | sP5(sK25, X0, X1) | ~ aElement0(X0)), inference(resolution, [], [f209, f215])).
fof(f209, plain, ! [X10, X8, X9] : (~ aElementOf0(X9, sdtpldt1(xI, xJ)) | ~ aElement0(X10) | sP5(X9, X10, X8) | ~ aElementOf0(X8, sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f120])).
fof(f814, plain, ! [X2, X0, X1] : (~ sP5(X0, X1, X2) | aElementOf0(sdtasdt0(X1, X2), sdtpldt1(xI, xJ))), inference(resolution, [], [f187, f199])).
fof(f199, plain, ! [X4, X2, X0, X3, X1] : (~ sP4(X0, X1, X2, X3, X4) | aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f105])).
fof(f105, plain, ! [X0, X1, X2, X3, X4] : ((aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & sP2(X2, X0) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, sK19(X0, X1, X2, X3, X4)), sdtpldt0(X3, sK20(X0, X1, X2, X3, X4)))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, sK20(X0, X1, X2, X3, X4)), xJ) & aElementOf0(sdtpldt0(X4, sK19(X0, X1, X2, X3, X4)), xI) & (sdtpldt0(sK19(X0, X1, X2, X3, X4), sK20(X0, X1, X2, X3, X4)) = X2) & aElementOf0(sK20(X0, X1, X2, X3, X4), xJ) & aElementOf0(sK19(X0, X1, X2, X3, X4), xI)) | ~ sP4(X0, X1, X2, X3, X4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19, sK20])], [f103, f104])).
fof(f104, plain, ! [X4, X3, X2, X1, X0] : (? [X5, X6] : (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & sP2(X2, X0) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, X5), sdtpldt0(X3, X6))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, X6), xJ) & aElementOf0(sdtpldt0(X4, X5), xI) & (sdtpldt0(X5, X6) = X2) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) => (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & sP2(X2, X0) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, sK19(X0, X1, X2, X3, X4)), sdtpldt0(X3, sK20(X0, X1, X2, X3, X4)))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, sK20(X0, X1, X2, X3, X4)), xJ) & aElementOf0(sdtpldt0(X4, sK19(X0, X1, X2, X3, X4)), xI) & (sdtpldt0(sK19(X0, X1, X2, X3, X4), sK20(X0, X1, X2, X3, X4)) = X2) & aElementOf0(sK20(X0, X1, X2, X3, X4), xJ) & aElementOf0(sK19(X0, X1, X2, X3, X4), xI))), introduced(choice_axiom, [])).
fof(f103, plain, ! [X0, X1, X2, X3, X4] : (? [X5, X6] : (aElementOf0(sdtasdt0(X1, X0), sdtpldt1(xI, xJ)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ)) & sP2(X2, X0) & (sdtpldt0(X0, X2) = sdtpldt0(sdtpldt0(X4, X5), sdtpldt0(X3, X6))) & aElementOf0(sdtasdt0(X1, X3), xJ) & aElementOf0(sdtasdt0(X1, X4), xI) & aElementOf0(sdtpldt0(X3, X6), xJ) & aElementOf0(sdtpldt0(X4, X5), xI) & (sdtpldt0(X5, X6) = X2) & aElementOf0(X6, xJ) & aElementOf0(X5, xI)) | ~ sP4(X0, X1, X2, X3, X4)), inference(rectify, [], [f102])).
fof(f102, plain, ! [X0, X2, X1, X8, X7] : (? [X9, X10] : (aElementOf0(sdtasdt0(X2, X0), sdtpldt1(xI, xJ)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), sdtpldt1(xI, xJ)) & sP2(X1, X0) & (sdtpldt0(X0, X1) = sdtpldt0(sdtpldt0(X7, X9), sdtpldt0(X8, X10))) & aElementOf0(sdtasdt0(X2, X8), xJ) & aElementOf0(sdtasdt0(X2, X7), xI) & aElementOf0(sdtpldt0(X8, X10), xJ) & aElementOf0(sdtpldt0(X7, X9), xI) & (sdtpldt0(X9, X10) = X1) & aElementOf0(X10, xJ) & aElementOf0(X9, xI)) | ~ sP4(X0, X2, X1, X8, X7)), inference(nnf_transformation, [], [f72])).
fof(f187, plain, ! [X2, X0, X1] : (sP4(X2, X1, X0, sK18(X0, X1, X2), sK17(X0, X1, X2)) | ~ sP5(X0, X1, X2)), inference(cnf_transformation, [], [f101])).
fof(f101, plain, ! [X0, X1, X2] : ((sP4(X2, X1, X0, sK18(X0, X1, X2), sK17(X0, X1, X2)) & (sdtpldt0(sK17(X0, X1, X2), sK18(X0, X1, X2)) = X2) & aElementOf0(sK18(X0, X1, X2), xJ) & aElementOf0(sK17(X0, X1, X2), xI)) | ~ sP5(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18])], [f99, f100])).
fof(f100, plain, ! [X2, X1, X0] : (? [X3, X4] : (sP4(X2, X1, X0, X4, X3) & (sdtpldt0(X3, X4) = X2) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) => (sP4(X2, X1, X0, sK18(X0, X1, X2), sK17(X0, X1, X2)) & (sdtpldt0(sK17(X0, X1, X2), sK18(X0, X1, X2)) = X2) & aElementOf0(sK18(X0, X1, X2), xJ) & aElementOf0(sK17(X0, X1, X2), xI))), introduced(choice_axiom, [])).
fof(f99, plain, ! [X0, X1, X2] : (? [X3, X4] : (sP4(X2, X1, X0, X4, X3) & (sdtpldt0(X3, X4) = X2) & aElementOf0(X4, xJ) & aElementOf0(X3, xI)) | ~ sP5(X0, X1, X2)), inference(rectify, [], [f98])).
fof(f98, plain, ! [X1, X2, X0] : (? [X7, X8] : (sP4(X0, X2, X1, X8, X7) & (sdtpldt0(X7, X8) = X0) & aElementOf0(X8, xJ) & aElementOf0(X7, xI)) | ~ sP5(X1, X2, X0)), inference(nnf_transformation, [], [f73])).
fof(f16550, plain, (spl30_1 | ~ spl30_3 | ~ spl30_85), inference(avatar_contradiction_clause, [], [f16549])).
fof(f16549, plain, ($false | (spl30_1 | ~ spl30_3 | ~ spl30_85)), inference(subsumption_resolution, [], [f16547, f235])).
fof(f235, plain, (~ aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ)) | spl30_1), inference(avatar_component_clause, [], [f233])).
fof(f233, plain, (spl30_1 <=> aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl30_1])])).
fof(f16547, plain, (aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ)) | (~ spl30_3 | ~ spl30_85)), inference(backward_demodulation, [], [f4014, f16515])).
fof(f16515, plain, ((sdtpldt0(sK25, sK27) = sdtpldt0(sK27, sK25)) | ~ spl30_3), inference(resolution, [], [f16463, f515])).
fof(f515, plain, ! [X10] : (~ aElement0(X10) | (sdtpldt0(X10, sK25) = sdtpldt0(sK25, X10))), inference(resolution, [], [f126, f274])).
fof(f126, plain, ! [X0, X1] : (~ aElement0(X1) | (sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aElement0(X0)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f39])).
fof(f39, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | (~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => (sdtpldt0(X0, X1) = sdtpldt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+2.p', mAddComm)).
fof(f16463, plain, (aElement0(sK27) | ~ spl30_3), inference(subsumption_resolution, [], [f16461, f210])).
fof(f16461, plain, (aElement0(sK27) | ~ aSet0(sdtpldt1(xI, xJ)) | ~ spl30_3), inference(resolution, [], [f244, f144])).
fof(f244, plain, (aElementOf0(sK27, sdtpldt1(xI, xJ)) | ~ spl30_3), inference(avatar_component_clause, [], [f242])).
fof(f242, plain, (spl30_3 <=> aElementOf0(sK27, sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl30_3])])).
fof(f4014, plain, (aElementOf0(sdtpldt0(sK27, sK25), sdtpldt1(xI, xJ)) | ~ spl30_85), inference(avatar_component_clause, [], [f4012])).
fof(f4012, plain, (spl30_85 <=> aElementOf0(sdtpldt0(sK27, sK25), sdtpldt1(xI, xJ))), introduced(avatar_definition, [new_symbols(naming, [spl30_85])])).
fof(f4015, plain, (spl30_85 | spl30_28 | ~ spl30_3), inference(avatar_split_clause, [], [f4006, f242, f1478, f4012])).
fof(f1478, plain, (spl30_28 <=> ! [X12] : ~ aElement0(X12)), introduced(avatar_definition, [new_symbols(naming, [spl30_28])])).
fof(f4006, plain, (! [X0] : (~ aElement0(X0) | aElementOf0(sdtpldt0(sK27, sK25), sdtpldt1(xI, xJ))) | ~ spl30_3), inference(resolution, [], [f3523, f815])).
fof(f815, plain, ! [X4, X5, X3] : (~ sP5(X3, X4, X5) | aElementOf0(sdtpldt0(X5, X3), sdtpldt1(xI, xJ))), inference(resolution, [], [f187, f197])).
fof(f197, plain, ! [X4, X2, X0, X3, X1] : (~ sP4(X0, X1, X2, X3, X4) | aElementOf0(sdtpldt0(X0, X2), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f105])).
fof(f3523, plain, (! [X0] : (sP5(sK25, X0, sK27) | ~ aElement0(X0)) | ~ spl30_3), inference(resolution, [], [f244, f800])).
fof(f1542, plain, ~ spl30_28, inference(avatar_contradiction_clause, [], [f1507])).
fof(f1507, plain, ($false | ~ spl30_28), inference(resolution, [], [f1479, f121])).
fof(f121, plain, aElement0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aElement0(sz00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG092+2.p', mSortsC)).
fof(f1479, plain, (! [X12] : ~ aElement0(X12) | ~ spl30_28), inference(avatar_component_clause, [], [f1478])).
fof(f251, plain, (spl30_3 | spl30_4), inference(avatar_split_clause, [], [f216, f247, f242])).
fof(f216, plain, (aElement0(sK26) | aElementOf0(sK27, sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f120])).
fof(f250, plain, (~ spl30_1 | spl30_4), inference(avatar_split_clause, [], [f217, f247, f233])).
fof(f217, plain, (aElement0(sK26) | ~ aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f120])).
fof(f245, plain, (spl30_3 | ~ spl30_2), inference(avatar_split_clause, [], [f218, f237, f242])).
fof(f218, plain, (~ aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ)) | aElementOf0(sK27, sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f120])).
fof(f240, plain, (~ spl30_1 | ~ spl30_2), inference(avatar_split_clause, [], [f219, f237, f233])).
fof(f219, plain, (~ aElementOf0(sdtasdt0(sK26, sK25), sdtpldt1(xI, xJ)) | ~ aElementOf0(sdtpldt0(sK25, sK27), sdtpldt1(xI, xJ))), inference(cnf_transformation, [], [f120])).