fof(f2565, plain, $false, inference(avatar_sat_refutation, [], [f272, f273, f282, f1031, f1928, f2050, f2176, f2551])).
fof(f2551, plain, (spl21_6 | ~ spl21_7 | ~ spl21_28 | ~ spl21_48 | ~ spl21_56), inference(avatar_contradiction_clause, [], [f2550])).
fof(f2550, plain, ($false | (spl21_6 | ~ spl21_7 | ~ spl21_28 | ~ spl21_48 | ~ spl21_56)), inference(subsumption_resolution, [], [f2549, f1017])).
fof(f1017, plain, (sdtlseqdt0(sz0z00, sz0z00) | ~ spl21_28), inference(avatar_component_clause, [], [f1016])).
fof(f1016, plain, (spl21_28 <=> sdtlseqdt0(sz0z00, sz0z00)), introduced(avatar_definition, [new_symbols(naming, [spl21_28])])).
fof(f2549, plain, (~ sdtlseqdt0(sz0z00, sz0z00) | (spl21_6 | ~ spl21_7 | ~ spl21_48 | ~ spl21_56)), inference(forward_demodulation, [], [f2491, f316])).
fof(f316, plain, (sz0z00 = sdtasdt0(sz0z00, sz0z00)), inference(resolution, [], [f160, f154])).
fof(f154, plain, aScalar0(sz0z00), inference(cnf_transformation, [], [f9])).
fof(f9, plain, aScalar0(sz0z00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', mSZeroSc)).
fof(f160, plain, ! [X0] : (~ aScalar0(X0) | (sz0z00 = sdtasdt0(X0, sz0z00))), inference(cnf_transformation, [], [f60])).
fof(f60, plain, ! [X0] : (((sz0z00 = smndt0(sz0z00)) & (smndt0(smndt0(X0)) = X0) & (sz0z00 = sdtpldt0(smndt0(X0), X0)) & (sz0z00 = sdtpldt0(X0, smndt0(X0))) & (sz0z00 = sdtasdt0(sz0z00, X0)) & (sz0z00 = sdtasdt0(X0, sz0z00)) & (sdtpldt0(sz0z00, X0) = X0) & (sdtpldt0(X0, sz0z00) = X0)) | ~ aScalar0(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (aScalar0(X0) => ((sz0z00 = smndt0(sz0z00)) & (smndt0(smndt0(X0)) = X0) & (sz0z00 = sdtpldt0(smndt0(X0), X0)) & (sz0z00 = sdtpldt0(X0, smndt0(X0))) & (sz0z00 = sdtasdt0(sz0z00, X0)) & (sz0z00 = sdtasdt0(X0, sz0z00)) & (sdtpldt0(sz0z00, X0) = X0) & (sdtpldt0(X0, sz0z00) = X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', mScZero)).
fof(f2491, plain, (~ sdtlseqdt0(sdtasdt0(sz0z00, sz0z00), sz0z00) | (spl21_6 | ~ spl21_7 | ~ spl21_48 | ~ spl21_56)), inference(backward_demodulation, [], [f2107, f2463])).
fof(f2463, plain, ((sz0z00 = sdtasasdt0(xs, xt)) | (~ spl21_7 | ~ spl21_56)), inference(subsumption_resolution, [], [f2461, f199])).
fof(f199, plain, aVector0(xt), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (aVector0(xt) & aVector0(xs)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', m__1678)).
fof(f2461, plain, ((sz0z00 = sdtasasdt0(xs, xt)) | ~ aVector0(xt) | (~ spl21_7 | ~ spl21_56)), inference(trivial_inequality_removal, [], [f2460])).
fof(f2460, plain, (~ (sz00 = sz00) | (sz0z00 = sdtasasdt0(xs, xt)) | ~ aVector0(xt) | (~ spl21_7 | ~ spl21_56)), inference(superposition, [], [f2199, f1460])).
fof(f1460, plain, ((sz00 = aDimensionOf0(xt)) | ~ spl21_56), inference(avatar_component_clause, [], [f1458])).
fof(f1458, plain, (spl21_56 <=> (sz00 = aDimensionOf0(xt))), introduced(avatar_definition, [new_symbols(naming, [spl21_56])])).
fof(f2199, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6)) | ~ spl21_7), inference(subsumption_resolution, [], [f2195, f198])).
fof(f198, plain, aVector0(xs), inference(cnf_transformation, [], [f38])).
fof(f2195, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl21_7), inference(duplicate_literal_removal, [], [f2190])).
fof(f2190, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | ~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl21_7), inference(superposition, [], [f195, f277])).
fof(f277, plain, ((sz00 = aDimensionOf0(xs)) | ~ spl21_7), inference(avatar_component_clause, [], [f275])).
fof(f275, plain, (spl21_7 <=> (sz00 = aDimensionOf0(xs))), introduced(avatar_definition, [new_symbols(naming, [spl21_7])])).
fof(f195, plain, ! [X0, X1] : (~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ (sz00 = aDimensionOf0(X1)) | (sz0z00 = sdtasasdt0(X0, X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(cnf_transformation, [], [f97])).
fof(f97, plain, ! [X0, X1] : ((sz0z00 = sdtasasdt0(X0, X1)) | ~ (sz00 = aDimensionOf0(X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(flattening, [], [f96])).
fof(f96, plain, ! [X0, X1] : (((sz0z00 = sdtasasdt0(X0, X1)) | (~ (sz00 = aDimensionOf0(X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)))) | (~ aVector0(X1) | ~ aVector0(X0))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0, X1] : ((aVector0(X1) & aVector0(X0)) => (((sz00 = aDimensionOf0(X1)) & (aDimensionOf0(X0) = aDimensionOf0(X1))) => (sz0z00 = sdtasasdt0(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', mDefSPZ)).
fof(f2107, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sz0z00) | (spl21_6 | ~ spl21_48)), inference(forward_demodulation, [], [f2051, f888])).
fof(f888, plain, (sz0z00 = sdtasdt0(sz0z00, sdtasasdt0(xt, xt))), inference(resolution, [], [f621, f199])).
fof(f621, plain, ! [X3] : (~ aVector0(X3) | (sz0z00 = sdtasdt0(sz0z00, sdtasasdt0(X3, X3)))), inference(resolution, [], [f567, f161])).
fof(f161, plain, ! [X0] : (~ aScalar0(X0) | (sz0z00 = sdtasdt0(sz0z00, X0))), inference(cnf_transformation, [], [f60])).
fof(f567, plain, ! [X0] : (aScalar0(sdtasasdt0(X0, X0)) | ~ aVector0(X0)), inference(duplicate_literal_removal, [], [f566])).
fof(f566, plain, ! [X0] : (aScalar0(sdtasasdt0(X0, X0)) | ~ aVector0(X0) | ~ aVector0(X0)), inference(equality_resolution, [], [f194])).
fof(f194, plain, ! [X0, X1] : (~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | aScalar0(sdtasasdt0(X0, X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(cnf_transformation, [], [f95])).
fof(f95, plain, ! [X0, X1] : (aScalar0(sdtasasdt0(X0, X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(flattening, [], [f94])).
fof(f94, plain, ! [X0, X1] : ((aScalar0(sdtasasdt0(X0, X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1))) | (~ aVector0(X1) | ~ aVector0(X0))), inference(ennf_transformation, [], [f34])).
fof(f34, plain, ! [X0, X1] : ((aVector0(X1) & aVector0(X0)) => ((aDimensionOf0(X0) = aDimensionOf0(X1)) => aScalar0(sdtasasdt0(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', mScPr)).
fof(f2051, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sz0z00, sdtasasdt0(xt, xt))) | (spl21_6 | ~ spl21_48)), inference(backward_demodulation, [], [f270, f1308])).
fof(f1308, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ spl21_48), inference(avatar_component_clause, [], [f1306])).
fof(f1306, plain, (spl21_48 <=> (sz0z00 = sdtasasdt0(xs, xs))), introduced(avatar_definition, [new_symbols(naming, [spl21_48])])).
fof(f270, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | spl21_6), inference(avatar_component_clause, [], [f269])).
fof(f269, plain, (spl21_6 <=> sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), introduced(avatar_definition, [new_symbols(naming, [spl21_6])])).
fof(f2176, plain, (~ spl21_5 | ~ spl21_8), inference(avatar_contradiction_clause, [], [f2175])).
fof(f2175, plain, ($false | (~ spl21_5 | ~ spl21_8)), inference(subsumption_resolution, [], [f281, f1576])).
fof(f1576, plain, (! [X0] : ~ sP4(X0) | ~ spl21_5), inference(resolution, [], [f1567, f208])).
fof(f208, plain, ! [X0] : (sP3(X0, sK7(X0), sK8(X0)) | ~ sP4(X0)), inference(cnf_transformation, [], [f121])).
fof(f121, plain, ! [X0] : (((sP3(X0, sK7(X0), sK8(X0)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = sK8(X0)) & aScalar0(sK8(X0))) & (sziznziztdt0(xt) = sK7(X0)) & ! [X3] : ((sdtlbdtrb0(xt, X3) = sdtlbdtrb0(sK7(X0), X3)) | ~ aNaturalNumber0(X3)) & (aDimensionOf0(xt) = szszuzczcdt0(aDimensionOf0(sK7(X0)))) & aVector0(sK7(X0))) | ~ sP4(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f118, f120, f119])).
fof(f119, plain, ! [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X3] : ((sdtlbdtrb0(X1, X3) = sdtlbdtrb0(xt, X3)) | ~ aNaturalNumber0(X3)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) => (? [X2] : (sP3(X0, sK7(X0), X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = sK7(X0)) & ! [X3] : ((sdtlbdtrb0(xt, X3) = sdtlbdtrb0(sK7(X0), X3)) | ~ aNaturalNumber0(X3)) & (aDimensionOf0(xt) = szszuzczcdt0(aDimensionOf0(sK7(X0)))) & aVector0(sK7(X0)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (? [X2] : (sP3(X0, sK7(X0), X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) => (sP3(X0, sK7(X0), sK8(X0)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = sK8(X0)) & aScalar0(sK8(X0)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X3] : ((sdtlbdtrb0(X1, X3) = sdtlbdtrb0(xt, X3)) | ~ aNaturalNumber0(X3)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) | ~ sP4(X0)), inference(rectify, [], [f117])).
fof(f117, plain, ! [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X14] : ((sdtlbdtrb0(X1, X14) = sdtlbdtrb0(xt, X14)) | ~ aNaturalNumber0(X14)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) | ~ sP4(X0)), inference(nnf_transformation, [], [f108])).
fof(f108, plain, ! [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X14] : ((sdtlbdtrb0(X1, X14) = sdtlbdtrb0(xt, X14)) | ~ aNaturalNumber0(X14)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) | ~ sP4(X0)), inference(usedef, [], [e108])).
fof(e108, plain, ! [X0] : (sP4(X0) <=> ? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X14] : ((sdtlbdtrb0(X1, X14) = sdtlbdtrb0(xt, X14)) | ~ aNaturalNumber0(X14)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f1567, plain, (! [X2, X0, X1] : ~ sP3(X0, X1, X2) | ~ spl21_5), inference(resolution, [], [f1348, f215])).
fof(f215, plain, ! [X2, X0, X1] : (sP2(X2, sK9(X0, X1, X2), sK11(X0, X1, X2), sK10(X0, X1, X2), X1, X0) | ~ sP3(X0, X1, X2)), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ! [X0, X1, X2] : ((((sP2(X2, sK9(X0, X1, X2), sK11(X0, X1, X2), sK10(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = sK11(X0, X1, X2)) & aScalar0(sK11(X0, X1, X2))) & (sdtasasdt0(X0, X0) = sK10(X0, X1, X2)) & aScalar0(sK10(X0, X1, X2))) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = sK9(X0, X1, X2)) & aScalar0(sK9(X0, X1, X2))) | ~ sP3(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f122, f125, f124, f123])).
fof(f123, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) => (? [X4] : (? [X5] : (sP2(X2, sK9(X0, X1, X2), X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = sK9(X0, X1, X2)) & aScalar0(sK9(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f124, plain, ! [X2, X1, X0] : (? [X4] : (? [X5] : (sP2(X2, sK9(X0, X1, X2), X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) => (? [X5] : (sP2(X2, sK9(X0, X1, X2), X5, sK10(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = sK10(X0, X1, X2)) & aScalar0(sK10(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f125, plain, ! [X2, X1, X0] : (? [X5] : (sP2(X2, sK9(X0, X1, X2), X5, sK10(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) => (sP2(X2, sK9(X0, X1, X2), sK11(X0, X1, X2), sK10(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = sK11(X0, X1, X2)) & aScalar0(sK11(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f122, plain, ! [X0, X1, X2] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) | ~ sP3(X0, X1, X2)), inference(nnf_transformation, [], [f107])).
fof(f107, plain, ! [X0, X1, X2] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) | ~ sP3(X0, X1, X2)), inference(usedef, [], [e107])).
fof(e107, plain, ! [X0, X1, X2] : (sP3(X0, X1, X2) <=> ? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f1348, plain, (! [X4, X2, X0, X5, X3, X1] : ~ sP2(X0, X1, X2, X3, X4, X5) | ~ spl21_5), inference(resolution, [], [f1345, f222])).
fof(f222, plain, ! [X4, X2, X0, X5, X3, X1] : (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, sK13(X0, X1, X2, X3, X4, X5), X2, sK14(X0, X1, X2, X3, X4, X5), X1, X0) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(cnf_transformation, [], [f132])).
fof(f132, plain, ! [X0, X1, X2, X3, X4, X5] : ((((sP1(sK12(X0, X1, X2, X3, X4, X5), X3, sK13(X0, X1, X2, X3, X4, X5), X2, sK14(X0, X1, X2, X3, X4, X5), X1, X0) & (sdtasdt0(X1, X1) = sK14(X0, X1, X2, X3, X4, X5)) & aScalar0(sK14(X0, X1, X2, X3, X4, X5))) & (sdtasdt0(X0, X0) = sK13(X0, X1, X2, X3, X4, X5)) & aScalar0(sK13(X0, X1, X2, X3, X4, X5))) & (sdtasasdt0(X5, X4) = sK12(X0, X1, X2, X3, X4, X5)) & aScalar0(sK12(X0, X1, X2, X3, X4, X5))) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12, sK13, sK14])], [f128, f131, f130, f129])).
fof(f129, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = X6) & aScalar0(X6)) => (? [X7] : (? [X8] : (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = sK12(X0, X1, X2, X3, X4, X5)) & aScalar0(sK12(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f130, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X7] : (? [X8] : (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) => (? [X8] : (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, sK13(X0, X1, X2, X3, X4, X5), X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = sK13(X0, X1, X2, X3, X4, X5)) & aScalar0(sK13(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f131, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X8] : (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, sK13(X0, X1, X2, X3, X4, X5), X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) => (sP1(sK12(X0, X1, X2, X3, X4, X5), X3, sK13(X0, X1, X2, X3, X4, X5), X2, sK14(X0, X1, X2, X3, X4, X5), X1, X0) & (sdtasdt0(X1, X1) = sK14(X0, X1, X2, X3, X4, X5)) & aScalar0(sK14(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X0, X1, X2, X3, X4, X5] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = X6) & aScalar0(X6)) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(rectify, [], [f127])).
fof(f127, plain, ! [X2, X3, X5, X4, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) | ~ sP2(X2, X3, X5, X4, X1, X0)), inference(nnf_transformation, [], [f106])).
fof(f106, plain, ! [X2, X3, X5, X4, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) | ~ sP2(X2, X3, X5, X4, X1, X0)), inference(usedef, [], [e106])).
fof(e106, plain, ! [X2, X3, X5, X4, X1, X0] : (sP2(X2, X3, X5, X4, X1, X0) <=> ? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f1345, plain, (! [X30, X28, X33, X31, X29, X34, X32] : ~ sP1(X28, X29, X30, X31, X32, X33, X34) | ~ spl21_5), inference(resolution, [], [f229, f267])).
fof(f267, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7) | ~ spl21_5), inference(avatar_component_clause, [], [f266])).
fof(f266, plain, (spl21_5 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), introduced(avatar_definition, [new_symbols(naming, [spl21_5])])).
fof(f229, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, sK16(X0, X1, X2, X3, X4, X5, X6), sK17(X0, X1, X2, X3, X4, X5, X6)) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(cnf_transformation, [], [f138])).
fof(f138, plain, ! [X0, X1, X2, X3, X4, X5, X6] : ((((sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, sK16(X0, X1, X2, X3, X4, X5, X6), sK17(X0, X1, X2, X3, X4, X5, X6)) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = sK17(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK17(X0, X1, X2, X3, X4, X5, X6))) & (sdtasdt0(X1, X4) = sK16(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK16(X0, X1, X2, X3, X4, X5, X6))) & (sdtasdt0(X6, X5) = sK15(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK15(X0, X1, X2, X3, X4, X5, X6))) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17])], [f134, f137, f136, f135])).
fof(f135, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X7] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, X7, X0, X8, X9) & (sdtasdt0(X0, X7) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = X7) & aScalar0(X7)) => (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, X8, X9) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = sK15(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK15(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f136, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, X8, X9) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) => (? [X9] : (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, sK16(X0, X1, X2, X3, X4, X5, X6), X9) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = sK16(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK16(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f137, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X9] : (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, sK16(X0, X1, X2, X3, X4, X5, X6), X9) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) => (sP0(X4, X3, X2, X1, sK15(X0, X1, X2, X3, X4, X5, X6), X0, sK16(X0, X1, X2, X3, X4, X5, X6), sK17(X0, X1, X2, X3, X4, X5, X6)) & (sdtasdt0(X0, sK15(X0, X1, X2, X3, X4, X5, X6)) = sK17(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK17(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f134, plain, ! [X0, X1, X2, X3, X4, X5, X6] : (? [X7] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, X7, X0, X8, X9) & (sdtasdt0(X0, X7) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = X7) & aScalar0(X7)) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(rectify, [], [f133])).
fof(f133, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) | ~ sP1(X6, X4, X7, X5, X8, X3, X2)), inference(nnf_transformation, [], [f105])).
fof(f105, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) | ~ sP1(X6, X4, X7, X5, X8, X3, X2)), inference(usedef, [], [e105])).
fof(e105, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (sP1(X6, X4, X7, X5, X8, X3, X2) <=> ? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f281, plain, (sP4(sK20) | ~ spl21_8), inference(avatar_component_clause, [], [f279])).
fof(f279, plain, (spl21_8 <=> sP4(sK20)), introduced(avatar_definition, [new_symbols(naming, [spl21_8])])).
fof(f2050, plain, (spl21_56 | ~ spl21_7), inference(avatar_split_clause, [], [f2049, f275, f1458])).
fof(f2049, plain, ((sz00 = aDimensionOf0(xt)) | ~ spl21_7), inference(forward_demodulation, [], [f201, f277])).
fof(f201, plain, (aDimensionOf0(xs) = aDimensionOf0(xt)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (aDimensionOf0(xs) = aDimensionOf0(xt)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', m__1678_01)).
fof(f1928, plain, (spl21_48 | ~ spl21_7), inference(avatar_split_clause, [], [f1927, f275, f1306])).
fof(f1927, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ spl21_7), inference(subsumption_resolution, [], [f1926, f198])).
fof(f1926, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ aVector0(xs) | ~ spl21_7), inference(trivial_inequality_removal, [], [f1923])).
fof(f1923, plain, (~ (sz00 = sz00) | (sz0z00 = sdtasasdt0(xs, xs)) | ~ aVector0(xs) | ~ spl21_7), inference(superposition, [], [f1529, f277])).
fof(f1529, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6)) | ~ spl21_7), inference(subsumption_resolution, [], [f1525, f198])).
fof(f1525, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl21_7), inference(duplicate_literal_removal, [], [f1520])).
fof(f1520, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | ~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl21_7), inference(superposition, [], [f195, f277])).
fof(f1031, plain, spl21_28, inference(avatar_contradiction_clause, [], [f1030])).
fof(f1030, plain, ($false | spl21_28), inference(subsumption_resolution, [], [f1025, f154])).
fof(f1025, plain, (~ aScalar0(sz0z00) | spl21_28), inference(resolution, [], [f1018, f176])).
fof(f176, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aScalar0(X0)), inference(cnf_transformation, [], [f71])).
fof(f71, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aScalar0(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (aScalar0(X0) => sdtlseqdt0(X0, X0)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', mLERef)).
fof(f1018, plain, (~ sdtlseqdt0(sz0z00, sz0z00) | spl21_28), inference(avatar_component_clause, [], [f1016])).
fof(f282, plain, (spl21_7 | spl21_8), inference(avatar_split_clause, [], [f242, f279, f275])).
fof(f242, plain, (sP4(sK20) | (sz00 = aDimensionOf0(xs))), inference(cnf_transformation, [], [f146])).
fof(f146, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & ((sP4(sK20) & (sziznziztdt0(xs) = sK20) & ! [X1] : ((sdtlbdtrb0(xs, X1) = sdtlbdtrb0(sK20, X1)) | ~ aNaturalNumber0(X1)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(sK20))) & aVector0(sK20)) | (sz00 = aDimensionOf0(xs)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK20])], [f144, f145])).
fof(f145, plain, (? [X0] : (sP4(X0) & (sziznziztdt0(xs) = X0) & ! [X1] : ((sdtlbdtrb0(X0, X1) = sdtlbdtrb0(xs, X1)) | ~ aNaturalNumber0(X1)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0)) => (sP4(sK20) & (sziznziztdt0(xs) = sK20) & ! [X1] : ((sdtlbdtrb0(xs, X1) = sdtlbdtrb0(sK20, X1)) | ~ aNaturalNumber0(X1)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(sK20))) & aVector0(sK20))), introduced(choice_axiom, [])).
fof(f144, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & (? [X0] : (sP4(X0) & (sziznziztdt0(xs) = X0) & ! [X1] : ((sdtlbdtrb0(X0, X1) = sdtlbdtrb0(xs, X1)) | ~ aNaturalNumber0(X1)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0)) | (sz00 = aDimensionOf0(xs)))), inference(rectify, [], [f109])).
fof(f109, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & (? [X0] : (sP4(X0) & (sziznziztdt0(xs) = X0) & ! [X15] : ((sdtlbdtrb0(X0, X15) = sdtlbdtrb0(xs, X15)) | ~ aNaturalNumber0(X15)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0)) | (sz00 = aDimensionOf0(xs)))), inference(definition_folding, [], [f103, e108, e107, e106, e105, e104])).
fof(f104, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) | ~ sP0(X8, X5, X7, X4, X9, X6, X10, X11)), inference(usedef, [], [e104])).
fof(e104, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) <=> ? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f103, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X14] : ((sdtlbdtrb0(X1, X14) = sdtlbdtrb0(xt, X14)) | ~ aNaturalNumber0(X14)) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & ! [X15] : ((sdtlbdtrb0(X0, X15) = sdtlbdtrb0(xs, X15)) | ~ aNaturalNumber0(X15)) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0)) | (sz00 = aDimensionOf0(xs)))), inference(ennf_transformation, [], [f48])).
fof(f48, plain, ~ ((~ (sz00 = aDimensionOf0(xs)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X14] : (aNaturalNumber0(X14) => (sdtlbdtrb0(X1, X14) = sdtlbdtrb0(xt, X14))) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & ! [X15] : (aNaturalNumber0(X15) => (sdtlbdtrb0(X0, X15) = sdtlbdtrb0(xs, X15))) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0))) => sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), inference(rectify, [], [f42])).
fof(f42, plain, ~ ((~ (sz00 = aDimensionOf0(xs)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X2] : (aNaturalNumber0(X2) => (sdtlbdtrb0(X1, X2) = sdtlbdtrb0(xt, X2))) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & ! [X1] : (aNaturalNumber0(X1) => (sdtlbdtrb0(X0, X1) = sdtlbdtrb0(xs, X1))) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0))) => sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), inference(negated_conjecture, [], [f41])).
fof(f41, plain, ~ ((~ (sz00 = aDimensionOf0(xs)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & ! [X2] : (aNaturalNumber0(X2) => (sdtlbdtrb0(X1, X2) = sdtlbdtrb0(xt, X2))) & (szszuzczcdt0(aDimensionOf0(X1)) = aDimensionOf0(xt)) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & ! [X1] : (aNaturalNumber0(X1) => (sdtlbdtrb0(X0, X1) = sdtlbdtrb0(xs, X1))) & (aDimensionOf0(xs) = szszuzczcdt0(aDimensionOf0(X0))) & aVector0(X0))) => sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+2.p', m__)).
fof(f273, plain, ~ spl21_6, inference(avatar_split_clause, [], [f243, f269])).
fof(f243, plain, ~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))), inference(cnf_transformation, [], [f146])).
fof(f272, plain, (spl21_5 | spl21_6), inference(avatar_split_clause, [], [f237, f269, f266])).
fof(f237, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(cnf_transformation, [], [f143])).
fof(f143, plain, ! [X0, X1, X2, X3, X4, X5, X6, X7] : (((sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7)) = sK19(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK19(X0, X1, X2, X3, X4, X5, X6, X7))) & (sdtasdt0(X2, X1) = sK18(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK18(X0, X1, X2, X3, X4, X5, X6, X7))) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18, sK19])], [f140, f142, f141])).
fof(f141, plain, ! [X7, X6, X5, X4, X3, X2, X1, X0] : (? [X8] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, X8)) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, X8) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = X8) & aScalar0(X8)) => (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7)) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = sK18(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK18(X0, X1, X2, X3, X4, X5, X6, X7)))), introduced(choice_axiom, [])).
fof(f142, plain, ! [X7, X6, X5, X4, X3, X2, X1, X0] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7)) = X9) & aScalar0(X9)) => (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK18(X0, X1, X2, X3, X4, X5, X6, X7)) = sK19(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK19(X0, X1, X2, X3, X4, X5, X6, X7)))), introduced(choice_axiom, [])).
fof(f140, plain, ! [X0, X1, X2, X3, X4, X5, X6, X7] : (? [X8] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, X8)) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, X8) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = X8) & aScalar0(X8)) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(rectify, [], [f139])).
fof(f139, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) | ~ sP0(X8, X5, X7, X4, X9, X6, X10, X11)), inference(nnf_transformation, [], [f104])).