fof(f2450, plain, $false, inference(avatar_sat_refutation, [], [f260, f261, f270, f1024, f1815, f2211, f2433, f2446])).
fof(f2446, plain, (spl20_61 | ~ spl20_7), inference(avatar_split_clause, [], [f2445, f263, f1448])).
fof(f1448, plain, (spl20_61 <=> (sz00 = aDimensionOf0(xt))), introduced(avatar_definition, [new_symbols(naming, [spl20_61])])).
fof(f263, plain, (spl20_7 <=> (sz00 = aDimensionOf0(xs))), introduced(avatar_definition, [new_symbols(naming, [spl20_7])])).
fof(f2445, plain, ((sz00 = aDimensionOf0(xt)) | ~ spl20_7), inference(forward_demodulation, [], [f195, f265])).
fof(f265, plain, ((sz00 = aDimensionOf0(xs)) | ~ spl20_7), inference(avatar_component_clause, [], [f263])).
fof(f195, plain, (aDimensionOf0(xs) = aDimensionOf0(xt)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, (aDimensionOf0(xs) = aDimensionOf0(xt)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', m__1678_01)).
fof(f2433, plain, (spl20_6 | ~ spl20_7 | ~ spl20_32 | ~ spl20_53 | ~ spl20_61), inference(avatar_contradiction_clause, [], [f2432])).
fof(f2432, plain, ($false | (spl20_6 | ~ spl20_7 | ~ spl20_32 | ~ spl20_53 | ~ spl20_61)), inference(subsumption_resolution, [], [f2431, f1010])).
fof(f1010, plain, (sdtlseqdt0(sz0z00, sz0z00) | ~ spl20_32), inference(avatar_component_clause, [], [f1009])).
fof(f1009, plain, (spl20_32 <=> sdtlseqdt0(sz0z00, sz0z00)), introduced(avatar_definition, [new_symbols(naming, [spl20_32])])).
fof(f2431, plain, (~ sdtlseqdt0(sz0z00, sz0z00) | (spl20_6 | ~ spl20_7 | ~ spl20_53 | ~ spl20_61)), inference(forward_demodulation, [], [f2381, f764])).
fof(f764, plain, (sz0z00 = sdtasdt0(sz0z00, sdtasasdt0(xt, xt))), inference(resolution, [], [f558, f193])).
fof(f193, plain, aVector0(xt), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (aVector0(xt) & aVector0(xs)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', m__1678)).
fof(f558, plain, ! [X3] : (~ aVector0(X3) | (sz0z00 = sdtasdt0(sz0z00, sdtasasdt0(X3, X3)))), inference(resolution, [], [f550, f155])).
fof(f155, plain, ! [X0] : (~ aScalar0(X0) | (sz0z00 = sdtasdt0(sz0z00, X0))), inference(cnf_transformation, [], [f59])).
fof(f59, plain, ! [X0] : (((sz0z00 = smndt0(sz0z00)) & (smndt0(smndt0(X0)) = X0) & (sz0z00 = sdtpldt0(smndt0(X0), X0)) & (sz0z00 = sdtpldt0(X0, smndt0(X0))) & (sz0z00 = sdtasdt0(sz0z00, X0)) & (sz0z00 = sdtasdt0(X0, sz0z00)) & (sdtpldt0(sz0z00, X0) = X0) & (sdtpldt0(X0, sz0z00) = X0)) | ~ aScalar0(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (aScalar0(X0) => ((sz0z00 = smndt0(sz0z00)) & (smndt0(smndt0(X0)) = X0) & (sz0z00 = sdtpldt0(smndt0(X0), X0)) & (sz0z00 = sdtpldt0(X0, smndt0(X0))) & (sz0z00 = sdtasdt0(sz0z00, X0)) & (sz0z00 = sdtasdt0(X0, sz0z00)) & (sdtpldt0(sz0z00, X0) = X0) & (sdtpldt0(X0, sz0z00) = X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', mScZero)).
fof(f550, plain, ! [X0] : (aScalar0(sdtasasdt0(X0, X0)) | ~ aVector0(X0)), inference(duplicate_literal_removal, [], [f549])).
fof(f549, plain, ! [X0] : (aScalar0(sdtasasdt0(X0, X0)) | ~ aVector0(X0) | ~ aVector0(X0)), inference(equality_resolution, [], [f188])).
fof(f188, plain, ! [X0, X1] : (~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | aScalar0(sdtasasdt0(X0, X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(cnf_transformation, [], [f94])).
fof(f94, plain, ! [X0, X1] : (aScalar0(sdtasasdt0(X0, X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(flattening, [], [f93])).
fof(f93, plain, ! [X0, X1] : ((aScalar0(sdtasasdt0(X0, X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1))) | (~ aVector0(X1) | ~ aVector0(X0))), inference(ennf_transformation, [], [f34])).
fof(f34, plain, ! [X0, X1] : ((aVector0(X1) & aVector0(X0)) => ((aDimensionOf0(X0) = aDimensionOf0(X1)) => aScalar0(sdtasasdt0(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', mScPr)).
fof(f2381, plain, (~ sdtlseqdt0(sz0z00, sdtasdt0(sz0z00, sdtasasdt0(xt, xt))) | (spl20_6 | ~ spl20_7 | ~ spl20_53 | ~ spl20_61)), inference(backward_demodulation, [], [f2269, f1273])).
fof(f1273, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ spl20_53), inference(avatar_component_clause, [], [f1271])).
fof(f1271, plain, (spl20_53 <=> (sz0z00 = sdtasasdt0(xs, xs))), introduced(avatar_definition, [new_symbols(naming, [spl20_53])])).
fof(f2269, plain, (~ sdtlseqdt0(sz0z00, sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | (spl20_6 | ~ spl20_7 | ~ spl20_61)), inference(forward_demodulation, [], [f2213, f311])).
fof(f311, plain, (sz0z00 = sdtasdt0(sz0z00, sz0z00)), inference(resolution, [], [f154, f148])).
fof(f148, plain, aScalar0(sz0z00), inference(cnf_transformation, [], [f9])).
fof(f9, plain, aScalar0(sz0z00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', mSZeroSc)).
fof(f154, plain, ! [X0] : (~ aScalar0(X0) | (sz0z00 = sdtasdt0(X0, sz0z00))), inference(cnf_transformation, [], [f59])).
fof(f2213, plain, (~ sdtlseqdt0(sdtasdt0(sz0z00, sz0z00), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | (spl20_6 | ~ spl20_7 | ~ spl20_61)), inference(backward_demodulation, [], [f258, f2212])).
fof(f2212, plain, ((sz0z00 = sdtasasdt0(xs, xt)) | (~ spl20_7 | ~ spl20_61)), inference(subsumption_resolution, [], [f2208, f193])).
fof(f2208, plain, ((sz0z00 = sdtasasdt0(xs, xt)) | ~ aVector0(xt) | (~ spl20_7 | ~ spl20_61)), inference(trivial_inequality_removal, [], [f2207])).
fof(f2207, plain, (~ (sz00 = sz00) | (sz0z00 = sdtasasdt0(xs, xt)) | ~ aVector0(xt) | (~ spl20_7 | ~ spl20_61)), inference(superposition, [], [f1553, f1450])).
fof(f1450, plain, ((sz00 = aDimensionOf0(xt)) | ~ spl20_61), inference(avatar_component_clause, [], [f1448])).
fof(f1553, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6)) | ~ spl20_7), inference(subsumption_resolution, [], [f1549, f192])).
fof(f192, plain, aVector0(xs), inference(cnf_transformation, [], [f38])).
fof(f1549, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl20_7), inference(duplicate_literal_removal, [], [f1544])).
fof(f1544, plain, (! [X6] : (~ (sz00 = aDimensionOf0(X6)) | ~ (sz00 = aDimensionOf0(X6)) | (sz0z00 = sdtasasdt0(xs, X6)) | ~ aVector0(X6) | ~ aVector0(xs)) | ~ spl20_7), inference(superposition, [], [f189, f265])).
fof(f189, plain, ! [X0, X1] : (~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ (sz00 = aDimensionOf0(X1)) | (sz0z00 = sdtasasdt0(X0, X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ! [X0, X1] : ((sz0z00 = sdtasasdt0(X0, X1)) | ~ (sz00 = aDimensionOf0(X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)) | ~ aVector0(X1) | ~ aVector0(X0)), inference(flattening, [], [f95])).
fof(f95, plain, ! [X0, X1] : (((sz0z00 = sdtasasdt0(X0, X1)) | (~ (sz00 = aDimensionOf0(X1)) | ~ (aDimensionOf0(X0) = aDimensionOf0(X1)))) | (~ aVector0(X1) | ~ aVector0(X0))), inference(ennf_transformation, [], [f35])).
fof(f35, plain, ! [X0, X1] : ((aVector0(X1) & aVector0(X0)) => (((sz00 = aDimensionOf0(X1)) & (aDimensionOf0(X0) = aDimensionOf0(X1))) => (sz0z00 = sdtasasdt0(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', mDefSPZ)).
fof(f258, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | spl20_6), inference(avatar_component_clause, [], [f257])).
fof(f257, plain, (spl20_6 <=> sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), introduced(avatar_definition, [new_symbols(naming, [spl20_6])])).
fof(f2211, plain, (spl20_53 | ~ spl20_7), inference(avatar_split_clause, [], [f2210, f263, f1271])).
fof(f2210, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ spl20_7), inference(subsumption_resolution, [], [f2209, f192])).
fof(f2209, plain, ((sz0z00 = sdtasasdt0(xs, xs)) | ~ aVector0(xs) | ~ spl20_7), inference(trivial_inequality_removal, [], [f2206])).
fof(f2206, plain, (~ (sz00 = sz00) | (sz0z00 = sdtasasdt0(xs, xs)) | ~ aVector0(xs) | ~ spl20_7), inference(superposition, [], [f1553, f265])).
fof(f1815, plain, (~ spl20_5 | ~ spl20_8), inference(avatar_contradiction_clause, [], [f1814])).
fof(f1814, plain, ($false | (~ spl20_5 | ~ spl20_8)), inference(subsumption_resolution, [], [f1813, f269])).
fof(f269, plain, (sP3(sK17, sK18, sK19) | ~ spl20_8), inference(avatar_component_clause, [], [f267])).
fof(f267, plain, (spl20_8 <=> sP3(sK17, sK18, sK19)), introduced(avatar_definition, [new_symbols(naming, [spl20_8])])).
fof(f1813, plain, (~ sP3(sK17, sK18, sK19) | (~ spl20_5 | ~ spl20_8)), inference(subsumption_resolution, [], [f1811, f1334])).
fof(f1334, plain, (! [X4, X2, X0, X5, X3, X1] : ~ sP2(X0, X1, X2, X3, X4, X5) | ~ spl20_5), inference(resolution, [], [f1331, f209])).
fof(f209, plain, ! [X4, X2, X0, X5, X3, X1] : (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, sK10(X0, X1, X2, X3, X4, X5), X2, sK11(X0, X1, X2, X3, X4, X5), X1, X0) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ! [X0, X1, X2, X3, X4, X5] : ((((sP1(sK9(X0, X1, X2, X3, X4, X5), X3, sK10(X0, X1, X2, X3, X4, X5), X2, sK11(X0, X1, X2, X3, X4, X5), X1, X0) & (sdtasdt0(X1, X1) = sK11(X0, X1, X2, X3, X4, X5)) & aScalar0(sK11(X0, X1, X2, X3, X4, X5))) & (sdtasdt0(X0, X0) = sK10(X0, X1, X2, X3, X4, X5)) & aScalar0(sK10(X0, X1, X2, X3, X4, X5))) & (sdtasasdt0(X5, X4) = sK9(X0, X1, X2, X3, X4, X5)) & aScalar0(sK9(X0, X1, X2, X3, X4, X5))) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f121, f124, f123, f122])).
fof(f122, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = X6) & aScalar0(X6)) => (? [X7] : (? [X8] : (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = sK9(X0, X1, X2, X3, X4, X5)) & aScalar0(sK9(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f123, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X7] : (? [X8] : (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) => (? [X8] : (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, sK10(X0, X1, X2, X3, X4, X5), X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = sK10(X0, X1, X2, X3, X4, X5)) & aScalar0(sK10(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f124, plain, ! [X5, X4, X3, X2, X1, X0] : (? [X8] : (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, sK10(X0, X1, X2, X3, X4, X5), X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) => (sP1(sK9(X0, X1, X2, X3, X4, X5), X3, sK10(X0, X1, X2, X3, X4, X5), X2, sK11(X0, X1, X2, X3, X4, X5), X1, X0) & (sdtasdt0(X1, X1) = sK11(X0, X1, X2, X3, X4, X5)) & aScalar0(sK11(X0, X1, X2, X3, X4, X5)))), introduced(choice_axiom, [])).
fof(f121, plain, ! [X0, X1, X2, X3, X4, X5] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X3, X7, X2, X8, X1, X0) & (sdtasdt0(X1, X1) = X8) & aScalar0(X8)) & (sdtasdt0(X0, X0) = X7) & aScalar0(X7)) & (sdtasasdt0(X5, X4) = X6) & aScalar0(X6)) | ~ sP2(X0, X1, X2, X3, X4, X5)), inference(rectify, [], [f120])).
fof(f120, plain, ! [X2, X3, X5, X4, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) | ~ sP2(X2, X3, X5, X4, X1, X0)), inference(nnf_transformation, [], [f105])).
fof(f105, plain, ! [X2, X3, X5, X4, X1, X0] : (? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) | ~ sP2(X2, X3, X5, X4, X1, X0)), inference(usedef, [], [e105])).
fof(e105, plain, ! [X2, X3, X5, X4, X1, X0] : (sP2(X2, X3, X5, X4, X1, X0) <=> ? [X6] : (? [X7] : (? [X8] : (sP1(X6, X4, X7, X5, X8, X3, X2) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f1331, plain, (! [X30, X28, X33, X31, X29, X34, X32] : ~ sP1(X28, X29, X30, X31, X32, X33, X34) | ~ spl20_5), inference(resolution, [], [f216, f255])).
fof(f255, plain, (! [X6, X4, X2, X0, X7, X5, X3, X1] : ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7) | ~ spl20_5), inference(avatar_component_clause, [], [f254])).
fof(f254, plain, (spl20_5 <=> ! [X1, X3, X5, X7, X0, X2, X4, X6] : ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), introduced(avatar_definition, [new_symbols(naming, [spl20_5])])).
fof(f216, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, sK13(X0, X1, X2, X3, X4, X5, X6), sK14(X0, X1, X2, X3, X4, X5, X6)) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(cnf_transformation, [], [f131])).
fof(f131, plain, ! [X0, X1, X2, X3, X4, X5, X6] : ((((sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, sK13(X0, X1, X2, X3, X4, X5, X6), sK14(X0, X1, X2, X3, X4, X5, X6)) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = sK14(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK14(X0, X1, X2, X3, X4, X5, X6))) & (sdtasdt0(X1, X4) = sK13(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK13(X0, X1, X2, X3, X4, X5, X6))) & (sdtasdt0(X6, X5) = sK12(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK12(X0, X1, X2, X3, X4, X5, X6))) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK12, sK13, sK14])], [f127, f130, f129, f128])).
fof(f128, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X7] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, X7, X0, X8, X9) & (sdtasdt0(X0, X7) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = X7) & aScalar0(X7)) => (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, X8, X9) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = sK12(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK12(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f129, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, X8, X9) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) => (? [X9] : (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, sK13(X0, X1, X2, X3, X4, X5, X6), X9) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = sK13(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK13(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f130, plain, ! [X6, X5, X4, X3, X2, X1, X0] : (? [X9] : (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, sK13(X0, X1, X2, X3, X4, X5, X6), X9) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = X9) & aScalar0(X9)) => (sP0(X4, X3, X2, X1, sK12(X0, X1, X2, X3, X4, X5, X6), X0, sK13(X0, X1, X2, X3, X4, X5, X6), sK14(X0, X1, X2, X3, X4, X5, X6)) & (sdtasdt0(X0, sK12(X0, X1, X2, X3, X4, X5, X6)) = sK14(X0, X1, X2, X3, X4, X5, X6)) & aScalar0(sK14(X0, X1, X2, X3, X4, X5, X6)))), introduced(choice_axiom, [])).
fof(f127, plain, ! [X0, X1, X2, X3, X4, X5, X6] : (? [X7] : (? [X8] : (? [X9] : (sP0(X4, X3, X2, X1, X7, X0, X8, X9) & (sdtasdt0(X0, X7) = X9) & aScalar0(X9)) & (sdtasdt0(X1, X4) = X8) & aScalar0(X8)) & (sdtasdt0(X6, X5) = X7) & aScalar0(X7)) | ~ sP1(X0, X1, X2, X3, X4, X5, X6)), inference(rectify, [], [f126])).
fof(f126, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) | ~ sP1(X6, X4, X7, X5, X8, X3, X2)), inference(nnf_transformation, [], [f104])).
fof(f104, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) | ~ sP1(X6, X4, X7, X5, X8, X3, X2)), inference(usedef, [], [e104])).
fof(e104, plain, ! [X6, X4, X7, X5, X8, X3, X2] : (sP1(X6, X4, X7, X5, X8, X3, X2) <=> ? [X9] : (? [X10] : (? [X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f1811, plain, (sP2(sK19, sK6(sK17, sK18, sK19), sdtasasdt0(sK18, sK18), sK7(sK17, sK18, sK19), sK18, sK17) | ~ sP3(sK17, sK18, sK19) | ~ spl20_8), inference(superposition, [], [f202, f1555])).
fof(f1555, plain, ((sdtasasdt0(sK18, sK18) = sK8(sK17, sK18, sK19)) | ~ spl20_8), inference(resolution, [], [f269, f201])).
fof(f201, plain, ! [X2, X0, X1] : (~ sP3(X0, X1, X2) | (sdtasasdt0(X1, X1) = sK8(X0, X1, X2))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0, X1, X2] : ((((sP2(X2, sK6(X0, X1, X2), sK8(X0, X1, X2), sK7(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = sK8(X0, X1, X2)) & aScalar0(sK8(X0, X1, X2))) & (sdtasasdt0(X0, X0) = sK7(X0, X1, X2)) & aScalar0(sK7(X0, X1, X2))) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = sK6(X0, X1, X2)) & aScalar0(sK6(X0, X1, X2))) | ~ sP3(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6, sK7, sK8])], [f115, f118, f117, f116])).
fof(f116, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) => (? [X4] : (? [X5] : (sP2(X2, sK6(X0, X1, X2), X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = sK6(X0, X1, X2)) & aScalar0(sK6(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f117, plain, ! [X2, X1, X0] : (? [X4] : (? [X5] : (sP2(X2, sK6(X0, X1, X2), X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) => (? [X5] : (sP2(X2, sK6(X0, X1, X2), X5, sK7(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = sK7(X0, X1, X2)) & aScalar0(sK7(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X2, X1, X0] : (? [X5] : (sP2(X2, sK6(X0, X1, X2), X5, sK7(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) => (sP2(X2, sK6(X0, X1, X2), sK8(X0, X1, X2), sK7(X0, X1, X2), X1, X0) & (sdtasasdt0(X1, X1) = sK8(X0, X1, X2)) & aScalar0(sK8(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f115, plain, ! [X0, X1, X2] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) | ~ sP3(X0, X1, X2)), inference(nnf_transformation, [], [f106])).
fof(f106, plain, ! [X0, X1, X2] : (? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) | ~ sP3(X0, X1, X2)), inference(usedef, [], [e106])).
fof(e106, plain, ! [X0, X1, X2] : (sP3(X0, X1, X2) <=> ? [X3] : (? [X4] : (? [X5] : (sP2(X2, X3, X5, X4, X1, X0) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f202, plain, ! [X2, X0, X1] : (sP2(X2, sK6(X0, X1, X2), sK8(X0, X1, X2), sK7(X0, X1, X2), X1, X0) | ~ sP3(X0, X1, X2)), inference(cnf_transformation, [], [f119])).
fof(f1024, plain, spl20_32, inference(avatar_contradiction_clause, [], [f1023])).
fof(f1023, plain, ($false | spl20_32), inference(subsumption_resolution, [], [f1018, f148])).
fof(f1018, plain, (~ aScalar0(sz0z00) | spl20_32), inference(resolution, [], [f1011, f170])).
fof(f170, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aScalar0(X0)), inference(cnf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (sdtlseqdt0(X0, X0) | ~ aScalar0(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (aScalar0(X0) => sdtlseqdt0(X0, X0)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', mLERef)).
fof(f1011, plain, (~ sdtlseqdt0(sz0z00, sz0z00) | spl20_32), inference(avatar_component_clause, [], [f1009])).
fof(f270, plain, (spl20_7 | spl20_8), inference(avatar_split_clause, [], [f231, f267, f263])).
fof(f231, plain, (sP3(sK17, sK18, sK19) | (sz00 = aDimensionOf0(xs))), inference(cnf_transformation, [], [f140])).
fof(f140, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & ((((sP3(sK17, sK18, sK19) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = sK19) & aScalar0(sK19)) & (sziznziztdt0(xt) = sK18) & aVector0(sK18)) & (sziznziztdt0(xs) = sK17) & aVector0(sK17)) | (sz00 = aDimensionOf0(xs)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18, sK19])], [f107, f139, f138, f137])).
fof(f137, plain, (? [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & aVector0(X0)) => (? [X1] : (? [X2] : (sP3(sK17, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = sK17) & aVector0(sK17))), introduced(choice_axiom, [])).
fof(f138, plain, (? [X1] : (? [X2] : (sP3(sK17, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) => (? [X2] : (sP3(sK17, sK18, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = sK18) & aVector0(sK18))), introduced(choice_axiom, [])).
fof(f139, plain, (? [X2] : (sP3(sK17, sK18, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) => (sP3(sK17, sK18, sK19) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = sK19) & aScalar0(sK19))), introduced(choice_axiom, [])).
fof(f107, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & (? [X0] : (? [X1] : (? [X2] : (sP3(X0, X1, X2) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & aVector0(X0)) | (sz00 = aDimensionOf0(xs)))), inference(definition_folding, [], [f102, e106, e105, e104, e103])).
fof(f103, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) | ~ sP0(X8, X5, X7, X4, X9, X6, X10, X11)), inference(usedef, [], [e103])).
fof(e103, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (sP0(X8, X5, X7, X4, X9, X6, X10, X11) <=> ? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f102, plain, (~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & (? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & aVector0(X0)) | (sz00 = aDimensionOf0(xs)))), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ~ ((~ (sz00 = aDimensionOf0(xs)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & aVector0(X0))) => sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), inference(negated_conjecture, [], [f41])).
fof(f41, plain, ~ ((~ (sz00 = aDimensionOf0(xs)) => ? [X0] : (? [X1] : (? [X2] : (? [X3] : (? [X4] : (? [X5] : (? [X6] : (? [X7] : (? [X8] : (? [X9] : (? [X10] : (? [X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) & (sdtasdt0(X6, X9) = X11) & aScalar0(X11)) & (sdtasdt0(X4, X8) = X10) & aScalar0(X10)) & (sdtasdt0(X2, X3) = X9) & aScalar0(X9)) & (sdtasdt0(X3, X3) = X8) & aScalar0(X8)) & (sdtasdt0(X2, X2) = X7) & aScalar0(X7)) & (sdtasasdt0(X0, X1) = X6) & aScalar0(X6)) & (sdtasasdt0(X1, X1) = X5) & aScalar0(X5)) & (sdtasasdt0(X0, X0) = X4) & aScalar0(X4)) & (sdtlbdtrb0(xt, aDimensionOf0(xt)) = X3) & aScalar0(X3)) & (sdtlbdtrb0(xs, aDimensionOf0(xs)) = X2) & aScalar0(X2)) & (sziznziztdt0(xt) = X1) & aVector0(X1)) & (sziznziztdt0(xs) = X0) & aVector0(X0))) => sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG081+1.p', m__)).
fof(f261, plain, ~ spl20_6, inference(avatar_split_clause, [], [f232, f257])).
fof(f232, plain, ~ sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))), inference(cnf_transformation, [], [f140])).
fof(f260, plain, (spl20_5 | spl20_6), inference(avatar_split_clause, [], [f224, f257, f254])).
fof(f224, plain, ! [X6, X4, X2, X0, X7, X5, X3, X1] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(cnf_transformation, [], [f136])).
fof(f136, plain, ! [X0, X1, X2, X3, X4, X5, X6, X7] : (((sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7)) = sK16(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK16(X0, X1, X2, X3, X4, X5, X6, X7))) & (sdtasdt0(X2, X1) = sK15(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK15(X0, X1, X2, X3, X4, X5, X6, X7))) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16])], [f133, f135, f134])).
fof(f134, plain, ! [X7, X6, X5, X4, X3, X2, X1, X0] : (? [X8] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, X8)) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, X8) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = X8) & aScalar0(X8)) => (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7)) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = sK15(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK15(X0, X1, X2, X3, X4, X5, X6, X7)))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X7, X6, X5, X4, X3, X2, X1, X0] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7)) = X9) & aScalar0(X9)) => (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7))) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, sK15(X0, X1, X2, X3, X4, X5, X6, X7)) = sK16(X0, X1, X2, X3, X4, X5, X6, X7)) & aScalar0(sK16(X0, X1, X2, X3, X4, X5, X6, X7)))), introduced(choice_axiom, [])).
fof(f133, plain, ! [X0, X1, X2, X3, X4, X5, X6, X7] : (? [X8] : (? [X9] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X5, X4), sdtpldt0(X5, X4)), sdtasdt0(sdtpldt0(X3, X2), sdtpldt0(X1, X0))) & sdtlseqdt0(sdtpldt0(X7, X7), sdtpldt0(X6, X8)) & sdtlseqdt0(sdtasdt0(X5, X5), sdtasdt0(X3, X1)) & (sdtasdt0(X6, X8) = X9) & aScalar0(X9)) & (sdtasdt0(X2, X1) = X8) & aScalar0(X8)) | ~ sP0(X0, X1, X2, X3, X4, X5, X6, X7)), inference(rectify, [], [f132])).
fof(f132, plain, ! [X8, X5, X7, X4, X9, X6, X10, X11] : (? [X12] : (? [X13] : (sdtlseqdt0(sdtasdt0(sdtasasdt0(xs, xt), sdtasasdt0(xs, xt)), sdtasdt0(sdtasasdt0(xs, xs), sdtasasdt0(xt, xt))) & sdtlseqdt0(sdtasdt0(sdtpldt0(X6, X9), sdtpldt0(X6, X9)), sdtasdt0(sdtpldt0(X4, X7), sdtpldt0(X5, X8))) & sdtlseqdt0(sdtpldt0(X11, X11), sdtpldt0(X10, X12)) & sdtlseqdt0(sdtasdt0(X6, X6), sdtasdt0(X4, X5)) & (sdtasdt0(X10, X12) = X13) & aScalar0(X13)) & (sdtasdt0(X7, X5) = X12) & aScalar0(X12)) | ~ sP0(X8, X5, X7, X4, X9, X6, X10, X11)), inference(nnf_transformation, [], [f103])).