fof(f436487, plain, $false, inference(avatar_sat_refutation, [], [f326, f331, f336, f337, f2746, f177424, f177474, f430245, f436485])).
fof(f436485, plain, (spl34_58 | spl34_1 | ~ spl34_3 | ~ spl34_1190), inference(avatar_split_clause, [], [f436484, f144369, f328, f319, f2680])).
fof(f2680, plain, (spl34_58 <=> ! [X5] : ~ aElement0(X5)), introduced(avatar_definition, [new_symbols(naming, [spl34_58])])).
fof(f319, plain, (spl34_1 <=> aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc))), introduced(avatar_definition, [new_symbols(naming, [spl34_1])])).
fof(f328, plain, (spl34_3 <=> aElementOf0(sK32, slsdtgt0(xc))), introduced(avatar_definition, [new_symbols(naming, [spl34_3])])).
fof(f144369, plain, (spl34_1190 <=> aElement0(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl34_1190])])).
fof(f436484, plain, (! [X16] : ~ aElement0(X16) | (spl34_1 | ~ spl34_3 | ~ spl34_1190)), inference(subsumption_resolution, [], [f436483, f321])).
fof(f321, plain, (~ aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc)) | spl34_1), inference(avatar_component_clause, [], [f319])).
fof(f436483, plain, (! [X16] : (aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc)) | ~ aElement0(X16)) | (~ spl34_3 | ~ spl34_1190)), inference(forward_demodulation, [], [f436454, f430815])).
fof(f430815, plain, ((sdtpldt0(sK30, sK32) = sdtpldt0(sK32, sK30)) | ~ spl34_1190), inference(resolution, [], [f144370, f653])).
fof(f653, plain, ! [X21] : (~ aElement0(X21) | (sdtpldt0(X21, sK30) = sdtpldt0(sK30, X21))), inference(resolution, [], [f186, f378])).
fof(f378, plain, aElement0(sK30), inference(subsumption_resolution, [], [f376, f290])).
fof(f290, plain, aSet0(slsdtgt0(xc)), inference(cnf_transformation, [], [f180])).
fof(f180, plain, (~ aIdeal0(slsdtgt0(xc)) & (((~ aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc)) & aElement0(sK31)) | (~ aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc)) & aElementOf0(sK32, slsdtgt0(xc)))) & aElementOf0(sK30, slsdtgt0(xc))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xc)) | ! [X4] : (~ (sdtasdt0(xc, X4) = X3) | ~ aElement0(X4))) & (((sdtasdt0(xc, sK33(X3)) = X3) & aElement0(sK33(X3))) | ~ aElementOf0(X3, slsdtgt0(xc)))) & aSet0(slsdtgt0(xc)) & ! [X6, X7, X8] : (sP5(X7, X8, X6) | ~ aElement0(X8) | (~ aElementOf0(X7, slsdtgt0(xc)) & ! [X9] : (~ (sdtasdt0(xc, X9) = X7) | ~ aElement0(X9))) | (~ aElementOf0(X6, slsdtgt0(xc)) & ! [X10] : (~ (sdtasdt0(xc, X10) = X6) | ~ aElement0(X10))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31, sK32, sK33])], [f175, f179, f178, f177, f176])).
fof(f176, plain, (? [X0] : ((? [X1] : (~ aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & aElementOf0(X2, slsdtgt0(xc)))) & aElementOf0(X0, slsdtgt0(xc))) => ((? [X1] : (~ aElementOf0(sdtasdt0(X1, sK30), slsdtgt0(xc)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(sK30, X2), slsdtgt0(xc)) & aElementOf0(X2, slsdtgt0(xc)))) & aElementOf0(sK30, slsdtgt0(xc)))), introduced(choice_axiom, [])).
fof(f177, plain, (? [X1] : (~ aElementOf0(sdtasdt0(X1, sK30), slsdtgt0(xc)) & aElement0(X1)) => (~ aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc)) & aElement0(sK31))), introduced(choice_axiom, [])).
fof(f178, plain, (? [X2] : (~ aElementOf0(sdtpldt0(sK30, X2), slsdtgt0(xc)) & aElementOf0(X2, slsdtgt0(xc))) => (~ aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc)) & aElementOf0(sK32, slsdtgt0(xc)))), introduced(choice_axiom, [])).
fof(f179, plain, ! [X3] : (? [X5] : ((sdtasdt0(xc, X5) = X3) & aElement0(X5)) => ((sdtasdt0(xc, sK33(X3)) = X3) & aElement0(sK33(X3)))), introduced(choice_axiom, [])).
fof(f175, plain, (~ aIdeal0(slsdtgt0(xc)) & ? [X0] : ((? [X1] : (~ aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & aElement0(X1)) | ? [X2] : (~ aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & aElementOf0(X2, slsdtgt0(xc)))) & aElementOf0(X0, slsdtgt0(xc))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xc)) | ! [X4] : (~ (sdtasdt0(xc, X4) = X3) | ~ aElement0(X4))) & (? [X5] : ((sdtasdt0(xc, X5) = X3) & aElement0(X5)) | ~ aElementOf0(X3, slsdtgt0(xc)))) & aSet0(slsdtgt0(xc)) & ! [X6, X7, X8] : (sP5(X7, X8, X6) | ~ aElement0(X8) | (~ aElementOf0(X7, slsdtgt0(xc)) & ! [X9] : (~ (sdtasdt0(xc, X9) = X7) | ~ aElement0(X9))) | (~ aElementOf0(X6, slsdtgt0(xc)) & ! [X10] : (~ (sdtasdt0(xc, X10) = X6) | ~ aElement0(X10))))), inference(rectify, [], [f174])).
fof(f174, plain, (~ aIdeal0(slsdtgt0(xc)) & ? [X11] : ((? [X12] : (~ aElementOf0(sdtasdt0(X12, X11), slsdtgt0(xc)) & aElement0(X12)) | ? [X13] : (~ aElementOf0(sdtpldt0(X11, X13), slsdtgt0(xc)) & aElementOf0(X13, slsdtgt0(xc)))) & aElementOf0(X11, slsdtgt0(xc))) & ! [X9] : ((aElementOf0(X9, slsdtgt0(xc)) | ! [X10] : (~ (sdtasdt0(xc, X10) = X9) | ~ aElement0(X10))) & (? [X10] : ((sdtasdt0(xc, X10) = X9) & aElement0(X10)) | ~ aElementOf0(X9, slsdtgt0(xc)))) & aSet0(slsdtgt0(xc)) & ! [X0, X1, X2] : (sP5(X1, X2, X0) | ~ aElement0(X2) | (~ aElementOf0(X1, slsdtgt0(xc)) & ! [X3] : (~ (sdtasdt0(xc, X3) = X1) | ~ aElement0(X3))) | (~ aElementOf0(X0, slsdtgt0(xc)) & ! [X4] : (~ (sdtasdt0(xc, X4) = X0) | ~ aElement0(X4))))), inference(nnf_transformation, [], [f108])).
fof(f108, plain, (~ aIdeal0(slsdtgt0(xc)) & ? [X11] : ((? [X12] : (~ aElementOf0(sdtasdt0(X12, X11), slsdtgt0(xc)) & aElement0(X12)) | ? [X13] : (~ aElementOf0(sdtpldt0(X11, X13), slsdtgt0(xc)) & aElementOf0(X13, slsdtgt0(xc)))) & aElementOf0(X11, slsdtgt0(xc))) & ! [X9] : (aElementOf0(X9, slsdtgt0(xc)) <=> ? [X10] : ((sdtasdt0(xc, X10) = X9) & aElement0(X10))) & aSet0(slsdtgt0(xc)) & ! [X0, X1, X2] : (sP5(X1, X2, X0) | ~ aElement0(X2) | (~ aElementOf0(X1, slsdtgt0(xc)) & ! [X3] : (~ (sdtasdt0(xc, X3) = X1) | ~ aElement0(X3))) | (~ aElementOf0(X0, slsdtgt0(xc)) & ! [X4] : (~ (sdtasdt0(xc, X4) = X0) | ~ aElement0(X4))))), inference(definition_folding, [], [f100, e107, e106, e105, e104])).
fof(f104, plain, ! [X1, X0] : (? [X8] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X8)) & aElement0(X8)) | ~ sP2(X1, X0)), inference(usedef, [], [e104])).
fof(e104, plain, ! [X1, X0] : (sP2(X1, X0) <=> ? [X8] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X8)) & aElement0(X8))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f105, plain, ! [X0, X2] : (? [X7] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X7)) & aElement0(X7)) | ~ sP3(X0, X2)), inference(usedef, [], [e105])).
fof(e105, plain, ! [X0, X2] : (sP3(X0, X2) <=> ? [X7] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X7)) & aElement0(X7))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f106, plain, ! [X0, X2, X1, X5] : (? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & sP2(X1, X0) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6)) | ~ sP4(X0, X2, X1, X5)), inference(usedef, [], [e106])).
fof(e106, plain, ! [X0, X2, X1, X5] : (sP4(X0, X2, X1, X5) <=> ? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & sP2(X1, X0) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f107, plain, ! [X1, X2, X0] : (? [X5] : (sP4(X0, X2, X1, X5) & (sdtasdt0(xc, X5) = X0) & aElement0(X5)) | ~ sP5(X1, X2, X0)), inference(usedef, [], [e107])).
fof(e107, plain, ! [X1, X2, X0] : (sP5(X1, X2, X0) <=> ? [X5] : (sP4(X0, X2, X1, X5) & (sdtasdt0(xc, X5) = X0) & aElement0(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f100, plain, (~ aIdeal0(slsdtgt0(xc)) & ? [X11] : ((? [X12] : (~ aElementOf0(sdtasdt0(X12, X11), slsdtgt0(xc)) & aElement0(X12)) | ? [X13] : (~ aElementOf0(sdtpldt0(X11, X13), slsdtgt0(xc)) & aElementOf0(X13, slsdtgt0(xc)))) & aElementOf0(X11, slsdtgt0(xc))) & ! [X9] : (aElementOf0(X9, slsdtgt0(xc)) <=> ? [X10] : ((sdtasdt0(xc, X10) = X9) & aElement0(X10))) & aSet0(slsdtgt0(xc)) & ! [X0, X1, X2] : (? [X5] : (? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & ? [X7] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X7)) & aElement0(X7)) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & ? [X8] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X8)) & aElement0(X8)) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6)) & (sdtasdt0(xc, X5) = X0) & aElement0(X5)) | ~ aElement0(X2) | (~ aElementOf0(X1, slsdtgt0(xc)) & ! [X3] : (~ (sdtasdt0(xc, X3) = X1) | ~ aElement0(X3))) | (~ aElementOf0(X0, slsdtgt0(xc)) & ! [X4] : (~ (sdtasdt0(xc, X4) = X0) | ~ aElement0(X4))))), inference(flattening, [], [f99])).
fof(f99, plain, (((~ aIdeal0(slsdtgt0(xc)) & ? [X11] : ((? [X12] : (~ aElementOf0(sdtasdt0(X12, X11), slsdtgt0(xc)) & aElement0(X12)) | ? [X13] : (~ aElementOf0(sdtpldt0(X11, X13), slsdtgt0(xc)) & aElementOf0(X13, slsdtgt0(xc)))) & aElementOf0(X11, slsdtgt0(xc)))) & (! [X9] : (aElementOf0(X9, slsdtgt0(xc)) <=> ? [X10] : ((sdtasdt0(xc, X10) = X9) & aElement0(X10))) & aSet0(slsdtgt0(xc)))) & ! [X0, X1, X2] : (? [X5] : (? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & ? [X7] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X7)) & aElement0(X7)) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & ? [X8] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X8)) & aElement0(X8)) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6)) & (sdtasdt0(xc, X5) = X0) & aElement0(X5)) | (~ aElement0(X2) | (~ aElementOf0(X1, slsdtgt0(xc)) & ! [X3] : (~ (sdtasdt0(xc, X3) = X1) | ~ aElement0(X3))) | (~ aElementOf0(X0, slsdtgt0(xc)) & ! [X4] : (~ (sdtasdt0(xc, X4) = X0) | ~ aElement0(X4)))))), inference(ennf_transformation, [], [f48])).
fof(f48, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, slsdtgt0(xc)) | ? [X3] : ((sdtasdt0(xc, X3) = X1) & aElement0(X3))) & (aElementOf0(X0, slsdtgt0(xc)) | ? [X4] : ((sdtasdt0(xc, X4) = X0) & aElement0(X4)))) => ? [X5] : (? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & ? [X7] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X7)) & aElement0(X7)) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & ? [X8] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X8)) & aElement0(X8)) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6)) & (sdtasdt0(xc, X5) = X0) & aElement0(X5))) => ((! [X9] : (aElementOf0(X9, slsdtgt0(xc)) <=> ? [X10] : ((sdtasdt0(xc, X10) = X9) & aElement0(X10))) & aSet0(slsdtgt0(xc))) => (aIdeal0(slsdtgt0(xc)) | ! [X11] : (aElementOf0(X11, slsdtgt0(xc)) => (! [X12] : (aElement0(X12) => aElementOf0(sdtasdt0(X12, X11), slsdtgt0(xc))) & ! [X13] : (aElementOf0(X13, slsdtgt0(xc)) => aElementOf0(sdtpldt0(X11, X13), slsdtgt0(xc)))))))), inference(rectify, [], [f40])).
fof(f40, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, slsdtgt0(xc)) | ? [X3] : ((sdtasdt0(xc, X3) = X1) & aElement0(X3))) & (aElementOf0(X0, slsdtgt0(xc)) | ? [X3] : ((sdtasdt0(xc, X3) = X0) & aElement0(X3)))) => ? [X3] : (? [X4] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & ? [X5] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X5)) & aElement0(X5)) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & ? [X5] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X5)) & aElement0(X5)) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X3, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X3, X4))) & (sdtasdt0(xc, X4) = X1) & aElement0(X4)) & (sdtasdt0(xc, X3) = X0) & aElement0(X3))) => ((! [X0] : (aElementOf0(X0, slsdtgt0(xc)) <=> ? [X1] : ((sdtasdt0(xc, X1) = X0) & aElement0(X1))) & aSet0(slsdtgt0(xc))) => (aIdeal0(slsdtgt0(xc)) | ! [X0] : (aElementOf0(X0, slsdtgt0(xc)) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc))) & ! [X1] : (aElementOf0(X1, slsdtgt0(xc)) => aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)))))))), inference(negated_conjecture, [], [f39])).
fof(f39, plain, ~ (! [X0, X1, X2] : ((aElement0(X2) & (aElementOf0(X1, slsdtgt0(xc)) | ? [X3] : ((sdtasdt0(xc, X3) = X1) & aElement0(X3))) & (aElementOf0(X0, slsdtgt0(xc)) | ? [X3] : ((sdtasdt0(xc, X3) = X0) & aElement0(X3)))) => ? [X3] : (? [X4] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & ? [X5] : ((sdtasdt0(X2, X0) = sdtasdt0(xc, X5)) & aElement0(X5)) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & ? [X5] : ((sdtpldt0(X0, X1) = sdtasdt0(xc, X5)) & aElement0(X5)) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X3, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X3, X4))) & (sdtasdt0(xc, X4) = X1) & aElement0(X4)) & (sdtasdt0(xc, X3) = X0) & aElement0(X3))) => ((! [X0] : (aElementOf0(X0, slsdtgt0(xc)) <=> ? [X1] : ((sdtasdt0(xc, X1) = X0) & aElement0(X1))) & aSet0(slsdtgt0(xc))) => (aIdeal0(slsdtgt0(xc)) | ! [X0] : (aElementOf0(X0, slsdtgt0(xc)) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc))) & ! [X1] : (aElementOf0(X1, slsdtgt0(xc)) => aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)))))))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG106+2.p', m__)).
fof(f376, plain, (aElement0(sK30) | ~ aSet0(slsdtgt0(xc))), inference(resolution, [], [f204, f294])).
fof(f294, plain, aElementOf0(sK30, slsdtgt0(xc)), inference(cnf_transformation, [], [f180])).
fof(f204, plain, ! [X0, X1] : (~ aElementOf0(X1, X0) | aElement0(X1) | ~ aSet0(X0)), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ! [X0] : (! [X1] : (aElement0(X1) | ~ aElementOf0(X1, X0)) | ~ aSet0(X0)), inference(ennf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1, X0) => aElement0(X1))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG106+2.p', mEOfElem)).
fof(f186, plain, ! [X0, X1] : (~ aElement0(X1) | (sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aElement0(X0)), inference(cnf_transformation, [], [f57])).
fof(f57, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f56])).
fof(f56, plain, ! [X0, X1] : ((sdtpldt0(X0, X1) = sdtpldt0(X1, X0)) | (~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => (sdtpldt0(X0, X1) = sdtpldt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG106+2.p', mAddComm)).
fof(f144370, plain, (aElement0(sK32) | ~ spl34_1190), inference(avatar_component_clause, [], [f144369])).
fof(f436454, plain, (! [X16] : (~ aElement0(X16) | aElementOf0(sdtpldt0(sK32, sK30), slsdtgt0(xc))) | ~ spl34_3), inference(resolution, [], [f431254, f845])).
fof(f845, plain, ! [X4, X5, X3] : (~ sP5(X3, X4, X5) | aElementOf0(sdtpldt0(X5, X3), slsdtgt0(xc))), inference(resolution, [], [f273, f279])).
fof(f279, plain, ! [X2, X0, X3, X1] : (~ sP4(X0, X1, X2, X3) | aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc))), inference(cnf_transformation, [], [f165])).
fof(f165, plain, ! [X0, X1, X2, X3] : ((aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & sP2(X2, X0) & (sdtasdt0(X1, X0) = sdtasdt0(xc, sdtasdt0(X3, X1))) & (sdtpldt0(X0, X2) = sdtasdt0(xc, sdtpldt0(X3, sK27(X0, X1, X2, X3)))) & (sdtasdt0(xc, sK27(X0, X1, X2, X3)) = X2) & aElement0(sK27(X0, X1, X2, X3))) | ~ sP4(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK27])], [f163, f164])).
fof(f164, plain, ! [X3, X2, X1, X0] : (? [X4] : (aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & sP2(X2, X0) & (sdtasdt0(X1, X0) = sdtasdt0(xc, sdtasdt0(X3, X1))) & (sdtasdt0(xc, sdtpldt0(X3, X4)) = sdtpldt0(X0, X2)) & (sdtasdt0(xc, X4) = X2) & aElement0(X4)) => (aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & sP2(X2, X0) & (sdtasdt0(X1, X0) = sdtasdt0(xc, sdtasdt0(X3, X1))) & (sdtpldt0(X0, X2) = sdtasdt0(xc, sdtpldt0(X3, sK27(X0, X1, X2, X3)))) & (sdtasdt0(xc, sK27(X0, X1, X2, X3)) = X2) & aElement0(sK27(X0, X1, X2, X3)))), introduced(choice_axiom, [])).
fof(f163, plain, ! [X0, X1, X2, X3] : (? [X4] : (aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc)) & sP3(X0, X1) & aElementOf0(sdtpldt0(X0, X2), slsdtgt0(xc)) & sP2(X2, X0) & (sdtasdt0(X1, X0) = sdtasdt0(xc, sdtasdt0(X3, X1))) & (sdtasdt0(xc, sdtpldt0(X3, X4)) = sdtpldt0(X0, X2)) & (sdtasdt0(xc, X4) = X2) & aElement0(X4)) | ~ sP4(X0, X1, X2, X3)), inference(rectify, [], [f162])).
fof(f162, plain, ! [X0, X2, X1, X5] : (? [X6] : (aElementOf0(sdtasdt0(X2, X0), slsdtgt0(xc)) & sP3(X0, X2) & aElementOf0(sdtpldt0(X0, X1), slsdtgt0(xc)) & sP2(X1, X0) & (sdtasdt0(X2, X0) = sdtasdt0(xc, sdtasdt0(X5, X2))) & (sdtpldt0(X0, X1) = sdtasdt0(xc, sdtpldt0(X5, X6))) & (sdtasdt0(xc, X6) = X1) & aElement0(X6)) | ~ sP4(X0, X2, X1, X5)), inference(nnf_transformation, [], [f106])).
fof(f273, plain, ! [X2, X0, X1] : (sP4(X2, X1, X0, sK26(X0, X1, X2)) | ~ sP5(X0, X1, X2)), inference(cnf_transformation, [], [f161])).
fof(f161, plain, ! [X0, X1, X2] : ((sP4(X2, X1, X0, sK26(X0, X1, X2)) & (sdtasdt0(xc, sK26(X0, X1, X2)) = X2) & aElement0(sK26(X0, X1, X2))) | ~ sP5(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26])], [f159, f160])).
fof(f160, plain, ! [X2, X1, X0] : (? [X3] : (sP4(X2, X1, X0, X3) & (sdtasdt0(xc, X3) = X2) & aElement0(X3)) => (sP4(X2, X1, X0, sK26(X0, X1, X2)) & (sdtasdt0(xc, sK26(X0, X1, X2)) = X2) & aElement0(sK26(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f159, plain, ! [X0, X1, X2] : (? [X3] : (sP4(X2, X1, X0, X3) & (sdtasdt0(xc, X3) = X2) & aElement0(X3)) | ~ sP5(X0, X1, X2)), inference(rectify, [], [f158])).
fof(f158, plain, ! [X1, X2, X0] : (? [X5] : (sP4(X0, X2, X1, X5) & (sdtasdt0(xc, X5) = X0) & aElement0(X5)) | ~ sP5(X1, X2, X0)), inference(nnf_transformation, [], [f107])).
fof(f431254, plain, (! [X2] : (sP5(sK30, X2, sK32) | ~ aElement0(X2)) | ~ spl34_3), inference(resolution, [], [f330, f955])).
fof(f955, plain, ! [X8, X7] : (~ aElementOf0(X8, slsdtgt0(xc)) | sP5(sK30, X7, X8) | ~ aElement0(X7)), inference(resolution, [], [f289, f294])).
fof(f289, plain, ! [X6, X8, X7] : (~ aElementOf0(X7, slsdtgt0(xc)) | ~ aElement0(X8) | sP5(X7, X8, X6) | ~ aElementOf0(X6, slsdtgt0(xc))), inference(cnf_transformation, [], [f180])).
fof(f330, plain, (aElementOf0(sK32, slsdtgt0(xc)) | ~ spl34_3), inference(avatar_component_clause, [], [f328])).
fof(f430245, plain, (spl34_2 | ~ spl34_4 | ~ spl34_1279), inference(avatar_contradiction_clause, [], [f430244])).
fof(f430244, plain, ($false | (spl34_2 | ~ spl34_4 | ~ spl34_1279)), inference(subsumption_resolution, [], [f430243, f335])).
fof(f335, plain, (aElement0(sK31) | ~ spl34_4), inference(avatar_component_clause, [], [f333])).
fof(f333, plain, (spl34_4 <=> aElement0(sK31)), introduced(avatar_definition, [new_symbols(naming, [spl34_4])])).
fof(f430243, plain, (~ aElement0(sK31) | (spl34_2 | ~ spl34_4 | ~ spl34_1279)), inference(subsumption_resolution, [], [f430239, f180428])).
fof(f180428, plain, (~ aElementOf0(sdtasdt0(sK30, sK31), slsdtgt0(xc)) | (spl34_2 | ~ spl34_4 | ~ spl34_1279)), inference(backward_demodulation, [], [f325, f180415])).
fof(f180415, plain, ((sdtasdt0(sK31, sK30) = sdtasdt0(sK30, sK31)) | (~ spl34_4 | ~ spl34_1279)), inference(resolution, [], [f177494, f176246])).
fof(f176246, plain, (aElement0(sK30) | ~ spl34_1279), inference(avatar_component_clause, [], [f176245])).
fof(f176245, plain, (spl34_1279 <=> aElement0(sK30)), introduced(avatar_definition, [new_symbols(naming, [spl34_1279])])).
fof(f177494, plain, (! [X3] : (~ aElement0(X3) | (sdtasdt0(X3, sK31) = sdtasdt0(sK31, X3))) | ~ spl34_4), inference(resolution, [], [f335, f192])).
fof(f192, plain, ! [X0, X1] : (~ aElement0(X1) | (sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aElement0(X0)), inference(cnf_transformation, [], [f63])).
fof(f63, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | ~ aElement0(X1) | ~ aElement0(X0)), inference(flattening, [], [f62])).
fof(f62, plain, ! [X0, X1] : ((sdtasdt0(X0, X1) = sdtasdt0(X1, X0)) | (~ aElement0(X1) | ~ aElement0(X0))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ! [X0, X1] : ((aElement0(X1) & aElement0(X0)) => (sdtasdt0(X0, X1) = sdtasdt0(X1, X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG106+2.p', mMulComm)).
fof(f325, plain, (~ aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc)) | spl34_2), inference(avatar_component_clause, [], [f323])).
fof(f323, plain, (spl34_2 <=> aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc))), introduced(avatar_definition, [new_symbols(naming, [spl34_2])])).
fof(f430239, plain, (aElementOf0(sdtasdt0(sK30, sK31), slsdtgt0(xc)) | ~ aElement0(sK31) | (~ spl34_4 | ~ spl34_1279)), inference(superposition, [], [f6564, f180415])).
fof(f6564, plain, ! [X3] : (aElementOf0(sdtasdt0(X3, sK30), slsdtgt0(xc)) | ~ aElement0(X3)), inference(resolution, [], [f3727, f844])).
fof(f844, plain, ! [X2, X0, X1] : (~ sP5(X0, X1, X2) | aElementOf0(sdtasdt0(X1, X2), slsdtgt0(xc))), inference(resolution, [], [f273, f281])).
fof(f281, plain, ! [X2, X0, X3, X1] : (~ sP4(X0, X1, X2, X3) | aElementOf0(sdtasdt0(X1, X0), slsdtgt0(xc))), inference(cnf_transformation, [], [f165])).
fof(f3727, plain, ! [X5] : (sP5(sK30, X5, sK30) | ~ aElement0(X5)), inference(resolution, [], [f955, f294])).
fof(f177474, plain, (spl34_1190 | ~ spl34_3), inference(avatar_split_clause, [], [f177473, f328, f144369])).
fof(f177473, plain, (aElement0(sK32) | ~ spl34_3), inference(subsumption_resolution, [], [f175512, f290])).
fof(f175512, plain, (aElement0(sK32) | ~ aSet0(slsdtgt0(xc)) | ~ spl34_3), inference(resolution, [], [f330, f204])).
fof(f177424, plain, spl34_1279, inference(avatar_split_clause, [], [f378, f176245])).
fof(f2746, plain, ~ spl34_58, inference(avatar_contradiction_clause, [], [f2687])).
fof(f2687, plain, ($false | ~ spl34_58), inference(resolution, [], [f2681, f181])).
fof(f181, plain, aElement0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aElement0(sz00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG106+2.p', mSortsC)).
fof(f2681, plain, (! [X5] : ~ aElement0(X5) | ~ spl34_58), inference(avatar_component_clause, [], [f2680])).
fof(f337, plain, (spl34_3 | spl34_4), inference(avatar_split_clause, [], [f295, f333, f328])).
fof(f295, plain, (aElement0(sK31) | aElementOf0(sK32, slsdtgt0(xc))), inference(cnf_transformation, [], [f180])).
fof(f336, plain, (~ spl34_1 | spl34_4), inference(avatar_split_clause, [], [f296, f333, f319])).
fof(f296, plain, (aElement0(sK31) | ~ aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc))), inference(cnf_transformation, [], [f180])).
fof(f331, plain, (spl34_3 | ~ spl34_2), inference(avatar_split_clause, [], [f297, f323, f328])).
fof(f297, plain, (~ aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc)) | aElementOf0(sK32, slsdtgt0(xc))), inference(cnf_transformation, [], [f180])).
fof(f326, plain, (~ spl34_1 | ~ spl34_2), inference(avatar_split_clause, [], [f298, f323, f319])).
fof(f298, plain, (~ aElementOf0(sdtasdt0(sK31, sK30), slsdtgt0(xc)) | ~ aElementOf0(sdtpldt0(sK30, sK32), slsdtgt0(xc))), inference(cnf_transformation, [], [f180])).