fof(f1565, plain, $false, inference(avatar_sat_refutation, [], [f420, f429, f434, f453, f754, f1510, f1513, f1564])).
fof(f1564, plain, (spl48_6 | ~ spl48_7 | ~ spl48_48), inference(avatar_contradiction_clause, [], [f1563])).
fof(f1563, plain, ($false | (spl48_6 | ~ spl48_7 | ~ spl48_48)), inference(subsumption_resolution, [], [f1560, f714])).
fof(f714, plain, (sP5(sK42) | (spl48_6 | ~ spl48_7)), inference(subsumption_resolution, [], [f711, f428])).
fof(f428, plain, (~ (sz00 = sK42) | spl48_6), inference(avatar_component_clause, [], [f426])).
fof(f426, plain, (spl48_6 <=> (sz00 = sK42)), introduced(avatar_definition, [new_symbols(naming, [spl48_6])])).
fof(f711, plain, ((sz00 = sK42) | sP5(sK42) | ~ spl48_7), inference(resolution, [], [f381, f433])).
fof(f433, plain, (aElementOf0(sK42, xI) | ~ spl48_7), inference(avatar_component_clause, [], [f431])).
fof(f431, plain, (spl48_7 <=> aElementOf0(sK42, xI)), introduced(avatar_definition, [new_symbols(naming, [spl48_7])])).
fof(f381, plain, ! [X0] : (~ aElementOf0(X0, xI) | (sz00 = X0) | sP5(X0)), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ! [X0] : (sP5(X0) | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa))))), inference(rectify, [], [f127])).
fof(f127, plain, ! [X0] : (sP5(X0) | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X0) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))), inference(definition_folding, [], [f117, e126])).
fof(f126, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) | ~ sP5(X0)), inference(usedef, [], [e126])).
fof(e126, plain, ! [X0] : (sP5(X0) <=> ? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f117, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X0) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))), inference(flattening, [], [f116])).
fof(f116, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & (~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))) | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X0) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ~ ? [X0] : (! [X1] : ((~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0))) & ~ (sz00 = X0) & (aElementOf0(X0, xI) | ? [X4, X5] : ((sdtpldt0(X4, X5) = X0) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))), inference(rectify, [], [f47])).
fof(f47, plain, ~ ? [X0] : (! [X1] : ((~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0))) & ~ (sz00 = X0) & (aElementOf0(X0, xI) | ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))))), inference(negated_conjecture, [], [f46])).
fof(f46, plain, ~ ? [X0] : (! [X1] : ((~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0))) & ~ (sz00 = X0) & (aElementOf0(X0, xI) | ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG112+4.p', m__)).
fof(f1560, plain, (~ sP5(sK42) | ~ spl48_48), inference(trivial_inequality_removal, [], [f1558])).
fof(f1558, plain, (~ (sz00 = sz00) | ~ sP5(sK42) | ~ spl48_48), inference(superposition, [], [f378, f1505])).
fof(f1505, plain, ((sz00 = sK45(sK42)) | ~ spl48_48), inference(avatar_component_clause, [], [f1503])).
fof(f1503, plain, (spl48_48 <=> (sz00 = sK45(sK42))), introduced(avatar_definition, [new_symbols(naming, [spl48_48])])).
fof(f378, plain, ! [X0] : (~ (sz00 = sK45(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f211])).
fof(f211, plain, ! [X0] : ((iLess0(sbrdtbr0(sK45(X0)), sbrdtbr0(X0)) & ~ (sz00 = sK45(X0)) & aElementOf0(sK45(X0), xI) & ((sK45(X0) = sdtpldt0(sK46(X0), sK47(X0))) & aElementOf0(sK47(X0), slsdtgt0(xb)) & aElementOf0(sK46(X0), slsdtgt0(xa)))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45, sK46, sK47])], [f208, f210, f209])).
fof(f209, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) => (iLess0(sbrdtbr0(sK45(X0)), sbrdtbr0(X0)) & ~ (sz00 = sK45(X0)) & aElementOf0(sK45(X0), xI) & ? [X3, X2] : ((sdtpldt0(X2, X3) = sK45(X0)) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))), introduced(choice_axiom, [])).
fof(f210, plain, ! [X0] : (? [X3, X2] : ((sdtpldt0(X2, X3) = sK45(X0)) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))) => ((sK45(X0) = sdtpldt0(sK46(X0), sK47(X0))) & aElementOf0(sK47(X0), slsdtgt0(xb)) & aElementOf0(sK46(X0), slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f208, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) | ~ sP5(X0)), inference(nnf_transformation, [], [f126])).
fof(f1513, plain, (spl48_6 | ~ spl48_7 | spl48_49), inference(avatar_contradiction_clause, [], [f1512])).
fof(f1512, plain, ($false | (spl48_6 | ~ spl48_7 | spl48_49)), inference(subsumption_resolution, [], [f1511, f714])).
fof(f1511, plain, (~ sP5(sK42) | spl48_49), inference(resolution, [], [f1509, f377])).
fof(f377, plain, ! [X0] : (aElementOf0(sK45(X0), xI) | ~ sP5(X0)), inference(cnf_transformation, [], [f211])).
fof(f1509, plain, (~ aElementOf0(sK45(sK42), xI) | spl48_49), inference(avatar_component_clause, [], [f1507])).
fof(f1507, plain, (spl48_49 <=> aElementOf0(sK45(sK42), xI)), introduced(avatar_definition, [new_symbols(naming, [spl48_49])])).
fof(f1510, plain, (spl48_48 | ~ spl48_49 | ~ spl48_4 | spl48_6 | ~ spl48_7), inference(avatar_split_clause, [], [f1501, f431, f426, f418, f1507, f1503])).
fof(f418, plain, (spl48_4 <=> ! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK42)) | ~ aElementOf0(X1, xI) | (sz00 = X1))), introduced(avatar_definition, [new_symbols(naming, [spl48_4])])).
fof(f1501, plain, (~ aElementOf0(sK45(sK42), xI) | (sz00 = sK45(sK42)) | (~ spl48_4 | spl48_6 | ~ spl48_7)), inference(subsumption_resolution, [], [f1500, f714])).
fof(f1500, plain, (~ aElementOf0(sK45(sK42), xI) | (sz00 = sK45(sK42)) | ~ sP5(sK42) | ~ spl48_4), inference(resolution, [], [f419, f379])).
fof(f379, plain, ! [X0] : (iLess0(sbrdtbr0(sK45(X0)), sbrdtbr0(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f211])).
fof(f419, plain, (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK42)) | ~ aElementOf0(X1, xI) | (sz00 = X1)) | ~ spl48_4), inference(avatar_component_clause, [], [f418])).
fof(f754, plain, ~ spl48_11, inference(avatar_contradiction_clause, [], [f753])).
fof(f753, plain, ($false | ~ spl48_11), inference(subsumption_resolution, [], [f750, f364])).
fof(f364, plain, ~ (sz00 = sK37), inference(cnf_transformation, [], [f202])).
fof(f202, plain, (~ (sz00 = sK37) & aElementOf0(sK37, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ((sK37 = sdtpldt0(sK38, sK39)) & aElementOf0(sK39, slsdtgt0(xb)) & aElementOf0(sK38, slsdtgt0(xa))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (((sdtasdt0(xb, sK40(X3)) = X3) & aElement0(sK40(X3))) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X6] : ((aElementOf0(X6, slsdtgt0(xa)) | ! [X7] : (~ (sdtasdt0(xa, X7) = X6) | ~ aElement0(X7))) & (((sdtasdt0(xa, sK41(X6)) = X6) & aElement0(sK41(X6))) | ~ aElementOf0(X6, slsdtgt0(xa))))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK37, sK38, sK39, sK40, sK41])], [f197, f201, f200, f199, f198])).
fof(f198, plain, (? [X0] : (~ (sz00 = X0) & aElementOf0(X0, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X5] : ((sdtasdt0(xb, X5) = X3) & aElement0(X5)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X6] : ((aElementOf0(X6, slsdtgt0(xa)) | ! [X7] : (~ (sdtasdt0(xa, X7) = X6) | ~ aElement0(X7))) & (? [X8] : ((sdtasdt0(xa, X8) = X6) & aElement0(X8)) | ~ aElementOf0(X6, slsdtgt0(xa))))) => (~ (sz00 = sK37) & aElementOf0(sK37, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X2, X1] : ((sdtpldt0(X1, X2) = sK37) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X5] : ((sdtasdt0(xb, X5) = X3) & aElement0(X5)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X6] : ((aElementOf0(X6, slsdtgt0(xa)) | ! [X7] : (~ (sdtasdt0(xa, X7) = X6) | ~ aElement0(X7))) & (? [X8] : ((sdtasdt0(xa, X8) = X6) & aElement0(X8)) | ~ aElementOf0(X6, slsdtgt0(xa)))))), introduced(choice_axiom, [])).
fof(f199, plain, (? [X2, X1] : ((sdtpldt0(X1, X2) = sK37) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) => ((sK37 = sdtpldt0(sK38, sK39)) & aElementOf0(sK39, slsdtgt0(xb)) & aElementOf0(sK38, slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f200, plain, ! [X3] : (? [X5] : ((sdtasdt0(xb, X5) = X3) & aElement0(X5)) => ((sdtasdt0(xb, sK40(X3)) = X3) & aElement0(sK40(X3)))), introduced(choice_axiom, [])).
fof(f201, plain, ! [X6] : (? [X8] : ((sdtasdt0(xa, X8) = X6) & aElement0(X8)) => ((sdtasdt0(xa, sK41(X6)) = X6) & aElement0(sK41(X6)))), introduced(choice_axiom, [])).
fof(f197, plain, ? [X0] : (~ (sz00 = X0) & aElementOf0(X0, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X5] : ((sdtasdt0(xb, X5) = X3) & aElement0(X5)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X6] : ((aElementOf0(X6, slsdtgt0(xa)) | ! [X7] : (~ (sdtasdt0(xa, X7) = X6) | ~ aElement0(X7))) & (? [X8] : ((sdtasdt0(xa, X8) = X6) & aElement0(X8)) | ~ aElementOf0(X6, slsdtgt0(xa))))), inference(rectify, [], [f196])).
fof(f196, plain, ? [X0] : (~ (sz00 = X0) & aElementOf0(X0, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xa)) | ! [X6] : (~ (sdtasdt0(xa, X6) = X5) | ~ aElement0(X6))) & (? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6)) | ~ aElementOf0(X5, slsdtgt0(xa))))), inference(nnf_transformation, [], [f58])).
fof(f58, plain, ? [X0] : (~ (sz00 = X0) & aElementOf0(X0, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X3] : (aElementOf0(X3, slsdtgt0(xb)) <=> ? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4))) & ! [X5] : (aElementOf0(X5, slsdtgt0(xa)) <=> ? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6)))), inference(rectify, [], [f44])).
fof(f44, plain, ? [X0] : (~ (sz00 = X0) & aElementOf0(X0, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) & ! [X1] : (aElementOf0(X1, slsdtgt0(xb)) <=> ? [X2] : ((sdtasdt0(xb, X2) = X1) & aElement0(X2))) & ! [X1] : (aElementOf0(X1, slsdtgt0(xa)) <=> ? [X2] : ((sdtasdt0(xa, X2) = X1) & aElement0(X2)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG112+4.p', m__2228)).
fof(f750, plain, ((sz00 = sK37) | ~ spl48_11), inference(resolution, [], [f452, f412])).
fof(f412, plain, aElementOf0(sK37, xI), inference(forward_demodulation, [], [f363, f341])).
fof(f341, plain, (xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))), inference(cnf_transformation, [], [f190])).
fof(f190, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (((sdtpldt0(sK29(X0), sK30(X0)) = X0) & aElementOf0(sK30(X0), slsdtgt0(xb)) & aElementOf0(sK29(X0), slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xb)) | ! [X6] : (~ (sdtasdt0(xb, X6) = X5) | ~ aElement0(X6))) & (((sdtasdt0(xb, sK31(X5)) = X5) & aElement0(sK31(X5))) | ~ aElementOf0(X5, slsdtgt0(xb)))) & ! [X8] : ((aElementOf0(X8, slsdtgt0(xa)) | ! [X9] : (~ (sdtasdt0(xa, X9) = X8) | ~ aElement0(X9))) & (((sdtasdt0(xa, sK32(X8)) = X8) & aElement0(sK32(X8))) | ~ aElementOf0(X8, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X11] : ((! [X12] : (aElementOf0(sdtasdt0(X12, X11), xI) | ~ aElement0(X12)) & ! [X13] : (aElementOf0(sdtpldt0(X11, X13), xI) | ~ aElementOf0(X13, xI))) | ~ aElementOf0(X11, xI)) & aSet0(xI)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30, sK31, sK32])], [f186, f189, f188, f187])).
fof(f187, plain, ! [X0] : (? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa))) => ((sdtpldt0(sK29(X0), sK30(X0)) = X0) & aElementOf0(sK30(X0), slsdtgt0(xb)) & aElementOf0(sK29(X0), slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f188, plain, ! [X5] : (? [X7] : ((sdtasdt0(xb, X7) = X5) & aElement0(X7)) => ((sdtasdt0(xb, sK31(X5)) = X5) & aElement0(sK31(X5)))), introduced(choice_axiom, [])).
fof(f189, plain, ! [X8] : (? [X10] : ((sdtasdt0(xa, X10) = X8) & aElement0(X10)) => ((sdtasdt0(xa, sK32(X8)) = X8) & aElement0(sK32(X8)))), introduced(choice_axiom, [])).
fof(f186, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xb)) | ! [X6] : (~ (sdtasdt0(xb, X6) = X5) | ~ aElement0(X6))) & (? [X7] : ((sdtasdt0(xb, X7) = X5) & aElement0(X7)) | ~ aElementOf0(X5, slsdtgt0(xb)))) & ! [X8] : ((aElementOf0(X8, slsdtgt0(xa)) | ! [X9] : (~ (sdtasdt0(xa, X9) = X8) | ~ aElement0(X9))) & (? [X10] : ((sdtasdt0(xa, X10) = X8) & aElement0(X10)) | ~ aElementOf0(X8, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X11] : ((! [X12] : (aElementOf0(sdtasdt0(X12, X11), xI) | ~ aElement0(X12)) & ! [X13] : (aElementOf0(sdtpldt0(X11, X13), xI) | ~ aElementOf0(X13, xI))) | ~ aElementOf0(X11, xI)) & aSet0(xI)), inference(rectify, [], [f185])).
fof(f185, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xa)) | ! [X6] : (~ (sdtasdt0(xa, X6) = X5) | ~ aElement0(X6))) & (? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6)) | ~ aElementOf0(X5, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X7] : ((! [X8] : (aElementOf0(sdtasdt0(X8, X7), xI) | ~ aElement0(X8)) & ! [X9] : (aElementOf0(sdtpldt0(X7, X9), xI) | ~ aElementOf0(X9, xI))) | ~ aElementOf0(X7, xI)) & aSet0(xI)), inference(nnf_transformation, [], [f113])).
fof(f113, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X3] : (aElementOf0(X3, slsdtgt0(xb)) <=> ? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4))) & ! [X5] : (aElementOf0(X5, slsdtgt0(xa)) <=> ? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6))) & aIdeal0(xI) & ! [X7] : ((! [X8] : (aElementOf0(sdtasdt0(X8, X7), xI) | ~ aElement0(X8)) & ! [X9] : (aElementOf0(sdtpldt0(X7, X9), xI) | ~ aElementOf0(X9, xI))) | ~ aElementOf0(X7, xI)) & aSet0(xI)), inference(ennf_transformation, [], [f56])).
fof(f56, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X3] : (aElementOf0(X3, slsdtgt0(xb)) <=> ? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4))) & ! [X5] : (aElementOf0(X5, slsdtgt0(xa)) <=> ? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6))) & aIdeal0(xI) & ! [X7] : (aElementOf0(X7, xI) => (! [X8] : (aElement0(X8) => aElementOf0(sdtasdt0(X8, X7), xI)) & ! [X9] : (aElementOf0(X9, xI) => aElementOf0(sdtpldt0(X7, X9), xI)))) & aSet0(xI)), inference(rectify, [], [f42])).
fof(f42, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X0] : (aElementOf0(X0, slsdtgt0(xb)) <=> ? [X1] : ((sdtasdt0(xb, X1) = X0) & aElement0(X1))) & ! [X0] : (aElementOf0(X0, slsdtgt0(xa)) <=> ? [X1] : ((sdtasdt0(xa, X1) = X0) & aElement0(X1))) & aIdeal0(xI) & ! [X0] : (aElementOf0(X0, xI) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), xI)) & ! [X1] : (aElementOf0(X1, xI) => aElementOf0(sdtpldt0(X0, X1), xI)))) & aSet0(xI)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG112+4.p', m__2174)).
fof(f363, plain, aElementOf0(sK37, sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))), inference(cnf_transformation, [], [f202])).
fof(f452, plain, (! [X0] : (~ aElementOf0(X0, xI) | (sz00 = X0)) | ~ spl48_11), inference(avatar_component_clause, [], [f451])).
fof(f451, plain, (spl48_11 <=> ! [X0] : ((sz00 = X0) | ~ aElementOf0(X0, xI))), introduced(avatar_definition, [new_symbols(naming, [spl48_11])])).
fof(f453, plain, (spl48_11 | spl48_3), inference(avatar_split_clause, [], [f373, f414, f451])).
fof(f414, plain, (spl48_3 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl48_3])])).
fof(f373, plain, ! [X0] : (sP4 | (sz00 = X0) | ~ aElementOf0(X0, xI)), inference(cnf_transformation, [], [f125])).
fof(f125, plain, ! [X0] : (sP4 | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa))))), inference(definition_folding, [], [f115, e124])).
fof(f124, plain, (? [X3] : (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))) | ~ sP4), inference(usedef, [], [e124])).
fof(e124, plain, (sP4 <=> ? [X3] : (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f115, plain, ! [X0] : (? [X3] : (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))) | (sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa))))), inference(flattening, [], [f114])).
fof(f114, plain, ! [X0] : (? [X3] : (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3)) | ((sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa)))))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))) | ((sz00 = X0) | (~ aElementOf0(X0, xI) & ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))))), inference(ennf_transformation, [], [f59])).
fof(f59, plain, ! [X0] : ((~ (sz00 = X0) & (aElementOf0(X0, xI) | ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))))) => ? [X3] : (! [X4] : ((~ (sz00 = X4) & (aElementOf0(X4, xI) | ? [X5, X6] : ((sdtpldt0(X5, X6) = X4) & aElementOf0(X6, slsdtgt0(xb)) & aElementOf0(X5, slsdtgt0(xa))))) => ~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa))))), inference(rectify, [], [f45])).
fof(f45, plain, ! [X0] : ((~ (sz00 = X0) & (aElementOf0(X0, xI) | ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))))) => ? [X1] : (! [X2] : ((~ (sz00 = X2) & (aElementOf0(X2, xI) | ? [X3, X4] : ((sdtpldt0(X3, X4) = X2) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa))))) => ~ iLess0(sbrdtbr0(X2), sbrdtbr0(X1))) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG112+4.p', m__2351)).
fof(f434, plain, (~ spl48_3 | spl48_7), inference(avatar_split_clause, [], [f368, f431, f414])).
fof(f368, plain, (aElementOf0(sK42, xI) | ~ sP4), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK42)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = sK42) & aElementOf0(sK42, xI) & ((sK42 = sdtpldt0(sK43, sK44)) & aElementOf0(sK44, slsdtgt0(xb)) & aElementOf0(sK43, slsdtgt0(xa)))) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42, sK43, sK44])], [f204, f206, f205])).
fof(f205, plain, (? [X0] : (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X4, X5] : ((sdtpldt0(X4, X5) = X0) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa)))) => (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK42)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = sK42) & aElementOf0(sK42, xI) & ? [X5, X4] : ((sdtpldt0(X4, X5) = sK42) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))), introduced(choice_axiom, [])).
fof(f206, plain, (? [X5, X4] : ((sdtpldt0(X4, X5) = sK42) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))) => ((sK42 = sdtpldt0(sK43, sK44)) & aElementOf0(sK44, slsdtgt0(xb)) & aElementOf0(sK43, slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f204, plain, (? [X0] : (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X4, X5] : ((sdtpldt0(X4, X5) = X0) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa)))) | ~ sP4), inference(rectify, [], [f203])).
fof(f203, plain, (? [X3] : (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(X3)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X3) & aElementOf0(X3, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X3) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))) | ~ sP4), inference(nnf_transformation, [], [f124])).
fof(f429, plain, (~ spl48_3 | ~ spl48_6), inference(avatar_split_clause, [], [f369, f426, f414])).
fof(f369, plain, (~ (sz00 = sK42) | ~ sP4), inference(cnf_transformation, [], [f207])).
fof(f420, plain, (~ spl48_3 | spl48_4), inference(avatar_split_clause, [], [f371, f418, f414])).
fof(f371, plain, ! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK42)) | (sz00 = X1) | ~ aElementOf0(X1, xI) | ~ sP4), inference(cnf_transformation, [], [f207])).