fof(f1714, plain, $false, inference(avatar_sat_refutation, [], [f424, f433, f438, f457, f1325, f1494, f1624, f1631, f1703, f1713])).
fof(f1713, plain, (spl51_6 | ~ spl51_7 | ~ spl51_54), inference(avatar_contradiction_clause, [], [f1712])).
fof(f1712, plain, ($false | (spl51_6 | ~ spl51_7 | ~ spl51_54)), inference(subsumption_resolution, [], [f1710, f732])).
fof(f732, plain, (sP5(sK45) | (spl51_6 | ~ spl51_7)), inference(subsumption_resolution, [], [f728, f432])).
fof(f432, plain, (~ (sz00 = sK45) | spl51_6), inference(avatar_component_clause, [], [f430])).
fof(f430, plain, (spl51_6 <=> (sz00 = sK45)), introduced(avatar_definition, [new_symbols(naming, [spl51_6])])).
fof(f728, plain, ((sz00 = sK45) | sP5(sK45) | ~ spl51_7), inference(resolution, [], [f385, f437])).
fof(f437, plain, (aElementOf0(sK45, xI) | ~ spl51_7), inference(avatar_component_clause, [], [f435])).
fof(f435, plain, (spl51_7 <=> aElementOf0(sK45, xI)), introduced(avatar_definition, [new_symbols(naming, [spl51_7])])).
fof(f385, plain, ! [X1] : (~ aElementOf0(X1, xI) | (sz00 = X1) | sP5(X1)), inference(cnf_transformation, [], [f211])).
fof(f211, plain, (! [X1] : (sP5(X1) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ! [X4] : (sP4 | ~ iLess0(sbrdtbr0(X4), sbrdtbr0(sK48)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = sK48) & aElementOf0(sK48, xI) & ((sK48 = sdtpldt0(sK49, sK50)) & aElementOf0(sK50, slsdtgt0(xb)) & aElementOf0(sK49, slsdtgt0(xa)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48, sK49, sK50])], [f208, f210, f209])).
fof(f209, plain, (? [X0] : (! [X1] : (sP5(X1) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ! [X4] : (sP4 | ~ iLess0(sbrdtbr0(X4), sbrdtbr0(X0)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X0) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))) => (! [X1] : (sP5(X1) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ! [X4] : (sP4 | ~ iLess0(sbrdtbr0(X4), sbrdtbr0(sK48)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = sK48) & aElementOf0(sK48, xI) & ? [X8, X7] : ((sdtpldt0(X7, X8) = sK48) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa))))), introduced(choice_axiom, [])).
fof(f210, plain, (? [X8, X7] : ((sdtpldt0(X7, X8) = sK48) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa))) => ((sK48 = sdtpldt0(sK49, sK50)) & aElementOf0(sK50, slsdtgt0(xb)) & aElementOf0(sK49, slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f208, plain, ? [X0] : (! [X1] : (sP5(X1) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ! [X4] : (sP4 | ~ iLess0(sbrdtbr0(X4), sbrdtbr0(X0)) | (sz00 = X4) | (~ aElementOf0(X4, xI) & ! [X5, X6] : (~ (sdtpldt0(X5, X6) = X4) | ~ aElementOf0(X6, slsdtgt0(xb)) | ~ aElementOf0(X5, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X7, X8] : ((sdtpldt0(X7, X8) = X0) & aElementOf0(X8, slsdtgt0(xb)) & aElementOf0(X7, slsdtgt0(xa)))), inference(rectify, [], [f122])).
fof(f122, plain, ? [X0] : (! [X12] : (sP5(X12) | (sz00 = X12) | (~ aElementOf0(X12, xI) & ! [X16, X17] : (~ (sdtpldt0(X16, X17) = X12) | ~ aElementOf0(X17, slsdtgt0(xb)) | ~ aElementOf0(X16, slsdtgt0(xa))))) & ! [X3] : (sP4 | ~ iLess0(sbrdtbr0(X3), sbrdtbr0(X0)) | (sz00 = X3) | (~ aElementOf0(X3, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))), inference(definition_folding, [], [f113, e121, e120])).
fof(f120, plain, (? [X6] : (! [X7] : (~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6)) | (sz00 = X7) | (~ aElementOf0(X7, xI) & ! [X8, X9] : (~ (sdtpldt0(X8, X9) = X7) | ~ aElementOf0(X9, slsdtgt0(xb)) | ~ aElementOf0(X8, slsdtgt0(xa))))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa)))) | ~ sP4), inference(usedef, [], [e120])).
fof(e120, plain, (sP4 <=> ? [X6] : (! [X7] : (~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6)) | (sz00 = X7) | (~ aElementOf0(X7, xI) & ! [X8, X9] : (~ (sdtpldt0(X8, X9) = X7) | ~ aElementOf0(X9, slsdtgt0(xb)) | ~ aElementOf0(X8, slsdtgt0(xa))))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f121, plain, ! [X12] : (? [X13] : (iLess0(sbrdtbr0(X13), sbrdtbr0(X12)) & ~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa)))) | ~ sP5(X12)), inference(usedef, [], [e121])).
fof(e121, plain, ! [X12] : (sP5(X12) <=> ? [X13] : (iLess0(sbrdtbr0(X13), sbrdtbr0(X12)) & ~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f113, plain, ? [X0] : (! [X12] : (? [X13] : (iLess0(sbrdtbr0(X13), sbrdtbr0(X12)) & ~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa)))) | (sz00 = X12) | (~ aElementOf0(X12, xI) & ! [X16, X17] : (~ (sdtpldt0(X16, X17) = X12) | ~ aElementOf0(X17, slsdtgt0(xb)) | ~ aElementOf0(X16, slsdtgt0(xa))))) & ! [X3] : (? [X6] : (! [X7] : (~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6)) | (sz00 = X7) | (~ aElementOf0(X7, xI) & ! [X8, X9] : (~ (sdtpldt0(X8, X9) = X7) | ~ aElementOf0(X9, slsdtgt0(xb)) | ~ aElementOf0(X8, slsdtgt0(xa))))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa)))) | ~ iLess0(sbrdtbr0(X3), sbrdtbr0(X0)) | (sz00 = X3) | (~ aElementOf0(X3, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))), inference(flattening, [], [f112])).
fof(f112, plain, ? [X0] : ((! [X12] : (? [X13] : (iLess0(sbrdtbr0(X13), sbrdtbr0(X12)) & (~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa))))) | (sz00 = X12) | (~ aElementOf0(X12, xI) & ! [X16, X17] : (~ (sdtpldt0(X16, X17) = X12) | ~ aElementOf0(X17, slsdtgt0(xb)) | ~ aElementOf0(X16, slsdtgt0(xa))))) & ! [X3] : ((? [X6] : (! [X7] : (~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6)) | ((sz00 = X7) | (~ aElementOf0(X7, xI) & ! [X8, X9] : (~ (sdtpldt0(X8, X9) = X7) | ~ aElementOf0(X9, slsdtgt0(xb)) | ~ aElementOf0(X8, slsdtgt0(xa)))))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa)))) | ~ iLess0(sbrdtbr0(X3), sbrdtbr0(X0))) | ((sz00 = X3) | (~ aElementOf0(X3, xI) & ! [X4, X5] : (~ (sdtpldt0(X4, X5) = X3) | ~ aElementOf0(X5, slsdtgt0(xb)) | ~ aElementOf0(X4, slsdtgt0(xa))))))) & (~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))))), inference(ennf_transformation, [], [f58])).
fof(f58, plain, ~ ! [X0] : ((~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) => (! [X3] : ((~ (sz00 = X3) & (aElementOf0(X3, xI) | ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))) => (iLess0(sbrdtbr0(X3), sbrdtbr0(X0)) => ? [X6] : (! [X7] : ((~ (sz00 = X7) & (aElementOf0(X7, xI) | ? [X8, X9] : ((sdtpldt0(X8, X9) = X7) & aElementOf0(X9, slsdtgt0(xb)) & aElementOf0(X8, slsdtgt0(xa))))) => ~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa)))))) => ? [X12] : (! [X13] : ((~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X13), sbrdtbr0(X12))) & ~ (sz00 = X12) & (aElementOf0(X12, xI) | ? [X16, X17] : ((sdtpldt0(X16, X17) = X12) & aElementOf0(X17, slsdtgt0(xb)) & aElementOf0(X16, slsdtgt0(xa))))))), inference(rectify, [], [f46])).
fof(f46, plain, ~ ! [X0] : ((~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) => (! [X1] : ((~ (sz00 = X1) & (aElementOf0(X1, xI) | ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))) => (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) => ? [X2] : (! [X3] : ((~ (sz00 = X3) & (aElementOf0(X3, xI) | ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))) => ~ iLess0(sbrdtbr0(X3), sbrdtbr0(X2))) & ~ (sz00 = X2) & aElementOf0(X2, xI) & ? [X3, X4] : ((sdtpldt0(X3, X4) = X2) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa)))))) => ? [X1] : (! [X2] : ((~ (sz00 = X2) & aElementOf0(X2, xI) & ? [X3, X4] : ((sdtpldt0(X3, X4) = X2) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X2), sbrdtbr0(X1))) & ~ (sz00 = X1) & (aElementOf0(X1, xI) | ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))))), inference(negated_conjecture, [], [f45])).
fof(f45, plain, ~ ! [X0] : ((~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) => (! [X1] : ((~ (sz00 = X1) & (aElementOf0(X1, xI) | ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))) => (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) => ? [X2] : (! [X3] : ((~ (sz00 = X3) & (aElementOf0(X3, xI) | ? [X4, X5] : ((sdtpldt0(X4, X5) = X3) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))) => ~ iLess0(sbrdtbr0(X3), sbrdtbr0(X2))) & ~ (sz00 = X2) & aElementOf0(X2, xI) & ? [X3, X4] : ((sdtpldt0(X3, X4) = X2) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa)))))) => ? [X1] : (! [X2] : ((~ (sz00 = X2) & aElementOf0(X2, xI) & ? [X3, X4] : ((sdtpldt0(X3, X4) = X2) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa)))) => ~ iLess0(sbrdtbr0(X2), sbrdtbr0(X1))) & ~ (sz00 = X1) & (aElementOf0(X1, xI) | ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG111+4.p', m__)).
fof(f1710, plain, (~ sP5(sK45) | ~ spl51_54), inference(trivial_inequality_removal, [], [f1708])).
fof(f1708, plain, (~ (sz00 = sz00) | ~ sP5(sK45) | ~ spl51_54), inference(superposition, [], [f368, f1619])).
fof(f1619, plain, ((sz00 = sK42(sK45)) | ~ spl51_54), inference(avatar_component_clause, [], [f1617])).
fof(f1617, plain, (spl51_54 <=> (sz00 = sK42(sK45))), introduced(avatar_definition, [new_symbols(naming, [spl51_54])])).
fof(f368, plain, ! [X0] : (~ (sz00 = sK42(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ! [X0] : ((iLess0(sbrdtbr0(sK42(X0)), sbrdtbr0(X0)) & ~ (sz00 = sK42(X0)) & aElementOf0(sK42(X0), xI) & ((sK42(X0) = sdtpldt0(sK43(X0), sK44(X0))) & aElementOf0(sK44(X0), slsdtgt0(xb)) & aElementOf0(sK43(X0), slsdtgt0(xa)))) | ~ sP5(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42, sK43, sK44])], [f199, f201, f200])).
fof(f200, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) => (iLess0(sbrdtbr0(sK42(X0)), sbrdtbr0(X0)) & ~ (sz00 = sK42(X0)) & aElementOf0(sK42(X0), xI) & ? [X3, X2] : ((sdtpldt0(X2, X3) = sK42(X0)) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))))), introduced(choice_axiom, [])).
fof(f201, plain, ! [X0] : (? [X3, X2] : ((sdtpldt0(X2, X3) = sK42(X0)) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa))) => ((sK42(X0) = sdtpldt0(sK43(X0), sK44(X0))) & aElementOf0(sK44(X0), slsdtgt0(xb)) & aElementOf0(sK43(X0), slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f199, plain, ! [X0] : (? [X1] : (iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) & ~ (sz00 = X1) & aElementOf0(X1, xI) & ? [X2, X3] : ((sdtpldt0(X2, X3) = X1) & aElementOf0(X3, slsdtgt0(xb)) & aElementOf0(X2, slsdtgt0(xa)))) | ~ sP5(X0)), inference(rectify, [], [f198])).
fof(f198, plain, ! [X12] : (? [X13] : (iLess0(sbrdtbr0(X13), sbrdtbr0(X12)) & ~ (sz00 = X13) & aElementOf0(X13, xI) & ? [X14, X15] : ((sdtpldt0(X14, X15) = X13) & aElementOf0(X15, slsdtgt0(xb)) & aElementOf0(X14, slsdtgt0(xa)))) | ~ sP5(X12)), inference(nnf_transformation, [], [f121])).
fof(f1703, plain, (spl51_6 | ~ spl51_7 | spl51_55), inference(avatar_contradiction_clause, [], [f1702])).
fof(f1702, plain, ($false | (spl51_6 | ~ spl51_7 | spl51_55)), inference(subsumption_resolution, [], [f1701, f732])).
fof(f1701, plain, (~ sP5(sK45) | spl51_55), inference(resolution, [], [f1623, f367])).
fof(f367, plain, ! [X0] : (aElementOf0(sK42(X0), xI) | ~ sP5(X0)), inference(cnf_transformation, [], [f202])).
fof(f1623, plain, (~ aElementOf0(sK42(sK45), xI) | spl51_55), inference(avatar_component_clause, [], [f1621])).
fof(f1621, plain, (spl51_55 <=> aElementOf0(sK42(sK45), xI)), introduced(avatar_definition, [new_symbols(naming, [spl51_55])])).
fof(f1631, plain, ~ spl51_51, inference(avatar_contradiction_clause, [], [f1630])).
fof(f1630, plain, ($false | ~ spl51_51), inference(subsumption_resolution, [], [f1628, f733])).
fof(f733, plain, sP5(sK48), inference(subsumption_resolution, [], [f729, f381])).
fof(f381, plain, ~ (sz00 = sK48), inference(cnf_transformation, [], [f211])).
fof(f729, plain, ((sz00 = sK48) | sP5(sK48)), inference(resolution, [], [f385, f380])).
fof(f380, plain, aElementOf0(sK48, xI), inference(cnf_transformation, [], [f211])).
fof(f1628, plain, (~ sP5(sK48) | ~ spl51_51), inference(trivial_inequality_removal, [], [f1626])).
fof(f1626, plain, (~ (sz00 = sz00) | ~ sP5(sK48) | ~ spl51_51), inference(superposition, [], [f368, f1320])).
fof(f1320, plain, ((sz00 = sK42(sK48)) | ~ spl51_51), inference(avatar_component_clause, [], [f1318])).
fof(f1318, plain, (spl51_51 <=> (sz00 = sK42(sK48))), introduced(avatar_definition, [new_symbols(naming, [spl51_51])])).
fof(f1624, plain, (spl51_54 | ~ spl51_55 | ~ spl51_4 | spl51_6 | ~ spl51_7), inference(avatar_split_clause, [], [f1615, f435, f430, f422, f1621, f1617])).
fof(f422, plain, (spl51_4 <=> ! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK45)) | ~ aElementOf0(X1, xI) | (sz00 = X1))), introduced(avatar_definition, [new_symbols(naming, [spl51_4])])).
fof(f1615, plain, (~ aElementOf0(sK42(sK45), xI) | (sz00 = sK42(sK45)) | (~ spl51_4 | spl51_6 | ~ spl51_7)), inference(subsumption_resolution, [], [f1614, f732])).
fof(f1614, plain, (~ aElementOf0(sK42(sK45), xI) | (sz00 = sK42(sK45)) | ~ sP5(sK45) | ~ spl51_4), inference(resolution, [], [f423, f369])).
fof(f369, plain, ! [X0] : (iLess0(sbrdtbr0(sK42(X0)), sbrdtbr0(X0)) | ~ sP5(X0)), inference(cnf_transformation, [], [f202])).
fof(f423, plain, (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK45)) | ~ aElementOf0(X1, xI) | (sz00 = X1)) | ~ spl51_4), inference(avatar_component_clause, [], [f422])).
fof(f1494, plain, spl51_52, inference(avatar_contradiction_clause, [], [f1493])).
fof(f1493, plain, ($false | spl51_52), inference(subsumption_resolution, [], [f1492, f733])).
fof(f1492, plain, (~ sP5(sK48) | spl51_52), inference(resolution, [], [f1324, f367])).
fof(f1324, plain, (~ aElementOf0(sK42(sK48), xI) | spl51_52), inference(avatar_component_clause, [], [f1322])).
fof(f1322, plain, (spl51_52 <=> aElementOf0(sK42(sK48), xI)), introduced(avatar_definition, [new_symbols(naming, [spl51_52])])).
fof(f1325, plain, (spl51_51 | ~ spl51_52 | ~ spl51_11), inference(avatar_split_clause, [], [f1316, f455, f1322, f1318])).
fof(f455, plain, (spl51_11 <=> ! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(sK48)) | ~ aElementOf0(X4, xI) | (sz00 = X4))), introduced(avatar_definition, [new_symbols(naming, [spl51_11])])).
fof(f1316, plain, (~ aElementOf0(sK42(sK48), xI) | (sz00 = sK42(sK48)) | ~ spl51_11), inference(subsumption_resolution, [], [f1315, f733])).
fof(f1315, plain, (~ aElementOf0(sK42(sK48), xI) | (sz00 = sK42(sK48)) | ~ sP5(sK48) | ~ spl51_11), inference(resolution, [], [f456, f369])).
fof(f456, plain, (! [X4] : (~ iLess0(sbrdtbr0(X4), sbrdtbr0(sK48)) | ~ aElementOf0(X4, xI) | (sz00 = X4)) | ~ spl51_11), inference(avatar_component_clause, [], [f455])).
fof(f457, plain, (spl51_11 | spl51_3), inference(avatar_split_clause, [], [f383, f418, f455])).
fof(f418, plain, (spl51_3 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl51_3])])).
fof(f383, plain, ! [X4] : (sP4 | ~ iLess0(sbrdtbr0(X4), sbrdtbr0(sK48)) | (sz00 = X4) | ~ aElementOf0(X4, xI)), inference(cnf_transformation, [], [f211])).
fof(f438, plain, (~ spl51_3 | spl51_7), inference(avatar_split_clause, [], [f373, f435, f418])).
fof(f373, plain, (aElementOf0(sK45, xI) | ~ sP4), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ((! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK45)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = sK45) & aElementOf0(sK45, xI) & ((sK45 = sdtpldt0(sK46, sK47)) & aElementOf0(sK47, slsdtgt0(xb)) & aElementOf0(sK46, slsdtgt0(xa)))) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK45, sK46, sK47])], [f204, f206, f205])).
fof(f205, plain, (? [X0] : (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X4, X5] : ((sdtpldt0(X4, X5) = X0) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa)))) => (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK45)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = sK45) & aElementOf0(sK45, xI) & ? [X5, X4] : ((sdtpldt0(X4, X5) = sK45) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))))), introduced(choice_axiom, [])).
fof(f206, plain, (? [X5, X4] : ((sdtpldt0(X4, X5) = sK45) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa))) => ((sK45 = sdtpldt0(sK46, sK47)) & aElementOf0(sK47, slsdtgt0(xb)) & aElementOf0(sK46, slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f204, plain, (? [X0] : (! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(X0)) | (sz00 = X1) | (~ aElementOf0(X1, xI) & ! [X2, X3] : (~ (sdtpldt0(X2, X3) = X1) | ~ aElementOf0(X3, slsdtgt0(xb)) | ~ aElementOf0(X2, slsdtgt0(xa))))) & ~ (sz00 = X0) & aElementOf0(X0, xI) & ? [X4, X5] : ((sdtpldt0(X4, X5) = X0) & aElementOf0(X5, slsdtgt0(xb)) & aElementOf0(X4, slsdtgt0(xa)))) | ~ sP4), inference(rectify, [], [f203])).
fof(f203, plain, (? [X6] : (! [X7] : (~ iLess0(sbrdtbr0(X7), sbrdtbr0(X6)) | (sz00 = X7) | (~ aElementOf0(X7, xI) & ! [X8, X9] : (~ (sdtpldt0(X8, X9) = X7) | ~ aElementOf0(X9, slsdtgt0(xb)) | ~ aElementOf0(X8, slsdtgt0(xa))))) & ~ (sz00 = X6) & aElementOf0(X6, xI) & ? [X10, X11] : ((sdtpldt0(X10, X11) = X6) & aElementOf0(X11, slsdtgt0(xb)) & aElementOf0(X10, slsdtgt0(xa)))) | ~ sP4), inference(nnf_transformation, [], [f120])).
fof(f433, plain, (~ spl51_3 | ~ spl51_6), inference(avatar_split_clause, [], [f374, f430, f418])).
fof(f374, plain, (~ (sz00 = sK45) | ~ sP4), inference(cnf_transformation, [], [f207])).
fof(f424, plain, (~ spl51_3 | spl51_4), inference(avatar_split_clause, [], [f376, f422, f418])).
fof(f376, plain, ! [X1] : (~ iLess0(sbrdtbr0(X1), sbrdtbr0(sK45)) | (sz00 = X1) | ~ aElementOf0(X1, xI) | ~ sP4), inference(cnf_transformation, [], [f207])).