fof(f630, plain, $false, inference(avatar_sat_refutation, [], [f358, f379, f480, f501, f607, f626])).
fof(f626, plain, ~ spl32_10, inference(avatar_contradiction_clause, [], [f625])).
fof(f625, plain, ($false | ~ spl32_10), inference(subsumption_resolution, [], [f624, f183])).
fof(f183, plain, aElement0(sz00), inference(cnf_transformation, [], [f2])).
fof(f2, plain, aElement0(sz00), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', mSortsC)).
fof(f624, plain, (~ aElement0(sz00) | ~ spl32_10), inference(trivial_inequality_removal, [], [f621])).
fof(f621, plain, (~ (sz00 = sz00) | ~ aElement0(sz00) | ~ spl32_10), inference(superposition, [], [f378, f433])).
fof(f433, plain, (sz00 = sdtasdt0(xb, sz00)), inference(resolution, [], [f202, f274])).
fof(f274, plain, aElement0(xb), inference(cnf_transformation, [], [f39])).
fof(f39, plain, (aElement0(xb) & aElement0(xa)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', m__2091)).
fof(f202, plain, ! [X0] : (~ aElement0(X0) | (sz00 = sdtasdt0(X0, sz00))), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ! [X0] : (((sz00 = sdtasdt0(sz00, X0)) & (sz00 = sdtasdt0(X0, sz00))) | ~ aElement0(X0)), inference(ennf_transformation, [], [f16])).
fof(f16, plain, ! [X0] : (aElement0(X0) => ((sz00 = sdtasdt0(sz00, X0)) & (sz00 = sdtasdt0(X0, sz00)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', mMulZero)).
fof(f378, plain, (! [X1] : (~ (sz00 = sdtasdt0(xb, X1)) | ~ aElement0(X1)) | ~ spl32_10), inference(avatar_component_clause, [], [f377])).
fof(f377, plain, (spl32_10 <=> ! [X1] : (~ (sz00 = sdtasdt0(xb, X1)) | ~ aElement0(X1))), introduced(avatar_definition, [new_symbols(naming, [spl32_10])])).
fof(f607, plain, ~ spl32_5, inference(avatar_contradiction_clause, [], [f606])).
fof(f606, plain, ($false | ~ spl32_5), inference(subsumption_resolution, [], [f602, f183])).
fof(f602, plain, (~ aElement0(sz00) | ~ spl32_5), inference(trivial_inequality_removal, [], [f600])).
fof(f600, plain, (~ (sz00 = sz00) | ~ aElement0(sz00) | ~ spl32_5), inference(superposition, [], [f357, f432])).
fof(f432, plain, (sz00 = sdtasdt0(xa, sz00)), inference(resolution, [], [f202, f273])).
fof(f273, plain, aElement0(xa), inference(cnf_transformation, [], [f39])).
fof(f357, plain, (! [X0] : (~ (sz00 = sdtasdt0(xa, X0)) | ~ aElement0(X0)) | ~ spl32_5), inference(avatar_component_clause, [], [f356])).
fof(f356, plain, (spl32_5 <=> ! [X0] : (~ (sz00 = sdtasdt0(xa, X0)) | ~ aElement0(X0))), introduced(avatar_definition, [new_symbols(naming, [spl32_5])])).
fof(f501, plain, spl32_6, inference(avatar_split_clause, [], [f500, f360])).
fof(f360, plain, (spl32_6 <=> aElementOf0(xa, slsdtgt0(xa))), introduced(avatar_definition, [new_symbols(naming, [spl32_6])])).
fof(f500, plain, aElementOf0(xa, slsdtgt0(xa)), inference(subsumption_resolution, [], [f498, f184])).
fof(f184, plain, aElement0(sz10), inference(cnf_transformation, [], [f3])).
fof(f3, plain, aElement0(sz10), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', mSortsC_01)).
fof(f498, plain, (aElementOf0(xa, slsdtgt0(xa)) | ~ aElement0(sz10)), inference(superposition, [], [f335, f414])).
fof(f414, plain, (xa = sdtasdt0(xa, sz10)), inference(resolution, [], [f196, f273])).
fof(f196, plain, ! [X0] : (~ aElement0(X0) | (sdtasdt0(X0, sz10) = X0)), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ! [X0] : (((sdtasdt0(sz10, X0) = X0) & (sdtasdt0(X0, sz10) = X0)) | ~ aElement0(X0)), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ! [X0] : (aElement0(X0) => ((sdtasdt0(sz10, X0) = X0) & (sdtasdt0(X0, sz10) = X0))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', mMulUnit)).
fof(f335, plain, ! [X9] : (aElementOf0(sdtasdt0(xa, X9), slsdtgt0(xa)) | ~ aElement0(X9)), inference(equality_resolution, [], [f302])).
fof(f302, plain, ! [X8, X9] : (aElementOf0(X8, slsdtgt0(xa)) | ~ (sdtasdt0(xa, X9) = X8) | ~ aElement0(X9)), inference(cnf_transformation, [], [f180])).
fof(f180, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (((sdtpldt0(sK28(X0), sK29(X0)) = X0) & aElementOf0(sK29(X0), slsdtgt0(xb)) & aElementOf0(sK28(X0), slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xb)) | ! [X6] : (~ (sdtasdt0(xb, X6) = X5) | ~ aElement0(X6))) & (((sdtasdt0(xb, sK30(X5)) = X5) & aElement0(sK30(X5))) | ~ aElementOf0(X5, slsdtgt0(xb)))) & ! [X8] : ((aElementOf0(X8, slsdtgt0(xa)) | ! [X9] : (~ (sdtasdt0(xa, X9) = X8) | ~ aElement0(X9))) & (((sdtasdt0(xa, sK31(X8)) = X8) & aElement0(sK31(X8))) | ~ aElementOf0(X8, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X11] : ((! [X12] : (aElementOf0(sdtasdt0(X12, X11), xI) | ~ aElement0(X12)) & ! [X13] : (aElementOf0(sdtpldt0(X11, X13), xI) | ~ aElementOf0(X13, xI))) | ~ aElementOf0(X11, xI)) & aSet0(xI)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29, sK30, sK31])], [f176, f179, f178, f177])).
fof(f177, plain, ! [X0] : (? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa))) => ((sdtpldt0(sK28(X0), sK29(X0)) = X0) & aElementOf0(sK29(X0), slsdtgt0(xb)) & aElementOf0(sK28(X0), slsdtgt0(xa)))), introduced(choice_axiom, [])).
fof(f178, plain, ! [X5] : (? [X7] : ((sdtasdt0(xb, X7) = X5) & aElement0(X7)) => ((sdtasdt0(xb, sK30(X5)) = X5) & aElement0(sK30(X5)))), introduced(choice_axiom, [])).
fof(f179, plain, ! [X8] : (? [X10] : ((sdtasdt0(xa, X10) = X8) & aElement0(X10)) => ((sdtasdt0(xa, sK31(X8)) = X8) & aElement0(sK31(X8)))), introduced(choice_axiom, [])).
fof(f176, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (? [X3, X4] : ((sdtpldt0(X3, X4) = X0) & aElementOf0(X4, slsdtgt0(xb)) & aElementOf0(X3, slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xb)) | ! [X6] : (~ (sdtasdt0(xb, X6) = X5) | ~ aElement0(X6))) & (? [X7] : ((sdtasdt0(xb, X7) = X5) & aElement0(X7)) | ~ aElementOf0(X5, slsdtgt0(xb)))) & ! [X8] : ((aElementOf0(X8, slsdtgt0(xa)) | ! [X9] : (~ (sdtasdt0(xa, X9) = X8) | ~ aElement0(X9))) & (? [X10] : ((sdtasdt0(xa, X10) = X8) & aElement0(X10)) | ~ aElementOf0(X8, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X11] : ((! [X12] : (aElementOf0(sdtasdt0(X12, X11), xI) | ~ aElement0(X12)) & ! [X13] : (aElementOf0(sdtpldt0(X11, X13), xI) | ~ aElementOf0(X13, xI))) | ~ aElementOf0(X11, xI)) & aSet0(xI)), inference(rectify, [], [f175])).
fof(f175, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : ((aElementOf0(X0, xI) | ! [X1, X2] : (~ (sdtpldt0(X1, X2) = X0) | ~ aElementOf0(X2, slsdtgt0(xb)) | ~ aElementOf0(X1, slsdtgt0(xa)))) & (? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa))) | ~ aElementOf0(X0, xI))) & ! [X3] : ((aElementOf0(X3, slsdtgt0(xb)) | ! [X4] : (~ (sdtasdt0(xb, X4) = X3) | ~ aElement0(X4))) & (? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4)) | ~ aElementOf0(X3, slsdtgt0(xb)))) & ! [X5] : ((aElementOf0(X5, slsdtgt0(xa)) | ! [X6] : (~ (sdtasdt0(xa, X6) = X5) | ~ aElement0(X6))) & (? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6)) | ~ aElementOf0(X5, slsdtgt0(xa)))) & aIdeal0(xI) & ! [X7] : ((! [X8] : (aElementOf0(sdtasdt0(X8, X7), xI) | ~ aElement0(X8)) & ! [X9] : (aElementOf0(sdtpldt0(X7, X9), xI) | ~ aElementOf0(X9, xI))) | ~ aElementOf0(X7, xI)) & aSet0(xI)), inference(nnf_transformation, [], [f108])).
fof(f108, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X3] : (aElementOf0(X3, slsdtgt0(xb)) <=> ? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4))) & ! [X5] : (aElementOf0(X5, slsdtgt0(xa)) <=> ? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6))) & aIdeal0(xI) & ! [X7] : ((! [X8] : (aElementOf0(sdtasdt0(X8, X7), xI) | ~ aElement0(X8)) & ! [X9] : (aElementOf0(sdtpldt0(X7, X9), xI) | ~ aElementOf0(X9, xI))) | ~ aElementOf0(X7, xI)) & aSet0(xI)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X3] : (aElementOf0(X3, slsdtgt0(xb)) <=> ? [X4] : ((sdtasdt0(xb, X4) = X3) & aElement0(X4))) & ! [X5] : (aElementOf0(X5, slsdtgt0(xa)) <=> ? [X6] : ((sdtasdt0(xa, X6) = X5) & aElement0(X6))) & aIdeal0(xI) & ! [X7] : (aElementOf0(X7, xI) => (! [X8] : (aElement0(X8) => aElementOf0(sdtasdt0(X8, X7), xI)) & ! [X9] : (aElementOf0(X9, xI) => aElementOf0(sdtpldt0(X7, X9), xI)))) & aSet0(xI)), inference(rectify, [], [f42])).
fof(f42, plain, ((xI = sdtpldt1(slsdtgt0(xa), slsdtgt0(xb))) & ! [X0] : (aElementOf0(X0, xI) <=> ? [X1, X2] : ((sdtpldt0(X1, X2) = X0) & aElementOf0(X2, slsdtgt0(xb)) & aElementOf0(X1, slsdtgt0(xa)))) & ! [X0] : (aElementOf0(X0, slsdtgt0(xb)) <=> ? [X1] : ((sdtasdt0(xb, X1) = X0) & aElement0(X1))) & ! [X0] : (aElementOf0(X0, slsdtgt0(xa)) <=> ? [X1] : ((sdtasdt0(xa, X1) = X0) & aElement0(X1))) & aIdeal0(xI) & ! [X0] : (aElementOf0(X0, xI) => (! [X1] : (aElement0(X1) => aElementOf0(sdtasdt0(X1, X0), xI)) & ! [X1] : (aElementOf0(X1, xI) => aElementOf0(sdtpldt0(X0, X1), xI)))) & aSet0(xI)), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', m__2174)).
fof(f480, plain, spl32_8, inference(avatar_split_clause, [], [f479, f368])).
fof(f368, plain, (spl32_8 <=> aElementOf0(xb, slsdtgt0(xb))), introduced(avatar_definition, [new_symbols(naming, [spl32_8])])).
fof(f479, plain, aElementOf0(xb, slsdtgt0(xb)), inference(subsumption_resolution, [], [f478, f184])).
fof(f478, plain, (aElementOf0(xb, slsdtgt0(xb)) | ~ aElement0(sz10)), inference(superposition, [], [f334, f415])).
fof(f415, plain, (xb = sdtasdt0(xb, sz10)), inference(resolution, [], [f196, f274])).
fof(f334, plain, ! [X6] : (aElementOf0(sdtasdt0(xb, X6), slsdtgt0(xb)) | ~ aElement0(X6)), inference(equality_resolution, [], [f305])).
fof(f305, plain, ! [X6, X5] : (aElementOf0(X5, slsdtgt0(xb)) | ~ (sdtasdt0(xb, X6) = X5) | ~ aElement0(X6)), inference(cnf_transformation, [], [f180])).
fof(f379, plain, (spl32_3 | ~ spl32_6 | spl32_10 | ~ spl32_8), inference(avatar_split_clause, [], [f318, f368, f377, f360, f347])).
fof(f347, plain, (spl32_3 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl32_3])])).
fof(f318, plain, ! [X1] : (~ aElementOf0(xb, slsdtgt0(xb)) | ~ (sz00 = sdtasdt0(xb, X1)) | ~ aElement0(X1) | ~ aElementOf0(xa, slsdtgt0(xa)) | sP4), inference(cnf_transformation, [], [f117])).
fof(f117, plain, ((~ aElementOf0(xb, slsdtgt0(xb)) & ! [X0] : (~ (xb = sdtasdt0(xb, X0)) | ~ aElement0(X0))) | (~ aElementOf0(sz00, slsdtgt0(xb)) & ! [X1] : (~ (sz00 = sdtasdt0(xb, X1)) | ~ aElement0(X1))) | (~ aElementOf0(xa, slsdtgt0(xa)) & ! [X2] : (~ (xa = sdtasdt0(xa, X2)) | ~ aElement0(X2))) | sP4), inference(definition_folding, [], [f109, e116])).
fof(f116, plain, ((~ aElementOf0(sz00, slsdtgt0(xa)) & ! [X3] : (~ (sz00 = sdtasdt0(xa, X3)) | ~ aElement0(X3))) | ~ sP4), inference(usedef, [], [e116])).
fof(e116, plain, (sP4 <=> (~ aElementOf0(sz00, slsdtgt0(xa)) & ! [X3] : (~ (sz00 = sdtasdt0(xa, X3)) | ~ aElement0(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f109, plain, ((~ aElementOf0(xb, slsdtgt0(xb)) & ! [X0] : (~ (xb = sdtasdt0(xb, X0)) | ~ aElement0(X0))) | (~ aElementOf0(sz00, slsdtgt0(xb)) & ! [X1] : (~ (sz00 = sdtasdt0(xb, X1)) | ~ aElement0(X1))) | (~ aElementOf0(xa, slsdtgt0(xa)) & ! [X2] : (~ (xa = sdtasdt0(xa, X2)) | ~ aElement0(X2))) | (~ aElementOf0(sz00, slsdtgt0(xa)) & ! [X3] : (~ (sz00 = sdtasdt0(xa, X3)) | ~ aElement0(X3)))), inference(ennf_transformation, [], [f54])).
fof(f54, plain, ~ ((aElementOf0(xb, slsdtgt0(xb)) | ? [X0] : ((xb = sdtasdt0(xb, X0)) & aElement0(X0))) & (aElementOf0(sz00, slsdtgt0(xb)) | ? [X1] : ((sz00 = sdtasdt0(xb, X1)) & aElement0(X1))) & (aElementOf0(xa, slsdtgt0(xa)) | ? [X2] : ((xa = sdtasdt0(xa, X2)) & aElement0(X2))) & (aElementOf0(sz00, slsdtgt0(xa)) | ? [X3] : ((sz00 = sdtasdt0(xa, X3)) & aElement0(X3)))), inference(rectify, [], [f44])).
fof(f44, plain, ~ ((aElementOf0(xb, slsdtgt0(xb)) | ? [X0] : ((xb = sdtasdt0(xb, X0)) & aElement0(X0))) & (aElementOf0(sz00, slsdtgt0(xb)) | ? [X0] : ((sz00 = sdtasdt0(xb, X0)) & aElement0(X0))) & (aElementOf0(xa, slsdtgt0(xa)) | ? [X0] : ((xa = sdtasdt0(xa, X0)) & aElement0(X0))) & (aElementOf0(sz00, slsdtgt0(xa)) | ? [X0] : ((sz00 = sdtasdt0(xa, X0)) & aElement0(X0)))), inference(negated_conjecture, [], [f43])).
fof(f43, plain, ~ ((aElementOf0(xb, slsdtgt0(xb)) | ? [X0] : ((xb = sdtasdt0(xb, X0)) & aElement0(X0))) & (aElementOf0(sz00, slsdtgt0(xb)) | ? [X0] : ((sz00 = sdtasdt0(xb, X0)) & aElement0(X0))) & (aElementOf0(xa, slsdtgt0(xa)) | ? [X0] : ((xa = sdtasdt0(xa, X0)) & aElement0(X0))) & (aElementOf0(sz00, slsdtgt0(xa)) | ? [X0] : ((sz00 = sdtasdt0(xa, X0)) & aElement0(X0)))), file('/home/ubuntu/library/tptp/Problems/RNG/RNG108+4.p', m__)).
fof(f358, plain, (~ spl32_3 | spl32_5), inference(avatar_split_clause, [], [f311, f356, f347])).
fof(f311, plain, ! [X0] : (~ (sz00 = sdtasdt0(xa, X0)) | ~ aElement0(X0) | ~ sP4), inference(cnf_transformation, [], [f182])).
fof(f182, plain, ((~ aElementOf0(sz00, slsdtgt0(xa)) & ! [X0] : (~ (sz00 = sdtasdt0(xa, X0)) | ~ aElement0(X0))) | ~ sP4), inference(rectify, [], [f181])).
fof(f181, plain, ((~ aElementOf0(sz00, slsdtgt0(xa)) & ! [X3] : (~ (sz00 = sdtasdt0(xa, X3)) | ~ aElement0(X3))) | ~ sP4), inference(nnf_transformation, [], [f116])).