fof(f12658, plain, $false, inference(avatar_sat_refutation, [], [f1977, f1987, f2029, f2264, f12657])).
fof(f12657, plain, (~ spl19_18 | ~ spl19_20 | ~ spl19_21 | ~ spl19_23), inference(avatar_contradiction_clause, [], [f12656])).
fof(f12656, plain, ($false | (~ spl19_18 | ~ spl19_20 | ~ spl19_21 | ~ spl19_23)), inference(subsumption_resolution, [], [f12655, f1971])).
fof(f1971, plain, (relation(relation_rng_restriction(sK17, sK18)) | ~ spl19_20), inference(avatar_component_clause, [], [f1970])).
fof(f1970, plain, (spl19_20 <=> relation(relation_rng_restriction(sK17, sK18))), introduced(avatar_definition, [new_symbols(naming, [spl19_20])])).
fof(f12655, plain, (~ relation(relation_rng_restriction(sK17, sK18)) | (~ spl19_18 | ~ spl19_21 | ~ spl19_23)), inference(subsumption_resolution, [], [f12654, f1975])).
fof(f1975, plain, (function(relation_rng_restriction(sK17, sK18)) | ~ spl19_21), inference(avatar_component_clause, [], [f1974])).
fof(f1974, plain, (spl19_21 <=> function(relation_rng_restriction(sK17, sK18))), introduced(avatar_definition, [new_symbols(naming, [spl19_21])])).
fof(f12654, plain, (~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f12653, f177])).
fof(f177, plain, ~ one_to_one(relation_rng_restriction(sK17, sK18)), inference(cnf_transformation, [], [f108])).
fof(f108, plain, (~ one_to_one(relation_rng_restriction(sK17, sK18)) & one_to_one(sK18) & function(sK18) & relation(sK18)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18])], [f69, f107])).
fof(f107, plain, (? [X0, X1] : (~ one_to_one(relation_rng_restriction(X0, X1)) & one_to_one(X1) & function(X1) & relation(X1)) => (~ one_to_one(relation_rng_restriction(sK17, sK18)) & one_to_one(sK18) & function(sK18) & relation(sK18))), introduced(choice_axiom, [])).
fof(f69, plain, ? [X0, X1] : (~ one_to_one(relation_rng_restriction(X0, X1)) & one_to_one(X1) & function(X1) & relation(X1)), inference(flattening, [], [f68])).
fof(f68, plain, ? [X0, X1] : ((~ one_to_one(relation_rng_restriction(X0, X1)) & one_to_one(X1)) & (function(X1) & relation(X1))), inference(ennf_transformation, [], [f36])).
fof(f36, plain, ~ ! [X0, X1] : ((function(X1) & relation(X1)) => (one_to_one(X1) => one_to_one(relation_rng_restriction(X0, X1)))), inference(negated_conjecture, [], [f35])).
fof(f35, plain, ~ ! [X0, X1] : ((function(X1) & relation(X1)) => (one_to_one(X1) => one_to_one(relation_rng_restriction(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU048+1.p', t99_funct_1)).
fof(f12653, plain, (one_to_one(relation_rng_restriction(sK17, sK18)) | ~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | (~ spl19_18 | ~ spl19_23)), inference(trivial_inequality_removal, [], [f12651])).
fof(f12651, plain, (~ (sK2(relation_rng_restriction(sK17, sK18)) = sK2(relation_rng_restriction(sK17, sK18))) | one_to_one(relation_rng_restriction(sK17, sK18)) | ~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | (~ spl19_18 | ~ spl19_23)), inference(superposition, [], [f119, f12623])).
fof(f12623, plain, ((sK3(relation_rng_restriction(sK17, sK18)) = sK2(relation_rng_restriction(sK17, sK18))) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f12622, f2297])).
fof(f2297, plain, (in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_23), inference(subsumption_resolution, [], [f2296, f175])).
fof(f175, plain, function(sK18), inference(cnf_transformation, [], [f108])).
fof(f2296, plain, (~ function(sK18) | in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_23), inference(subsumption_resolution, [], [f2283, f174])).
fof(f174, plain, relation(sK18), inference(cnf_transformation, [], [f108])).
fof(f2283, plain, (~ relation(sK18) | ~ function(sK18) | in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_23), inference(resolution, [], [f2135, f450])).
fof(f450, plain, ! [X6, X8, X7] : (~ in(X7, relation_dom(relation_rng_restriction(X8, X6))) | ~ relation(X6) | ~ function(X6) | in(X7, relation_dom(X6))), inference(resolution, [], [f283, f162])).
fof(f162, plain, ! [X6, X2, X0, X1] : (~ sP0(X0, X1, X2) | ~ in(X6, relation_dom(X0)) | in(X6, relation_dom(X1))), inference(cnf_transformation, [], [f106])).
fof(f106, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | (~ (apply(X1, sK15(X0, X1)) = apply(X0, sK15(X0, X1))) & in(sK15(X0, X1), relation_dom(X0))) | ((~ in(apply(X1, sK16(X0, X1, X2)), X2) | ~ in(sK16(X0, X1, X2), relation_dom(X1)) | ~ in(sK16(X0, X1, X2), relation_dom(X0))) & ((in(apply(X1, sK16(X0, X1, X2)), X2) & in(sK16(X0, X1, X2), relation_dom(X1))) | in(sK16(X0, X1, X2), relation_dom(X0))))) & ((! [X5] : ((apply(X1, X5) = apply(X0, X5)) | ~ in(X5, relation_dom(X0))) & ! [X6] : ((in(X6, relation_dom(X0)) | ~ in(apply(X1, X6), X2) | ~ in(X6, relation_dom(X1))) & ((in(apply(X1, X6), X2) & in(X6, relation_dom(X1))) | ~ in(X6, relation_dom(X0))))) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16])], [f103, f105, f104])).
fof(f104, plain, ! [X1, X0] : (? [X3] : (~ (apply(X1, X3) = apply(X0, X3)) & in(X3, relation_dom(X0))) => (~ (apply(X1, sK15(X0, X1)) = apply(X0, sK15(X0, X1))) & in(sK15(X0, X1), relation_dom(X0)))), introduced(choice_axiom, [])).
fof(f105, plain, ! [X2, X1, X0] : (? [X4] : ((~ in(apply(X1, X4), X2) | ~ in(X4, relation_dom(X1)) | ~ in(X4, relation_dom(X0))) & ((in(apply(X1, X4), X2) & in(X4, relation_dom(X1))) | in(X4, relation_dom(X0)))) => ((~ in(apply(X1, sK16(X0, X1, X2)), X2) | ~ in(sK16(X0, X1, X2), relation_dom(X1)) | ~ in(sK16(X0, X1, X2), relation_dom(X0))) & ((in(apply(X1, sK16(X0, X1, X2)), X2) & in(sK16(X0, X1, X2), relation_dom(X1))) | in(sK16(X0, X1, X2), relation_dom(X0))))), introduced(choice_axiom, [])).
fof(f103, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ? [X3] : (~ (apply(X1, X3) = apply(X0, X3)) & in(X3, relation_dom(X0))) | ? [X4] : ((~ in(apply(X1, X4), X2) | ~ in(X4, relation_dom(X1)) | ~ in(X4, relation_dom(X0))) & ((in(apply(X1, X4), X2) & in(X4, relation_dom(X1))) | in(X4, relation_dom(X0))))) & ((! [X5] : ((apply(X1, X5) = apply(X0, X5)) | ~ in(X5, relation_dom(X0))) & ! [X6] : ((in(X6, relation_dom(X0)) | ~ in(apply(X1, X6), X2) | ~ in(X6, relation_dom(X1))) & ((in(apply(X1, X6), X2) & in(X6, relation_dom(X1))) | ~ in(X6, relation_dom(X0))))) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f102])).
fof(f102, plain, ! [X1, X2, X0] : ((sP0(X1, X2, X0) | ? [X3] : (~ (apply(X2, X3) = apply(X1, X3)) & in(X3, relation_dom(X1))) | ? [X4] : ((~ in(apply(X2, X4), X0) | ~ in(X4, relation_dom(X2)) | ~ in(X4, relation_dom(X1))) & ((in(apply(X2, X4), X0) & in(X4, relation_dom(X2))) | in(X4, relation_dom(X1))))) & ((! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : ((in(X4, relation_dom(X1)) | ~ in(apply(X2, X4), X0) | ~ in(X4, relation_dom(X2))) & ((in(apply(X2, X4), X0) & in(X4, relation_dom(X2))) | ~ in(X4, relation_dom(X1))))) | ~ sP0(X1, X2, X0))), inference(flattening, [], [f101])).
fof(f101, plain, ! [X1, X2, X0] : ((sP0(X1, X2, X0) | (? [X3] : (~ (apply(X2, X3) = apply(X1, X3)) & in(X3, relation_dom(X1))) | ? [X4] : (((~ in(apply(X2, X4), X0) | ~ in(X4, relation_dom(X2))) | ~ in(X4, relation_dom(X1))) & ((in(apply(X2, X4), X0) & in(X4, relation_dom(X2))) | in(X4, relation_dom(X1)))))) & ((! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : ((in(X4, relation_dom(X1)) | (~ in(apply(X2, X4), X0) | ~ in(X4, relation_dom(X2)))) & ((in(apply(X2, X4), X0) & in(X4, relation_dom(X2))) | ~ in(X4, relation_dom(X1))))) | ~ sP0(X1, X2, X0))), inference(nnf_transformation, [], [f70])).
fof(f70, plain, ! [X1, X2, X0] : (sP0(X1, X2, X0) <=> (! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : (in(X4, relation_dom(X1)) <=> (in(apply(X2, X4), X0) & in(X4, relation_dom(X2)))))), inference(usedef, [], [e70])).
fof(e70, plain, ! [X1, X2, X0] : (sP0(X1, X2, X0) <=> (! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : (in(X4, relation_dom(X1)) <=> (in(apply(X2, X4), X0) & in(X4, relation_dom(X2)))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f283, plain, ! [X0, X1] : (sP0(relation_rng_restriction(X0, X1), X1, X0) | ~ function(X1) | ~ relation(X1)), inference(subsumption_resolution, [], [f282, f120])).
fof(f120, plain, ! [X0, X1] : (relation(relation_rng_restriction(X0, X1)) | ~ relation(X1)), inference(cnf_transformation, [], [f49])).
fof(f49, plain, ! [X0, X1] : (relation(relation_rng_restriction(X0, X1)) | ~ relation(X1)), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1] : (relation(X1) => relation(relation_rng_restriction(X0, X1))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU048+1.p', dt_k8_relat_1)).
fof(f282, plain, ! [X0, X1] : (sP0(relation_rng_restriction(X0, X1), X1, X0) | ~ function(X1) | ~ relation(X1) | ~ relation(relation_rng_restriction(X0, X1))), inference(subsumption_resolution, [], [f281, f129])).
fof(f129, plain, ! [X0, X1] : (function(relation_rng_restriction(X0, X1)) | ~ function(X1) | ~ relation(X1)), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ! [X0, X1] : ((function(relation_rng_restriction(X0, X1)) & relation(relation_rng_restriction(X0, X1))) | ~ function(X1) | ~ relation(X1)), inference(flattening, [], [f50])).
fof(f50, plain, ! [X0, X1] : ((function(relation_rng_restriction(X0, X1)) & relation(relation_rng_restriction(X0, X1))) | (~ function(X1) | ~ relation(X1))), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0, X1] : ((function(X1) & relation(X1)) => (function(relation_rng_restriction(X0, X1)) & relation(relation_rng_restriction(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU048+1.p', fc5_funct_1)).
fof(f281, plain, ! [X0, X1] : (sP0(relation_rng_restriction(X0, X1), X1, X0) | ~ function(X1) | ~ relation(X1) | ~ function(relation_rng_restriction(X0, X1)) | ~ relation(relation_rng_restriction(X0, X1))), inference(resolution, [], [f178, f172])).
fof(f172, plain, ! [X2, X0, X1] : (sP1(X0, X2, X1) | ~ function(X2) | ~ relation(X2) | ~ function(X1) | ~ relation(X1)), inference(cnf_transformation, [], [f72])).
fof(f72, plain, ! [X0, X1] : (! [X2] : (sP1(X0, X2, X1) | ~ function(X2) | ~ relation(X2)) | ~ function(X1) | ~ relation(X1)), inference(definition_folding, [], [f66, e71, e70])).
fof(f71, plain, ! [X0, X2, X1] : (((relation_rng_restriction(X0, X2) = X1) <=> sP0(X1, X2, X0)) | ~ sP1(X0, X2, X1)), inference(usedef, [], [e71])).
fof(e71, plain, ! [X0, X2, X1] : (sP1(X0, X2, X1) <=> ((relation_rng_restriction(X0, X2) = X1) <=> sP0(X1, X2, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f66, plain, ! [X0, X1] : (! [X2] : (((relation_rng_restriction(X0, X2) = X1) <=> (! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : (in(X4, relation_dom(X1)) <=> (in(apply(X2, X4), X0) & in(X4, relation_dom(X2)))))) | ~ function(X2) | ~ relation(X2)) | ~ function(X1) | ~ relation(X1)), inference(flattening, [], [f65])).
fof(f65, plain, ! [X0, X1] : (! [X2] : (((relation_rng_restriction(X0, X2) = X1) <=> (! [X3] : ((apply(X2, X3) = apply(X1, X3)) | ~ in(X3, relation_dom(X1))) & ! [X4] : (in(X4, relation_dom(X1)) <=> (in(apply(X2, X4), X0) & in(X4, relation_dom(X2)))))) | (~ function(X2) | ~ relation(X2))) | (~ function(X1) | ~ relation(X1))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ! [X0, X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => ((relation_rng_restriction(X0, X2) = X1) <=> (! [X3] : (in(X3, relation_dom(X1)) => (apply(X2, X3) = apply(X1, X3))) & ! [X4] : (in(X4, relation_dom(X1)) <=> (in(apply(X2, X4), X0) & in(X4, relation_dom(X2)))))))), inference(rectify, [], [f33])).
fof(f33, plain, ! [X0, X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => ((relation_rng_restriction(X0, X2) = X1) <=> (! [X3] : (in(X3, relation_dom(X1)) => (apply(X2, X3) = apply(X1, X3))) & ! [X3] : (in(X3, relation_dom(X1)) <=> (in(apply(X2, X3), X0) & in(X3, relation_dom(X2)))))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU048+1.p', t85_funct_1)).
fof(f178, plain, ! [X0, X1] : (~ sP1(X0, X1, relation_rng_restriction(X0, X1)) | sP0(relation_rng_restriction(X0, X1), X1, X0)), inference(equality_resolution, [], [f160])).
fof(f160, plain, ! [X2, X0, X1] : (sP0(X2, X1, X0) | ~ (relation_rng_restriction(X0, X1) = X2) | ~ sP1(X0, X1, X2)), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ! [X0, X1, X2] : ((((relation_rng_restriction(X0, X1) = X2) | ~ sP0(X2, X1, X0)) & (sP0(X2, X1, X0) | ~ (relation_rng_restriction(X0, X1) = X2))) | ~ sP1(X0, X1, X2)), inference(rectify, [], [f99])).
fof(f99, plain, ! [X0, X2, X1] : ((((relation_rng_restriction(X0, X2) = X1) | ~ sP0(X1, X2, X0)) & (sP0(X1, X2, X0) | ~ (relation_rng_restriction(X0, X2) = X1))) | ~ sP1(X0, X2, X1)), inference(nnf_transformation, [], [f71])).
fof(f2135, plain, (in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18))) | ~ spl19_23), inference(avatar_component_clause, [], [f2134])).
fof(f2134, plain, (spl19_23 <=> in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl19_23])])).
fof(f12622, plain, ((sK3(relation_rng_restriction(sK17, sK18)) = sK2(relation_rng_restriction(sK17, sK18))) | ~ in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | (~ spl19_18 | ~ spl19_23)), inference(equality_resolution, [], [f2305])).
fof(f2305, plain, (! [X0] : (~ (apply(sK18, X0) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (sK3(relation_rng_restriction(sK17, sK18)) = X0) | ~ in(X0, relation_dom(sK18))) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f2304, f174])).
fof(f2304, plain, (! [X0] : (~ (apply(sK18, X0) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (sK3(relation_rng_restriction(sK17, sK18)) = X0) | ~ in(X0, relation_dom(sK18)) | ~ relation(sK18)) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f2303, f175])).
fof(f2303, plain, (! [X0] : (~ (apply(sK18, X0) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (sK3(relation_rng_restriction(sK17, sK18)) = X0) | ~ in(X0, relation_dom(sK18)) | ~ function(sK18) | ~ relation(sK18)) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f2302, f176])).
fof(f176, plain, one_to_one(sK18), inference(cnf_transformation, [], [f108])).
fof(f2302, plain, (! [X0] : (~ (apply(sK18, X0) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (sK3(relation_rng_restriction(sK17, sK18)) = X0) | ~ in(X0, relation_dom(sK18)) | ~ one_to_one(sK18) | ~ function(sK18) | ~ relation(sK18)) | (~ spl19_18 | ~ spl19_23)), inference(subsumption_resolution, [], [f2299, f2129])).
fof(f2129, plain, (in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_18), inference(subsumption_resolution, [], [f2128, f175])).
fof(f2128, plain, (~ function(sK18) | in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_18), inference(subsumption_resolution, [], [f2118, f174])).
fof(f2118, plain, (~ relation(sK18) | ~ function(sK18) | in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ spl19_18), inference(resolution, [], [f1961, f450])).
fof(f1961, plain, (in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18))) | ~ spl19_18), inference(avatar_component_clause, [], [f1960])).
fof(f1960, plain, (spl19_18 <=> in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18)))), introduced(avatar_definition, [new_symbols(naming, [spl19_18])])).
fof(f2299, plain, (! [X0] : (~ (apply(sK18, X0) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (sK3(relation_rng_restriction(sK17, sK18)) = X0) | ~ in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(sK18)) | ~ in(X0, relation_dom(sK18)) | ~ one_to_one(sK18) | ~ function(sK18) | ~ relation(sK18)) | (~ spl19_18 | ~ spl19_23)), inference(superposition, [], [f115, f2289])).
fof(f2289, plain, ((apply(sK18, sK3(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | (~ spl19_18 | ~ spl19_23)), inference(backward_demodulation, [], [f2124, f2288])).
fof(f2288, plain, ((apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | ~ spl19_23), inference(subsumption_resolution, [], [f2287, f175])).
fof(f2287, plain, (~ function(sK18) | (apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | ~ spl19_23), inference(subsumption_resolution, [], [f2281, f174])).
fof(f2281, plain, (~ relation(sK18) | ~ function(sK18) | (apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK2(relation_rng_restriction(sK17, sK18)))) | ~ spl19_23), inference(resolution, [], [f2135, f448])).
fof(f448, plain, ! [X2, X0, X1] : (~ in(X1, relation_dom(relation_rng_restriction(X2, X0))) | ~ relation(X0) | ~ function(X0) | (apply(X0, X1) = apply(relation_rng_restriction(X2, X0), X1))), inference(resolution, [], [f283, f165])).
fof(f165, plain, ! [X2, X0, X5, X1] : (~ sP0(X0, X1, X2) | ~ in(X5, relation_dom(X0)) | (apply(X1, X5) = apply(X0, X5))), inference(cnf_transformation, [], [f106])).
fof(f2124, plain, ((apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK3(relation_rng_restriction(sK17, sK18)))) | ~ spl19_18), inference(backward_demodulation, [], [f1417, f2123])).
fof(f2123, plain, ((apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK3(relation_rng_restriction(sK17, sK18)))) | ~ spl19_18), inference(subsumption_resolution, [], [f2122, f175])).
fof(f2122, plain, (~ function(sK18) | (apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK3(relation_rng_restriction(sK17, sK18)))) | ~ spl19_18), inference(subsumption_resolution, [], [f2116, f174])).
fof(f2116, plain, (~ relation(sK18) | ~ function(sK18) | (apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(sK18, sK3(relation_rng_restriction(sK17, sK18)))) | ~ spl19_18), inference(resolution, [], [f1961, f448])).
fof(f1417, plain, (apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18)))), inference(subsumption_resolution, [], [f1416, f175])).
fof(f1416, plain, ((apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18)))) | ~ function(sK18)), inference(subsumption_resolution, [], [f1411, f174])).
fof(f1411, plain, ((apply(relation_rng_restriction(sK17, sK18), sK3(relation_rng_restriction(sK17, sK18))) = apply(relation_rng_restriction(sK17, sK18), sK2(relation_rng_restriction(sK17, sK18)))) | ~ relation(sK18) | ~ function(sK18)), inference(resolution, [], [f708, f177])).
fof(f708, plain, ! [X0, X1] : (one_to_one(relation_rng_restriction(X0, X1)) | (apply(relation_rng_restriction(X0, X1), sK3(relation_rng_restriction(X0, X1))) = apply(relation_rng_restriction(X0, X1), sK2(relation_rng_restriction(X0, X1)))) | ~ relation(X1) | ~ function(X1)), inference(duplicate_literal_removal, [], [f706])).
fof(f706, plain, ! [X0, X1] : ((apply(relation_rng_restriction(X0, X1), sK3(relation_rng_restriction(X0, X1))) = apply(relation_rng_restriction(X0, X1), sK2(relation_rng_restriction(X0, X1)))) | one_to_one(relation_rng_restriction(X0, X1)) | ~ relation(X1) | ~ function(X1) | ~ relation(X1)), inference(resolution, [], [f294, f129])).
fof(f294, plain, ! [X2, X1] : (~ function(relation_rng_restriction(X1, X2)) | (apply(relation_rng_restriction(X1, X2), sK2(relation_rng_restriction(X1, X2))) = apply(relation_rng_restriction(X1, X2), sK3(relation_rng_restriction(X1, X2)))) | one_to_one(relation_rng_restriction(X1, X2)) | ~ relation(X2)), inference(resolution, [], [f118, f120])).
fof(f118, plain, ! [X0] : (~ relation(X0) | (apply(X0, sK2(X0)) = apply(X0, sK3(X0))) | ~ function(X0) | one_to_one(X0)), inference(cnf_transformation, [], [f76])).
fof(f76, plain, ! [X0] : (((one_to_one(X0) | (~ (sK2(X0) = sK3(X0)) & (apply(X0, sK2(X0)) = apply(X0, sK3(X0))) & in(sK3(X0), relation_dom(X0)) & in(sK2(X0), relation_dom(X0)))) & (! [X3, X4] : ((X3 = X4) | ~ (apply(X0, X3) = apply(X0, X4)) | ~ in(X4, relation_dom(X0)) | ~ in(X3, relation_dom(X0))) | ~ one_to_one(X0))) | ~ function(X0) | ~ relation(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3])], [f74, f75])).
fof(f75, plain, ! [X0] : (? [X1, X2] : (~ (X1 = X2) & (apply(X0, X1) = apply(X0, X2)) & in(X2, relation_dom(X0)) & in(X1, relation_dom(X0))) => (~ (sK2(X0) = sK3(X0)) & (apply(X0, sK2(X0)) = apply(X0, sK3(X0))) & in(sK3(X0), relation_dom(X0)) & in(sK2(X0), relation_dom(X0)))), introduced(choice_axiom, [])).
fof(f74, plain, ! [X0] : (((one_to_one(X0) | ? [X1, X2] : (~ (X1 = X2) & (apply(X0, X1) = apply(X0, X2)) & in(X2, relation_dom(X0)) & in(X1, relation_dom(X0)))) & (! [X3, X4] : ((X3 = X4) | ~ (apply(X0, X3) = apply(X0, X4)) | ~ in(X4, relation_dom(X0)) | ~ in(X3, relation_dom(X0))) | ~ one_to_one(X0))) | ~ function(X0) | ~ relation(X0)), inference(rectify, [], [f73])).
fof(f73, plain, ! [X0] : (((one_to_one(X0) | ? [X1, X2] : (~ (X1 = X2) & (apply(X0, X1) = apply(X0, X2)) & in(X2, relation_dom(X0)) & in(X1, relation_dom(X0)))) & (! [X1, X2] : ((X1 = X2) | ~ (apply(X0, X1) = apply(X0, X2)) | ~ in(X2, relation_dom(X0)) | ~ in(X1, relation_dom(X0))) | ~ one_to_one(X0))) | ~ function(X0) | ~ relation(X0)), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ! [X0] : ((one_to_one(X0) <=> ! [X1, X2] : ((X1 = X2) | ~ (apply(X0, X1) = apply(X0, X2)) | ~ in(X2, relation_dom(X0)) | ~ in(X1, relation_dom(X0)))) | ~ function(X0) | ~ relation(X0)), inference(flattening, [], [f47])).
fof(f47, plain, ! [X0] : ((one_to_one(X0) <=> ! [X1, X2] : ((X1 = X2) | (~ (apply(X0, X1) = apply(X0, X2)) | ~ in(X2, relation_dom(X0)) | ~ in(X1, relation_dom(X0))))) | (~ function(X0) | ~ relation(X0))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0] : ((function(X0) & relation(X0)) => (one_to_one(X0) <=> ! [X1, X2] : (((apply(X0, X1) = apply(X0, X2)) & in(X2, relation_dom(X0)) & in(X1, relation_dom(X0))) => (X1 = X2)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU048+1.p', d8_funct_1)).
fof(f115, plain, ! [X4, X0, X3] : (~ (apply(X0, X3) = apply(X0, X4)) | (X3 = X4) | ~ in(X4, relation_dom(X0)) | ~ in(X3, relation_dom(X0)) | ~ one_to_one(X0) | ~ function(X0) | ~ relation(X0)), inference(cnf_transformation, [], [f76])).
fof(f119, plain, ! [X0] : (~ (sK2(X0) = sK3(X0)) | one_to_one(X0) | ~ function(X0) | ~ relation(X0)), inference(cnf_transformation, [], [f76])).
fof(f2264, plain, (~ spl19_20 | ~ spl19_21 | spl19_23), inference(avatar_contradiction_clause, [], [f2263])).
fof(f2263, plain, ($false | (~ spl19_20 | ~ spl19_21 | spl19_23)), inference(subsumption_resolution, [], [f2262, f1971])).
fof(f2262, plain, (~ relation(relation_rng_restriction(sK17, sK18)) | (~ spl19_21 | spl19_23)), inference(subsumption_resolution, [], [f2261, f1975])).
fof(f2261, plain, (~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | spl19_23), inference(subsumption_resolution, [], [f2260, f177])).
fof(f2260, plain, (one_to_one(relation_rng_restriction(sK17, sK18)) | ~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | spl19_23), inference(resolution, [], [f2136, f116])).
fof(f116, plain, ! [X0] : (in(sK2(X0), relation_dom(X0)) | one_to_one(X0) | ~ function(X0) | ~ relation(X0)), inference(cnf_transformation, [], [f76])).
fof(f2136, plain, (~ in(sK2(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18))) | spl19_23), inference(avatar_component_clause, [], [f2134])).
fof(f2029, plain, spl19_21, inference(avatar_contradiction_clause, [], [f2028])).
fof(f2028, plain, ($false | spl19_21), inference(subsumption_resolution, [], [f2027, f174])).
fof(f2027, plain, (~ relation(sK18) | spl19_21), inference(subsumption_resolution, [], [f2026, f175])).
fof(f2026, plain, (~ function(sK18) | ~ relation(sK18) | spl19_21), inference(resolution, [], [f1976, f129])).
fof(f1976, plain, (~ function(relation_rng_restriction(sK17, sK18)) | spl19_21), inference(avatar_component_clause, [], [f1974])).
fof(f1987, plain, spl19_20, inference(avatar_contradiction_clause, [], [f1986])).
fof(f1986, plain, ($false | spl19_20), inference(subsumption_resolution, [], [f1985, f174])).
fof(f1985, plain, (~ relation(sK18) | spl19_20), inference(resolution, [], [f1972, f120])).
fof(f1972, plain, (~ relation(relation_rng_restriction(sK17, sK18)) | spl19_20), inference(avatar_component_clause, [], [f1970])).
fof(f1977, plain, (~ spl19_20 | ~ spl19_21 | spl19_18), inference(avatar_split_clause, [], [f1968, f1960, f1974, f1970])).
fof(f1968, plain, (~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | spl19_18), inference(subsumption_resolution, [], [f1967, f177])).
fof(f1967, plain, (one_to_one(relation_rng_restriction(sK17, sK18)) | ~ function(relation_rng_restriction(sK17, sK18)) | ~ relation(relation_rng_restriction(sK17, sK18)) | spl19_18), inference(resolution, [], [f1962, f117])).
fof(f117, plain, ! [X0] : (in(sK3(X0), relation_dom(X0)) | one_to_one(X0) | ~ function(X0) | ~ relation(X0)), inference(cnf_transformation, [], [f76])).
fof(f1962, plain, (~ in(sK3(relation_rng_restriction(sK17, sK18)), relation_dom(relation_rng_restriction(sK17, sK18))) | spl19_18), inference(avatar_component_clause, [], [f1960])).