fof(f814, plain, $false, inference(avatar_sat_refutation, [], [f387, f406, f411, f445, f465, f467, f622, f739, f756, f757, f759, f785, f807, f813])).
fof(f813, plain, (spl29_8 | spl29_41), inference(avatar_contradiction_clause, [], [f812])).
fof(f812, plain, ($false | (spl29_8 | spl29_41)), inference(subsumption_resolution, [], [f809, f434])).
fof(f434, plain, (~ (sK3(sK26(sK1, sK2)) = set_difference(sK4(sK26(sK1, sK2)), singleton(sK1))) | spl29_8), inference(avatar_component_clause, [], [f433])).
fof(f433, plain, (spl29_8 <=> (sK3(sK26(sK1, sK2)) = set_difference(sK4(sK26(sK1, sK2)), singleton(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl29_8])])).
fof(f809, plain, ((sK3(sK26(sK1, sK2)) = set_difference(sK4(sK26(sK1, sK2)), singleton(sK1))) | spl29_41), inference(resolution, [], [f733, f131])).
fof(f131, plain, ! [X2] : (in(sK3(X2), X2) | (sK3(X2) = set_difference(sK4(X2), singleton(sK1)))), inference(cnf_transformation, [], [f83])).
fof(f83, plain, (! [X2] : ((! [X4] : (~ (set_difference(X4, singleton(sK1)) = sK3(X2)) | ~ in(X4, sK2)) | ~ in(sK3(X2), powerset(sK1)) | ~ in(sK3(X2), X2)) & ((((sK3(X2) = set_difference(sK4(X2), singleton(sK1))) & in(sK4(X2), sK2)) & in(sK3(X2), powerset(sK1))) | in(sK3(X2), X2))) & element(sK2, powerset(powerset(succ(sK1)))) & ordinal(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3, sK4])], [f79, f82, f81, f80])).
fof(f80, plain, (? [X0, X1] : (! [X2] : ? [X3] : ((! [X4] : (~ (set_difference(X4, singleton(X0)) = X3) | ~ in(X4, X1)) | ~ in(X3, powerset(X0)) | ~ in(X3, X2)) & ((? [X5] : ((set_difference(X5, singleton(X0)) = X3) & in(X5, X1)) & in(X3, powerset(X0))) | in(X3, X2))) & element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)) => (! [X2] : ? [X3] : ((! [X4] : (~ (set_difference(X4, singleton(sK1)) = X3) | ~ in(X4, sK2)) | ~ in(X3, powerset(sK1)) | ~ in(X3, X2)) & ((? [X5] : ((set_difference(X5, singleton(sK1)) = X3) & in(X5, sK2)) & in(X3, powerset(sK1))) | in(X3, X2))) & element(sK2, powerset(powerset(succ(sK1)))) & ordinal(sK1))), introduced(choice_axiom, [])).
fof(f81, plain, ! [X2] : (? [X3] : ((! [X4] : (~ (set_difference(X4, singleton(sK1)) = X3) | ~ in(X4, sK2)) | ~ in(X3, powerset(sK1)) | ~ in(X3, X2)) & ((? [X5] : ((set_difference(X5, singleton(sK1)) = X3) & in(X5, sK2)) & in(X3, powerset(sK1))) | in(X3, X2))) => ((! [X4] : (~ (set_difference(X4, singleton(sK1)) = sK3(X2)) | ~ in(X4, sK2)) | ~ in(sK3(X2), powerset(sK1)) | ~ in(sK3(X2), X2)) & ((? [X5] : ((set_difference(X5, singleton(sK1)) = sK3(X2)) & in(X5, sK2)) & in(sK3(X2), powerset(sK1))) | in(sK3(X2), X2)))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X2] : (? [X5] : ((set_difference(X5, singleton(sK1)) = sK3(X2)) & in(X5, sK2)) => ((sK3(X2) = set_difference(sK4(X2), singleton(sK1))) & in(sK4(X2), sK2))), introduced(choice_axiom, [])).
fof(f79, plain, ? [X0, X1] : (! [X2] : ? [X3] : ((! [X4] : (~ (set_difference(X4, singleton(X0)) = X3) | ~ in(X4, X1)) | ~ in(X3, powerset(X0)) | ~ in(X3, X2)) & ((? [X5] : ((set_difference(X5, singleton(X0)) = X3) & in(X5, X1)) & in(X3, powerset(X0))) | in(X3, X2))) & element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)), inference(rectify, [], [f78])).
fof(f78, plain, ? [X0, X1] : (! [X2] : ? [X3] : ((! [X4] : (~ (set_difference(X4, singleton(X0)) = X3) | ~ in(X4, X1)) | ~ in(X3, powerset(X0)) | ~ in(X3, X2)) & ((? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0))) | in(X3, X2))) & element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)), inference(flattening, [], [f77])).
fof(f77, plain, ? [X0, X1] : (! [X2] : ? [X3] : (((! [X4] : (~ (set_difference(X4, singleton(X0)) = X3) | ~ in(X4, X1)) | ~ in(X3, powerset(X0))) | ~ in(X3, X2)) & ((? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0))) | in(X3, X2))) & element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)), inference(nnf_transformation, [], [f50])).
fof(f50, plain, ? [X0, X1] : (! [X2] : ? [X3] : ~ (in(X3, X2) <=> (? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0)))) & element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)), inference(flattening, [], [f49])).
fof(f49, plain, ? [X0, X1] : (! [X2] : ? [X3] : ~ (in(X3, X2) <=> (? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0)))) & (element(X1, powerset(powerset(succ(X0)))) & ordinal(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1] : ((element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)) => ? [X2] : ! [X3] : (in(X3, X2) <=> (? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1] : ((element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)) => ? [X2] : ! [X3] : (in(X3, X2) <=> (? [X4] : ((set_difference(X4, singleton(X0)) = X3) & in(X4, X1)) & in(X3, powerset(X0))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU298+1.p', s1_xboole_0__e4_27_3_1__finset_1)).
fof(f733, plain, (~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | spl29_41), inference(avatar_component_clause, [], [f731])).
fof(f731, plain, (spl29_41 <=> in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2))), introduced(avatar_definition, [new_symbols(naming, [spl29_41])])).
fof(f807, plain, (~ spl29_10 | ~ spl29_41 | ~ spl29_44), inference(avatar_contradiction_clause, [], [f806])).
fof(f806, plain, ($false | (~ spl29_10 | ~ spl29_41 | ~ spl29_44)), inference(subsumption_resolution, [], [f805, f444])).
fof(f444, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ spl29_10), inference(avatar_component_clause, [], [f442])).
fof(f442, plain, (spl29_10 <=> in(sK3(sK26(sK1, sK2)), powerset(sK1))), introduced(avatar_definition, [new_symbols(naming, [spl29_10])])).
fof(f805, plain, (~ in(sK3(sK26(sK1, sK2)), powerset(sK1)) | (~ spl29_41 | ~ spl29_44)), inference(subsumption_resolution, [], [f804, f732])).
fof(f732, plain, (in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | ~ spl29_41), inference(avatar_component_clause, [], [f731])).
fof(f804, plain, (~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | ~ in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ spl29_44), inference(equality_resolution, [], [f755])).
fof(f755, plain, (! [X0] : (~ (sK3(X0) = sK3(sK26(sK1, sK2))) | ~ in(sK3(X0), X0) | ~ in(sK3(X0), powerset(sK1))) | ~ spl29_44), inference(avatar_component_clause, [], [f754])).
fof(f754, plain, (spl29_44 <=> ! [X0] : (~ (sK3(X0) = sK3(sK26(sK1, sK2))) | ~ in(sK3(X0), X0) | ~ in(sK3(X0), powerset(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl29_44])])).
fof(f785, plain, (~ spl29_7 | ~ spl29_41 | spl29_43), inference(avatar_contradiction_clause, [], [f784])).
fof(f784, plain, ($false | (~ spl29_7 | ~ spl29_41 | spl29_43)), inference(subsumption_resolution, [], [f783, f732])).
fof(f783, plain, (~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | (~ spl29_7 | spl29_43)), inference(resolution, [], [f752, f410])).
fof(f410, plain, (! [X0] : (in(sK28(sK1, sK2, X0), sK2) | ~ in(X0, sK26(sK1, sK2))) | ~ spl29_7), inference(avatar_component_clause, [], [f409])).
fof(f409, plain, (spl29_7 <=> ! [X0] : (~ in(X0, sK26(sK1, sK2)) | in(sK28(sK1, sK2, X0), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl29_7])])).
fof(f752, plain, (~ in(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), sK2) | spl29_43), inference(avatar_component_clause, [], [f750])).
fof(f750, plain, (spl29_43 <=> in(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), sK2)), introduced(avatar_definition, [new_symbols(naming, [spl29_43])])).
fof(f759, plain, (spl29_41 | spl29_11), inference(avatar_split_clause, [], [f758, f447, f731])).
fof(f447, plain, (spl29_11 <=> in(sK4(sK26(sK1, sK2)), sK2)), introduced(avatar_definition, [new_symbols(naming, [spl29_11])])).
fof(f758, plain, (in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | spl29_11), inference(resolution, [], [f448, f130])).
fof(f130, plain, ! [X2] : (in(sK4(X2), sK2) | in(sK3(X2), X2)), inference(cnf_transformation, [], [f83])).
fof(f448, plain, (~ in(sK4(sK26(sK1, sK2)), sK2) | spl29_11), inference(avatar_component_clause, [], [f447])).
fof(f757, plain, (~ spl29_11 | spl29_44 | ~ spl29_8), inference(avatar_split_clause, [], [f493, f433, f754, f447])).
fof(f493, plain, (! [X0] : (~ (sK3(X0) = sK3(sK26(sK1, sK2))) | ~ in(sK4(sK26(sK1, sK2)), sK2) | ~ in(sK3(X0), powerset(sK1)) | ~ in(sK3(X0), X0)) | ~ spl29_8), inference(superposition, [], [f132, f435])).
fof(f435, plain, ((sK3(sK26(sK1, sK2)) = set_difference(sK4(sK26(sK1, sK2)), singleton(sK1))) | ~ spl29_8), inference(avatar_component_clause, [], [f433])).
fof(f132, plain, ! [X4, X2] : (~ (set_difference(X4, singleton(sK1)) = sK3(X2)) | ~ in(X4, sK2) | ~ in(sK3(X2), powerset(sK1)) | ~ in(sK3(X2), X2)), inference(cnf_transformation, [], [f83])).
fof(f756, plain, (~ spl29_43 | spl29_44 | ~ spl29_12), inference(avatar_split_clause, [], [f580, f462, f754, f750])).
fof(f462, plain, (spl29_12 <=> (sK3(sK26(sK1, sK2)) = set_difference(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), singleton(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl29_12])])).
fof(f580, plain, (! [X0] : (~ (sK3(X0) = sK3(sK26(sK1, sK2))) | ~ in(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), sK2) | ~ in(sK3(X0), powerset(sK1)) | ~ in(sK3(X0), X0)) | ~ spl29_12), inference(superposition, [], [f132, f464])).
fof(f464, plain, ((sK3(sK26(sK1, sK2)) = set_difference(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), singleton(sK1))) | ~ spl29_12), inference(avatar_component_clause, [], [f462])).
fof(f739, plain, (spl29_10 | spl29_5 | ~ spl29_9), inference(avatar_split_clause, [], [f738, f437, f381, f442])).
fof(f381, plain, (spl29_5 <=> sP0(sK1, sK2)), introduced(avatar_definition, [new_symbols(naming, [spl29_5])])).
fof(f437, plain, (spl29_9 <=> (sK3(sK26(sK1, sK2)) = sK27(sK1, sK2, sK3(sK26(sK1, sK2))))), introduced(avatar_definition, [new_symbols(naming, [spl29_9])])).
fof(f738, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | (spl29_5 | ~ spl29_9)), inference(subsumption_resolution, [], [f737, f129])).
fof(f129, plain, ! [X2] : (in(sK3(X2), powerset(sK1)) | in(sK3(X2), X2)), inference(cnf_transformation, [], [f83])).
fof(f737, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | (spl29_5 | ~ spl29_9)), inference(subsumption_resolution, [], [f736, f127])).
fof(f127, plain, ordinal(sK1), inference(cnf_transformation, [], [f83])).
fof(f736, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | ~ ordinal(sK1) | (spl29_5 | ~ spl29_9)), inference(subsumption_resolution, [], [f735, f128])).
fof(f128, plain, element(sK2, powerset(powerset(succ(sK1)))), inference(cnf_transformation, [], [f83])).
fof(f735, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | ~ element(sK2, powerset(powerset(succ(sK1)))) | ~ ordinal(sK1) | (spl29_5 | ~ spl29_9)), inference(subsumption_resolution, [], [f532, f382])).
fof(f382, plain, (~ sP0(sK1, sK2) | spl29_5), inference(avatar_component_clause, [], [f381])).
fof(f532, plain, (in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | sP0(sK1, sK2) | ~ element(sK2, powerset(powerset(succ(sK1)))) | ~ ordinal(sK1) | ~ spl29_9), inference(superposition, [], [f225, f439])).
fof(f439, plain, ((sK3(sK26(sK1, sK2)) = sK27(sK1, sK2, sK3(sK26(sK1, sK2)))) | ~ spl29_9), inference(avatar_component_clause, [], [f437])).
fof(f225, plain, ! [X0, X3, X1] : (in(sK27(X0, X1, X3), powerset(X0)) | ~ in(X3, sK26(X0, X1)) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(cnf_transformation, [], [f126])).
fof(f126, plain, ! [X0, X1] : (! [X3] : ((in(X3, sK26(X0, X1)) | ! [X4] : (! [X5] : (~ (set_difference(X5, singleton(X0)) = X3) | ~ in(X5, X1)) | ~ (X3 = X4) | ~ in(X4, powerset(X0)))) & ((((set_difference(sK28(X0, X1, X3), singleton(X0)) = X3) & in(sK28(X0, X1, X3), X1)) & (sK27(X0, X1, X3) = X3) & in(sK27(X0, X1, X3), powerset(X0))) | ~ in(X3, sK26(X0, X1)))) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27, sK28])], [f122, f125, f124, f123])).
fof(f123, plain, ! [X1, X0] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (! [X5] : (~ (set_difference(X5, singleton(X0)) = X3) | ~ in(X5, X1)) | ~ (X3 = X4) | ~ in(X4, powerset(X0)))) & (? [X6] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (X3 = X6) & in(X6, powerset(X0))) | ~ in(X3, X2))) => ! [X3] : ((in(X3, sK26(X0, X1)) | ! [X4] : (! [X5] : (~ (set_difference(X5, singleton(X0)) = X3) | ~ in(X5, X1)) | ~ (X3 = X4) | ~ in(X4, powerset(X0)))) & (? [X6] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (X3 = X6) & in(X6, powerset(X0))) | ~ in(X3, sK26(X0, X1))))), introduced(choice_axiom, [])).
fof(f124, plain, ! [X3, X1, X0] : (? [X6] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (X3 = X6) & in(X6, powerset(X0))) => (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (sK27(X0, X1, X3) = X3) & in(sK27(X0, X1, X3), powerset(X0)))), introduced(choice_axiom, [])).
fof(f125, plain, ! [X3, X1, X0] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) => ((set_difference(sK28(X0, X1, X3), singleton(X0)) = X3) & in(sK28(X0, X1, X3), X1))), introduced(choice_axiom, [])).
fof(f122, plain, ! [X0, X1] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (! [X5] : (~ (set_difference(X5, singleton(X0)) = X3) | ~ in(X5, X1)) | ~ (X3 = X4) | ~ in(X4, powerset(X0)))) & (? [X6] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (X3 = X6) & in(X6, powerset(X0))) | ~ in(X3, X2))) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(rectify, [], [f121])).
fof(f121, plain, ! [X0, X1] : (? [X7] : ! [X8] : ((in(X8, X7) | ! [X9] : (! [X10] : (~ (set_difference(X10, singleton(X0)) = X8) | ~ in(X10, X1)) | ~ (X8 = X9) | ~ in(X9, powerset(X0)))) & (? [X9] : (? [X10] : ((set_difference(X10, singleton(X0)) = X8) & in(X10, X1)) & (X8 = X9) & in(X9, powerset(X0))) | ~ in(X8, X7))) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(nnf_transformation, [], [f76])).
fof(f76, plain, ! [X0, X1] : (? [X7] : ! [X8] : (in(X8, X7) <=> ? [X9] : (? [X10] : ((set_difference(X10, singleton(X0)) = X8) & in(X10, X1)) & (X8 = X9) & in(X9, powerset(X0)))) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(definition_folding, [], [f74, e75])).
fof(f75, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)) | ~ sP0(X0, X1)), inference(usedef, [], [e75])).
fof(e75, plain, ! [X0, X1] : (sP0(X0, X1) <=> ? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f74, plain, ! [X0, X1] : (? [X7] : ! [X8] : (in(X8, X7) <=> ? [X9] : (? [X10] : ((set_difference(X10, singleton(X0)) = X8) & in(X10, X1)) & (X8 = X9) & in(X9, powerset(X0)))) | ? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(flattening, [], [f73])).
fof(f73, plain, ! [X0, X1] : ((? [X7] : ! [X8] : (in(X8, X7) <=> ? [X9] : (? [X10] : ((set_difference(X10, singleton(X0)) = X8) & in(X10, X1)) & (X8 = X9) & in(X9, powerset(X0)))) | ? [X2, X3, X4] : (~ (X3 = X4) & (? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)))) | (~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0))), inference(ennf_transformation, [], [f44])).
fof(f44, plain, ! [X0, X1] : ((element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)) => (! [X2, X3, X4] : ((? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)) => (X3 = X4)) => ? [X7] : ! [X8] : (in(X8, X7) <=> ? [X9] : (? [X10] : ((set_difference(X10, singleton(X0)) = X8) & in(X10, X1)) & (X8 = X9) & in(X9, powerset(X0)))))), inference(rectify, [], [f43])).
fof(f43, plain, ! [X0, X1] : ((element(X1, powerset(powerset(succ(X0)))) & ordinal(X0)) => (! [X2, X3, X4] : ((? [X6] : ((set_difference(X6, singleton(X0)) = X4) & in(X6, X1)) & (X2 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X3) & in(X5, X1)) & (X2 = X3)) => (X3 = X4)) => ? [X2] : ! [X3] : (in(X3, X2) <=> ? [X4] : (? [X7] : ((set_difference(X7, singleton(X0)) = X3) & in(X7, X1)) & (X3 = X4) & in(X4, powerset(X0)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU298+1.p', s1_tarski__e4_27_3_1__finset_1__1)).
fof(f622, plain, (spl29_5 | ~ spl29_8 | ~ spl29_10 | ~ spl29_11), inference(avatar_contradiction_clause, [], [f621])).
fof(f621, plain, ($false | (spl29_5 | ~ spl29_8 | ~ spl29_10 | ~ spl29_11)), inference(subsumption_resolution, [], [f620, f128])).
fof(f620, plain, (~ element(sK2, powerset(powerset(succ(sK1)))) | (spl29_5 | ~ spl29_8 | ~ spl29_10 | ~ spl29_11)), inference(subsumption_resolution, [], [f619, f382])).
fof(f619, plain, (sP0(sK1, sK2) | ~ element(sK2, powerset(powerset(succ(sK1)))) | (~ spl29_8 | ~ spl29_10 | ~ spl29_11)), inference(subsumption_resolution, [], [f599, f449])).
fof(f449, plain, (in(sK4(sK26(sK1, sK2)), sK2) | ~ spl29_11), inference(avatar_component_clause, [], [f447])).
fof(f599, plain, (~ in(sK4(sK26(sK1, sK2)), sK2) | sP0(sK1, sK2) | ~ element(sK2, powerset(powerset(succ(sK1)))) | (~ spl29_8 | ~ spl29_10 | ~ spl29_11)), inference(resolution, [], [f499, f535])).
fof(f535, plain, (~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | (~ spl29_8 | ~ spl29_10 | ~ spl29_11)), inference(subsumption_resolution, [], [f534, f444])).
fof(f534, plain, (~ in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ in(sK3(sK26(sK1, sK2)), sK26(sK1, sK2)) | (~ spl29_8 | ~ spl29_11)), inference(equality_resolution, [], [f497])).
fof(f497, plain, (! [X0] : (~ (sK3(X0) = sK3(sK26(sK1, sK2))) | ~ in(sK3(X0), powerset(sK1)) | ~ in(sK3(X0), X0)) | (~ spl29_8 | ~ spl29_11)), inference(subsumption_resolution, [], [f493, f449])).
fof(f499, plain, (! [X1] : (in(sK3(sK26(sK1, sK2)), sK26(sK1, X1)) | ~ in(sK4(sK26(sK1, sK2)), X1) | sP0(sK1, X1) | ~ element(X1, powerset(powerset(succ(sK1))))) | (~ spl29_8 | ~ spl29_10)), inference(subsumption_resolution, [], [f498, f127])).
fof(f498, plain, (! [X1] : (in(sK3(sK26(sK1, sK2)), sK26(sK1, X1)) | ~ in(sK4(sK26(sK1, sK2)), X1) | sP0(sK1, X1) | ~ element(X1, powerset(powerset(succ(sK1)))) | ~ ordinal(sK1)) | (~ spl29_8 | ~ spl29_10)), inference(subsumption_resolution, [], [f494, f444])).
fof(f494, plain, (! [X1] : (in(sK3(sK26(sK1, sK2)), sK26(sK1, X1)) | ~ in(sK4(sK26(sK1, sK2)), X1) | ~ in(sK3(sK26(sK1, sK2)), powerset(sK1)) | sP0(sK1, X1) | ~ element(X1, powerset(powerset(succ(sK1)))) | ~ ordinal(sK1)) | ~ spl29_8), inference(superposition, [], [f231, f435])).
fof(f231, plain, ! [X0, X5, X1] : (in(set_difference(X5, singleton(X0)), sK26(X0, X1)) | ~ in(X5, X1) | ~ in(set_difference(X5, singleton(X0)), powerset(X0)) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(equality_resolution, [], [f230])).
fof(f230, plain, ! [X4, X0, X5, X1] : (in(set_difference(X5, singleton(X0)), sK26(X0, X1)) | ~ in(X5, X1) | ~ (set_difference(X5, singleton(X0)) = X4) | ~ in(X4, powerset(X0)) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(equality_resolution, [], [f229])).
fof(f229, plain, ! [X4, X0, X5, X3, X1] : (in(X3, sK26(X0, X1)) | ~ (set_difference(X5, singleton(X0)) = X3) | ~ in(X5, X1) | ~ (X3 = X4) | ~ in(X4, powerset(X0)) | sP0(X0, X1) | ~ element(X1, powerset(powerset(succ(X0)))) | ~ ordinal(X0)), inference(cnf_transformation, [], [f126])).
fof(f467, plain, (spl29_11 | spl29_12 | spl29_5), inference(avatar_split_clause, [], [f455, f381, f462, f447])).
fof(f455, plain, ((sK3(sK26(sK1, sK2)) = set_difference(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), singleton(sK1))) | in(sK4(sK26(sK1, sK2)), sK2) | spl29_5), inference(resolution, [], [f419, f130])).
fof(f419, plain, (! [X0] : (~ in(X0, sK26(sK1, sK2)) | (set_difference(sK28(sK1, sK2, X0), singleton(sK1)) = X0)) | spl29_5), inference(subsumption_resolution, [], [f418, f127])).
fof(f418, plain, (! [X0] : (~ in(X0, sK26(sK1, sK2)) | (set_difference(sK28(sK1, sK2, X0), singleton(sK1)) = X0) | ~ ordinal(sK1)) | spl29_5), inference(subsumption_resolution, [], [f413, f382])).
fof(f413, plain, ! [X0] : (~ in(X0, sK26(sK1, sK2)) | sP0(sK1, sK2) | (set_difference(sK28(sK1, sK2, X0), singleton(sK1)) = X0) | ~ ordinal(sK1)), inference(resolution, [], [f228, f128])).
fof(f228, plain, ! [X0, X3, X1] : (~ element(X1, powerset(powerset(succ(X0)))) | ~ in(X3, sK26(X0, X1)) | sP0(X0, X1) | (set_difference(sK28(X0, X1, X3), singleton(X0)) = X3) | ~ ordinal(X0)), inference(cnf_transformation, [], [f126])).
fof(f465, plain, (spl29_8 | spl29_12 | spl29_5), inference(avatar_split_clause, [], [f453, f381, f462, f433])).
fof(f453, plain, ((sK3(sK26(sK1, sK2)) = set_difference(sK28(sK1, sK2, sK3(sK26(sK1, sK2))), singleton(sK1))) | (sK3(sK26(sK1, sK2)) = set_difference(sK4(sK26(sK1, sK2)), singleton(sK1))) | spl29_5), inference(resolution, [], [f419, f131])).
fof(f445, plain, (spl29_10 | spl29_9 | ~ spl29_6), inference(avatar_split_clause, [], [f425, f385, f437, f442])).
fof(f385, plain, (spl29_6 <=> ! [X0] : (~ in(X0, sK26(sK1, sK2)) | (sK27(sK1, sK2, X0) = X0))), introduced(avatar_definition, [new_symbols(naming, [spl29_6])])).
fof(f425, plain, ((sK3(sK26(sK1, sK2)) = sK27(sK1, sK2, sK3(sK26(sK1, sK2)))) | in(sK3(sK26(sK1, sK2)), powerset(sK1)) | ~ spl29_6), inference(resolution, [], [f386, f129])).
fof(f386, plain, (! [X0] : (~ in(X0, sK26(sK1, sK2)) | (sK27(sK1, sK2, X0) = X0)) | ~ spl29_6), inference(avatar_component_clause, [], [f385])).
fof(f411, plain, (spl29_5 | spl29_7), inference(avatar_split_clause, [], [f407, f409, f381])).
fof(f407, plain, ! [X0] : (~ in(X0, sK26(sK1, sK2)) | sP0(sK1, sK2) | in(sK28(sK1, sK2, X0), sK2)), inference(subsumption_resolution, [], [f396, f127])).
fof(f396, plain, ! [X0] : (~ in(X0, sK26(sK1, sK2)) | sP0(sK1, sK2) | in(sK28(sK1, sK2, X0), sK2) | ~ ordinal(sK1)), inference(resolution, [], [f227, f128])).
fof(f227, plain, ! [X0, X3, X1] : (~ element(X1, powerset(powerset(succ(X0)))) | ~ in(X3, sK26(X0, X1)) | sP0(X0, X1) | in(sK28(X0, X1, X3), X1) | ~ ordinal(X0)), inference(cnf_transformation, [], [f126])).
fof(f406, plain, ~ spl29_5, inference(avatar_contradiction_clause, [], [f405])).
fof(f405, plain, ($false | ~ spl29_5), inference(subsumption_resolution, [], [f404, f383])).
fof(f383, plain, (sP0(sK1, sK2) | ~ spl29_5), inference(avatar_component_clause, [], [f381])).
fof(f404, plain, (~ sP0(sK1, sK2) | ~ spl29_5), inference(subsumption_resolution, [], [f403, f393])).
fof(f393, plain, ((sK22(sK1, sK2) = sK21(sK1, sK2)) | ~ spl29_5), inference(resolution, [], [f383, f218])).
fof(f218, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK21(X0, X1) = sK22(X0, X1))), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ! [X0, X1] : ((~ (sK22(X0, X1) = sK23(X0, X1)) & ((sK23(X0, X1) = set_difference(sK24(X0, X1), singleton(X0))) & in(sK24(X0, X1), X1)) & (sK21(X0, X1) = sK23(X0, X1)) & ((sK22(X0, X1) = set_difference(sK25(X0, X1), singleton(X0))) & in(sK25(X0, X1), X1)) & (sK21(X0, X1) = sK22(X0, X1))) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22, sK23, sK24, sK25])], [f116, f119, f118, f117])).
fof(f117, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)) => (~ (sK22(X0, X1) = sK23(X0, X1)) & ? [X5] : ((set_difference(X5, singleton(X0)) = sK23(X0, X1)) & in(X5, X1)) & (sK21(X0, X1) = sK23(X0, X1)) & ? [X6] : ((set_difference(X6, singleton(X0)) = sK22(X0, X1)) & in(X6, X1)) & (sK21(X0, X1) = sK22(X0, X1)))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X1, X0] : (? [X5] : ((set_difference(X5, singleton(X0)) = sK23(X0, X1)) & in(X5, X1)) => ((sK23(X0, X1) = set_difference(sK24(X0, X1), singleton(X0))) & in(sK24(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f119, plain, ! [X1, X0] : (? [X6] : ((set_difference(X6, singleton(X0)) = sK22(X0, X1)) & in(X6, X1)) => ((sK22(X0, X1) = set_difference(sK25(X0, X1), singleton(X0))) & in(sK25(X0, X1), X1))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : ((set_difference(X5, singleton(X0)) = X4) & in(X5, X1)) & (X2 = X4) & ? [X6] : ((set_difference(X6, singleton(X0)) = X3) & in(X6, X1)) & (X2 = X3)) | ~ sP0(X0, X1)), inference(nnf_transformation, [], [f75])).
fof(f403, plain, (~ (sK22(sK1, sK2) = sK21(sK1, sK2)) | ~ sP0(sK1, sK2) | ~ spl29_5), inference(superposition, [], [f224, f392])).
fof(f392, plain, ((sK23(sK1, sK2) = sK21(sK1, sK2)) | ~ spl29_5), inference(resolution, [], [f383, f221])).
fof(f221, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK21(X0, X1) = sK23(X0, X1))), inference(cnf_transformation, [], [f120])).
fof(f224, plain, ! [X0, X1] : (~ (sK22(X0, X1) = sK23(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f120])).
fof(f387, plain, (spl29_5 | spl29_6), inference(avatar_split_clause, [], [f379, f385, f381])).
fof(f379, plain, ! [X0] : (~ in(X0, sK26(sK1, sK2)) | sP0(sK1, sK2) | (sK27(sK1, sK2, X0) = X0)), inference(subsumption_resolution, [], [f374, f127])).
fof(f374, plain, ! [X0] : (~ in(X0, sK26(sK1, sK2)) | sP0(sK1, sK2) | (sK27(sK1, sK2, X0) = X0) | ~ ordinal(sK1)), inference(resolution, [], [f226, f128])).
fof(f226, plain, ! [X0, X3, X1] : (~ element(X1, powerset(powerset(succ(X0)))) | ~ in(X3, sK26(X0, X1)) | sP0(X0, X1) | (sK27(X0, X1, X3) = X3) | ~ ordinal(X0)), inference(cnf_transformation, [], [f126])).