fof(f2488, plain, $false, inference(avatar_sat_refutation, [], [f1741, f1743, f1925, f1932, f1941, f1943, f1945, f2126, f2241, f2333, f2336, f2356, f2487])).
fof(f2487, plain, (spl22_11 | spl22_30), inference(avatar_contradiction_clause, [], [f2486])).
fof(f2486, plain, ($false | (spl22_11 | spl22_30)), inference(subsumption_resolution, [], [f2485, f1920])).
fof(f1920, plain, (~ in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | spl22_30), inference(avatar_component_clause, [], [f1918])).
fof(f1918, plain, (spl22_30 <=> in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl22_30])])).
fof(f2485, plain, (in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | spl22_11), inference(resolution, [], [f930, f62])).
fof(f62, plain, ! [X2] : (in(sK5(X2), sK2) | in(sK4(X2), X2)), inference(cnf_transformation, [], [f36])).
fof(f36, plain, ! [X2] : ((! [X4, X5] : (~ (singleton(X4) = X5) | ~ in(X4, sK2) | ~ (ordered_pair(X4, X5) = sK4(X2))) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)) & ((((sK6(X2) = singleton(sK5(X2))) & in(sK5(X2), sK2) & (sK4(X2) = ordered_pair(sK5(X2), sK6(X2)))) & in(sK4(X2), cartesian_product2(sK2, sK3))) | in(sK4(X2), X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3, sK4, sK5, sK6])], [f32, f35, f34, f33])).
fof(f33, plain, (? [X0, X1] : ! [X2] : ? [X3] : ((! [X4, X5] : (~ (singleton(X4) = X5) | ~ in(X4, X0) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(X0, X1)) | ~ in(X3, X2)) & ((? [X6, X7] : ((singleton(X6) = X7) & in(X6, X0) & (ordered_pair(X6, X7) = X3)) & in(X3, cartesian_product2(X0, X1))) | in(X3, X2))) => ! [X2] : ? [X3] : ((! [X5, X4] : (~ (singleton(X4) = X5) | ~ in(X4, sK2) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(sK2, sK3)) | ~ in(X3, X2)) & ((? [X7, X6] : ((singleton(X6) = X7) & in(X6, sK2) & (ordered_pair(X6, X7) = X3)) & in(X3, cartesian_product2(sK2, sK3))) | in(X3, X2)))), introduced(choice_axiom, [])).
fof(f34, plain, ! [X2] : (? [X3] : ((! [X5, X4] : (~ (singleton(X4) = X5) | ~ in(X4, sK2) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(sK2, sK3)) | ~ in(X3, X2)) & ((? [X7, X6] : ((singleton(X6) = X7) & in(X6, sK2) & (ordered_pair(X6, X7) = X3)) & in(X3, cartesian_product2(sK2, sK3))) | in(X3, X2))) => ((! [X5, X4] : (~ (singleton(X4) = X5) | ~ in(X4, sK2) | ~ (ordered_pair(X4, X5) = sK4(X2))) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)) & ((? [X7, X6] : ((singleton(X6) = X7) & in(X6, sK2) & (ordered_pair(X6, X7) = sK4(X2))) & in(sK4(X2), cartesian_product2(sK2, sK3))) | in(sK4(X2), X2)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X2] : (? [X7, X6] : ((singleton(X6) = X7) & in(X6, sK2) & (ordered_pair(X6, X7) = sK4(X2))) => ((sK6(X2) = singleton(sK5(X2))) & in(sK5(X2), sK2) & (sK4(X2) = ordered_pair(sK5(X2), sK6(X2))))), introduced(choice_axiom, [])).
fof(f32, plain, ? [X0, X1] : ! [X2] : ? [X3] : ((! [X4, X5] : (~ (singleton(X4) = X5) | ~ in(X4, X0) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(X0, X1)) | ~ in(X3, X2)) & ((? [X6, X7] : ((singleton(X6) = X7) & in(X6, X0) & (ordered_pair(X6, X7) = X3)) & in(X3, cartesian_product2(X0, X1))) | in(X3, X2))), inference(rectify, [], [f31])).
fof(f31, plain, ? [X0, X1] : ! [X2] : ? [X3] : ((! [X4, X5] : (~ (singleton(X4) = X5) | ~ in(X4, X0) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(X0, X1)) | ~ in(X3, X2)) & ((? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & in(X3, cartesian_product2(X0, X1))) | in(X3, X2))), inference(flattening, [], [f30])).
fof(f30, plain, ? [X0, X1] : ! [X2] : ? [X3] : (((! [X4, X5] : (~ (singleton(X4) = X5) | ~ in(X4, X0) | ~ (ordered_pair(X4, X5) = X3)) | ~ in(X3, cartesian_product2(X0, X1))) | ~ in(X3, X2)) & ((? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & in(X3, cartesian_product2(X0, X1))) | in(X3, X2))), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ? [X0, X1] : ! [X2] : ? [X3] : ~ (in(X3, X2) <=> (? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & in(X3, cartesian_product2(X0, X1)))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1] : ? [X2] : ! [X3] : (in(X3, X2) <=> (? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & in(X3, cartesian_product2(X0, X1)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1] : ? [X2] : ! [X3] : (in(X3, X2) <=> (? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & in(X3, cartesian_product2(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU281+1.p', s1_xboole_0__e16_22__wellord2__1)).
fof(f930, plain, (~ in(sK5(sK18(sK2, sK3)), sK2) | spl22_11), inference(avatar_component_clause, [], [f929])).
fof(f929, plain, (spl22_11 <=> in(sK5(sK18(sK2, sK3)), sK2)), introduced(avatar_definition, [new_symbols(naming, [spl22_11])])).
fof(f2356, plain, (spl22_27 | spl22_31 | ~ spl22_21), inference(avatar_split_clause, [], [f2355, f1617, f1922, f1715])).
fof(f1715, plain, (spl22_27 <=> sP1(sK2)), introduced(avatar_definition, [new_symbols(naming, [spl22_27])])).
fof(f1922, plain, (spl22_31 <=> in(sK4(sK18(sK2, sK3)), cartesian_product2(sK2, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl22_31])])).
fof(f1617, plain, (spl22_21 <=> (sK4(sK18(sK2, sK3)) = sK19(sK2, sK3, sK4(sK18(sK2, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl22_21])])).
fof(f2355, plain, (in(sK4(sK18(sK2, sK3)), cartesian_product2(sK2, sK3)) | sP1(sK2) | ~ spl22_21), inference(subsumption_resolution, [], [f2196, f60])).
fof(f60, plain, ! [X2] : (in(sK4(X2), cartesian_product2(sK2, sK3)) | in(sK4(X2), X2)), inference(cnf_transformation, [], [f36])).
fof(f2196, plain, (in(sK4(sK18(sK2, sK3)), cartesian_product2(sK2, sK3)) | ~ in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | sP1(sK2) | ~ spl22_21), inference(superposition, [], [f92, f1619])).
fof(f1619, plain, ((sK4(sK18(sK2, sK3)) = sK19(sK2, sK3, sK4(sK18(sK2, sK3)))) | ~ spl22_21), inference(avatar_component_clause, [], [f1617])).
fof(f92, plain, ! [X0, X3, X1] : (in(sK19(X0, X1, X3), cartesian_product2(X0, X1)) | ~ in(X3, sK18(X0, X1)) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f59, plain, ! [X0, X1] : (! [X3] : ((in(X3, sK18(X0, X1)) | ! [X4] : (! [X5, X6] : (~ (singleton(X5) = X6) | ~ in(X5, X0) | ~ (ordered_pair(X5, X6) = X3)) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)))) & ((((sK21(X0, X3) = singleton(sK20(X0, X3))) & in(sK20(X0, X3), X0) & (ordered_pair(sK20(X0, X3), sK21(X0, X3)) = X3)) & (sK19(X0, X1, X3) = X3) & in(sK19(X0, X1, X3), cartesian_product2(X0, X1))) | ~ in(X3, sK18(X0, X1)))) | sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK18, sK19, sK20, sK21])], [f55, f58, f57, f56])).
fof(f56, plain, ! [X1, X0] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (! [X5, X6] : (~ (singleton(X5) = X6) | ~ in(X5, X0) | ~ (ordered_pair(X5, X6) = X3)) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)))) & (? [X7] : (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) & (X3 = X7) & in(X7, cartesian_product2(X0, X1))) | ~ in(X3, X2))) => ! [X3] : ((in(X3, sK18(X0, X1)) | ! [X4] : (! [X5, X6] : (~ (singleton(X5) = X6) | ~ in(X5, X0) | ~ (ordered_pair(X5, X6) = X3)) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)))) & (? [X7] : (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) & (X3 = X7) & in(X7, cartesian_product2(X0, X1))) | ~ in(X3, sK18(X0, X1))))), introduced(choice_axiom, [])).
fof(f57, plain, ! [X3, X1, X0] : (? [X7] : (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) & (X3 = X7) & in(X7, cartesian_product2(X0, X1))) => (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) & (sK19(X0, X1, X3) = X3) & in(sK19(X0, X1, X3), cartesian_product2(X0, X1)))), introduced(choice_axiom, [])).
fof(f58, plain, ! [X3, X0] : (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) => ((sK21(X0, X3) = singleton(sK20(X0, X3))) & in(sK20(X0, X3), X0) & (ordered_pair(sK20(X0, X3), sK21(X0, X3)) = X3))), introduced(choice_axiom, [])).
fof(f55, plain, ! [X0, X1] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (! [X5, X6] : (~ (singleton(X5) = X6) | ~ in(X5, X0) | ~ (ordered_pair(X5, X6) = X3)) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)))) & (? [X7] : (? [X8, X9] : ((singleton(X8) = X9) & in(X8, X0) & (ordered_pair(X8, X9) = X3)) & (X3 = X7) & in(X7, cartesian_product2(X0, X1))) | ~ in(X3, X2))) | sP1(X0)), inference(rectify, [], [f54])).
fof(f54, plain, ! [X0, X1] : (? [X9] : ! [X10] : ((in(X10, X9) | ! [X11] : (! [X12, X13] : (~ (singleton(X12) = X13) | ~ in(X12, X0) | ~ (ordered_pair(X12, X13) = X10)) | ~ (X10 = X11) | ~ in(X11, cartesian_product2(X0, X1)))) & (? [X11] : (? [X12, X13] : ((singleton(X12) = X13) & in(X12, X0) & (ordered_pair(X12, X13) = X10)) & (X10 = X11) & in(X11, cartesian_product2(X0, X1))) | ~ in(X10, X9))) | sP1(X0)), inference(nnf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1] : (? [X9] : ! [X10] : (in(X10, X9) <=> ? [X11] : (? [X12, X13] : ((singleton(X12) = X13) & in(X12, X0) & (ordered_pair(X12, X13) = X10)) & (X10 = X11) & in(X11, cartesian_product2(X0, X1)))) | sP1(X0)), inference(definition_folding, [], [f26, e28, e27])).
fof(f27, plain, ! [X0, X3] : (? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X3)) | ~ sP0(X0, X3)), inference(usedef, [], [e27])).
fof(e27, plain, ! [X0, X3] : (sP0(X0, X3) <=> ? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f28, plain, ! [X0] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & sP0(X0, X3) & (X2 = X3)) | ~ sP1(X0)), inference(usedef, [], [e28])).
fof(e28, plain, ! [X0] : (sP1(X0) <=> ? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & sP0(X0, X3) & (X2 = X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f26, plain, ! [X0, X1] : (? [X9] : ! [X10] : (in(X10, X9) <=> ? [X11] : (? [X12, X13] : ((singleton(X12) = X13) & in(X12, X0) & (ordered_pair(X12, X13) = X10)) & (X10 = X11) & in(X11, cartesian_product2(X0, X1)))) | ? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & ? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X3)) & (X2 = X3))), inference(flattening, [], [f25])).
fof(f25, plain, ! [X0, X1] : (? [X9] : ! [X10] : (in(X10, X9) <=> ? [X11] : (? [X12, X13] : ((singleton(X12) = X13) & in(X12, X0) & (ordered_pair(X12, X13) = X10)) & (X10 = X11) & in(X11, cartesian_product2(X0, X1)))) | ? [X2, X3, X4] : (~ (X3 = X4) & (? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & ? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X3)) & (X2 = X3)))), inference(ennf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1] : (! [X2, X3, X4] : ((? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & ? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X3)) & (X2 = X3)) => (X3 = X4)) => ? [X9] : ! [X10] : (in(X10, X9) <=> ? [X11] : (? [X12, X13] : ((singleton(X12) = X13) & in(X12, X0) & (ordered_pair(X12, X13) = X10)) & (X10 = X11) & in(X11, cartesian_product2(X0, X1))))), inference(rectify, [], [f16])).
fof(f16, plain, ! [X0, X1] : (! [X2, X3, X4] : ((? [X7, X8] : ((singleton(X7) = X8) & in(X7, X0) & (ordered_pair(X7, X8) = X4)) & (X2 = X4) & ? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X3)) & (X2 = X3)) => (X3 = X4)) => ? [X2] : ! [X3] : (in(X3, X2) <=> ? [X4] : (? [X9, X10] : ((singleton(X9) = X10) & in(X9, X0) & (ordered_pair(X9, X10) = X3)) & (X3 = X4) & in(X4, cartesian_product2(X0, X1))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU281+1.p', s1_tarski__e16_22__wellord2__2)).
fof(f2336, plain, (~ spl22_27 | ~ spl22_1 | ~ spl22_10), inference(avatar_split_clause, [], [f2335, f771, f697, f1715])).
fof(f697, plain, (spl22_1 <=> (sK11(sK2) = sK13(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl22_1])])).
fof(f771, plain, (spl22_10 <=> (sK11(sK2) = sK12(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl22_10])])).
fof(f2335, plain, (~ sP1(sK2) | (~ spl22_1 | ~ spl22_10)), inference(subsumption_resolution, [], [f2280, f773])).
fof(f773, plain, ((sK11(sK2) = sK12(sK2)) | ~ spl22_10), inference(avatar_component_clause, [], [f771])).
fof(f2280, plain, (~ (sK11(sK2) = sK12(sK2)) | ~ sP1(sK2) | ~ spl22_1), inference(superposition, [], [f88, f699])).
fof(f699, plain, ((sK11(sK2) = sK13(sK2)) | ~ spl22_1), inference(avatar_component_clause, [], [f697])).
fof(f88, plain, ! [X0] : (~ (sK12(X0) = sK13(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f49])).
fof(f49, plain, ! [X0] : ((~ (sK12(X0) = sK13(X0)) & ((sK15(X0) = singleton(sK14(X0))) & in(sK14(X0), X0) & (sK13(X0) = ordered_pair(sK14(X0), sK15(X0)))) & (sK11(X0) = sK13(X0)) & sP0(X0, sK12(X0)) & (sK11(X0) = sK12(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12, sK13, sK14, sK15])], [f46, f48, f47])).
fof(f47, plain, ! [X0] : (? [X1, X2, X3] : (~ (X2 = X3) & ? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & (X1 = X3) & sP0(X0, X2) & (X1 = X2)) => (~ (sK12(X0) = sK13(X0)) & ? [X5, X4] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = sK13(X0))) & (sK11(X0) = sK13(X0)) & sP0(X0, sK12(X0)) & (sK11(X0) = sK12(X0)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X0] : (? [X5, X4] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = sK13(X0))) => ((sK15(X0) = singleton(sK14(X0))) & in(sK14(X0), X0) & (sK13(X0) = ordered_pair(sK14(X0), sK15(X0))))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X0] : (? [X1, X2, X3] : (~ (X2 = X3) & ? [X4, X5] : ((singleton(X4) = X5) & in(X4, X0) & (ordered_pair(X4, X5) = X3)) & (X1 = X3) & sP0(X0, X2) & (X1 = X2)) | ~ sP1(X0)), inference(rectify, [], [f45])).
fof(f45, plain, ! [X0] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : ((singleton(X5) = X6) & in(X5, X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & sP0(X0, X3) & (X2 = X3)) | ~ sP1(X0)), inference(nnf_transformation, [], [f28])).
fof(f2333, plain, (spl22_6 | spl22_30), inference(avatar_contradiction_clause, [], [f2332])).
fof(f2332, plain, ($false | (spl22_6 | spl22_30)), inference(subsumption_resolution, [], [f2329, f752])).
fof(f752, plain, (~ (sK6(sK18(sK2, sK3)) = singleton(sK5(sK18(sK2, sK3)))) | spl22_6), inference(avatar_component_clause, [], [f751])).
fof(f751, plain, (spl22_6 <=> (sK6(sK18(sK2, sK3)) = singleton(sK5(sK18(sK2, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl22_6])])).
fof(f2329, plain, ((sK6(sK18(sK2, sK3)) = singleton(sK5(sK18(sK2, sK3)))) | spl22_30), inference(resolution, [], [f1920, f63])).
fof(f63, plain, ! [X2] : (in(sK4(X2), X2) | (sK6(X2) = singleton(sK5(X2)))), inference(cnf_transformation, [], [f36])).
fof(f2241, plain, (spl22_20 | ~ spl22_19 | spl22_27 | ~ spl22_28 | ~ spl22_30), inference(avatar_split_clause, [], [f2240, f1918, f1721, f1715, f1609, f1613])).
fof(f1613, plain, (spl22_20 <=> ! [X0] : (~ (sK4(X0) = sK4(sK18(sK2, sK3))) | ~ in(sK4(X0), cartesian_product2(sK2, sK3)) | ~ in(sK4(X0), X0))), introduced(avatar_definition, [new_symbols(naming, [spl22_20])])).
fof(f1609, plain, (spl22_19 <=> (sK21(sK2, sK4(sK18(sK2, sK3))) = singleton(sK20(sK2, sK4(sK18(sK2, sK3)))))), introduced(avatar_definition, [new_symbols(naming, [spl22_19])])).
fof(f1721, plain, (spl22_28 <=> (sK4(sK18(sK2, sK3)) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3)))))), introduced(avatar_definition, [new_symbols(naming, [spl22_28])])).
fof(f2240, plain, (! [X2] : (~ (sK4(X2) = sK4(sK18(sK2, sK3))) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)) | (~ spl22_19 | spl22_27 | ~ spl22_28 | ~ spl22_30)), inference(forward_demodulation, [], [f2239, f1723])).
fof(f1723, plain, ((sK4(sK18(sK2, sK3)) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3))))) | ~ spl22_28), inference(avatar_component_clause, [], [f1721])).
fof(f2239, plain, (! [X2] : (~ (sK4(X2) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3))))) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)) | (~ spl22_19 | spl22_27 | ~ spl22_30)), inference(subsumption_resolution, [], [f2236, f2150])).
fof(f2150, plain, (in(sK20(sK2, sK4(sK18(sK2, sK3))), sK2) | (spl22_27 | ~ spl22_30)), inference(subsumption_resolution, [], [f2148, f1716])).
fof(f1716, plain, (~ sP1(sK2) | spl22_27), inference(avatar_component_clause, [], [f1715])).
fof(f2148, plain, (in(sK20(sK2, sK4(sK18(sK2, sK3))), sK2) | sP1(sK2) | ~ spl22_30), inference(resolution, [], [f1919, f95])).
fof(f95, plain, ! [X0, X3, X1] : (~ in(X3, sK18(X0, X1)) | in(sK20(X0, X3), X0) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f1919, plain, (in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | ~ spl22_30), inference(avatar_component_clause, [], [f1918])).
fof(f2236, plain, (! [X2] : (~ (sK4(X2) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3))))) | ~ in(sK20(sK2, sK4(sK18(sK2, sK3))), sK2) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)) | ~ spl22_19), inference(superposition, [], [f98, f1611])).
fof(f1611, plain, ((sK21(sK2, sK4(sK18(sK2, sK3))) = singleton(sK20(sK2, sK4(sK18(sK2, sK3))))) | ~ spl22_19), inference(avatar_component_clause, [], [f1609])).
fof(f98, plain, ! [X4, X2] : (~ (sK4(X2) = ordered_pair(X4, singleton(X4))) | ~ in(X4, sK2) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)), inference(equality_resolution, [], [f64])).
fof(f64, plain, ! [X4, X2, X5] : (~ (singleton(X4) = X5) | ~ in(X4, sK2) | ~ (ordered_pair(X4, X5) = sK4(X2)) | ~ in(sK4(X2), cartesian_product2(sK2, sK3)) | ~ in(sK4(X2), X2)), inference(cnf_transformation, [], [f36])).
fof(f2126, plain, (~ spl22_6 | ~ spl22_11 | ~ spl22_16 | spl22_27 | spl22_30), inference(avatar_contradiction_clause, [], [f2125])).
fof(f2125, plain, ($false | (~ spl22_6 | ~ spl22_11 | ~ spl22_16 | spl22_27 | spl22_30)), inference(subsumption_resolution, [], [f2124, f1716])).
fof(f2124, plain, (sP1(sK2) | (~ spl22_6 | ~ spl22_11 | ~ spl22_16 | spl22_30)), inference(subsumption_resolution, [], [f2123, f931])).
fof(f931, plain, (in(sK5(sK18(sK2, sK3)), sK2) | ~ spl22_11), inference(avatar_component_clause, [], [f929])).
fof(f2123, plain, (~ in(sK5(sK18(sK2, sK3)), sK2) | sP1(sK2) | (~ spl22_6 | ~ spl22_16 | spl22_30)), inference(subsumption_resolution, [], [f2103, f1920])).
fof(f2103, plain, (in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | ~ in(sK5(sK18(sK2, sK3)), sK2) | sP1(sK2) | (~ spl22_6 | ~ spl22_16)), inference(duplicate_literal_removal, [], [f2102])).
fof(f2102, plain, (in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | ~ in(sK5(sK18(sK2, sK3)), sK2) | sP1(sK2) | in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | (~ spl22_6 | ~ spl22_16)), inference(resolution, [], [f1949, f60])).
fof(f1949, plain, (! [X0, X1] : (~ in(sK4(sK18(sK2, sK3)), cartesian_product2(X0, X1)) | in(sK4(sK18(sK2, sK3)), sK18(X0, X1)) | ~ in(sK5(sK18(sK2, sK3)), X0) | sP1(X0)) | (~ spl22_6 | ~ spl22_16)), inference(forward_demodulation, [], [f1948, f1518])).
fof(f1518, plain, ((ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))) = sK4(sK18(sK2, sK3))) | ~ spl22_16), inference(avatar_component_clause, [], [f1516])).
fof(f1516, plain, (spl22_16 <=> (ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))) = sK4(sK18(sK2, sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl22_16])])).
fof(f1948, plain, (! [X0, X1] : (in(sK4(sK18(sK2, sK3)), sK18(X0, X1)) | ~ in(ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))), cartesian_product2(X0, X1)) | ~ in(sK5(sK18(sK2, sK3)), X0) | sP1(X0)) | (~ spl22_6 | ~ spl22_16)), inference(backward_demodulation, [], [f1551, f1518])).
fof(f1551, plain, (! [X0, X1] : (~ in(ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))), cartesian_product2(X0, X1)) | ~ in(sK5(sK18(sK2, sK3)), X0) | in(ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))), sK18(X0, X1)) | sP1(X0)) | ~ spl22_6), inference(superposition, [], [f101, f753])).
fof(f753, plain, ((sK6(sK18(sK2, sK3)) = singleton(sK5(sK18(sK2, sK3)))) | ~ spl22_6), inference(avatar_component_clause, [], [f751])).
fof(f101, plain, ! [X0, X5, X1] : (~ in(ordered_pair(X5, singleton(X5)), cartesian_product2(X0, X1)) | ~ in(X5, X0) | in(ordered_pair(X5, singleton(X5)), sK18(X0, X1)) | sP1(X0)), inference(equality_resolution, [], [f100])).
fof(f100, plain, ! [X4, X0, X5, X1] : (in(ordered_pair(X5, singleton(X5)), sK18(X0, X1)) | ~ in(X5, X0) | ~ (ordered_pair(X5, singleton(X5)) = X4) | ~ in(X4, cartesian_product2(X0, X1)) | sP1(X0)), inference(equality_resolution, [], [f99])).
fof(f99, plain, ! [X4, X0, X5, X3, X1] : (in(X3, sK18(X0, X1)) | ~ in(X5, X0) | ~ (ordered_pair(X5, singleton(X5)) = X3) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)) | sP1(X0)), inference(equality_resolution, [], [f97])).
fof(f97, plain, ! [X6, X4, X0, X5, X3, X1] : (in(X3, sK18(X0, X1)) | ~ (singleton(X5) = X6) | ~ in(X5, X0) | ~ (ordered_pair(X5, X6) = X3) | ~ (X3 = X4) | ~ in(X4, cartesian_product2(X0, X1)) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f1945, plain, (spl22_21 | spl22_27 | ~ spl22_30), inference(avatar_split_clause, [], [f1944, f1918, f1715, f1617])).
fof(f1944, plain, ((sK4(sK18(sK2, sK3)) = sK19(sK2, sK3, sK4(sK18(sK2, sK3)))) | (spl22_27 | ~ spl22_30)), inference(subsumption_resolution, [], [f1937, f1716])).
fof(f1937, plain, ((sK4(sK18(sK2, sK3)) = sK19(sK2, sK3, sK4(sK18(sK2, sK3)))) | sP1(sK2) | ~ spl22_30), inference(resolution, [], [f1919, f93])).
fof(f93, plain, ! [X0, X3, X1] : (~ in(X3, sK18(X0, X1)) | (sK19(X0, X1, X3) = X3) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f1943, plain, (spl22_19 | spl22_27 | ~ spl22_30), inference(avatar_split_clause, [], [f1942, f1918, f1715, f1609])).
fof(f1942, plain, ((sK21(sK2, sK4(sK18(sK2, sK3))) = singleton(sK20(sK2, sK4(sK18(sK2, sK3))))) | (spl22_27 | ~ spl22_30)), inference(subsumption_resolution, [], [f1936, f1716])).
fof(f1936, plain, ((sK21(sK2, sK4(sK18(sK2, sK3))) = singleton(sK20(sK2, sK4(sK18(sK2, sK3))))) | sP1(sK2) | ~ spl22_30), inference(resolution, [], [f1919, f96])).
fof(f96, plain, ! [X0, X3, X1] : (~ in(X3, sK18(X0, X1)) | (sK21(X0, X3) = singleton(sK20(X0, X3))) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f1941, plain, (spl22_28 | spl22_27 | ~ spl22_30), inference(avatar_split_clause, [], [f1940, f1918, f1715, f1721])).
fof(f1940, plain, ((sK4(sK18(sK2, sK3)) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3))))) | (spl22_27 | ~ spl22_30)), inference(subsumption_resolution, [], [f1935, f1716])).
fof(f1935, plain, ((sK4(sK18(sK2, sK3)) = ordered_pair(sK20(sK2, sK4(sK18(sK2, sK3))), sK21(sK2, sK4(sK18(sK2, sK3))))) | sP1(sK2) | ~ spl22_30), inference(resolution, [], [f1919, f94])).
fof(f94, plain, ! [X0, X3, X1] : (~ in(X3, sK18(X0, X1)) | (ordered_pair(sK20(X0, X3), sK21(X0, X3)) = X3) | sP1(X0)), inference(cnf_transformation, [], [f59])).
fof(f1932, plain, (spl22_16 | spl22_30), inference(avatar_split_clause, [], [f1928, f1918, f1516])).
fof(f1928, plain, ((ordered_pair(sK5(sK18(sK2, sK3)), sK6(sK18(sK2, sK3))) = sK4(sK18(sK2, sK3))) | spl22_30), inference(resolution, [], [f1920, f61])).
fof(f61, plain, ! [X2] : (in(sK4(X2), X2) | (sK4(X2) = ordered_pair(sK5(X2), sK6(X2)))), inference(cnf_transformation, [], [f36])).
fof(f1925, plain, (~ spl22_30 | ~ spl22_31 | ~ spl22_20), inference(avatar_split_clause, [], [f1916, f1613, f1922, f1918])).
fof(f1916, plain, (~ in(sK4(sK18(sK2, sK3)), cartesian_product2(sK2, sK3)) | ~ in(sK4(sK18(sK2, sK3)), sK18(sK2, sK3)) | ~ spl22_20), inference(equality_resolution, [], [f1614])).
fof(f1614, plain, (! [X0] : (~ (sK4(X0) = sK4(sK18(sK2, sK3))) | ~ in(sK4(X0), cartesian_product2(sK2, sK3)) | ~ in(sK4(X0), X0)) | ~ spl22_20), inference(avatar_component_clause, [], [f1613])).
fof(f1743, plain, (spl22_10 | ~ spl22_27), inference(avatar_contradiction_clause, [], [f1742])).
fof(f1742, plain, ($false | (spl22_10 | ~ spl22_27)), inference(subsumption_resolution, [], [f1736, f772])).
fof(f772, plain, (~ (sK11(sK2) = sK12(sK2)) | spl22_10), inference(avatar_component_clause, [], [f771])).
fof(f1736, plain, ((sK11(sK2) = sK12(sK2)) | ~ spl22_27), inference(resolution, [], [f1717, f82])).
fof(f82, plain, ! [X0] : (~ sP1(X0) | (sK11(X0) = sK12(X0))), inference(cnf_transformation, [], [f49])).
fof(f1717, plain, (sP1(sK2) | ~ spl22_27), inference(avatar_component_clause, [], [f1715])).
fof(f1741, plain, (spl22_1 | ~ spl22_27), inference(avatar_contradiction_clause, [], [f1740])).
fof(f1740, plain, ($false | (spl22_1 | ~ spl22_27)), inference(subsumption_resolution, [], [f1735, f698])).
fof(f698, plain, (~ (sK11(sK2) = sK13(sK2)) | spl22_1), inference(avatar_component_clause, [], [f697])).
fof(f1735, plain, ((sK11(sK2) = sK13(sK2)) | ~ spl22_27), inference(resolution, [], [f1717, f84])).
fof(f84, plain, ! [X0] : (~ sP1(X0) | (sK11(X0) = sK13(X0))), inference(cnf_transformation, [], [f49])).