fof(f465, plain, $false, inference(avatar_sat_refutation, [], [f379, f381, f388, f426, f434, f442, f464])).
fof(f464, plain, (~ spl26_7 | ~ spl26_10 | ~ spl26_11 | ~ spl26_12), inference(avatar_contradiction_clause, [], [f463])).
fof(f463, plain, ($false | (~ spl26_7 | ~ spl26_10 | ~ spl26_11 | ~ spl26_12)), inference(subsumption_resolution, [], [f462, f400])).
fof(f400, plain, (in(sK4(sK24(sK2, sK3)), sK24(sK2, sK3)) | ~ spl26_7), inference(factoring, [], [f378])).
fof(f378, plain, (! [X0] : (in(sK4(X0), sK24(sK2, sK3)) | in(sK4(X0), X0)) | ~ spl26_7), inference(avatar_component_clause, [], [f377])).
fof(f377, plain, (spl26_7 <=> ! [X0] : (in(sK4(X0), sK24(sK2, sK3)) | in(sK4(X0), X0))), introduced(avatar_definition, [new_symbols(naming, [spl26_7])])).
fof(f462, plain, (~ in(sK4(sK24(sK2, sK3)), sK24(sK2, sK3)) | (~ spl26_10 | ~ spl26_11 | ~ spl26_12)), inference(subsumption_resolution, [], [f460, f441])).
fof(f441, plain, (in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | ~ spl26_12), inference(avatar_component_clause, [], [f439])).
fof(f439, plain, (spl26_12 <=> in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2)), introduced(avatar_definition, [new_symbols(naming, [spl26_12])])).
fof(f460, plain, (~ in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | ~ in(sK4(sK24(sK2, sK3)), sK24(sK2, sK3)) | (~ spl26_10 | ~ spl26_11)), inference(resolution, [], [f443, f125])).
fof(f125, plain, ! [X3] : (~ in(sK4(X3), powerset(relation_dom(sK3))) | ~ in(relation_image(sK3, sK4(X3)), sK2) | ~ in(sK4(X3), X3)), inference(cnf_transformation, [], [f78])).
fof(f78, plain, (! [X3] : ((~ in(relation_image(sK3, sK4(X3)), sK2) | ~ in(sK4(X3), powerset(relation_dom(sK3))) | ~ in(sK4(X3), X3)) & ((in(relation_image(sK3, sK4(X3)), sK2) & in(sK4(X3), powerset(relation_dom(sK3)))) | in(sK4(X3), X3))) & function(sK3) & relation(sK3) & element(sK2, powerset(powerset(sK1)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3, sK4])], [f75, f77, f76])).
fof(f76, plain, (? [X0, X1, X2] : (! [X3] : ? [X4] : ((~ in(relation_image(X2, X4), X1) | ~ in(X4, powerset(relation_dom(X2))) | ~ in(X4, X3)) & ((in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2)))) | in(X4, X3))) & function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))) => (! [X3] : ? [X4] : ((~ in(relation_image(sK3, X4), sK2) | ~ in(X4, powerset(relation_dom(sK3))) | ~ in(X4, X3)) & ((in(relation_image(sK3, X4), sK2) & in(X4, powerset(relation_dom(sK3)))) | in(X4, X3))) & function(sK3) & relation(sK3) & element(sK2, powerset(powerset(sK1))))), introduced(choice_axiom, [])).
fof(f77, plain, ! [X3] : (? [X4] : ((~ in(relation_image(sK3, X4), sK2) | ~ in(X4, powerset(relation_dom(sK3))) | ~ in(X4, X3)) & ((in(relation_image(sK3, X4), sK2) & in(X4, powerset(relation_dom(sK3)))) | in(X4, X3))) => ((~ in(relation_image(sK3, sK4(X3)), sK2) | ~ in(sK4(X3), powerset(relation_dom(sK3))) | ~ in(sK4(X3), X3)) & ((in(relation_image(sK3, sK4(X3)), sK2) & in(sK4(X3), powerset(relation_dom(sK3)))) | in(sK4(X3), X3)))), introduced(choice_axiom, [])).
fof(f75, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ((~ in(relation_image(X2, X4), X1) | ~ in(X4, powerset(relation_dom(X2))) | ~ in(X4, X3)) & ((in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2)))) | in(X4, X3))) & function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))), inference(flattening, [], [f74])).
fof(f74, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : (((~ in(relation_image(X2, X4), X1) | ~ in(X4, powerset(relation_dom(X2)))) | ~ in(X4, X3)) & ((in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2)))) | in(X4, X3))) & function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))), inference(nnf_transformation, [], [f48])).
fof(f48, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2))))) & function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))), inference(flattening, [], [f47])).
fof(f47, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2))))) & (function(X2) & relation(X2) & element(X1, powerset(powerset(X0))))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1, X2] : ((function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))) => ? [X3] : ! [X4] : (in(X4, X3) <=> (in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2)))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1, X2] : ((function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))) => ? [X3] : ! [X4] : (in(X4, X3) <=> (in(relation_image(X2, X4), X1) & in(X4, powerset(relation_dom(X2)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU297+1.p', s1_xboole_0__e6_27__finset_1)).
fof(f443, plain, (in(sK4(sK24(sK2, sK3)), powerset(relation_dom(sK3))) | (~ spl26_10 | ~ spl26_11)), inference(forward_demodulation, [], [f425, f433])).
fof(f433, plain, ((sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3)))) | ~ spl26_11), inference(avatar_component_clause, [], [f431])).
fof(f431, plain, (spl26_11 <=> (sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3))))), introduced(avatar_definition, [new_symbols(naming, [spl26_11])])).
fof(f425, plain, (in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3))) | ~ spl26_10), inference(avatar_component_clause, [], [f423])).
fof(f423, plain, (spl26_10 <=> in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl26_10])])).
fof(f442, plain, (spl26_5 | spl26_12 | spl26_6 | ~ spl26_7), inference(avatar_split_clause, [], [f437, f377, f373, f439, f370])).
fof(f370, plain, (spl26_5 <=> ! [X1] : ~ element(sK2, powerset(powerset(X1)))), introduced(avatar_definition, [new_symbols(naming, [spl26_5])])).
fof(f373, plain, (spl26_6 <=> sP0(sK2, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl26_6])])).
fof(f437, plain, (! [X2] : (in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | ~ element(sK2, powerset(powerset(X2)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f436, f121])).
fof(f121, plain, relation(sK3), inference(cnf_transformation, [], [f78])).
fof(f436, plain, (! [X2] : (in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X2)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f435, f122])).
fof(f122, plain, function(sK3), inference(cnf_transformation, [], [f78])).
fof(f435, plain, (! [X2] : (in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X2)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f417, f374])).
fof(f374, plain, (~ sP0(sK2, sK3) | spl26_6), inference(avatar_component_clause, [], [f373])).
fof(f417, plain, (! [X2] : (in(relation_image(sK3, sK4(sK24(sK2, sK3))), sK2) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X2)))) | ~ spl26_7), inference(resolution, [], [f400, f204])).
fof(f204, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X1, X2)) | in(relation_image(X2, X4), X1) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, ! [X0, X1, X2] : (! [X4] : ((in(X4, sK24(X1, X2)) | ! [X5] : (~ in(relation_image(X2, X4), X1) | ~ (X4 = X5) | ~ in(X5, powerset(relation_dom(X2))))) & ((in(relation_image(X2, X4), X1) & (sK25(X1, X2, X4) = X4) & in(sK25(X1, X2, X4), powerset(relation_dom(X2)))) | ~ in(X4, sK24(X1, X2)))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25])], [f116, f118, f117])).
fof(f117, plain, ! [X2, X1] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (~ in(relation_image(X2, X4), X1) | ~ (X4 = X5) | ~ in(X5, powerset(relation_dom(X2))))) & (? [X6] : (in(relation_image(X2, X4), X1) & (X4 = X6) & in(X6, powerset(relation_dom(X2)))) | ~ in(X4, X3))) => ! [X4] : ((in(X4, sK24(X1, X2)) | ! [X5] : (~ in(relation_image(X2, X4), X1) | ~ (X4 = X5) | ~ in(X5, powerset(relation_dom(X2))))) & (? [X6] : (in(relation_image(X2, X4), X1) & (X4 = X6) & in(X6, powerset(relation_dom(X2)))) | ~ in(X4, sK24(X1, X2))))), introduced(choice_axiom, [])).
fof(f118, plain, ! [X4, X2, X1] : (? [X6] : (in(relation_image(X2, X4), X1) & (X4 = X6) & in(X6, powerset(relation_dom(X2)))) => (in(relation_image(X2, X4), X1) & (sK25(X1, X2, X4) = X4) & in(sK25(X1, X2, X4), powerset(relation_dom(X2))))), introduced(choice_axiom, [])).
fof(f116, plain, ! [X0, X1, X2] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (~ in(relation_image(X2, X4), X1) | ~ (X4 = X5) | ~ in(X5, powerset(relation_dom(X2))))) & (? [X6] : (in(relation_image(X2, X4), X1) & (X4 = X6) & in(X6, powerset(relation_dom(X2)))) | ~ in(X4, X3))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(rectify, [], [f115])).
fof(f115, plain, ! [X0, X1, X2] : (? [X6] : ! [X7] : ((in(X7, X6) | ! [X8] : (~ in(relation_image(X2, X7), X1) | ~ (X7 = X8) | ~ in(X8, powerset(relation_dom(X2))))) & (? [X8] : (in(relation_image(X2, X7), X1) & (X7 = X8) & in(X8, powerset(relation_dom(X2)))) | ~ in(X7, X6))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(nnf_transformation, [], [f73])).
fof(f73, plain, ! [X0, X1, X2] : (? [X6] : ! [X7] : (in(X7, X6) <=> ? [X8] : (in(relation_image(X2, X7), X1) & (X7 = X8) & in(X8, powerset(relation_dom(X2))))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(definition_folding, [], [f71, e72])).
fof(f72, plain, ! [X1, X2] : (? [X3, X4, X5] : (~ (X4 = X5) & in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)) | ~ sP0(X1, X2)), inference(usedef, [], [e72])).
fof(e72, plain, ! [X1, X2] : (sP0(X1, X2) <=> ? [X3, X4, X5] : (~ (X4 = X5) & in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f71, plain, ! [X0, X1, X2] : (? [X6] : ! [X7] : (in(X7, X6) <=> ? [X8] : (in(relation_image(X2, X7), X1) & (X7 = X8) & in(X8, powerset(relation_dom(X2))))) | ? [X3, X4, X5] : (~ (X4 = X5) & in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(flattening, [], [f70])).
fof(f70, plain, ! [X0, X1, X2] : ((? [X6] : ! [X7] : (in(X7, X6) <=> ? [X8] : (in(relation_image(X2, X7), X1) & (X7 = X8) & in(X8, powerset(relation_dom(X2))))) | ? [X3, X4, X5] : (~ (X4 = X5) & (in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)))) | (~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0))))), inference(ennf_transformation, [], [f39])).
fof(f39, plain, ! [X0, X1, X2] : ((function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))) => (! [X3, X4, X5] : ((in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)) => (X4 = X5)) => ? [X6] : ! [X7] : (in(X7, X6) <=> ? [X8] : (in(relation_image(X2, X7), X1) & (X7 = X8) & in(X8, powerset(relation_dom(X2))))))), inference(rectify, [], [f38])).
fof(f38, plain, ! [X0, X1, X2] : ((function(X2) & relation(X2) & element(X1, powerset(powerset(X0)))) => (! [X3, X4, X5] : ((in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)) => (X4 = X5)) => ? [X3] : ! [X4] : (in(X4, X3) <=> ? [X5] : (in(relation_image(X2, X4), X1) & (X4 = X5) & in(X5, powerset(relation_dom(X2))))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU297+1.p', s1_tarski__e6_27__finset_1__1)).
fof(f434, plain, (spl26_5 | spl26_11 | spl26_6 | ~ spl26_7), inference(avatar_split_clause, [], [f429, f377, f373, f431, f370])).
fof(f429, plain, (! [X1] : ((sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3)))) | ~ element(sK2, powerset(powerset(X1)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f428, f121])).
fof(f428, plain, (! [X1] : ((sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3)))) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X1)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f427, f122])).
fof(f427, plain, (! [X1] : ((sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X1)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f416, f374])).
fof(f416, plain, (! [X1] : ((sK4(sK24(sK2, sK3)) = sK25(sK2, sK3, sK4(sK24(sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X1)))) | ~ spl26_7), inference(resolution, [], [f400, f203])).
fof(f203, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X1, X2)) | (sK25(X1, X2, X4) = X4) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(cnf_transformation, [], [f119])).
fof(f426, plain, (spl26_5 | spl26_10 | spl26_6 | ~ spl26_7), inference(avatar_split_clause, [], [f421, f377, f373, f423, f370])).
fof(f421, plain, (! [X0] : (in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3))) | ~ element(sK2, powerset(powerset(X0)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f420, f121])).
fof(f420, plain, (! [X0] : (in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3))) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X0)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f419, f122])).
fof(f419, plain, (! [X0] : (in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3))) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X0)))) | (spl26_6 | ~ spl26_7)), inference(subsumption_resolution, [], [f415, f374])).
fof(f415, plain, (! [X0] : (in(sK25(sK2, sK3, sK4(sK24(sK2, sK3))), powerset(relation_dom(sK3))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X0)))) | ~ spl26_7), inference(resolution, [], [f400, f202])).
fof(f202, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X1, X2)) | in(sK25(X1, X2, X4), powerset(relation_dom(X2))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(cnf_transformation, [], [f119])).
fof(f388, plain, ~ spl26_6, inference(avatar_contradiction_clause, [], [f387])).
fof(f387, plain, ($false | ~ spl26_6), inference(subsumption_resolution, [], [f386, f375])).
fof(f375, plain, (sP0(sK2, sK3) | ~ spl26_6), inference(avatar_component_clause, [], [f373])).
fof(f386, plain, (~ sP0(sK2, sK3) | ~ spl26_6), inference(subsumption_resolution, [], [f384, f383])).
fof(f383, plain, ((sK21(sK2, sK3) = sK22(sK2, sK3)) | ~ spl26_6), inference(resolution, [], [f375, f197])).
fof(f197, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK21(X0, X1) = sK22(X0, X1))), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0, X1] : ((~ (sK22(X0, X1) = sK23(X0, X1)) & in(relation_image(X1, sK23(X0, X1)), X0) & (sK21(X0, X1) = sK23(X0, X1)) & in(relation_image(X1, sK22(X0, X1)), X0) & (sK21(X0, X1) = sK22(X0, X1))) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22, sK23])], [f112, f113])).
fof(f113, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & in(relation_image(X1, X4), X0) & (X2 = X4) & in(relation_image(X1, X3), X0) & (X2 = X3)) => (~ (sK22(X0, X1) = sK23(X0, X1)) & in(relation_image(X1, sK23(X0, X1)), X0) & (sK21(X0, X1) = sK23(X0, X1)) & in(relation_image(X1, sK22(X0, X1)), X0) & (sK21(X0, X1) = sK22(X0, X1)))), introduced(choice_axiom, [])).
fof(f112, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & in(relation_image(X1, X4), X0) & (X2 = X4) & in(relation_image(X1, X3), X0) & (X2 = X3)) | ~ sP0(X0, X1)), inference(rectify, [], [f111])).
fof(f111, plain, ! [X1, X2] : (? [X3, X4, X5] : (~ (X4 = X5) & in(relation_image(X2, X5), X1) & (X3 = X5) & in(relation_image(X2, X4), X1) & (X3 = X4)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f72])).
fof(f384, plain, (~ (sK21(sK2, sK3) = sK22(sK2, sK3)) | ~ sP0(sK2, sK3) | ~ spl26_6), inference(superposition, [], [f201, f382])).
fof(f382, plain, ((sK21(sK2, sK3) = sK23(sK2, sK3)) | ~ spl26_6), inference(resolution, [], [f375, f199])).
fof(f199, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK21(X0, X1) = sK23(X0, X1))), inference(cnf_transformation, [], [f114])).
fof(f201, plain, ! [X0, X1] : (~ (sK22(X0, X1) = sK23(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f114])).
fof(f381, plain, ~ spl26_5, inference(avatar_contradiction_clause, [], [f380])).
fof(f380, plain, ($false | ~ spl26_5), inference(resolution, [], [f371, f120])).
fof(f120, plain, element(sK2, powerset(powerset(sK1))), inference(cnf_transformation, [], [f78])).
fof(f371, plain, (! [X1] : ~ element(sK2, powerset(powerset(X1))) | ~ spl26_5), inference(avatar_component_clause, [], [f370])).
fof(f379, plain, (spl26_5 | spl26_6 | spl26_7), inference(avatar_split_clause, [], [f368, f377, f373, f370])).
fof(f368, plain, ! [X0, X1] : (in(sK4(X0), sK24(sK2, sK3)) | sP0(sK2, sK3) | ~ element(sK2, powerset(powerset(X1))) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f367, f123])).
fof(f123, plain, ! [X3] : (in(sK4(X3), powerset(relation_dom(sK3))) | in(sK4(X3), X3)), inference(cnf_transformation, [], [f78])).
fof(f367, plain, ! [X0, X1] : (in(sK4(X0), sK24(sK2, sK3)) | ~ in(sK4(X0), powerset(relation_dom(sK3))) | sP0(sK2, sK3) | ~ element(sK2, powerset(powerset(X1))) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f366, f121])).
fof(f366, plain, ! [X0, X1] : (in(sK4(X0), sK24(sK2, sK3)) | ~ in(sK4(X0), powerset(relation_dom(sK3))) | sP0(sK2, sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X1))) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f363, f122])).
fof(f363, plain, ! [X0, X1] : (in(sK4(X0), sK24(sK2, sK3)) | ~ in(sK4(X0), powerset(relation_dom(sK3))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ element(sK2, powerset(powerset(X1))) | in(sK4(X0), X0)), inference(resolution, [], [f206, f124])).
fof(f124, plain, ! [X3] : (in(relation_image(sK3, sK4(X3)), sK2) | in(sK4(X3), X3)), inference(cnf_transformation, [], [f78])).
fof(f206, plain, ! [X2, X0, X5, X1] : (~ in(relation_image(X2, X5), X1) | in(X5, sK24(X1, X2)) | ~ in(X5, powerset(relation_dom(X2))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(equality_resolution, [], [f205])).
fof(f205, plain, ! [X4, X2, X0, X5, X1] : (in(X4, sK24(X1, X2)) | ~ in(relation_image(X2, X4), X1) | ~ (X4 = X5) | ~ in(X5, powerset(relation_dom(X2))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ element(X1, powerset(powerset(X0)))), inference(cnf_transformation, [], [f119])).