fof(f2510, plain, $false, inference(avatar_sat_refutation, [], [f190, f298, f307, f2271, f2332, f2378, f2414, f2460, f2508])).
fof(f2508, plain, (~ spl26_1 | spl26_5), inference(avatar_contradiction_clause, [], [f2507])).
fof(f2507, plain, ($false | (~ spl26_1 | spl26_5)), inference(subsumption_resolution, [], [f2503, f244])).
fof(f244, plain, (~ (sK15(sK2, sK3) = sK17(sK2, sK3)) | spl26_5), inference(avatar_component_clause, [], [f243])).
fof(f243, plain, (spl26_5 <=> (sK15(sK2, sK3) = sK17(sK2, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl26_5])])).
fof(f2503, plain, ((sK15(sK2, sK3) = sK17(sK2, sK3)) | ~ spl26_1), inference(resolution, [], [f186, f118])).
fof(f118, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK15(X0, X1) = sK17(X0, X1))), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ! [X0, X1] : ((~ (sK16(X0, X1) = sK17(X0, X1)) & (in(ordered_pair(apply(X1, sK18(X0, X1)), apply(X1, sK19(X0, X1))), X0) & (sK17(X0, X1) = ordered_pair(sK18(X0, X1), sK19(X0, X1)))) & (sK15(X0, X1) = sK17(X0, X1)) & (in(ordered_pair(apply(X1, sK20(X0, X1)), apply(X1, sK21(X0, X1))), X0) & (sK16(X0, X1) = ordered_pair(sK20(X0, X1), sK21(X0, X1)))) & (sK15(X0, X1) = sK16(X0, X1))) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK15, sK16, sK17, sK18, sK19, sK20, sK21])], [f64, f67, f66, f65])).
fof(f65, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : (in(ordered_pair(apply(X1, X5), apply(X1, X6)), X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & ? [X7, X8] : (in(ordered_pair(apply(X1, X7), apply(X1, X8)), X0) & (ordered_pair(X7, X8) = X3)) & (X2 = X3)) => (~ (sK16(X0, X1) = sK17(X0, X1)) & ? [X6, X5] : (in(ordered_pair(apply(X1, X5), apply(X1, X6)), X0) & (ordered_pair(X5, X6) = sK17(X0, X1))) & (sK15(X0, X1) = sK17(X0, X1)) & ? [X8, X7] : (in(ordered_pair(apply(X1, X7), apply(X1, X8)), X0) & (ordered_pair(X7, X8) = sK16(X0, X1))) & (sK15(X0, X1) = sK16(X0, X1)))), introduced(choice_axiom, [])).
fof(f66, plain, ! [X1, X0] : (? [X6, X5] : (in(ordered_pair(apply(X1, X5), apply(X1, X6)), X0) & (ordered_pair(X5, X6) = sK17(X0, X1))) => (in(ordered_pair(apply(X1, sK18(X0, X1)), apply(X1, sK19(X0, X1))), X0) & (sK17(X0, X1) = ordered_pair(sK18(X0, X1), sK19(X0, X1))))), introduced(choice_axiom, [])).
fof(f67, plain, ! [X1, X0] : (? [X8, X7] : (in(ordered_pair(apply(X1, X7), apply(X1, X8)), X0) & (ordered_pair(X7, X8) = sK16(X0, X1))) => (in(ordered_pair(apply(X1, sK20(X0, X1)), apply(X1, sK21(X0, X1))), X0) & (sK16(X0, X1) = ordered_pair(sK20(X0, X1), sK21(X0, X1))))), introduced(choice_axiom, [])).
fof(f64, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5, X6] : (in(ordered_pair(apply(X1, X5), apply(X1, X6)), X0) & (ordered_pair(X5, X6) = X4)) & (X2 = X4) & ? [X7, X8] : (in(ordered_pair(apply(X1, X7), apply(X1, X8)), X0) & (ordered_pair(X7, X8) = X3)) & (X2 = X3)) | ~ sP0(X0, X1)), inference(rectify, [], [f63])).
fof(f63, plain, ! [X1, X2] : (? [X3, X4, X5] : (~ (X4 = X5) & ? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4)) | ~ sP0(X1, X2)), inference(nnf_transformation, [], [f38])).
fof(f38, plain, ! [X1, X2] : (? [X3, X4, X5] : (~ (X4 = X5) & ? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4)) | ~ sP0(X1, X2)), inference(usedef, [], [e38])).
fof(e38, plain, ! [X1, X2] : (sP0(X1, X2) <=> ? [X3, X4, X5] : (~ (X4 = X5) & ? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f186, plain, (sP0(sK2, sK3) | ~ spl26_1), inference(avatar_component_clause, [], [f184])).
fof(f184, plain, (spl26_1 <=> sP0(sK2, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl26_1])])).
fof(f2460, plain, (spl26_131 | spl26_1 | ~ spl26_128 | ~ spl26_132), inference(avatar_split_clause, [], [f2459, f2384, f2265, f184, f2277])).
fof(f2277, plain, (spl26_131 <=> ! [X1] : (~ (sK4(X1) = sK4(sK22(sK1, sK2, sK3))) | ~ in(sK4(X1), X1) | ~ in(sK4(X1), cartesian_product2(sK1, sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl26_131])])).
fof(f2265, plain, (spl26_128 <=> in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl26_128])])).
fof(f2384, plain, (spl26_132 <=> (sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))))), introduced(avatar_definition, [new_symbols(naming, [spl26_132])])).
fof(f2459, plain, (! [X2] : (~ (sK4(X2) = sK4(sK22(sK1, sK2, sK3))) | ~ in(sK4(X2), cartesian_product2(sK1, sK1)) | ~ in(sK4(X2), X2)) | (spl26_1 | ~ spl26_128 | ~ spl26_132)), inference(subsumption_resolution, [], [f2457, f2409])).
fof(f2409, plain, (in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2408, f75])).
fof(f75, plain, relation(sK2), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (! [X3] : ((! [X5, X6] : (~ in(ordered_pair(apply(sK3, X5), apply(sK3, X6)), sK2) | ~ (ordered_pair(X5, X6) = sK4(X3))) | ~ in(sK4(X3), cartesian_product2(sK1, sK1)) | ~ in(sK4(X3), X3)) & (((in(ordered_pair(apply(sK3, sK5(X3)), apply(sK3, sK6(X3))), sK2) & (sK4(X3) = ordered_pair(sK5(X3), sK6(X3)))) & in(sK4(X3), cartesian_product2(sK1, sK1))) | in(sK4(X3), X3))) & function(sK3) & relation(sK3) & relation(sK2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3, sK4, sK5, sK6])], [f42, f45, f44, f43])).
fof(f43, plain, (? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5, X6] : (~ in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(X0, X0)) | ~ in(X4, X3)) & ((? [X7, X8] : (in(ordered_pair(apply(X2, X7), apply(X2, X8)), X1) & (ordered_pair(X7, X8) = X4)) & in(X4, cartesian_product2(X0, X0))) | in(X4, X3))) & function(X2) & relation(X2) & relation(X1)) => (! [X3] : ? [X4] : ((! [X6, X5] : (~ in(ordered_pair(apply(sK3, X5), apply(sK3, X6)), sK2) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(sK1, sK1)) | ~ in(X4, X3)) & ((? [X8, X7] : (in(ordered_pair(apply(sK3, X7), apply(sK3, X8)), sK2) & (ordered_pair(X7, X8) = X4)) & in(X4, cartesian_product2(sK1, sK1))) | in(X4, X3))) & function(sK3) & relation(sK3) & relation(sK2))), introduced(choice_axiom, [])).
fof(f44, plain, ! [X3] : (? [X4] : ((! [X6, X5] : (~ in(ordered_pair(apply(sK3, X5), apply(sK3, X6)), sK2) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(sK1, sK1)) | ~ in(X4, X3)) & ((? [X8, X7] : (in(ordered_pair(apply(sK3, X7), apply(sK3, X8)), sK2) & (ordered_pair(X7, X8) = X4)) & in(X4, cartesian_product2(sK1, sK1))) | in(X4, X3))) => ((! [X6, X5] : (~ in(ordered_pair(apply(sK3, X5), apply(sK3, X6)), sK2) | ~ (ordered_pair(X5, X6) = sK4(X3))) | ~ in(sK4(X3), cartesian_product2(sK1, sK1)) | ~ in(sK4(X3), X3)) & ((? [X8, X7] : (in(ordered_pair(apply(sK3, X7), apply(sK3, X8)), sK2) & (ordered_pair(X7, X8) = sK4(X3))) & in(sK4(X3), cartesian_product2(sK1, sK1))) | in(sK4(X3), X3)))), introduced(choice_axiom, [])).
fof(f45, plain, ! [X3] : (? [X8, X7] : (in(ordered_pair(apply(sK3, X7), apply(sK3, X8)), sK2) & (ordered_pair(X7, X8) = sK4(X3))) => (in(ordered_pair(apply(sK3, sK5(X3)), apply(sK3, sK6(X3))), sK2) & (sK4(X3) = ordered_pair(sK5(X3), sK6(X3))))), introduced(choice_axiom, [])).
fof(f42, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5, X6] : (~ in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(X0, X0)) | ~ in(X4, X3)) & ((? [X7, X8] : (in(ordered_pair(apply(X2, X7), apply(X2, X8)), X1) & (ordered_pair(X7, X8) = X4)) & in(X4, cartesian_product2(X0, X0))) | in(X4, X3))) & function(X2) & relation(X2) & relation(X1)), inference(rectify, [], [f41])).
fof(f41, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5, X6] : (~ in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(X0, X0)) | ~ in(X4, X3)) & ((? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0))) | in(X4, X3))) & function(X2) & relation(X2) & relation(X1)), inference(flattening, [], [f40])).
fof(f40, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : (((! [X5, X6] : (~ in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) | ~ (ordered_pair(X5, X6) = X4)) | ~ in(X4, cartesian_product2(X0, X0))) | ~ in(X4, X3)) & ((? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0))) | in(X4, X3))) & function(X2) & relation(X2) & relation(X1)), inference(nnf_transformation, [], [f27])).
fof(f27, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0)))) & function(X2) & relation(X2) & relation(X1)), inference(flattening, [], [f26])).
fof(f26, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0)))) & (function(X2) & relation(X2) & relation(X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1, X2] : ((function(X2) & relation(X2) & relation(X1)) => ? [X3] : ! [X4] : (in(X4, X3) <=> (? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1, X2] : ((function(X2) & relation(X2) & relation(X1)) => ? [X3] : ! [X4] : (in(X4, X3) <=> (? [X5, X6] : (in(ordered_pair(apply(X2, X5), apply(X2, X6)), X1) & (ordered_pair(X5, X6) = X4)) & in(X4, cartesian_product2(X0, X0))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU277+1.p', s1_xboole_0__e6_21__wellord2__1)).
fof(f2408, plain, (in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2407, f76])).
fof(f76, plain, relation(sK3), inference(cnf_transformation, [], [f46])).
fof(f2407, plain, (in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2406, f77])).
fof(f77, plain, function(sK3), inference(cnf_transformation, [], [f46])).
fof(f2406, plain, (in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2402, f185])).
fof(f185, plain, (~ sP0(sK2, sK3) | spl26_1), inference(avatar_component_clause, [], [f184])).
fof(f2402, plain, (in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | ~ spl26_128), inference(resolution, [], [f2267, f125])).
fof(f125, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK22(X0, X1, X2)) | in(ordered_pair(apply(X2, sK24(X1, X2, X4)), apply(X2, sK25(X1, X2, X4))), X1) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ! [X0, X1, X2] : (! [X4] : ((in(X4, sK22(X0, X1, X2)) | ! [X5] : (! [X6, X7] : (~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X4)) | ~ (X4 = X5) | ~ in(X5, cartesian_product2(X0, X0)))) & (((in(ordered_pair(apply(X2, sK24(X1, X2, X4)), apply(X2, sK25(X1, X2, X4))), X1) & (ordered_pair(sK24(X1, X2, X4), sK25(X1, X2, X4)) = X4)) & (sK23(X0, X1, X2, X4) = X4) & in(sK23(X0, X1, X2, X4), cartesian_product2(X0, X0))) | ~ in(X4, sK22(X0, X1, X2)))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK22, sK23, sK24, sK25])], [f70, f73, f72, f71])).
fof(f71, plain, ! [X2, X1, X0] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6, X7] : (~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X4)) | ~ (X4 = X5) | ~ in(X5, cartesian_product2(X0, X0)))) & (? [X8] : (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) & (X4 = X8) & in(X8, cartesian_product2(X0, X0))) | ~ in(X4, X3))) => ! [X4] : ((in(X4, sK22(X0, X1, X2)) | ! [X5] : (! [X6, X7] : (~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X4)) | ~ (X4 = X5) | ~ in(X5, cartesian_product2(X0, X0)))) & (? [X8] : (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) & (X4 = X8) & in(X8, cartesian_product2(X0, X0))) | ~ in(X4, sK22(X0, X1, X2))))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X4, X2, X1, X0] : (? [X8] : (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) & (X4 = X8) & in(X8, cartesian_product2(X0, X0))) => (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) & (sK23(X0, X1, X2, X4) = X4) & in(sK23(X0, X1, X2, X4), cartesian_product2(X0, X0)))), introduced(choice_axiom, [])).
fof(f73, plain, ! [X4, X2, X1] : (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) => (in(ordered_pair(apply(X2, sK24(X1, X2, X4)), apply(X2, sK25(X1, X2, X4))), X1) & (ordered_pair(sK24(X1, X2, X4), sK25(X1, X2, X4)) = X4))), introduced(choice_axiom, [])).
fof(f70, plain, ! [X0, X1, X2] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6, X7] : (~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X4)) | ~ (X4 = X5) | ~ in(X5, cartesian_product2(X0, X0)))) & (? [X8] : (? [X9, X10] : (in(ordered_pair(apply(X2, X9), apply(X2, X10)), X1) & (ordered_pair(X9, X10) = X4)) & (X4 = X8) & in(X8, cartesian_product2(X0, X0))) | ~ in(X4, X3))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(rectify, [], [f69])).
fof(f69, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : ((in(X11, X10) | ! [X12] : (! [X13, X14] : (~ in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) | ~ (ordered_pair(X13, X14) = X11)) | ~ (X11 = X12) | ~ in(X12, cartesian_product2(X0, X0)))) & (? [X12] : (? [X13, X14] : (in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) & (ordered_pair(X13, X14) = X11)) & (X11 = X12) & in(X12, cartesian_product2(X0, X0))) | ~ in(X11, X10))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(nnf_transformation, [], [f39])).
fof(f39, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13, X14] : (in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) & (ordered_pair(X13, X14) = X11)) & (X11 = X12) & in(X12, cartesian_product2(X0, X0)))) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(definition_folding, [], [f37, e38])).
fof(f37, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13, X14] : (in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) & (ordered_pair(X13, X14) = X11)) & (X11 = X12) & in(X12, cartesian_product2(X0, X0)))) | ? [X3, X4, X5] : (~ (X4 = X5) & ? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4)) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(flattening, [], [f36])).
fof(f36, plain, ! [X0, X1, X2] : ((? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13, X14] : (in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) & (ordered_pair(X13, X14) = X11)) & (X11 = X12) & in(X12, cartesian_product2(X0, X0)))) | ? [X3, X4, X5] : (~ (X4 = X5) & (? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4)))) | (~ function(X2) | ~ relation(X2) | ~ relation(X1))), inference(ennf_transformation, [], [f22])).
fof(f22, plain, ! [X0, X1, X2] : ((function(X2) & relation(X2) & relation(X1)) => (! [X3, X4, X5] : ((? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X5)) & (X3 = X5) & ? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X4)) & (X3 = X4)) => (X4 = X5)) => ? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13, X14] : (in(ordered_pair(apply(X2, X13), apply(X2, X14)), X1) & (ordered_pair(X13, X14) = X11)) & (X11 = X12) & in(X12, cartesian_product2(X0, X0)))))), inference(rectify, [], [f21])).
fof(f21, plain, ! [X0, X1, X2] : ((function(X2) & relation(X2) & relation(X1)) => (! [X3, X4, X5] : ((? [X8, X9] : (in(ordered_pair(apply(X2, X8), apply(X2, X9)), X1) & (ordered_pair(X8, X9) = X5)) & (X3 = X5) & ? [X6, X7] : (in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) & (ordered_pair(X6, X7) = X4)) & (X3 = X4)) => (X4 = X5)) => ? [X3] : ! [X4] : (in(X4, X3) <=> ? [X5] : (? [X10, X11] : (in(ordered_pair(apply(X2, X10), apply(X2, X11)), X1) & (ordered_pair(X10, X11) = X4)) & (X4 = X5) & in(X5, cartesian_product2(X0, X0)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU277+1.p', s1_tarski__e6_21__wellord2__1)).
fof(f2267, plain, (in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ spl26_128), inference(avatar_component_clause, [], [f2265])).
fof(f2457, plain, (! [X2] : (~ (sK4(X2) = sK4(sK22(sK1, sK2, sK3))) | ~ in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))), sK2) | ~ in(sK4(X2), cartesian_product2(sK1, sK1)) | ~ in(sK4(X2), X2)) | ~ spl26_132), inference(superposition, [], [f81, f2386])).
fof(f2386, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | ~ spl26_132), inference(avatar_component_clause, [], [f2384])).
fof(f81, plain, ! [X6, X5, X3] : (~ (ordered_pair(X5, X6) = sK4(X3)) | ~ in(ordered_pair(apply(sK3, X5), apply(sK3, X6)), sK2) | ~ in(sK4(X3), cartesian_product2(sK1, sK1)) | ~ in(sK4(X3), X3)), inference(cnf_transformation, [], [f46])).
fof(f2414, plain, (spl26_132 | spl26_1 | ~ spl26_128), inference(avatar_split_clause, [], [f2413, f2265, f184, f2384])).
fof(f2413, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2412, f75])).
fof(f2412, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2411, f76])).
fof(f2411, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2410, f77])).
fof(f2410, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_128)), inference(subsumption_resolution, [], [f2403, f185])).
fof(f2403, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(sK1, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(sK1, sK2, sK3))))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | ~ spl26_128), inference(resolution, [], [f2267, f124])).
fof(f124, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK22(X0, X1, X2)) | (ordered_pair(sK24(X1, X2, X4), sK25(X1, X2, X4)) = X4) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(cnf_transformation, [], [f74])).
fof(f2378, plain, (spl26_128 | ~ spl26_129), inference(avatar_split_clause, [], [f2368, f2269, f2265])).
fof(f2269, plain, (spl26_129 <=> ! [X0] : (~ in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(X0, X0)) | in(sK4(sK22(sK1, sK2, sK3)), sK22(X0, sK2, sK3)))), introduced(avatar_definition, [new_symbols(naming, [spl26_129])])).
fof(f2368, plain, (in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ spl26_129), inference(duplicate_literal_removal, [], [f2367])).
fof(f2367, plain, (in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ spl26_129), inference(resolution, [], [f2270, f78])).
fof(f78, plain, ! [X3] : (in(sK4(X3), cartesian_product2(sK1, sK1)) | in(sK4(X3), X3)), inference(cnf_transformation, [], [f46])).
fof(f2270, plain, (! [X0] : (~ in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(X0, X0)) | in(sK4(sK22(sK1, sK2, sK3)), sK22(X0, sK2, sK3))) | ~ spl26_129), inference(avatar_component_clause, [], [f2269])).
fof(f2332, plain, (~ spl26_128 | spl26_1 | ~ spl26_2 | ~ spl26_6 | ~ spl26_131), inference(avatar_split_clause, [], [f2331, f2277, f247, f188, f184, f2265])).
fof(f188, plain, (spl26_2 <=> ! [X1, X0] : (in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3)) | in(sK4(X0), X0) | ~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)))), introduced(avatar_definition, [new_symbols(naming, [spl26_2])])).
fof(f247, plain, (spl26_6 <=> ! [X0] : ((sK4(sK22(X0, sK2, sK3)) = sK23(X0, sK2, sK3, sK4(sK22(X0, sK2, sK3)))) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))))), introduced(avatar_definition, [new_symbols(naming, [spl26_6])])).
fof(f2331, plain, (~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (spl26_1 | ~ spl26_2 | ~ spl26_6 | ~ spl26_131)), inference(subsumption_resolution, [], [f2330, f2228])).
fof(f2228, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2227, f78])).
fof(f2227, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2226, f75])).
fof(f2226, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ relation(sK2) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2225, f76])).
fof(f2225, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2224, f77])).
fof(f2224, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2218, f185])).
fof(f2218, plain, (in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(superposition, [], [f122, f2194])).
fof(f2194, plain, ((sK4(sK22(sK1, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(sK1, sK2, sK3)))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(trivial_inequality_removal, [], [f2191])).
fof(f2191, plain, (~ (sK4(sK22(sK1, sK2, sK3)) = sK4(sK22(sK1, sK2, sK3))) | (sK4(sK22(sK1, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(sK1, sK2, sK3)))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(equality_factoring, [], [f2157])).
fof(f2157, plain, (! [X2] : ((sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(X2, sK2, sK3))))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2156, f75])).
fof(f2156, plain, (! [X2] : ((sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2155, f76])).
fof(f2155, plain, (! [X2] : ((sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2154, f77])).
fof(f2154, plain, (! [X2] : ((sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2144, f185])).
fof(f2144, plain, (! [X2] : ((sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(sK1, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(resolution, [], [f2141, f123])).
fof(f123, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK22(X0, X1, X2)) | (sK23(X0, X1, X2, X4) = X4) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(cnf_transformation, [], [f74])).
fof(f2141, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3))))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2140, f75])).
fof(f2140, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2139, f76])).
fof(f2139, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2138, f77])).
fof(f2138, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2117, f185])).
fof(f2117, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(duplicate_literal_removal, [], [f2114])).
fof(f2114, plain, (! [X2] : (in(sK4(sK22(X2, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | (sK4(sK22(X2, sK2, sK3)) = sK23(X2, sK2, sK3, sK4(sK22(X2, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(resolution, [], [f2107, f123])).
fof(f2107, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3))))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2106, f75])).
fof(f2106, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2105, f76])).
fof(f2105, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2104, f77])).
fof(f2104, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(subsumption_resolution, [], [f2099, f185])).
fof(f2099, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | (~ spl26_2 | ~ spl26_6)), inference(duplicate_literal_removal, [], [f2096])).
fof(f2096, plain, (! [X1] : (in(sK4(sK22(X1, sK2, sK3)), sK22(X1, sK2, sK3)) | in(sK4(sK22(X1, sK2, sK3)), sK22(sK1, sK2, sK3)) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3))))) | (~ spl26_2 | ~ spl26_6)), inference(resolution, [], [f441, f165])).
fof(f165, plain, ! [X4, X5, X3] : (in(sK4(sK22(X3, X4, X5)), cartesian_product2(sK1, sK1)) | sP0(X4, X5) | ~ function(X5) | ~ relation(X5) | ~ relation(X4) | (sK4(sK22(X3, X4, X5)) = sK23(X3, X4, X5, sK4(sK22(X3, X4, X5))))), inference(resolution, [], [f123, f78])).
fof(f441, plain, (! [X0, X1] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(X1, X1)) | in(sK4(sK22(X0, sK2, sK3)), sK22(X0, sK2, sK3)) | in(sK4(sK22(X0, sK2, sK3)), sK22(X1, sK2, sK3)) | (sK4(sK22(X0, sK2, sK3)) = sK23(X0, sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | (~ spl26_2 | ~ spl26_6)), inference(superposition, [], [f189, f248])).
fof(f248, plain, (! [X0] : ((sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | (sK4(sK22(X0, sK2, sK3)) = sK23(X0, sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | ~ spl26_6), inference(avatar_component_clause, [], [f247])).
fof(f189, plain, (! [X0, X1] : (~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)) | in(sK4(X0), X0) | in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3))) | ~ spl26_2), inference(avatar_component_clause, [], [f188])).
fof(f122, plain, ! [X4, X2, X0, X1] : (in(sK23(X0, X1, X2, X4), cartesian_product2(X0, X0)) | ~ in(X4, sK22(X0, X1, X2)) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(cnf_transformation, [], [f74])).
fof(f2330, plain, (~ in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | ~ in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ spl26_131), inference(equality_resolution, [], [f2278])).
fof(f2278, plain, (! [X1] : (~ (sK4(X1) = sK4(sK22(sK1, sK2, sK3))) | ~ in(sK4(X1), X1) | ~ in(sK4(X1), cartesian_product2(sK1, sK1))) | ~ spl26_131), inference(avatar_component_clause, [], [f2277])).
fof(f2271, plain, (spl26_128 | spl26_129 | spl26_1 | ~ spl26_2 | ~ spl26_6), inference(avatar_split_clause, [], [f2261, f247, f188, f184, f2269, f2265])).
fof(f2261, plain, (! [X0] : (~ in(sK4(sK22(sK1, sK2, sK3)), cartesian_product2(X0, X0)) | in(sK4(sK22(sK1, sK2, sK3)), sK22(sK1, sK2, sK3)) | in(sK4(sK22(sK1, sK2, sK3)), sK22(X0, sK2, sK3))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(superposition, [], [f189, f2249])).
fof(f2249, plain, ((sK4(sK22(sK1, sK2, sK3)) = ordered_pair(sK5(sK22(sK1, sK2, sK3)), sK6(sK22(sK1, sK2, sK3)))) | (spl26_1 | ~ spl26_2 | ~ spl26_6)), inference(resolution, [], [f2228, f1826])).
fof(f1826, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f1825, f75])).
fof(f1825, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | ~ relation(sK2)) | spl26_1), inference(subsumption_resolution, [], [f1824, f76])).
fof(f1824, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | ~ relation(sK3) | ~ relation(sK2)) | spl26_1), inference(subsumption_resolution, [], [f1823, f77])).
fof(f1823, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | spl26_1), inference(subsumption_resolution, [], [f1822, f185])).
fof(f1822, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2)) | spl26_1), inference(duplicate_literal_removal, [], [f1820])).
fof(f1820, plain, (! [X0] : (~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3)))) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(resolution, [], [f1432, f173])).
fof(f173, plain, ! [X2, X0, X1] : (in(ordered_pair(apply(X0, sK24(X1, X0, sK4(sK22(X2, X1, X0)))), apply(X0, sK25(X1, X0, sK4(sK22(X2, X1, X0))))), X1) | sP0(X1, X0) | ~ function(X0) | ~ relation(X0) | ~ relation(X1) | (sK4(sK22(X2, X1, X0)) = ordered_pair(sK5(sK22(X2, X1, X0)), sK6(sK22(X2, X1, X0))))), inference(resolution, [], [f125, f79])).
fof(f79, plain, ! [X3] : (in(sK4(X3), X3) | (sK4(X3) = ordered_pair(sK5(X3), sK6(X3)))), inference(cnf_transformation, [], [f46])).
fof(f1432, plain, (! [X0] : (~ in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))), sK2) | ~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f1431, f79])).
fof(f1431, plain, (! [X0] : (~ in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))), sK2) | ~ in(sK4(sK22(X0, sK2, sK3)), cartesian_product2(sK1, sK1)) | ~ in(sK4(sK22(X0, sK2, sK3)), sK22(X0, sK2, sK3)) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(equality_resolution, [], [f340])).
fof(f340, plain, (! [X0, X1] : (~ (sK4(sK22(X0, sK2, sK3)) = sK4(X1)) | ~ in(ordered_pair(apply(sK3, sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3)))), apply(sK3, sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))), sK2) | ~ in(sK4(X1), cartesian_product2(sK1, sK1)) | ~ in(sK4(X1), X1) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(superposition, [], [f81, f303])).
fof(f303, plain, (! [X0] : ((sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f302, f75])).
fof(f302, plain, (! [X0] : ((sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | ~ relation(sK2) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f301, f76])).
fof(f301, plain, (! [X0] : ((sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f299, f77])).
fof(f299, plain, (! [X0] : ((sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK24(sK2, sK3, sK4(sK22(X0, sK2, sK3))), sK25(sK2, sK3, sK4(sK22(X0, sK2, sK3))))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X0, sK2, sK3)) = ordered_pair(sK5(sK22(X0, sK2, sK3)), sK6(sK22(X0, sK2, sK3))))) | spl26_1), inference(resolution, [], [f185, f169])).
fof(f169, plain, ! [X2, X0, X1] : (sP0(X0, X1) | (sK4(sK22(X2, X0, X1)) = ordered_pair(sK24(X0, X1, sK4(sK22(X2, X0, X1))), sK25(X0, X1, sK4(sK22(X2, X0, X1))))) | ~ function(X1) | ~ relation(X1) | ~ relation(X0) | (sK4(sK22(X2, X0, X1)) = ordered_pair(sK5(sK22(X2, X0, X1)), sK6(sK22(X2, X0, X1))))), inference(resolution, [], [f124, f79])).
fof(f307, plain, (spl26_6 | spl26_1), inference(avatar_split_clause, [], [f306, f184, f247])).
fof(f306, plain, (! [X1] : ((sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | (sK4(sK22(X1, sK2, sK3)) = ordered_pair(sK5(sK22(X1, sK2, sK3)), sK6(sK22(X1, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f305, f75])).
fof(f305, plain, (! [X1] : ((sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ relation(sK2) | (sK4(sK22(X1, sK2, sK3)) = ordered_pair(sK5(sK22(X1, sK2, sK3)), sK6(sK22(X1, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f304, f76])).
fof(f304, plain, (! [X1] : ((sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X1, sK2, sK3)) = ordered_pair(sK5(sK22(X1, sK2, sK3)), sK6(sK22(X1, sK2, sK3))))) | spl26_1), inference(subsumption_resolution, [], [f300, f77])).
fof(f300, plain, (! [X1] : ((sK4(sK22(X1, sK2, sK3)) = sK23(X1, sK2, sK3, sK4(sK22(X1, sK2, sK3)))) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | (sK4(sK22(X1, sK2, sK3)) = ordered_pair(sK5(sK22(X1, sK2, sK3)), sK6(sK22(X1, sK2, sK3))))) | spl26_1), inference(resolution, [], [f185, f164])).
fof(f164, plain, ! [X2, X0, X1] : (sP0(X1, X2) | (sK4(sK22(X0, X1, X2)) = sK23(X0, X1, X2, sK4(sK22(X0, X1, X2)))) | ~ function(X2) | ~ relation(X2) | ~ relation(X1) | (sK4(sK22(X0, X1, X2)) = ordered_pair(sK5(sK22(X0, X1, X2)), sK6(sK22(X0, X1, X2))))), inference(resolution, [], [f123, f79])).
fof(f298, plain, (~ spl26_1 | ~ spl26_5), inference(avatar_split_clause, [], [f297, f243, f184])).
fof(f297, plain, (~ sP0(sK2, sK3) | ~ spl26_5), inference(subsumption_resolution, [], [f296, f115])).
fof(f115, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK15(X0, X1) = sK16(X0, X1))), inference(cnf_transformation, [], [f68])).
fof(f296, plain, (~ (sK15(sK2, sK3) = sK16(sK2, sK3)) | ~ sP0(sK2, sK3) | ~ spl26_5), inference(superposition, [], [f121, f245])).
fof(f245, plain, ((sK15(sK2, sK3) = sK17(sK2, sK3)) | ~ spl26_5), inference(avatar_component_clause, [], [f243])).
fof(f121, plain, ! [X0, X1] : (~ (sK16(X0, X1) = sK17(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f68])).
fof(f190, plain, (spl26_1 | spl26_2), inference(avatar_split_clause, [], [f182, f188, f184])).
fof(f182, plain, ! [X0, X1] : (in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3)) | ~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)) | sP0(sK2, sK3) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f181, f75])).
fof(f181, plain, ! [X0, X1] : (in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3)) | ~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)) | sP0(sK2, sK3) | ~ relation(sK2) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f180, f76])).
fof(f180, plain, ! [X0, X1] : (in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3)) | ~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)) | sP0(sK2, sK3) | ~ relation(sK3) | ~ relation(sK2) | in(sK4(X0), X0)), inference(subsumption_resolution, [], [f177, f77])).
fof(f177, plain, ! [X0, X1] : (in(ordered_pair(sK5(X0), sK6(X0)), sK22(X1, sK2, sK3)) | ~ in(ordered_pair(sK5(X0), sK6(X0)), cartesian_product2(X1, X1)) | sP0(sK2, sK3) | ~ function(sK3) | ~ relation(sK3) | ~ relation(sK2) | in(sK4(X0), X0)), inference(resolution, [], [f128, f80])).
fof(f80, plain, ! [X3] : (in(ordered_pair(apply(sK3, sK5(X3)), apply(sK3, sK6(X3))), sK2) | in(sK4(X3), X3)), inference(cnf_transformation, [], [f46])).
fof(f128, plain, ! [X6, X2, X0, X7, X1] : (~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | in(ordered_pair(X6, X7), sK22(X0, X1, X2)) | ~ in(ordered_pair(X6, X7), cartesian_product2(X0, X0)) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(equality_resolution, [], [f127])).
fof(f127, plain, ! [X6, X2, X0, X7, X5, X1] : (in(ordered_pair(X6, X7), sK22(X0, X1, X2)) | ~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X5) | ~ in(X5, cartesian_product2(X0, X0)) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(equality_resolution, [], [f126])).
fof(f126, plain, ! [X6, X4, X2, X0, X7, X5, X1] : (in(X4, sK22(X0, X1, X2)) | ~ in(ordered_pair(apply(X2, X6), apply(X2, X7)), X1) | ~ (ordered_pair(X6, X7) = X4) | ~ (X4 = X5) | ~ in(X5, cartesian_product2(X0, X0)) | sP0(X1, X2) | ~ function(X2) | ~ relation(X2) | ~ relation(X1)), inference(cnf_transformation, [], [f74])).