fof(f1637, plain, $false, inference(avatar_sat_refutation, [], [f228, f245, f484, f723, f731, f748, f1094, f1123, f1380, f1382, f1507, f1636])).
fof(f1636, plain, (spl28_3 | spl28_6 | ~ spl28_17), inference(avatar_contradiction_clause, [], [f1635])).
fof(f1635, plain, ($false | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1634, f86])).
fof(f86, plain, ~ empty_carrier(sK2), inference(cnf_transformation, [], [f49])).
fof(f49, plain, (! [X3] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(sK2, X5, X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2))) | ~ (sK5(X3) = X5)) | ~ in(sK5(X3), powerset(sK4)) | ~ in(sK5(X3), X3)) & ((((relstr_set_smaller(sK2, sK6(X3), sK7(X3)) & in(sK7(X3), sK3) & element(sK7(X3), the_carrier(sK2))) & (sK5(X3) = sK6(X3))) & in(sK5(X3), powerset(sK4))) | in(sK5(X3), X3))) & element(sK4, powerset(sK3)) & finite(sK4) & element(sK3, powerset(the_carrier(sK2))) & rel_str(sK2) & transitive_relstr(sK2) & ~ empty_carrier(sK2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2, sK3, sK4, sK5, sK6, sK7])], [f44, f48, f47, f46, f45])).
fof(f45, plain, (? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(X0, X5, X6) | ~ in(X6, X1) | ~ element(X6, the_carrier(X0))) | ~ (X4 = X5)) | ~ in(X4, powerset(X2)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : (relstr_set_smaller(X0, X7, X8) & in(X8, X1) & element(X8, the_carrier(X0))) & (X4 = X7)) & in(X4, powerset(X2))) | in(X4, X3))) & element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)) => (! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(sK2, X5, X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2))) | ~ (X4 = X5)) | ~ in(X4, powerset(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : (relstr_set_smaller(sK2, X7, X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) & (X4 = X7)) & in(X4, powerset(sK4))) | in(X4, X3))) & element(sK4, powerset(sK3)) & finite(sK4) & element(sK3, powerset(the_carrier(sK2))) & rel_str(sK2) & transitive_relstr(sK2) & ~ empty_carrier(sK2))), introduced(choice_axiom, [])).
fof(f46, plain, ! [X3] : (? [X4] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(sK2, X5, X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2))) | ~ (X4 = X5)) | ~ in(X4, powerset(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : (relstr_set_smaller(sK2, X7, X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) & (X4 = X7)) & in(X4, powerset(sK4))) | in(X4, X3))) => ((! [X5] : (! [X6] : (~ relstr_set_smaller(sK2, X5, X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2))) | ~ (sK5(X3) = X5)) | ~ in(sK5(X3), powerset(sK4)) | ~ in(sK5(X3), X3)) & ((? [X7] : (? [X8] : (relstr_set_smaller(sK2, X7, X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) & (sK5(X3) = X7)) & in(sK5(X3), powerset(sK4))) | in(sK5(X3), X3)))), introduced(choice_axiom, [])).
fof(f47, plain, ! [X3] : (? [X7] : (? [X8] : (relstr_set_smaller(sK2, X7, X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) & (sK5(X3) = X7)) => (? [X8] : (relstr_set_smaller(sK2, sK6(X3), X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) & (sK5(X3) = sK6(X3)))), introduced(choice_axiom, [])).
fof(f48, plain, ! [X3] : (? [X8] : (relstr_set_smaller(sK2, sK6(X3), X8) & in(X8, sK3) & element(X8, the_carrier(sK2))) => (relstr_set_smaller(sK2, sK6(X3), sK7(X3)) & in(sK7(X3), sK3) & element(sK7(X3), the_carrier(sK2)))), introduced(choice_axiom, [])).
fof(f44, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(X0, X5, X6) | ~ in(X6, X1) | ~ element(X6, the_carrier(X0))) | ~ (X4 = X5)) | ~ in(X4, powerset(X2)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : (relstr_set_smaller(X0, X7, X8) & in(X8, X1) & element(X8, the_carrier(X0))) & (X4 = X7)) & in(X4, powerset(X2))) | in(X4, X3))) & element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)), inference(rectify, [], [f43])).
fof(f43, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ relstr_set_smaller(X0, X5, X6) | ~ in(X6, X1) | ~ element(X6, the_carrier(X0))) | ~ (X4 = X5)) | ~ in(X4, powerset(X2)) | ~ in(X4, X3)) & ((? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2))) | in(X4, X3))) & element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)), inference(flattening, [], [f42])).
fof(f42, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : (((! [X5] : (! [X6] : (~ relstr_set_smaller(X0, X5, X6) | ~ in(X6, X1) | ~ element(X6, the_carrier(X0))) | ~ (X4 = X5)) | ~ in(X4, powerset(X2))) | ~ in(X4, X3)) & ((? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2))) | in(X4, X3))) & element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)), inference(nnf_transformation, [], [f25])).
fof(f25, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2)))) & element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)), inference(flattening, [], [f24])).
fof(f24, plain, ? [X0, X1, X2] : (! [X3] : ? [X4] : ~ (in(X4, X3) <=> (? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2)))) & (element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1, X2] : ((element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)) => ? [X3] : ! [X4] : (in(X4, X3) <=> (? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1, X2] : ((element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)) => ? [X3] : ! [X4] : (in(X4, X3) <=> (? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & in(X4, powerset(X2))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU364+1.p', s1_xboole_0__e11_2_1__waybel_0__1)).
fof(f1634, plain, (empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1633, f87])).
fof(f87, plain, transitive_relstr(sK2), inference(cnf_transformation, [], [f49])).
fof(f1633, plain, (~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1632, f88])).
fof(f88, plain, rel_str(sK2), inference(cnf_transformation, [], [f49])).
fof(f1632, plain, (~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1631, f89])).
fof(f89, plain, element(sK3, powerset(the_carrier(sK2))), inference(cnf_transformation, [], [f49])).
fof(f1631, plain, (~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1630, f90])).
fof(f90, plain, finite(sK4), inference(cnf_transformation, [], [f49])).
fof(f1630, plain, (~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1629, f91])).
fof(f91, plain, element(sK4, powerset(sK3)), inference(cnf_transformation, [], [f49])).
fof(f1629, plain, (~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1628, f223])).
fof(f223, plain, (~ sP1(sK2, sK3) | spl28_3), inference(avatar_component_clause, [], [f222])).
fof(f222, plain, (spl28_3 <=> sP1(sK2, sK3)), introduced(avatar_definition, [new_symbols(naming, [spl28_3])])).
fof(f1628, plain, (sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_6 | ~ spl28_17)), inference(subsumption_resolution, [], [f1619, f403])).
fof(f403, plain, (~ (sK5(sK24(sK2, sK3, sK4)) = sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | spl28_6), inference(avatar_component_clause, [], [f402])).
fof(f402, plain, (spl28_6 <=> (sK5(sK24(sK2, sK3, sK4)) = sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))))), introduced(avatar_definition, [new_symbols(naming, [spl28_6])])).
fof(f1619, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | ~ spl28_17), inference(resolution, [], [f480, f136])).
fof(f136, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X0, X1, X2)) | (sK26(X0, X1, X4) = X4) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f85, plain, ! [X0, X1, X2] : (! [X4] : ((in(X4, sK24(X0, X1, X2)) | ! [X5] : (! [X6] : (! [X7] : (~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0))) | ~ (X4 = X6)) | ~ (X4 = X5) | ~ in(X5, powerset(X2)))) & ((((relstr_set_smaller(X0, sK26(X0, X1, X4), sK27(X0, X1, X4)) & in(sK27(X0, X1, X4), X1) & element(sK27(X0, X1, X4), the_carrier(X0))) & (sK26(X0, X1, X4) = X4)) & (sK25(X0, X1, X2, X4) = X4) & in(sK25(X0, X1, X2, X4), powerset(X2))) | ~ in(X4, sK24(X0, X1, X2)))) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24, sK25, sK26, sK27])], [f80, f84, f83, f82, f81])).
fof(f81, plain, ! [X2, X1, X0] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6] : (! [X7] : (~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0))) | ~ (X4 = X6)) | ~ (X4 = X5) | ~ in(X5, powerset(X2)))) & (? [X8] : (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) & (X4 = X8) & in(X8, powerset(X2))) | ~ in(X4, X3))) => ! [X4] : ((in(X4, sK24(X0, X1, X2)) | ! [X5] : (! [X6] : (! [X7] : (~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0))) | ~ (X4 = X6)) | ~ (X4 = X5) | ~ in(X5, powerset(X2)))) & (? [X8] : (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) & (X4 = X8) & in(X8, powerset(X2))) | ~ in(X4, sK24(X0, X1, X2))))), introduced(choice_axiom, [])).
fof(f82, plain, ! [X4, X2, X1, X0] : (? [X8] : (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) & (X4 = X8) & in(X8, powerset(X2))) => (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) & (sK25(X0, X1, X2, X4) = X4) & in(sK25(X0, X1, X2, X4), powerset(X2)))), introduced(choice_axiom, [])).
fof(f83, plain, ! [X4, X1, X0] : (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) => (? [X10] : (relstr_set_smaller(X0, sK26(X0, X1, X4), X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (sK26(X0, X1, X4) = X4))), introduced(choice_axiom, [])).
fof(f84, plain, ! [X4, X1, X0] : (? [X10] : (relstr_set_smaller(X0, sK26(X0, X1, X4), X10) & in(X10, X1) & element(X10, the_carrier(X0))) => (relstr_set_smaller(X0, sK26(X0, X1, X4), sK27(X0, X1, X4)) & in(sK27(X0, X1, X4), X1) & element(sK27(X0, X1, X4), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f80, plain, ! [X0, X1, X2] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6] : (! [X7] : (~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0))) | ~ (X4 = X6)) | ~ (X4 = X5) | ~ in(X5, powerset(X2)))) & (? [X8] : (? [X9] : (? [X10] : (relstr_set_smaller(X0, X9, X10) & in(X10, X1) & element(X10, the_carrier(X0))) & (X4 = X9)) & (X4 = X8) & in(X8, powerset(X2))) | ~ in(X4, X3))) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(rectify, [], [f79])).
fof(f79, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : ((in(X11, X10) | ! [X12] : (! [X13] : (! [X14] : (~ relstr_set_smaller(X0, X13, X14) | ~ in(X14, X1) | ~ element(X14, the_carrier(X0))) | ~ (X11 = X13)) | ~ (X11 = X12) | ~ in(X12, powerset(X2)))) & (? [X12] : (? [X13] : (? [X14] : (relstr_set_smaller(X0, X13, X14) & in(X14, X1) & element(X14, the_carrier(X0))) & (X11 = X13)) & (X11 = X12) & in(X12, powerset(X2))) | ~ in(X11, X10))) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(nnf_transformation, [], [f41])).
fof(f41, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13] : (? [X14] : (relstr_set_smaller(X0, X13, X14) & in(X14, X1) & element(X14, the_carrier(X0))) & (X11 = X13)) & (X11 = X12) & in(X12, powerset(X2)))) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(definition_folding, [], [f38, e40, e39])).
fof(f39, plain, ! [X0, X1, X4] : (? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X4 = X8)) | ~ sP0(X0, X1, X4)), inference(usedef, [], [e39])).
fof(e39, plain, ! [X0, X1, X4] : (sP0(X0, X1, X4) <=> ? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X4 = X8))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f40, plain, ! [X0, X1] : (? [X3, X4, X5] : (~ (X4 = X5) & ? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & sP0(X0, X1, X4) & (X3 = X4)) | ~ sP1(X0, X1)), inference(usedef, [], [e40])).
fof(e40, plain, ! [X0, X1] : (sP1(X0, X1) <=> ? [X3, X4, X5] : (~ (X4 = X5) & ? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & sP0(X0, X1, X4) & (X3 = X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f38, plain, ! [X0, X1, X2] : (? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13] : (? [X14] : (relstr_set_smaller(X0, X13, X14) & in(X14, X1) & element(X14, the_carrier(X0))) & (X11 = X13)) & (X11 = X12) & in(X12, powerset(X2)))) | ? [X3, X4, X5] : (~ (X4 = X5) & ? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & ? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X4 = X8)) & (X3 = X4)) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(flattening, [], [f37])).
fof(f37, plain, ! [X0, X1, X2] : ((? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13] : (? [X14] : (relstr_set_smaller(X0, X13, X14) & in(X14, X1) & element(X14, the_carrier(X0))) & (X11 = X13)) & (X11 = X12) & in(X12, powerset(X2)))) | ? [X3, X4, X5] : (~ (X4 = X5) & (? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & ? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X4 = X8)) & (X3 = X4)))) | (~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0))), inference(ennf_transformation, [], [f23])).
fof(f23, plain, ! [X0, X1, X2] : ((element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)) => (! [X3, X4, X5] : ((? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & ? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X4 = X8)) & (X3 = X4)) => (X4 = X5)) => ? [X10] : ! [X11] : (in(X11, X10) <=> ? [X12] : (? [X13] : (? [X14] : (relstr_set_smaller(X0, X13, X14) & in(X14, X1) & element(X14, the_carrier(X0))) & (X11 = X13)) & (X11 = X12) & in(X12, powerset(X2)))))), inference(rectify, [], [f22])).
fof(f22, plain, ! [X0, X1, X2] : ((element(X2, powerset(X1)) & finite(X2) & element(X1, powerset(the_carrier(X0))) & rel_str(X0) & transitive_relstr(X0) & ~ empty_carrier(X0)) => (! [X3, X4, X5] : ((? [X8] : (? [X9] : (relstr_set_smaller(X0, X8, X9) & in(X9, X1) & element(X9, the_carrier(X0))) & (X5 = X8)) & (X3 = X5) & ? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X4 = X6)) & (X3 = X4)) => (X4 = X5)) => ? [X3] : ! [X4] : (in(X4, X3) <=> ? [X5] : (? [X10] : (? [X11] : (relstr_set_smaller(X0, X10, X11) & in(X11, X1) & element(X11, the_carrier(X0))) & (X4 = X10)) & (X4 = X5) & in(X5, powerset(X2)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU364+1.p', s1_tarski__e11_2_1__waybel_0__1)).
fof(f480, plain, (in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ spl28_17), inference(avatar_component_clause, [], [f478])).
fof(f478, plain, (spl28_17 <=> in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4))), introduced(avatar_definition, [new_symbols(naming, [spl28_17])])).
fof(f1507, plain, (spl28_3 | spl28_29 | spl28_32), inference(avatar_contradiction_clause, [], [f1506])).
fof(f1506, plain, ($false | (spl28_3 | spl28_29 | spl28_32)), inference(subsumption_resolution, [], [f1505, f663])).
fof(f663, plain, (~ (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | spl28_32), inference(avatar_component_clause, [], [f662])).
fof(f662, plain, (spl28_32 <=> (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4))))), introduced(avatar_definition, [new_symbols(naming, [spl28_32])])).
fof(f1505, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1504, f86])).
fof(f1504, plain, (empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1503, f87])).
fof(f1503, plain, (~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1502, f88])).
fof(f1502, plain, (~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1501, f89])).
fof(f1501, plain, (~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1500, f90])).
fof(f1500, plain, (~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1499, f91])).
fof(f1499, plain, (~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | spl28_29)), inference(subsumption_resolution, [], [f1496, f223])).
fof(f1496, plain, (sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | spl28_29), inference(resolution, [], [f605, f188])).
fof(f188, plain, ! [X6, X8, X7] : (in(sK5(sK24(X6, X7, X8)), powerset(sK4)) | sP1(X6, X7) | ~ element(X8, powerset(X7)) | ~ finite(X8) | ~ element(X7, powerset(the_carrier(X6))) | ~ rel_str(X6) | ~ transitive_relstr(X6) | empty_carrier(X6) | (sK5(sK24(X6, X7, X8)) = sK25(X6, X7, X8, sK5(sK24(X6, X7, X8))))), inference(resolution, [], [f135, f92])).
fof(f92, plain, ! [X3] : (in(sK5(X3), powerset(sK4)) | in(sK5(X3), X3)), inference(cnf_transformation, [], [f49])).
fof(f135, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X0, X1, X2)) | (sK25(X0, X1, X2, X4) = X4) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f605, plain, (~ in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | spl28_29), inference(avatar_component_clause, [], [f603])).
fof(f603, plain, (spl28_29 <=> in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl28_29])])).
fof(f1382, plain, (~ spl28_17 | ~ spl28_29 | ~ spl28_30 | ~ spl28_31 | ~ spl28_33), inference(avatar_split_clause, [], [f1368, f674, f611, f607, f603, f478])).
fof(f607, plain, (spl28_30 <=> element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2))), introduced(avatar_definition, [new_symbols(naming, [spl28_30])])).
fof(f611, plain, (spl28_31 <=> in(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK3)), introduced(avatar_definition, [new_symbols(naming, [spl28_31])])).
fof(f674, plain, (spl28_33 <=> relstr_set_smaller(sK2, sK5(sK24(sK2, sK3, sK4)), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))))), introduced(avatar_definition, [new_symbols(naming, [spl28_33])])).
fof(f1368, plain, (~ in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | (~ spl28_30 | ~ spl28_31 | ~ spl28_33)), inference(subsumption_resolution, [], [f1367, f608])).
fof(f608, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ spl28_30), inference(avatar_component_clause, [], [f607])).
fof(f1367, plain, (~ element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | (~ spl28_31 | ~ spl28_33)), inference(subsumption_resolution, [], [f1365, f612])).
fof(f612, plain, (in(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK3) | ~ spl28_31), inference(avatar_component_clause, [], [f611])).
fof(f1365, plain, (~ in(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK3) | ~ element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ spl28_33), inference(resolution, [], [f676, f141])).
fof(f141, plain, ! [X6, X3] : (~ relstr_set_smaller(sK2, sK5(X3), X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2)) | ~ in(sK5(X3), powerset(sK4)) | ~ in(sK5(X3), X3)), inference(equality_resolution, [], [f97])).
fof(f97, plain, ! [X6, X5, X3] : (~ relstr_set_smaller(sK2, X5, X6) | ~ in(X6, sK3) | ~ element(X6, the_carrier(sK2)) | ~ (sK5(X3) = X5) | ~ in(sK5(X3), powerset(sK4)) | ~ in(sK5(X3), X3)), inference(cnf_transformation, [], [f49])).
fof(f676, plain, (relstr_set_smaller(sK2, sK5(sK24(sK2, sK3, sK4)), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ spl28_33), inference(avatar_component_clause, [], [f674])).
fof(f1380, plain, (spl28_29 | spl28_3 | ~ spl28_32), inference(avatar_split_clause, [], [f1379, f662, f222, f603])).
fof(f1379, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1378, f92])).
fof(f1378, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1377, f86])).
fof(f1377, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1376, f87])).
fof(f1376, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1375, f88])).
fof(f1375, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1374, f89])).
fof(f1374, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1373, f90])).
fof(f1373, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1372, f91])).
fof(f1372, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_32)), inference(subsumption_resolution, [], [f1225, f223])).
fof(f1225, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | ~ spl28_32), inference(superposition, [], [f134, f664])).
fof(f664, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK25(sK2, sK3, sK4, sK5(sK24(sK2, sK3, sK4)))) | ~ spl28_32), inference(avatar_component_clause, [], [f662])).
fof(f134, plain, ! [X4, X2, X0, X1] : (in(sK25(X0, X1, X2, X4), powerset(X2)) | ~ in(X4, sK24(X0, X1, X2)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f1123, plain, (spl28_5 | spl28_17), inference(avatar_split_clause, [], [f764, f478, f398])).
fof(f398, plain, (spl28_5 <=> (sK5(sK24(sK2, sK3, sK4)) = sK6(sK24(sK2, sK3, sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl28_5])])).
fof(f764, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK6(sK24(sK2, sK3, sK4))) | spl28_17), inference(resolution, [], [f479, f93])).
fof(f93, plain, ! [X3] : (in(sK5(X3), X3) | (sK5(X3) = sK6(X3))), inference(cnf_transformation, [], [f49])).
fof(f479, plain, (~ in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4)) | spl28_17), inference(avatar_component_clause, [], [f478])).
fof(f1094, plain, (spl28_17 | ~ spl28_18 | ~ spl28_29), inference(avatar_contradiction_clause, [], [f1093])).
fof(f1093, plain, ($false | (spl28_17 | ~ spl28_18 | ~ spl28_29)), inference(subsumption_resolution, [], [f1092, f90])).
fof(f1092, plain, (~ finite(sK4) | (spl28_17 | ~ spl28_18 | ~ spl28_29)), inference(subsumption_resolution, [], [f1091, f91])).
fof(f1091, plain, (~ element(sK4, powerset(sK3)) | ~ finite(sK4) | (spl28_17 | ~ spl28_18 | ~ spl28_29)), inference(subsumption_resolution, [], [f1079, f604])).
fof(f604, plain, (in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ spl28_29), inference(avatar_component_clause, [], [f603])).
fof(f1079, plain, (~ in(sK5(sK24(sK2, sK3, sK4)), powerset(sK4)) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | (spl28_17 | ~ spl28_18)), inference(resolution, [], [f483, f479])).
fof(f483, plain, (! [X0] : (in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, X0)) | ~ in(sK5(sK24(sK2, sK3, sK4)), powerset(X0)) | ~ element(X0, powerset(sK3)) | ~ finite(X0)) | ~ spl28_18), inference(avatar_component_clause, [], [f482])).
fof(f482, plain, (spl28_18 <=> ! [X0] : (~ in(sK5(sK24(sK2, sK3, sK4)), powerset(X0)) | in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, X0)) | ~ element(X0, powerset(sK3)) | ~ finite(X0))), introduced(avatar_definition, [new_symbols(naming, [spl28_18])])).
fof(f748, plain, (spl28_3 | ~ spl28_17 | spl28_31), inference(avatar_contradiction_clause, [], [f747])).
fof(f747, plain, ($false | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f746, f86])).
fof(f746, plain, (empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f745, f87])).
fof(f745, plain, (~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f744, f88])).
fof(f744, plain, (~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f743, f89])).
fof(f743, plain, (~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f742, f90])).
fof(f742, plain, (~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f741, f91])).
fof(f741, plain, (~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f740, f223])).
fof(f740, plain, (sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (~ spl28_17 | spl28_31)), inference(subsumption_resolution, [], [f712, f613])).
fof(f613, plain, (~ in(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK3) | spl28_31), inference(avatar_component_clause, [], [f611])).
fof(f712, plain, (in(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK3) | sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | ~ spl28_17), inference(resolution, [], [f480, f138])).
fof(f138, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X0, X1, X2)) | in(sK27(X0, X1, X4), X1) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f731, plain, (spl28_30 | spl28_3 | ~ spl28_17), inference(avatar_split_clause, [], [f730, f478, f222, f607])).
fof(f730, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f729, f86])).
fof(f729, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f728, f87])).
fof(f728, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f727, f88])).
fof(f727, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f726, f89])).
fof(f726, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f725, f90])).
fof(f725, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f724, f91])).
fof(f724, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f710, f223])).
fof(f710, plain, (element(sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), the_carrier(sK2)) | sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | ~ spl28_17), inference(resolution, [], [f480, f137])).
fof(f137, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X0, X1, X2)) | element(sK27(X0, X1, X4), the_carrier(X0)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f723, plain, (spl28_33 | spl28_3 | ~ spl28_6 | ~ spl28_17), inference(avatar_split_clause, [], [f722, f478, f402, f222, f674])).
fof(f722, plain, (relstr_set_smaller(sK2, sK5(sK24(sK2, sK3, sK4)), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | ~ spl28_6 | ~ spl28_17)), inference(forward_demodulation, [], [f721, f404])).
fof(f404, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ spl28_6), inference(avatar_component_clause, [], [f402])).
fof(f721, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f720, f86])).
fof(f720, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f719, f87])).
fof(f719, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f718, f88])).
fof(f718, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f717, f89])).
fof(f717, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f716, f90])).
fof(f716, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f715, f91])).
fof(f715, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | (spl28_3 | ~ spl28_17)), inference(subsumption_resolution, [], [f709, f223])).
fof(f709, plain, (relstr_set_smaller(sK2, sK26(sK2, sK3, sK5(sK24(sK2, sK3, sK4))), sK27(sK2, sK3, sK5(sK24(sK2, sK3, sK4)))) | sP1(sK2, sK3) | ~ element(sK4, powerset(sK3)) | ~ finite(sK4) | ~ element(sK3, powerset(the_carrier(sK2))) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2) | ~ spl28_17), inference(resolution, [], [f480, f139])).
fof(f139, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK24(X0, X1, X2)) | relstr_set_smaller(X0, sK26(X0, X1, X4), sK27(X0, X1, X4)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).
fof(f484, plain, (spl28_17 | spl28_18 | ~ spl28_4 | ~ spl28_5), inference(avatar_split_clause, [], [f475, f398, f226, f482, f478])).
fof(f226, plain, (spl28_4 <=> ! [X1, X0, X2] : (~ relstr_set_smaller(sK2, X0, X1) | in(X0, sK24(sK2, sK3, X2)) | ~ finite(X2) | ~ element(X2, powerset(sK3)) | ~ in(X0, powerset(X2)) | ~ element(X1, the_carrier(sK2)) | ~ in(X1, sK3))), introduced(avatar_definition, [new_symbols(naming, [spl28_4])])).
fof(f475, plain, (! [X0] : (~ in(sK5(sK24(sK2, sK3, sK4)), powerset(X0)) | ~ finite(X0) | ~ element(X0, powerset(sK3)) | in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, X0)) | in(sK5(sK24(sK2, sK3, sK4)), sK24(sK2, sK3, sK4))) | (~ spl28_4 | ~ spl28_5)), inference(superposition, [], [f250, f400])).
fof(f400, plain, ((sK5(sK24(sK2, sK3, sK4)) = sK6(sK24(sK2, sK3, sK4))) | ~ spl28_5), inference(avatar_component_clause, [], [f398])).
fof(f250, plain, (! [X0, X1] : (~ in(sK6(X0), powerset(X1)) | ~ finite(X1) | ~ element(X1, powerset(sK3)) | in(sK6(X0), sK24(sK2, sK3, X1)) | in(sK5(X0), X0)) | ~ spl28_4), inference(subsumption_resolution, [], [f249, f95])).
fof(f95, plain, ! [X3] : (in(sK7(X3), sK3) | in(sK5(X3), X3)), inference(cnf_transformation, [], [f49])).
fof(f249, plain, (! [X0, X1] : (in(sK6(X0), sK24(sK2, sK3, X1)) | ~ finite(X1) | ~ element(X1, powerset(sK3)) | ~ in(sK6(X0), powerset(X1)) | ~ in(sK7(X0), sK3) | in(sK5(X0), X0)) | ~ spl28_4), inference(subsumption_resolution, [], [f246, f94])).
fof(f94, plain, ! [X3] : (element(sK7(X3), the_carrier(sK2)) | in(sK5(X3), X3)), inference(cnf_transformation, [], [f49])).
fof(f246, plain, (! [X0, X1] : (in(sK6(X0), sK24(sK2, sK3, X1)) | ~ finite(X1) | ~ element(X1, powerset(sK3)) | ~ in(sK6(X0), powerset(X1)) | ~ element(sK7(X0), the_carrier(sK2)) | ~ in(sK7(X0), sK3) | in(sK5(X0), X0)) | ~ spl28_4), inference(resolution, [], [f227, f96])).
fof(f96, plain, ! [X3] : (relstr_set_smaller(sK2, sK6(X3), sK7(X3)) | in(sK5(X3), X3)), inference(cnf_transformation, [], [f49])).
fof(f227, plain, (! [X2, X0, X1] : (~ relstr_set_smaller(sK2, X0, X1) | in(X0, sK24(sK2, sK3, X2)) | ~ finite(X2) | ~ element(X2, powerset(sK3)) | ~ in(X0, powerset(X2)) | ~ element(X1, the_carrier(sK2)) | ~ in(X1, sK3)) | ~ spl28_4), inference(avatar_component_clause, [], [f226])).
fof(f245, plain, ~ spl28_3, inference(avatar_contradiction_clause, [], [f244])).
fof(f244, plain, ($false | ~ spl28_3), inference(subsumption_resolution, [], [f243, f224])).
fof(f224, plain, (sP1(sK2, sK3) | ~ spl28_3), inference(avatar_component_clause, [], [f222])).
fof(f243, plain, (~ sP1(sK2, sK3) | ~ spl28_3), inference(subsumption_resolution, [], [f242, f233])).
fof(f233, plain, ((sK18(sK2, sK3) = sK17(sK2, sK3)) | ~ spl28_3), inference(resolution, [], [f224, f122])).
fof(f122, plain, ! [X0, X1] : (~ sP1(X0, X1) | (sK17(X0, X1) = sK18(X0, X1))), inference(cnf_transformation, [], [f73])).
fof(f73, plain, ! [X0, X1] : ((~ (sK18(X0, X1) = sK19(X0, X1)) & ((relstr_set_smaller(X0, sK20(X0, X1), sK21(X0, X1)) & in(sK21(X0, X1), X1) & element(sK21(X0, X1), the_carrier(X0))) & (sK19(X0, X1) = sK20(X0, X1))) & (sK17(X0, X1) = sK19(X0, X1)) & sP0(X0, X1, sK18(X0, X1)) & (sK17(X0, X1) = sK18(X0, X1))) | ~ sP1(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK17, sK18, sK19, sK20, sK21])], [f69, f72, f71, f70])).
fof(f70, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & (X2 = X4) & sP0(X0, X1, X3) & (X2 = X3)) => (~ (sK18(X0, X1) = sK19(X0, X1)) & ? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (sK19(X0, X1) = X5)) & (sK17(X0, X1) = sK19(X0, X1)) & sP0(X0, X1, sK18(X0, X1)) & (sK17(X0, X1) = sK18(X0, X1)))), introduced(choice_axiom, [])).
fof(f71, plain, ! [X1, X0] : (? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (sK19(X0, X1) = X5)) => (? [X6] : (relstr_set_smaller(X0, sK20(X0, X1), X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (sK19(X0, X1) = sK20(X0, X1)))), introduced(choice_axiom, [])).
fof(f72, plain, ! [X1, X0] : (? [X6] : (relstr_set_smaller(X0, sK20(X0, X1), X6) & in(X6, X1) & element(X6, the_carrier(X0))) => (relstr_set_smaller(X0, sK20(X0, X1), sK21(X0, X1)) & in(sK21(X0, X1), X1) & element(sK21(X0, X1), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f69, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & ? [X5] : (? [X6] : (relstr_set_smaller(X0, X5, X6) & in(X6, X1) & element(X6, the_carrier(X0))) & (X4 = X5)) & (X2 = X4) & sP0(X0, X1, X3) & (X2 = X3)) | ~ sP1(X0, X1)), inference(rectify, [], [f68])).
fof(f68, plain, ! [X0, X1] : (? [X3, X4, X5] : (~ (X4 = X5) & ? [X6] : (? [X7] : (relstr_set_smaller(X0, X6, X7) & in(X7, X1) & element(X7, the_carrier(X0))) & (X5 = X6)) & (X3 = X5) & sP0(X0, X1, X4) & (X3 = X4)) | ~ sP1(X0, X1)), inference(nnf_transformation, [], [f40])).
fof(f242, plain, (~ (sK18(sK2, sK3) = sK17(sK2, sK3)) | ~ sP1(sK2, sK3) | ~ spl28_3), inference(superposition, [], [f129, f232])).
fof(f232, plain, ((sK19(sK2, sK3) = sK17(sK2, sK3)) | ~ spl28_3), inference(resolution, [], [f224, f124])).
fof(f124, plain, ! [X0, X1] : (~ sP1(X0, X1) | (sK17(X0, X1) = sK19(X0, X1))), inference(cnf_transformation, [], [f73])).
fof(f129, plain, ! [X0, X1] : (~ (sK18(X0, X1) = sK19(X0, X1)) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f73])).
fof(f228, plain, (spl28_3 | spl28_4), inference(avatar_split_clause, [], [f220, f226, f222])).
fof(f220, plain, ! [X2, X0, X1] : (~ relstr_set_smaller(sK2, X0, X1) | ~ in(X1, sK3) | ~ element(X1, the_carrier(sK2)) | ~ in(X0, powerset(X2)) | sP1(sK2, sK3) | ~ element(X2, powerset(sK3)) | ~ finite(X2) | in(X0, sK24(sK2, sK3, X2))), inference(subsumption_resolution, [], [f219, f86])).
fof(f219, plain, ! [X2, X0, X1] : (~ relstr_set_smaller(sK2, X0, X1) | ~ in(X1, sK3) | ~ element(X1, the_carrier(sK2)) | ~ in(X0, powerset(X2)) | sP1(sK2, sK3) | ~ element(X2, powerset(sK3)) | ~ finite(X2) | in(X0, sK24(sK2, sK3, X2)) | empty_carrier(sK2)), inference(subsumption_resolution, [], [f218, f87])).
fof(f218, plain, ! [X2, X0, X1] : (~ relstr_set_smaller(sK2, X0, X1) | ~ in(X1, sK3) | ~ element(X1, the_carrier(sK2)) | ~ in(X0, powerset(X2)) | sP1(sK2, sK3) | ~ element(X2, powerset(sK3)) | ~ finite(X2) | in(X0, sK24(sK2, sK3, X2)) | ~ transitive_relstr(sK2) | empty_carrier(sK2)), inference(subsumption_resolution, [], [f211, f88])).
fof(f211, plain, ! [X2, X0, X1] : (~ relstr_set_smaller(sK2, X0, X1) | ~ in(X1, sK3) | ~ element(X1, the_carrier(sK2)) | ~ in(X0, powerset(X2)) | sP1(sK2, sK3) | ~ element(X2, powerset(sK3)) | ~ finite(X2) | in(X0, sK24(sK2, sK3, X2)) | ~ rel_str(sK2) | ~ transitive_relstr(sK2) | empty_carrier(sK2)), inference(resolution, [], [f143, f89])).
fof(f143, plain, ! [X6, X2, X0, X7, X1] : (~ element(X1, powerset(the_carrier(X0))) | ~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0)) | ~ in(X6, powerset(X2)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | in(X6, sK24(X0, X1, X2)) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(equality_resolution, [], [f142])).
fof(f142, plain, ! [X6, X2, X0, X7, X5, X1] : (in(X6, sK24(X0, X1, X2)) | ~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0)) | ~ (X5 = X6) | ~ in(X5, powerset(X2)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(equality_resolution, [], [f140])).
fof(f140, plain, ! [X6, X4, X2, X0, X7, X5, X1] : (in(X4, sK24(X0, X1, X2)) | ~ relstr_set_smaller(X0, X6, X7) | ~ in(X7, X1) | ~ element(X7, the_carrier(X0)) | ~ (X4 = X6) | ~ (X4 = X5) | ~ in(X5, powerset(X2)) | sP1(X0, X1) | ~ element(X2, powerset(X1)) | ~ finite(X2) | ~ element(X1, powerset(the_carrier(X0))) | ~ rel_str(X0) | ~ transitive_relstr(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f85])).