fof(f436, plain, $false, inference(avatar_sat_refutation, [], [f334, f347, f353, f387, f419, f435])).
fof(f435, plain, (spl13_7 | ~ spl13_9 | ~ spl13_12), inference(avatar_contradiction_clause, [], [f434])).
fof(f434, plain, ($false | (spl13_7 | ~ spl13_9 | ~ spl13_12)), inference(subsumption_resolution, [], [f433, f365])).
fof(f365, plain, (in(sK3(sK11(sK1, sK2)), sK11(sK1, sK2)) | spl13_7), inference(factoring, [], [f362])).
fof(f362, plain, (! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | in(sK3(X0), X0)) | spl13_7), inference(subsumption_resolution, [], [f361, f108])).
fof(f108, plain, ! [X2] : (in(sK3(X2), powerset(the_carrier(sK1))) | in(sK3(X2), X2)), inference(cnf_transformation, [], [f87])).
fof(f87, plain, (! [X2] : ((~ in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) | ~ in(sK3(X2), powerset(the_carrier(sK1))) | ~ in(sK3(X2), X2)) & ((in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) & in(sK3(X2), powerset(the_carrier(sK1)))) | in(sK3(X2), X2))) & element(sK2, powerset(powerset(the_carrier(sK1)))) & top_str(sK1) & topological_space(sK1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2, sK3])], [f84, f86, f85])).
fof(f85, plain, (? [X0, X1] : (! [X2] : ? [X3] : ((~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ in(X3, powerset(the_carrier(X0))) | ~ in(X3, X2)) & ((in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0)))) | in(X3, X2))) & element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)) => (! [X2] : ? [X3] : ((~ in(set_difference(cast_as_carrier_subset(sK1), X3), sK2) | ~ in(X3, powerset(the_carrier(sK1))) | ~ in(X3, X2)) & ((in(set_difference(cast_as_carrier_subset(sK1), X3), sK2) & in(X3, powerset(the_carrier(sK1)))) | in(X3, X2))) & element(sK2, powerset(powerset(the_carrier(sK1)))) & top_str(sK1) & topological_space(sK1))), introduced(choice_axiom, [])).
fof(f86, plain, ! [X2] : (? [X3] : ((~ in(set_difference(cast_as_carrier_subset(sK1), X3), sK2) | ~ in(X3, powerset(the_carrier(sK1))) | ~ in(X3, X2)) & ((in(set_difference(cast_as_carrier_subset(sK1), X3), sK2) & in(X3, powerset(the_carrier(sK1)))) | in(X3, X2))) => ((~ in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) | ~ in(sK3(X2), powerset(the_carrier(sK1))) | ~ in(sK3(X2), X2)) & ((in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) & in(sK3(X2), powerset(the_carrier(sK1)))) | in(sK3(X2), X2)))), introduced(choice_axiom, [])).
fof(f84, plain, ? [X0, X1] : (! [X2] : ? [X3] : ((~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ in(X3, powerset(the_carrier(X0))) | ~ in(X3, X2)) & ((in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0)))) | in(X3, X2))) & element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)), inference(flattening, [], [f83])).
fof(f83, plain, ? [X0, X1] : (! [X2] : ? [X3] : (((~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ in(X3, powerset(the_carrier(X0)))) | ~ in(X3, X2)) & ((in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0)))) | in(X3, X2))) & element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)), inference(nnf_transformation, [], [f57])).
fof(f57, plain, ? [X0, X1] : (! [X2] : ? [X3] : ~ (in(X3, X2) <=> (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0))))) & element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)), inference(flattening, [], [f56])).
fof(f56, plain, ? [X0, X1] : (! [X2] : ? [X3] : ~ (in(X3, X2) <=> (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0))))) & (element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0, X1] : ((element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)) => ? [X2] : ! [X3] : (in(X3, X2) <=> (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0)))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1] : ((element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)) => ? [X2] : ! [X3] : (in(X3, X2) <=> (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & in(X3, powerset(the_carrier(X0)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU310+1.p', s1_xboole_0__e2_37_1_1__pre_topc__1)).
fof(f361, plain, (! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | in(sK3(X0), X0)) | spl13_7), inference(subsumption_resolution, [], [f360, f105])).
fof(f105, plain, topological_space(sK1), inference(cnf_transformation, [], [f87])).
fof(f360, plain, (! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | ~ topological_space(sK1) | in(sK3(X0), X0)) | spl13_7), inference(subsumption_resolution, [], [f359, f106])).
fof(f106, plain, top_str(sK1), inference(cnf_transformation, [], [f87])).
fof(f359, plain, (! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | ~ top_str(sK1) | ~ topological_space(sK1) | in(sK3(X0), X0)) | spl13_7), inference(subsumption_resolution, [], [f358, f107])).
fof(f107, plain, element(sK2, powerset(powerset(the_carrier(sK1)))), inference(cnf_transformation, [], [f87])).
fof(f358, plain, (! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | ~ element(sK2, powerset(powerset(the_carrier(sK1)))) | ~ top_str(sK1) | ~ topological_space(sK1) | in(sK3(X0), X0)) | spl13_7), inference(subsumption_resolution, [], [f355, f329])).
fof(f329, plain, (~ sP0(sK2, sK1) | spl13_7), inference(avatar_component_clause, [], [f328])).
fof(f328, plain, (spl13_7 <=> sP0(sK2, sK1)), introduced(avatar_definition, [new_symbols(naming, [spl13_7])])).
fof(f355, plain, ! [X0] : (in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | sP0(sK2, sK1) | ~ element(sK2, powerset(powerset(the_carrier(sK1)))) | ~ top_str(sK1) | ~ topological_space(sK1) | in(sK3(X0), X0)), inference(resolution, [], [f174, f109])).
fof(f109, plain, ! [X2] : (in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) | in(sK3(X2), X2)), inference(cnf_transformation, [], [f87])).
fof(f174, plain, ! [X4, X0, X1] : (~ in(set_difference(cast_as_carrier_subset(X0), X4), X1) | in(X4, sK11(X0, X1)) | ~ in(X4, powerset(the_carrier(X0))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(equality_resolution, [], [f173])).
fof(f173, plain, ! [X4, X0, X3, X1] : (in(X3, sK11(X0, X1)) | ~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ (X3 = X4) | ~ in(X4, powerset(the_carrier(X0))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ! [X0, X1] : (! [X3] : ((in(X3, sK11(X0, X1)) | ! [X4] : (~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ (X3 = X4) | ~ in(X4, powerset(the_carrier(X0))))) & ((in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (sK12(X0, X1, X3) = X3) & in(sK12(X0, X1, X3), powerset(the_carrier(X0)))) | ~ in(X3, sK11(X0, X1)))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK11, sK12])], [f101, f103, f102])).
fof(f102, plain, ! [X1, X0] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ (X3 = X4) | ~ in(X4, powerset(the_carrier(X0))))) & (? [X5] : (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X3 = X5) & in(X5, powerset(the_carrier(X0)))) | ~ in(X3, X2))) => ! [X3] : ((in(X3, sK11(X0, X1)) | ! [X4] : (~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ (X3 = X4) | ~ in(X4, powerset(the_carrier(X0))))) & (? [X5] : (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X3 = X5) & in(X5, powerset(the_carrier(X0)))) | ~ in(X3, sK11(X0, X1))))), introduced(choice_axiom, [])).
fof(f103, plain, ! [X3, X1, X0] : (? [X5] : (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X3 = X5) & in(X5, powerset(the_carrier(X0)))) => (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (sK12(X0, X1, X3) = X3) & in(sK12(X0, X1, X3), powerset(the_carrier(X0))))), introduced(choice_axiom, [])).
fof(f101, plain, ! [X0, X1] : (? [X2] : ! [X3] : ((in(X3, X2) | ! [X4] : (~ in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ (X3 = X4) | ~ in(X4, powerset(the_carrier(X0))))) & (? [X5] : (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X3 = X5) & in(X5, powerset(the_carrier(X0)))) | ~ in(X3, X2))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(rectify, [], [f100])).
fof(f100, plain, ! [X0, X1] : (? [X5] : ! [X6] : ((in(X6, X5) | ! [X7] : (~ in(set_difference(cast_as_carrier_subset(X0), X6), X1) | ~ (X6 = X7) | ~ in(X7, powerset(the_carrier(X0))))) & (? [X7] : (in(set_difference(cast_as_carrier_subset(X0), X6), X1) & (X6 = X7) & in(X7, powerset(the_carrier(X0)))) | ~ in(X6, X5))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(nnf_transformation, [], [f82])).
fof(f82, plain, ! [X0, X1] : (? [X5] : ! [X6] : (in(X6, X5) <=> ? [X7] : (in(set_difference(cast_as_carrier_subset(X0), X6), X1) & (X6 = X7) & in(X7, powerset(the_carrier(X0))))) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(definition_folding, [], [f80, e81])).
fof(f81, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)) | ~ sP0(X1, X0)), inference(usedef, [], [e81])).
fof(e81, plain, ! [X1, X0] : (sP0(X1, X0) <=> ? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f80, plain, ! [X0, X1] : (? [X5] : ! [X6] : (in(X6, X5) <=> ? [X7] : (in(set_difference(cast_as_carrier_subset(X0), X6), X1) & (X6 = X7) & in(X7, powerset(the_carrier(X0))))) | ? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(flattening, [], [f79])).
fof(f79, plain, ! [X0, X1] : ((? [X5] : ! [X6] : (in(X6, X5) <=> ? [X7] : (in(set_difference(cast_as_carrier_subset(X0), X6), X1) & (X6 = X7) & in(X7, powerset(the_carrier(X0))))) | ? [X2, X3, X4] : (~ (X3 = X4) & (in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)))) | (~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0))), inference(ennf_transformation, [], [f38])).
fof(f38, plain, ! [X0, X1] : ((element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)) => (! [X2, X3, X4] : ((in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)) => (X3 = X4)) => ? [X5] : ! [X6] : (in(X6, X5) <=> ? [X7] : (in(set_difference(cast_as_carrier_subset(X0), X6), X1) & (X6 = X7) & in(X7, powerset(the_carrier(X0))))))), inference(rectify, [], [f37])).
fof(f37, plain, ! [X0, X1] : ((element(X1, powerset(powerset(the_carrier(X0)))) & top_str(X0) & topological_space(X0)) => (! [X2, X3, X4] : ((in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)) => (X3 = X4)) => ? [X2] : ! [X3] : (in(X3, X2) <=> ? [X4] : (in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X3 = X4) & in(X4, powerset(the_carrier(X0))))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU310+1.p', s1_tarski__e2_37_1_1__pre_topc__1)).
fof(f433, plain, (~ in(sK3(sK11(sK1, sK2)), sK11(sK1, sK2)) | (spl13_7 | ~ spl13_9 | ~ spl13_12)), inference(subsumption_resolution, [], [f432, f386])).
fof(f386, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ spl13_12), inference(avatar_component_clause, [], [f384])).
fof(f384, plain, (spl13_12 <=> in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1)))), introduced(avatar_definition, [new_symbols(naming, [spl13_12])])).
fof(f432, plain, (~ in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ in(sK3(sK11(sK1, sK2)), sK11(sK1, sK2)) | (spl13_7 | ~ spl13_9)), inference(resolution, [], [f389, f365])).
fof(f389, plain, (! [X0] : (~ in(sK3(X0), sK11(sK1, sK2)) | ~ in(sK3(X0), powerset(the_carrier(sK1))) | ~ in(sK3(X0), X0)) | ~ spl13_9), inference(resolution, [], [f352, f110])).
fof(f110, plain, ! [X2] : (~ in(set_difference(cast_as_carrier_subset(sK1), sK3(X2)), sK2) | ~ in(sK3(X2), powerset(the_carrier(sK1))) | ~ in(sK3(X2), X2)), inference(cnf_transformation, [], [f87])).
fof(f352, plain, (! [X0] : (in(set_difference(cast_as_carrier_subset(sK1), X0), sK2) | ~ in(X0, sK11(sK1, sK2))) | ~ spl13_9), inference(avatar_component_clause, [], [f351])).
fof(f351, plain, (spl13_9 <=> ! [X0] : (~ in(X0, sK11(sK1, sK2)) | in(set_difference(cast_as_carrier_subset(sK1), X0), sK2))), introduced(avatar_definition, [new_symbols(naming, [spl13_9])])).
fof(f419, plain, (spl13_12 | spl13_7 | ~ spl13_11), inference(avatar_split_clause, [], [f418, f379, f328, f384])).
fof(f379, plain, (spl13_11 <=> (sK3(sK11(sK1, sK2)) = sK12(sK1, sK2, sK3(sK11(sK1, sK2))))), introduced(avatar_definition, [new_symbols(naming, [spl13_11])])).
fof(f418, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | (spl13_7 | ~ spl13_11)), inference(subsumption_resolution, [], [f417, f105])).
fof(f417, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ topological_space(sK1) | (spl13_7 | ~ spl13_11)), inference(subsumption_resolution, [], [f416, f106])).
fof(f416, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ top_str(sK1) | ~ topological_space(sK1) | (spl13_7 | ~ spl13_11)), inference(subsumption_resolution, [], [f415, f107])).
fof(f415, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ element(sK2, powerset(powerset(the_carrier(sK1)))) | ~ top_str(sK1) | ~ topological_space(sK1) | (spl13_7 | ~ spl13_11)), inference(subsumption_resolution, [], [f414, f329])).
fof(f414, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | sP0(sK2, sK1) | ~ element(sK2, powerset(powerset(the_carrier(sK1)))) | ~ top_str(sK1) | ~ topological_space(sK1) | (spl13_7 | ~ spl13_11)), inference(subsumption_resolution, [], [f408, f365])).
fof(f408, plain, (in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ in(sK3(sK11(sK1, sK2)), sK11(sK1, sK2)) | sP0(sK2, sK1) | ~ element(sK2, powerset(powerset(the_carrier(sK1)))) | ~ top_str(sK1) | ~ topological_space(sK1) | ~ spl13_11), inference(superposition, [], [f170, f381])).
fof(f381, plain, ((sK3(sK11(sK1, sK2)) = sK12(sK1, sK2, sK3(sK11(sK1, sK2)))) | ~ spl13_11), inference(avatar_component_clause, [], [f379])).
fof(f170, plain, ! [X0, X3, X1] : (in(sK12(X0, X1, X3), powerset(the_carrier(X0))) | ~ in(X3, sK11(X0, X1)) | sP0(X1, X0) | ~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ top_str(X0) | ~ topological_space(X0)), inference(cnf_transformation, [], [f104])).
fof(f387, plain, (spl13_12 | spl13_11 | ~ spl13_8), inference(avatar_split_clause, [], [f370, f332, f379, f384])).
fof(f332, plain, (spl13_8 <=> ! [X0] : (~ in(X0, sK11(sK1, sK2)) | (sK12(sK1, sK2, X0) = X0))), introduced(avatar_definition, [new_symbols(naming, [spl13_8])])).
fof(f370, plain, ((sK3(sK11(sK1, sK2)) = sK12(sK1, sK2, sK3(sK11(sK1, sK2)))) | in(sK3(sK11(sK1, sK2)), powerset(the_carrier(sK1))) | ~ spl13_8), inference(resolution, [], [f333, f108])).
fof(f333, plain, (! [X0] : (~ in(X0, sK11(sK1, sK2)) | (sK12(sK1, sK2, X0) = X0)) | ~ spl13_8), inference(avatar_component_clause, [], [f332])).
fof(f353, plain, (spl13_7 | spl13_9), inference(avatar_split_clause, [], [f349, f351, f328])).
fof(f349, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | in(set_difference(cast_as_carrier_subset(sK1), X0), sK2)), inference(subsumption_resolution, [], [f348, f105])).
fof(f348, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | in(set_difference(cast_as_carrier_subset(sK1), X0), sK2) | ~ topological_space(sK1)), inference(subsumption_resolution, [], [f338, f106])).
fof(f338, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | in(set_difference(cast_as_carrier_subset(sK1), X0), sK2) | ~ top_str(sK1) | ~ topological_space(sK1)), inference(resolution, [], [f172, f107])).
fof(f172, plain, ! [X0, X3, X1] : (~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ in(X3, sK11(X0, X1)) | sP0(X1, X0) | in(set_difference(cast_as_carrier_subset(X0), X3), X1) | ~ top_str(X0) | ~ topological_space(X0)), inference(cnf_transformation, [], [f104])).
fof(f347, plain, ~ spl13_7, inference(avatar_contradiction_clause, [], [f346])).
fof(f346, plain, ($false | ~ spl13_7), inference(subsumption_resolution, [], [f345, f330])).
fof(f330, plain, (sP0(sK2, sK1) | ~ spl13_7), inference(avatar_component_clause, [], [f328])).
fof(f345, plain, (~ sP0(sK2, sK1) | ~ spl13_7), inference(subsumption_resolution, [], [f343, f337])).
fof(f337, plain, ((sK8(sK2, sK1) = sK9(sK2, sK1)) | ~ spl13_7), inference(resolution, [], [f330, f165])).
fof(f165, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK8(X0, X1) = sK9(X0, X1))), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ! [X0, X1] : ((~ (sK9(X0, X1) = sK10(X0, X1)) & in(set_difference(cast_as_carrier_subset(X1), sK10(X0, X1)), X0) & (sK8(X0, X1) = sK10(X0, X1)) & in(set_difference(cast_as_carrier_subset(X1), sK9(X0, X1)), X0) & (sK8(X0, X1) = sK9(X0, X1))) | ~ sP0(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8, sK9, sK10])], [f97, f98])).
fof(f98, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X1), X4), X0) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X1), X3), X0) & (X2 = X3)) => (~ (sK9(X0, X1) = sK10(X0, X1)) & in(set_difference(cast_as_carrier_subset(X1), sK10(X0, X1)), X0) & (sK8(X0, X1) = sK10(X0, X1)) & in(set_difference(cast_as_carrier_subset(X1), sK9(X0, X1)), X0) & (sK8(X0, X1) = sK9(X0, X1)))), introduced(choice_axiom, [])).
fof(f97, plain, ! [X0, X1] : (? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X1), X4), X0) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X1), X3), X0) & (X2 = X3)) | ~ sP0(X0, X1)), inference(rectify, [], [f96])).
fof(f96, plain, ! [X1, X0] : (? [X2, X3, X4] : (~ (X3 = X4) & in(set_difference(cast_as_carrier_subset(X0), X4), X1) & (X2 = X4) & in(set_difference(cast_as_carrier_subset(X0), X3), X1) & (X2 = X3)) | ~ sP0(X1, X0)), inference(nnf_transformation, [], [f81])).
fof(f343, plain, (~ (sK8(sK2, sK1) = sK9(sK2, sK1)) | ~ sP0(sK2, sK1) | ~ spl13_7), inference(superposition, [], [f169, f336])).
fof(f336, plain, ((sK8(sK2, sK1) = sK10(sK2, sK1)) | ~ spl13_7), inference(resolution, [], [f330, f167])).
fof(f167, plain, ! [X0, X1] : (~ sP0(X0, X1) | (sK8(X0, X1) = sK10(X0, X1))), inference(cnf_transformation, [], [f99])).
fof(f169, plain, ! [X0, X1] : (~ (sK9(X0, X1) = sK10(X0, X1)) | ~ sP0(X0, X1)), inference(cnf_transformation, [], [f99])).
fof(f334, plain, (spl13_7 | spl13_8), inference(avatar_split_clause, [], [f326, f332, f328])).
fof(f326, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | (sK12(sK1, sK2, X0) = X0)), inference(subsumption_resolution, [], [f325, f105])).
fof(f325, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | (sK12(sK1, sK2, X0) = X0) | ~ topological_space(sK1)), inference(subsumption_resolution, [], [f322, f106])).
fof(f322, plain, ! [X0] : (~ in(X0, sK11(sK1, sK2)) | sP0(sK2, sK1) | (sK12(sK1, sK2, X0) = X0) | ~ top_str(sK1) | ~ topological_space(sK1)), inference(resolution, [], [f171, f107])).
fof(f171, plain, ! [X0, X3, X1] : (~ element(X1, powerset(powerset(the_carrier(X0)))) | ~ in(X3, sK11(X0, X1)) | sP0(X1, X0) | (sK12(X0, X1, X3) = X3) | ~ top_str(X0) | ~ topological_space(X0)), inference(cnf_transformation, [], [f104])).