fof(f110997, plain, $false, inference(avatar_sat_refutation, [], [f68447, f68764, f68769, f68781, f70065, f70066, f70067, f70068, f70069, f70070, f80055, f80146, f106723, f107143, f107163, f107292, f107317, f107319, f109851, f110127, f110729, f110986])).
fof(f110986, plain, (spl40_1779 | ~ spl40_1781 | ~ spl40_2173), inference(avatar_contradiction_clause, [], [f110985])).
fof(f110985, plain, ($false | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110984, f251])).
fof(f251, plain, ~ empty_carrier(sK3), inference(cnf_transformation, [], [f189])).
fof(f189, plain, (! [X3] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (sK6(X3) = X6) | ~ (topstr_closure(sK3, X5) = sK5) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(sK6(X3), the_carrier(sK4)) | ~ in(sK6(X3), X3)) & (((((sK7(X3) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(X3))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(X3))))) & (sK6(X3) = sK8(X3)) & (sK5 = topstr_closure(sK3, sK7(X3))) & element(sK8(X3), the_carrier(sK4))) & netstr_induced_subset(sK7(X3), sK3, sK4)) & in(sK6(X3), the_carrier(sK4))) | in(sK6(X3), X3))) & net_str(sK4, sK3) & directed_relstr(sK4) & transitive_relstr(sK4) & ~ empty_carrier(sK4) & top_str(sK3) & topological_space(sK3) & ~ empty_carrier(sK3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5, sK6, sK7, sK8])], [f183, f188, f187, f186, f185, f184])).
fof(f184, plain, (? [X0, X1] : (? [X2] : ! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X6)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(X0, X5) = X2) | ~ element(X6, the_carrier(X1))) | ~ netstr_induced_subset(X5, X0, X1)) | ~ in(X4, the_carrier(X1)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X4 = X8) & (topstr_closure(X0, X7) = X2) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & in(X4, the_carrier(X1))) | in(X4, X3))) & net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => (? [X2] : ! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(sK3, X5) = X2) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(X4, the_carrier(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (X4 = X8) & (topstr_closure(sK3, X7) = X2) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) & in(X4, the_carrier(sK4))) | in(X4, X3))) & net_str(sK4, sK3) & directed_relstr(sK4) & transitive_relstr(sK4) & ~ empty_carrier(sK4) & top_str(sK3) & topological_space(sK3) & ~ empty_carrier(sK3))), introduced(choice_axiom, [])).
fof(f185, plain, (? [X2] : ! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(sK3, X5) = X2) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(X4, the_carrier(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (X4 = X8) & (topstr_closure(sK3, X7) = X2) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) & in(X4, the_carrier(sK4))) | in(X4, X3))) => ! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(sK3, X5) = sK5) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(X4, the_carrier(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (X4 = X8) & (topstr_closure(sK3, X7) = sK5) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) & in(X4, the_carrier(sK4))) | in(X4, X3)))), introduced(choice_axiom, [])).
fof(f186, plain, ! [X3] : (? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(sK3, X5) = sK5) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(X4, the_carrier(sK4)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (X4 = X8) & (topstr_closure(sK3, X7) = sK5) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) & in(X4, the_carrier(sK4))) | in(X4, X3))) => ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (sK6(X3) = X6) | ~ (topstr_closure(sK3, X5) = sK5) | ~ element(X6, the_carrier(sK4))) | ~ netstr_induced_subset(X5, sK3, sK4)) | ~ in(sK6(X3), the_carrier(sK4)) | ~ in(sK6(X3), X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (sK6(X3) = X8) & (topstr_closure(sK3, X7) = sK5) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) & in(sK6(X3), the_carrier(sK4))) | in(sK6(X3), X3)))), introduced(choice_axiom, [])).
fof(f187, plain, ! [X3] : (? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = X7) & (sK6(X3) = X8) & (topstr_closure(sK3, X7) = sK5) & element(X8, the_carrier(sK4))) & netstr_induced_subset(X7, sK3, sK4)) => (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = sK7(X3)) & (sK6(X3) = X8) & (sK5 = topstr_closure(sK3, sK7(X3))) & element(X8, the_carrier(sK4))) & netstr_induced_subset(sK7(X3), sK3, sK4))), introduced(choice_axiom, [])).
fof(f188, plain, ! [X3] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X8)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X8))) = sK7(X3)) & (sK6(X3) = X8) & (sK5 = topstr_closure(sK3, sK7(X3))) & element(X8, the_carrier(sK4))) => ((sK7(X3) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(X3))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(X3))))) & (sK6(X3) = sK8(X3)) & (sK5 = topstr_closure(sK3, sK7(X3))) & element(sK8(X3), the_carrier(sK4)))), introduced(choice_axiom, [])).
fof(f183, plain, ? [X0, X1] : (? [X2] : ! [X3] : ? [X4] : ((! [X5] : (! [X6] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X6)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X6))) = X5) | ~ (X4 = X6) | ~ (topstr_closure(X0, X5) = X2) | ~ element(X6, the_carrier(X1))) | ~ netstr_induced_subset(X5, X0, X1)) | ~ in(X4, the_carrier(X1)) | ~ in(X4, X3)) & ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X4 = X8) & (topstr_closure(X0, X7) = X2) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & in(X4, the_carrier(X1))) | in(X4, X3))) & net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)), inference(rectify, [], [f182])).
fof(f182, plain, ? [X0, X1] : (? [X3] : ! [X4] : ? [X5] : ((! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) | ~ (X5 = X7) | ~ (topstr_closure(X0, X6) = X3) | ~ element(X7, the_carrier(X1))) | ~ netstr_induced_subset(X6, X0, X1)) | ~ in(X5, the_carrier(X1)) | ~ in(X5, X4)) & ((? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1))) | in(X5, X4))) & net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)), inference(flattening, [], [f181])).
fof(f181, plain, ? [X0, X1] : (? [X3] : ! [X4] : ? [X5] : (((! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) | ~ (X5 = X7) | ~ (topstr_closure(X0, X6) = X3) | ~ element(X7, the_carrier(X1))) | ~ netstr_induced_subset(X6, X0, X1)) | ~ in(X5, the_carrier(X1))) | ~ in(X5, X4)) & ((? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1))) | in(X5, X4))) & net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)), inference(nnf_transformation, [], [f89])).
fof(f89, plain, ? [X0, X1] : (? [X3] : ! [X4] : ? [X5] : ~ (in(X5, X4) <=> (? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1)))) & net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)), inference(flattening, [], [f88])).
fof(f88, plain, ? [X0, X1] : (? [X3] : ! [X4] : ? [X5] : ~ (in(X5, X4) <=> (? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1)))) & (net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0))), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ~ ! [X0, X1] : ((net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => ! [X3] : ? [X4] : ! [X5] : (in(X5, X4) <=> (? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1))))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ! [X0, X1, X2] : ((net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => ! [X3] : ? [X4] : ! [X5] : (in(X5, X4) <=> (? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0, X1, X2] : ((net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => ! [X3] : ? [X4] : ! [X5] : (in(X5, X4) <=> (? [X6] : (? [X7] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X7)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X7))) = X6) & (X5 = X7) & (topstr_closure(X0, X6) = X3) & element(X7, the_carrier(X1))) & netstr_induced_subset(X6, X0, X1)) & in(X5, the_carrier(X1))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU399+1.p', s1_xboole_0__e6_39_3__yellow19__1)).
fof(f110984, plain, (empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110983, f252])).
fof(f252, plain, topological_space(sK3), inference(cnf_transformation, [], [f189])).
fof(f110983, plain, (~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110982, f253])).
fof(f253, plain, top_str(sK3), inference(cnf_transformation, [], [f189])).
fof(f110982, plain, (~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110981, f254])).
fof(f254, plain, ~ empty_carrier(sK4), inference(cnf_transformation, [], [f189])).
fof(f110981, plain, (empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110980, f255])).
fof(f255, plain, transitive_relstr(sK4), inference(cnf_transformation, [], [f189])).
fof(f110980, plain, (~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110979, f256])).
fof(f256, plain, directed_relstr(sK4), inference(cnf_transformation, [], [f189])).
fof(f110979, plain, (~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110978, f257])).
fof(f257, plain, net_str(sK4, sK3), inference(cnf_transformation, [], [f189])).
fof(f110978, plain, (~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | ~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110977, f67873])).
fof(f67873, plain, (~ sP2(sK4, sK3, sK5) | spl40_1779), inference(avatar_component_clause, [], [f67872])).
fof(f67872, plain, (spl40_1779 <=> sP2(sK4, sK3, sK5)), introduced(avatar_definition, [new_symbols(naming, [spl40_1779])])).
fof(f110977, plain, (sP2(sK4, sK3, sK5) | ~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (~ spl40_1781 | ~ spl40_2173)), inference(subsumption_resolution, [], [f110974, f67889])).
fof(f67889, plain, ((sK37(sK3, sK4, sK5) = sK38(sK3, sK4, sK5)) | ~ spl40_1781), inference(avatar_component_clause, [], [f67888])).
fof(f67888, plain, (spl40_1781 <=> (sK37(sK3, sK4, sK5) = sK38(sK3, sK4, sK5))), introduced(avatar_definition, [new_symbols(naming, [spl40_1781])])).
fof(f110974, plain, (~ (sK37(sK3, sK4, sK5) = sK38(sK3, sK4, sK5)) | sP2(sK4, sK3, sK5) | ~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | ~ spl40_2173), inference(superposition, [], [f410, f80141])).
fof(f80141, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ spl40_2173), inference(avatar_component_clause, [], [f80139])).
fof(f80139, plain, (spl40_2173 <=> (sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5))), introduced(avatar_definition, [new_symbols(naming, [spl40_2173])])).
fof(f410, plain, ! [X2, X0, X1] : (~ (sK38(X0, X1, X2) = sK39(X0, X1, X2)) | sP2(X1, X0, X2) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f250])).
fof(f250, plain, ! [X0, X1] : (! [X2] : (sP2(X1, X0, X2) | (~ (sK38(X0, X1, X2) = sK39(X0, X1, X2)) & sP1(X1, X0, sK39(X0, X1, X2), X2) & (sK37(X0, X1, X2) = sK39(X0, X1, X2)) & sP0(X1, X0, sK38(X0, X1, X2), X2) & (sK37(X0, X1, X2) = sK38(X0, X1, X2)))) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK37, sK38, sK39])], [f248, f249])).
fof(f249, plain, ! [X2, X1, X0] : (? [X3, X4, X5] : (~ (X4 = X5) & sP1(X1, X0, X5, X2) & (X3 = X5) & sP0(X1, X0, X4, X2) & (X3 = X4)) => (~ (sK38(X0, X1, X2) = sK39(X0, X1, X2)) & sP1(X1, X0, sK39(X0, X1, X2), X2) & (sK37(X0, X1, X2) = sK39(X0, X1, X2)) & sP0(X1, X0, sK38(X0, X1, X2), X2) & (sK37(X0, X1, X2) = sK38(X0, X1, X2)))), introduced(choice_axiom, [])).
fof(f248, plain, ! [X0, X1] : (! [X2] : (sP2(X1, X0, X2) | ? [X3, X4, X5] : (~ (X4 = X5) & sP1(X1, X0, X5, X2) & (X3 = X5) & sP0(X1, X0, X4, X2) & (X3 = X4))) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(rectify, [], [f180])).
fof(f180, plain, ! [X0, X1] : (! [X3] : (sP2(X1, X0, X3) | ? [X4, X5, X6] : (~ (X5 = X6) & sP1(X1, X0, X6, X3) & (X4 = X6) & sP0(X1, X0, X5, X3) & (X4 = X5))) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(definition_folding, [], [f176, e179, e178, e177])).
fof(f177, plain, ! [X1, X0, X5, X3] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X5 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1)) | ~ sP0(X1, X0, X5, X3)), inference(usedef, [], [e177])).
fof(e177, plain, ! [X1, X0, X5, X3] : (sP0(X1, X0, X5, X3) <=> ? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X5 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f178, plain, ! [X1, X0, X6, X3] : (? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X6 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) | ~ sP1(X1, X0, X6, X3)), inference(usedef, [], [e178])).
fof(e178, plain, ! [X1, X0, X6, X3] : (sP1(X1, X0, X6, X3) <=> ? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X6 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f179, plain, ! [X1, X0, X3] : (? [X11] : ! [X12] : (in(X12, X11) <=> ? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1)))) | ~ sP2(X1, X0, X3)), inference(usedef, [], [e179])).
fof(e179, plain, ! [X1, X0, X3] : (sP2(X1, X0, X3) <=> ? [X11] : ! [X12] : (in(X12, X11) <=> ? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f176, plain, ! [X0, X1] : (! [X3] : (? [X11] : ! [X12] : (in(X12, X11) <=> ? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1)))) | ? [X4, X5, X6] : (~ (X5 = X6) & ? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X6 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & (X4 = X6) & ? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X5 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1)) & (X4 = X5))) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(flattening, [], [f175])).
fof(f175, plain, ! [X0, X1] : (! [X3] : (? [X11] : ! [X12] : (in(X12, X11) <=> ? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1)))) | ? [X4, X5, X6] : (~ (X5 = X6) & (? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X6 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & (X4 = X6) & ? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X5 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1)) & (X4 = X5)))) | (~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0))), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ! [X0, X1] : ((net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => ! [X3] : (! [X4, X5, X6] : ((? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X6 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & (X4 = X6) & ? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X5 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1)) & (X4 = X5)) => (X5 = X6)) => ? [X11] : ! [X12] : (in(X12, X11) <=> ? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1)))))), inference(rectify, [], [f70])).
fof(f70, plain, ! [X0, X1, X2] : ((net_str(X1, X0) & directed_relstr(X1) & transitive_relstr(X1) & ~ empty_carrier(X1) & top_str(X0) & topological_space(X0) & ~ empty_carrier(X0)) => ! [X3] : (! [X4, X5, X6] : ((? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X10)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X10))) = X9) & (X6 = X10) & (topstr_closure(X0, X9) = X3) & element(X10, the_carrier(X1))) & netstr_induced_subset(X9, X0, X1)) & (X4 = X6) & ? [X7] : (? [X8] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X8)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X8))) = X7) & (X5 = X8) & (topstr_closure(X0, X7) = X3) & element(X8, the_carrier(X1))) & netstr_induced_subset(X7, X0, X1)) & (X4 = X5)) => (X5 = X6)) => ? [X4] : ! [X5] : (in(X5, X4) <=> ? [X6] : (? [X11] : (? [X12] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X12)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X12))) = X11) & (X5 = X12) & (topstr_closure(X0, X11) = X3) & element(X12, the_carrier(X1))) & netstr_induced_subset(X11, X0, X1)) & (X5 = X6) & in(X6, the_carrier(X1)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU399+1.p', s1_tarski__e6_39_3__yellow19__1)).
fof(f110729, plain, (spl40_2173 | spl40_1779), inference(avatar_split_clause, [], [f110728, f67872, f80139])).
fof(f110728, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | spl40_1779), inference(subsumption_resolution, [], [f110727, f251])).
fof(f110727, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110726, f252])).
fof(f110726, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110725, f253])).
fof(f110725, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110724, f254])).
fof(f110724, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110723, f255])).
fof(f110723, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110722, f256])).
fof(f110722, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(subsumption_resolution, [], [f110714, f257])).
fof(f110714, plain, ((sK37(sK3, sK4, sK5) = sK39(sK3, sK4, sK5)) | ~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(resolution, [], [f67873, f408])).
fof(f408, plain, ! [X2, X0, X1] : (sP2(X1, X0, X2) | (sK37(X0, X1, X2) = sK39(X0, X1, X2)) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f250])).
fof(f110127, plain, (~ spl40_1779 | spl40_1772 | ~ spl40_1778 | ~ spl40_1792), inference(avatar_split_clause, [], [f110124, f68761, f67502, f67452, f67872])).
fof(f67452, plain, (spl40_1772 <=> element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl40_1772])])).
fof(f67502, plain, (spl40_1778 <=> in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5))), introduced(avatar_definition, [new_symbols(naming, [spl40_1778])])).
fof(f68761, plain, (spl40_1792 <=> (sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1792])])).
fof(f110124, plain, (~ sP2(sK4, sK3, sK5) | (spl40_1772 | ~ spl40_1778 | ~ spl40_1792)), inference(subsumption_resolution, [], [f110123, f67504])).
fof(f67504, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ spl40_1778), inference(avatar_component_clause, [], [f67502])).
fof(f110123, plain, (~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ sP2(sK4, sK3, sK5) | (spl40_1772 | ~ spl40_1792)), inference(subsumption_resolution, [], [f110113, f67454])).
fof(f67454, plain, (~ element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | spl40_1772), inference(avatar_component_clause, [], [f67452])).
fof(f110113, plain, (element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1792), inference(superposition, [], [f391, f68763])).
fof(f68763, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | ~ spl40_1792), inference(avatar_component_clause, [], [f68761])).
fof(f391, plain, ! [X4, X2, X0, X1] : (element(sK32(X0, X1, X2, X4), the_carrier(X0)) | ~ in(X4, sK29(X0, X1, X2)) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ! [X0, X1, X2] : (! [X4] : ((in(X4, sK29(X0, X1, X2)) | ! [X5] : (! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))) = X6) | ~ (X4 = X7) | ~ (topstr_closure(X1, X6) = X2) | ~ element(X7, the_carrier(X0))) | ~ netstr_induced_subset(X6, X1, X0)) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)))) & (((((sK31(X0, X1, X2, X4) = relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))))) & (sK32(X0, X1, X2, X4) = X4) & (topstr_closure(X1, sK31(X0, X1, X2, X4)) = X2) & element(sK32(X0, X1, X2, X4), the_carrier(X0))) & netstr_induced_subset(sK31(X0, X1, X2, X4), X1, X0)) & (sK30(X0, X1, X2, X4) = X4) & in(sK30(X0, X1, X2, X4), the_carrier(X0))) | ~ in(X4, sK29(X0, X1, X2)))) | ~ sP2(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK29, sK30, sK31, sK32])], [f232, f236, f235, f234, f233])).
fof(f233, plain, ! [X2, X1, X0] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))) = X6) | ~ (X4 = X7) | ~ (topstr_closure(X1, X6) = X2) | ~ element(X7, the_carrier(X0))) | ~ netstr_induced_subset(X6, X1, X0)) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)))) & (? [X8] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) & (X4 = X8) & in(X8, the_carrier(X0))) | ~ in(X4, X3))) => ! [X4] : ((in(X4, sK29(X0, X1, X2)) | ! [X5] : (! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))) = X6) | ~ (X4 = X7) | ~ (topstr_closure(X1, X6) = X2) | ~ element(X7, the_carrier(X0))) | ~ netstr_induced_subset(X6, X1, X0)) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)))) & (? [X8] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) & (X4 = X8) & in(X8, the_carrier(X0))) | ~ in(X4, sK29(X0, X1, X2))))), introduced(choice_axiom, [])).
fof(f234, plain, ! [X4, X2, X1, X0] : (? [X8] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) & (X4 = X8) & in(X8, the_carrier(X0))) => (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) & (sK30(X0, X1, X2, X4) = X4) & in(sK30(X0, X1, X2, X4), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f235, plain, ! [X4, X2, X1, X0] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) => (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = sK31(X0, X1, X2, X4)) & (X4 = X10) & (topstr_closure(X1, sK31(X0, X1, X2, X4)) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(sK31(X0, X1, X2, X4), X1, X0))), introduced(choice_axiom, [])).
fof(f236, plain, ! [X4, X2, X1, X0] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = sK31(X0, X1, X2, X4)) & (X4 = X10) & (topstr_closure(X1, sK31(X0, X1, X2, X4)) = X2) & element(X10, the_carrier(X0))) => ((sK31(X0, X1, X2, X4) = relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))))) & (sK32(X0, X1, X2, X4) = X4) & (topstr_closure(X1, sK31(X0, X1, X2, X4)) = X2) & element(sK32(X0, X1, X2, X4), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f232, plain, ! [X0, X1, X2] : (? [X3] : ! [X4] : ((in(X4, X3) | ! [X5] : (! [X6] : (! [X7] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))) = X6) | ~ (X4 = X7) | ~ (topstr_closure(X1, X6) = X2) | ~ element(X7, the_carrier(X0))) | ~ netstr_induced_subset(X6, X1, X0)) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)))) & (? [X8] : (? [X9] : (? [X10] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X10)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X10))) = X9) & (X4 = X10) & (topstr_closure(X1, X9) = X2) & element(X10, the_carrier(X0))) & netstr_induced_subset(X9, X1, X0)) & (X4 = X8) & in(X8, the_carrier(X0))) | ~ in(X4, X3))) | ~ sP2(X0, X1, X2)), inference(rectify, [], [f231])).
fof(f231, plain, ! [X1, X0, X3] : (? [X11] : ! [X12] : ((in(X12, X11) | ! [X13] : (! [X14] : (! [X15] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) | ~ (X12 = X15) | ~ (topstr_closure(X0, X14) = X3) | ~ element(X15, the_carrier(X1))) | ~ netstr_induced_subset(X14, X0, X1)) | ~ (X12 = X13) | ~ in(X13, the_carrier(X1)))) & (? [X13] : (? [X14] : (? [X15] : ((relation_rng_as_subset(the_carrier(subnetstr_of_element(X0, X1, X15)), the_carrier(X0), the_mapping(X0, subnetstr_of_element(X0, X1, X15))) = X14) & (X12 = X15) & (topstr_closure(X0, X14) = X3) & element(X15, the_carrier(X1))) & netstr_induced_subset(X14, X0, X1)) & (X12 = X13) & in(X13, the_carrier(X1))) | ~ in(X12, X11))) | ~ sP2(X1, X0, X3)), inference(nnf_transformation, [], [f179])).
fof(f109851, plain, (spl40_1772 | ~ spl40_1793 | ~ spl40_1802), inference(avatar_split_clause, [], [f108621, f68825, f68766, f67452])).
fof(f68766, plain, (spl40_1793 <=> (sK6(sK29(sK4, sK3, sK5)) = sK8(sK29(sK4, sK3, sK5)))), introduced(avatar_definition, [new_symbols(naming, [spl40_1793])])).
fof(f68825, plain, (spl40_1802 <=> element(sK8(sK29(sK4, sK3, sK5)), the_carrier(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl40_1802])])).
fof(f108621, plain, (element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | (~ spl40_1793 | ~ spl40_1802)), inference(forward_demodulation, [], [f68827, f68768])).
fof(f68768, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK8(sK29(sK4, sK3, sK5))) | ~ spl40_1793), inference(avatar_component_clause, [], [f68766])).
fof(f68827, plain, (element(sK8(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ spl40_1802), inference(avatar_component_clause, [], [f68825])).
fof(f107319, plain, (spl40_1792 | ~ spl40_1778 | ~ spl40_1779), inference(avatar_split_clause, [], [f107318, f67872, f67502, f68761])).
fof(f107318, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (~ spl40_1778 | ~ spl40_1779)), inference(subsumption_resolution, [], [f80285, f67874])).
fof(f67874, plain, (sP2(sK4, sK3, sK5) | ~ spl40_1779), inference(avatar_component_clause, [], [f67872])).
fof(f80285, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1778), inference(resolution, [], [f67504, f393])).
fof(f393, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK29(X0, X1, X2)) | (sK32(X0, X1, X2, X4) = X4) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f107317, plain, (spl40_1796 | ~ spl40_1778 | ~ spl40_1779), inference(avatar_split_clause, [], [f80292, f67872, f67502, f68783])).
fof(f68783, plain, (spl40_1796 <=> (sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1796])])).
fof(f80292, plain, ((sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))))) | (~ spl40_1778 | ~ spl40_1779)), inference(subsumption_resolution, [], [f80286, f67874])).
fof(f80286, plain, ((sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))))) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1778), inference(resolution, [], [f67504, f394])).
fof(f394, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK29(X0, X1, X2)) | (sK31(X0, X1, X2, X4) = relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, sK32(X0, X1, X2, X4))))) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f107292, plain, (spl40_1794 | ~ spl40_1778 | ~ spl40_1779), inference(avatar_split_clause, [], [f107291, f67872, f67502, f68771])).
fof(f68771, plain, (spl40_1794 <=> (sK5 = topstr_closure(sK3, sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1794])])).
fof(f107291, plain, ((sK5 = topstr_closure(sK3, sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))) | (~ spl40_1778 | ~ spl40_1779)), inference(subsumption_resolution, [], [f80284, f67874])).
fof(f80284, plain, ((sK5 = topstr_closure(sK3, sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1778), inference(resolution, [], [f67504, f392])).
fof(f392, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK29(X0, X1, X2)) | (topstr_closure(X1, sK31(X0, X1, X2, X4)) = X2) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f107163, plain, (spl40_1750 | ~ spl40_1778 | ~ spl40_1779), inference(avatar_split_clause, [], [f107162, f67872, f67502, f63418])).
fof(f63418, plain, (spl40_1750 <=> (sK6(sK29(sK4, sK3, sK5)) = sK30(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1750])])).
fof(f107162, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK30(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (~ spl40_1778 | ~ spl40_1779)), inference(subsumption_resolution, [], [f80283, f67874])).
fof(f80283, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK30(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1778), inference(resolution, [], [f67504, f389])).
fof(f389, plain, ! [X4, X2, X0, X1] : (~ in(X4, sK29(X0, X1, X2)) | (sK30(X0, X1, X2, X4) = X4) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f107143, plain, (spl40_1800 | ~ spl40_1750 | ~ spl40_1778 | ~ spl40_1779), inference(avatar_split_clause, [], [f107142, f67872, f67502, f63418, f68810])).
fof(f68810, plain, (spl40_1800 <=> in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4))), introduced(avatar_definition, [new_symbols(naming, [spl40_1800])])).
fof(f107142, plain, (in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | (~ spl40_1750 | ~ spl40_1778 | ~ spl40_1779)), inference(subsumption_resolution, [], [f107141, f67874])).
fof(f107141, plain, (in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ sP2(sK4, sK3, sK5) | (~ spl40_1750 | ~ spl40_1778)), inference(subsumption_resolution, [], [f70553, f67504])).
fof(f70553, plain, (in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ sP2(sK4, sK3, sK5) | ~ spl40_1750), inference(superposition, [], [f388, f63420])).
fof(f63420, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK30(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | ~ spl40_1750), inference(avatar_component_clause, [], [f63418])).
fof(f388, plain, ! [X4, X2, X0, X1] : (in(sK30(X0, X1, X2, X4), the_carrier(X0)) | ~ in(X4, sK29(X0, X1, X2)) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f106723, plain, (~ spl40_1772 | ~ spl40_1778 | ~ spl40_1779 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167), inference(avatar_contradiction_clause, [], [f106722])).
fof(f106722, plain, ($false | (~ spl40_1772 | ~ spl40_1778 | ~ spl40_1779 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106721, f67874])).
fof(f106721, plain, (~ sP2(sK4, sK3, sK5) | (~ spl40_1772 | ~ spl40_1778 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106720, f67504])).
fof(f106720, plain, (~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ sP2(sK4, sK3, sK5) | (~ spl40_1772 | ~ spl40_1778 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167)), inference(resolution, [], [f106663, f390])).
fof(f390, plain, ! [X4, X2, X0, X1] : (netstr_induced_subset(sK31(X0, X1, X2, X4), X1, X0) | ~ in(X4, sK29(X0, X1, X2)) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f106663, plain, (~ netstr_induced_subset(sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))), sK3, sK4) | (~ spl40_1772 | ~ spl40_1778 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106662, f67504])).
fof(f106662, plain, (~ netstr_induced_subset(sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | (~ spl40_1772 | ~ spl40_1794 | ~ spl40_1800 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106661, f68812])).
fof(f68812, plain, (in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ spl40_1800), inference(avatar_component_clause, [], [f68810])).
fof(f106661, plain, (~ netstr_induced_subset(sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | (~ spl40_1772 | ~ spl40_1794 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106660, f67453])).
fof(f67453, plain, (element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ spl40_1772), inference(avatar_component_clause, [], [f67452])).
fof(f106660, plain, (~ element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ netstr_induced_subset(sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | (~ spl40_1794 | ~ spl40_2167)), inference(subsumption_resolution, [], [f106595, f68773])).
fof(f68773, plain, ((sK5 = topstr_closure(sK3, sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))) | ~ spl40_1794), inference(avatar_component_clause, [], [f68771])).
fof(f106595, plain, (~ (sK5 = topstr_closure(sK3, sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))) | ~ element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ netstr_induced_subset(sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ spl40_2167), inference(superposition, [], [f412, f80104])).
fof(f80104, plain, ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5))))) = sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | ~ spl40_2167), inference(avatar_component_clause, [], [f80102])).
fof(f80102, plain, (spl40_2167 <=> (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5))))) = sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), introduced(avatar_definition, [new_symbols(naming, [spl40_2167])])).
fof(f412, plain, ! [X3] : (~ (sK5 = topstr_closure(sK3, relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(X3))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(X3)))))) | ~ element(sK6(X3), the_carrier(sK4)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(X3))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(X3)))), sK3, sK4) | ~ in(sK6(X3), the_carrier(sK4)) | ~ in(sK6(X3), X3)), inference(equality_resolution, [], [f411])).
fof(f411, plain, ! [X6, X3] : (~ (sK6(X3) = X6) | ~ (sK5 = topstr_closure(sK3, relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))))) | ~ element(X6, the_carrier(sK4)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))), sK3, sK4) | ~ in(sK6(X3), the_carrier(sK4)) | ~ in(sK6(X3), X3)), inference(equality_resolution, [], [f264])).
fof(f264, plain, ! [X6, X5, X3] : (~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, X6)), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, X6))) = X5) | ~ (sK6(X3) = X6) | ~ (topstr_closure(sK3, X5) = sK5) | ~ element(X6, the_carrier(sK4)) | ~ netstr_induced_subset(X5, sK3, sK4) | ~ in(sK6(X3), the_carrier(sK4)) | ~ in(sK6(X3), X3)), inference(cnf_transformation, [], [f189])).
fof(f80146, plain, (spl40_2167 | ~ spl40_1792 | ~ spl40_1796), inference(avatar_split_clause, [], [f68882, f68783, f68761, f80102])).
fof(f68882, plain, ((relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5))))) = sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (~ spl40_1792 | ~ spl40_1796)), inference(backward_demodulation, [], [f68785, f68763])).
fof(f68785, plain, ((sK31(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5))))))) | ~ spl40_1796), inference(avatar_component_clause, [], [f68783])).
fof(f80055, plain, (spl40_1778 | ~ spl40_1772 | ~ spl40_1779 | ~ spl40_1782 | ~ spl40_1791 | ~ spl40_1793 | ~ spl40_1795 | ~ spl40_1800), inference(avatar_split_clause, [], [f80054, f68810, f68777, f68766, f68757, f67908, f67872, f67452, f67502])).
fof(f67908, plain, (spl40_1782 <=> netstr_induced_subset(sK7(sK29(sK4, sK3, sK5)), sK3, sK4)), introduced(avatar_definition, [new_symbols(naming, [spl40_1782])])).
fof(f68757, plain, (spl40_1791 <=> (sK5 = topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1791])])).
fof(f68777, plain, (spl40_1795 <=> (sK7(sK29(sK4, sK3, sK5)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5))))))), introduced(avatar_definition, [new_symbols(naming, [spl40_1795])])).
fof(f80054, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | (~ spl40_1772 | ~ spl40_1779 | ~ spl40_1782 | ~ spl40_1791 | ~ spl40_1793 | ~ spl40_1795 | ~ spl40_1800)), inference(subsumption_resolution, [], [f80053, f67874])).
fof(f80053, plain, (~ sP2(sK4, sK3, sK5) | in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | (~ spl40_1772 | ~ spl40_1782 | ~ spl40_1791 | ~ spl40_1793 | ~ spl40_1795 | ~ spl40_1800)), inference(forward_demodulation, [], [f79511, f68759])).
fof(f68759, plain, ((sK5 = topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | ~ spl40_1791), inference(avatar_component_clause, [], [f68757])).
fof(f79511, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | ~ sP2(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | (~ spl40_1772 | ~ spl40_1782 | ~ spl40_1791 | ~ spl40_1793 | ~ spl40_1795 | ~ spl40_1800)), inference(forward_demodulation, [], [f79510, f68759])).
fof(f79510, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5))))) | ~ sP2(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | (~ spl40_1772 | ~ spl40_1782 | ~ spl40_1793 | ~ spl40_1795 | ~ spl40_1800)), inference(subsumption_resolution, [], [f79509, f68812])).
fof(f79509, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5))))) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ sP2(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | (~ spl40_1772 | ~ spl40_1782 | ~ spl40_1793 | ~ spl40_1795)), inference(subsumption_resolution, [], [f79508, f67910])).
fof(f67910, plain, (netstr_induced_subset(sK7(sK29(sK4, sK3, sK5)), sK3, sK4) | ~ spl40_1782), inference(avatar_component_clause, [], [f67908])).
fof(f79508, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5))))) | ~ netstr_induced_subset(sK7(sK29(sK4, sK3, sK5)), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ sP2(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | (~ spl40_1772 | ~ spl40_1793 | ~ spl40_1795)), inference(subsumption_resolution, [], [f79444, f67453])).
fof(f79444, plain, (in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5))))) | ~ element(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ netstr_induced_subset(sK7(sK29(sK4, sK3, sK5)), sK3, sK4) | ~ in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | ~ sP2(sK4, sK3, topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | (~ spl40_1793 | ~ spl40_1795)), inference(superposition, [], [f416, f70075])).
fof(f70075, plain, ((sK7(sK29(sK4, sK3, sK5)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK6(sK29(sK4, sK3, sK5)))))) | (~ spl40_1793 | ~ spl40_1795)), inference(backward_demodulation, [], [f68779, f68768])).
fof(f68779, plain, ((sK7(sK29(sK4, sK3, sK5)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))))) | ~ spl40_1795), inference(avatar_component_clause, [], [f68777])).
fof(f416, plain, ! [X0, X7, X1] : (in(X7, sK29(X0, X1, topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))))) | ~ element(X7, the_carrier(X0)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))), X1, X0) | ~ in(X7, the_carrier(X0)) | ~ sP2(X0, X1, topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))))), inference(equality_resolution, [], [f415])).
fof(f415, plain, ! [X0, X7, X5, X1] : (in(X7, sK29(X0, X1, topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))))) | ~ element(X7, the_carrier(X0)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))), X1, X0) | ~ (X5 = X7) | ~ in(X5, the_carrier(X0)) | ~ sP2(X0, X1, topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))))), inference(equality_resolution, [], [f414])).
fof(f414, plain, ! [X2, X0, X7, X5, X1] : (in(X7, sK29(X0, X1, X2)) | ~ (topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))) = X2) | ~ element(X7, the_carrier(X0)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))), X1, X0) | ~ (X5 = X7) | ~ in(X5, the_carrier(X0)) | ~ sP2(X0, X1, X2)), inference(equality_resolution, [], [f413])).
fof(f413, plain, ! [X4, X2, X0, X7, X5, X1] : (in(X4, sK29(X0, X1, X2)) | ~ (X4 = X7) | ~ (topstr_closure(X1, relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7)))) = X2) | ~ element(X7, the_carrier(X0)) | ~ netstr_induced_subset(relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))), X1, X0) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)) | ~ sP2(X0, X1, X2)), inference(equality_resolution, [], [f395])).
fof(f395, plain, ! [X6, X4, X2, X0, X7, X5, X1] : (in(X4, sK29(X0, X1, X2)) | ~ (relation_rng_as_subset(the_carrier(subnetstr_of_element(X1, X0, X7)), the_carrier(X1), the_mapping(X1, subnetstr_of_element(X1, X0, X7))) = X6) | ~ (X4 = X7) | ~ (topstr_closure(X1, X6) = X2) | ~ element(X7, the_carrier(X0)) | ~ netstr_induced_subset(X6, X1, X0) | ~ (X4 = X5) | ~ in(X5, the_carrier(X0)) | ~ sP2(X0, X1, X2)), inference(cnf_transformation, [], [f237])).
fof(f70070, plain, (spl40_1795 | spl40_1778), inference(avatar_split_clause, [], [f70055, f67502, f68777])).
fof(f70055, plain, ((sK7(sK29(sK4, sK3, sK5)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))))) | spl40_1778), inference(resolution, [], [f67503, f263])).
fof(f263, plain, ! [X3] : (in(sK6(X3), X3) | (sK7(X3) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(X3))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(X3)))))), inference(cnf_transformation, [], [f189])).
fof(f67503, plain, (~ in(sK6(sK29(sK4, sK3, sK5)), sK29(sK4, sK3, sK5)) | spl40_1778), inference(avatar_component_clause, [], [f67502])).
fof(f70069, plain, (spl40_1793 | spl40_1778), inference(avatar_split_clause, [], [f70054, f67502, f68766])).
fof(f70054, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK8(sK29(sK4, sK3, sK5))) | spl40_1778), inference(resolution, [], [f67503, f262])).
fof(f262, plain, ! [X3] : (in(sK6(X3), X3) | (sK6(X3) = sK8(X3))), inference(cnf_transformation, [], [f189])).
fof(f70068, plain, (spl40_1791 | spl40_1778), inference(avatar_split_clause, [], [f70053, f67502, f68757])).
fof(f70053, plain, ((sK5 = topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | spl40_1778), inference(resolution, [], [f67503, f261])).
fof(f261, plain, ! [X3] : (in(sK6(X3), X3) | (sK5 = topstr_closure(sK3, sK7(X3)))), inference(cnf_transformation, [], [f189])).
fof(f70067, plain, (spl40_1802 | spl40_1778), inference(avatar_split_clause, [], [f70052, f67502, f68825])).
fof(f70052, plain, (element(sK8(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | spl40_1778), inference(resolution, [], [f67503, f260])).
fof(f260, plain, ! [X3] : (element(sK8(X3), the_carrier(sK4)) | in(sK6(X3), X3)), inference(cnf_transformation, [], [f189])).
fof(f70066, plain, (spl40_1782 | spl40_1778), inference(avatar_split_clause, [], [f70051, f67502, f67908])).
fof(f70051, plain, (netstr_induced_subset(sK7(sK29(sK4, sK3, sK5)), sK3, sK4) | spl40_1778), inference(resolution, [], [f67503, f259])).
fof(f259, plain, ! [X3] : (netstr_induced_subset(sK7(X3), sK3, sK4) | in(sK6(X3), X3)), inference(cnf_transformation, [], [f189])).
fof(f70065, plain, (spl40_1800 | spl40_1778), inference(avatar_split_clause, [], [f70050, f67502, f68810])).
fof(f70050, plain, (in(sK6(sK29(sK4, sK3, sK5)), the_carrier(sK4)) | spl40_1778), inference(resolution, [], [f67503, f258])).
fof(f258, plain, ! [X3] : (in(sK6(X3), the_carrier(sK4)) | in(sK6(X3), X3)), inference(cnf_transformation, [], [f189])).
fof(f68781, plain, (spl40_1795 | spl40_1792 | ~ spl40_1779), inference(avatar_split_clause, [], [f68717, f67872, f68761, f68777])).
fof(f68717, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (sK7(sK29(sK4, sK3, sK5)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(sK29(sK4, sK3, sK5)))))) | ~ spl40_1779), inference(resolution, [], [f67874, f808])).
fof(f808, plain, ! [X6, X4, X5] : (~ sP2(X4, X5, X6) | (sK6(sK29(X4, X5, X6)) = sK32(X4, X5, X6, sK6(sK29(X4, X5, X6)))) | (sK7(sK29(X4, X5, X6)) = relation_rng_as_subset(the_carrier(subnetstr_of_element(sK3, sK4, sK8(sK29(X4, X5, X6)))), the_carrier(sK3), the_mapping(sK3, subnetstr_of_element(sK3, sK4, sK8(sK29(X4, X5, X6))))))), inference(resolution, [], [f263, f393])).
fof(f68769, plain, (spl40_1793 | spl40_1792 | ~ spl40_1779), inference(avatar_split_clause, [], [f68713, f67872, f68761, f68766])).
fof(f68713, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (sK6(sK29(sK4, sK3, sK5)) = sK8(sK29(sK4, sK3, sK5))) | ~ spl40_1779), inference(resolution, [], [f67874, f654])).
fof(f654, plain, ! [X4, X5, X3] : (~ sP2(X3, X4, X5) | (sK6(sK29(X3, X4, X5)) = sK32(X3, X4, X5, sK6(sK29(X3, X4, X5)))) | (sK6(sK29(X3, X4, X5)) = sK8(sK29(X3, X4, X5)))), inference(resolution, [], [f393, f262])).
fof(f68764, plain, (spl40_1791 | spl40_1792 | ~ spl40_1779), inference(avatar_split_clause, [], [f68712, f67872, f68761, f68757])).
fof(f68712, plain, ((sK6(sK29(sK4, sK3, sK5)) = sK32(sK4, sK3, sK5, sK6(sK29(sK4, sK3, sK5)))) | (sK5 = topstr_closure(sK3, sK7(sK29(sK4, sK3, sK5)))) | ~ spl40_1779), inference(resolution, [], [f67874, f653])).
fof(f653, plain, ! [X2, X0, X1] : (~ sP2(X0, X1, X2) | (sK6(sK29(X0, X1, X2)) = sK32(X0, X1, X2, sK6(sK29(X0, X1, X2)))) | (sK5 = topstr_closure(sK3, sK7(sK29(X0, X1, X2))))), inference(resolution, [], [f393, f261])).
fof(f68447, plain, (spl40_1779 | spl40_1781), inference(avatar_contradiction_clause, [], [f68446])).
fof(f68446, plain, ($false | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68445, f251])).
fof(f68445, plain, (empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68444, f252])).
fof(f68444, plain, (~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68443, f253])).
fof(f68443, plain, (~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68442, f254])).
fof(f68442, plain, (empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68441, f255])).
fof(f68441, plain, (~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68440, f256])).
fof(f68440, plain, (~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68439, f257])).
fof(f68439, plain, (~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | (spl40_1779 | spl40_1781)), inference(subsumption_resolution, [], [f68430, f67890])).
fof(f67890, plain, (~ (sK37(sK3, sK4, sK5) = sK38(sK3, sK4, sK5)) | spl40_1781), inference(avatar_component_clause, [], [f67888])).
fof(f68430, plain, ((sK37(sK3, sK4, sK5) = sK38(sK3, sK4, sK5)) | ~ net_str(sK4, sK3) | ~ directed_relstr(sK4) | ~ transitive_relstr(sK4) | empty_carrier(sK4) | ~ top_str(sK3) | ~ topological_space(sK3) | empty_carrier(sK3) | spl40_1779), inference(resolution, [], [f67873, f406])).
fof(f406, plain, ! [X2, X0, X1] : (sP2(X1, X0, X2) | (sK37(X0, X1, X2) = sK38(X0, X1, X2)) | ~ net_str(X1, X0) | ~ directed_relstr(X1) | ~ transitive_relstr(X1) | empty_carrier(X1) | ~ top_str(X0) | ~ topological_space(X0) | empty_carrier(X0)), inference(cnf_transformation, [], [f250])).