fof(f6602, plain, $false, inference(avatar_sat_refutation, [], [f378, f382, f391, f396, f400, f404, f409, f413, f417, f1417, f3491, f3495, f3718, f4287, f4335, f4875, f4878, f4882, f5053, f5134, f5175, f5409, f6601])).
fof(f6601, plain, (~ spl41_5 | ~ spl41_6 | ~ spl41_28 | ~ spl41_81 | spl41_82), inference(avatar_contradiction_clause, [], [f6600])).
fof(f6600, plain, ($false | (~ spl41_5 | ~ spl41_6 | ~ spl41_28 | ~ spl41_81 | spl41_82)), inference(subsumption_resolution, [], [f6599, f390])).
fof(f390, plain, (sP2(sK34) | ~ spl41_5), inference(avatar_component_clause, [], [f388])).
fof(f388, plain, (spl41_5 <=> sP2(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl41_5])])).
fof(f6599, plain, (~ sP2(sK34) | (~ spl41_6 | ~ spl41_28 | ~ spl41_81 | spl41_82)), inference(subsumption_resolution, [], [f6598, f1390])).
fof(f1390, plain, (ordinal(sK40(sK34)) | ~ spl41_81), inference(avatar_component_clause, [], [f1389])).
fof(f1389, plain, (spl41_81 <=> ordinal(sK40(sK34))), introduced(avatar_definition, [new_symbols(naming, [spl41_81])])).
fof(f6598, plain, (~ ordinal(sK40(sK34)) | ~ sP2(sK34) | (~ spl41_6 | ~ spl41_28 | spl41_82)), inference(subsumption_resolution, [], [f6597, f1394])).
fof(f1394, plain, (~ sP0(sK40(sK34)) | spl41_82), inference(avatar_component_clause, [], [f1393])).
fof(f1393, plain, (spl41_82 <=> sP0(sK40(sK34))), introduced(avatar_definition, [new_symbols(naming, [spl41_82])])).
fof(f6597, plain, (sP0(sK40(sK34)) | ~ ordinal(sK40(sK34)) | ~ sP2(sK34) | (~ spl41_6 | ~ spl41_28)), inference(subsumption_resolution, [], [f6596, f395])).
fof(f395, plain, (in(sK34, omega) | ~ spl41_6), inference(avatar_component_clause, [], [f393])).
fof(f393, plain, (spl41_6 <=> in(sK34, omega)), introduced(avatar_definition, [new_symbols(naming, [spl41_6])])).
fof(f6596, plain, (~ in(sK34, omega) | sP0(sK40(sK34)) | ~ ordinal(sK40(sK34)) | ~ sP2(sK34) | ~ spl41_28), inference(superposition, [], [f6168, f843])).
fof(f843, plain, ((sK34 = set_union2(sK40(sK34), singleton(sK40(sK34)))) | ~ spl41_28), inference(avatar_component_clause, [], [f841])).
fof(f841, plain, (spl41_28 <=> (sK34 = set_union2(sK40(sK34), singleton(sK40(sK34))))), introduced(avatar_definition, [new_symbols(naming, [spl41_28])])).
fof(f6168, plain, ! [X0] : (~ in(set_union2(X0, singleton(X0)), omega) | sP0(X0) | ~ ordinal(X0) | ~ sP2(set_union2(X0, singleton(X0)))), inference(subsumption_resolution, [], [f6167, f936])).
fof(f936, plain, ! [X5] : (in(sK13(sK36(set_union2(X5, singleton(X5)))), sK36(set_union2(X5, singleton(X5)))) | ~ in(set_union2(X5, singleton(X5)), omega) | sP0(X5) | ~ ordinal(X5) | ~ sP2(set_union2(X5, singleton(X5)))), inference(subsumption_resolution, [], [f929, f332])).
fof(f332, plain, ! [X0] : (~ (empty_set = sK36(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f184])).
fof(f184, plain, ! [X0] : ((! [X2] : ((~ (sK37(X0, X2) = X2) & subset(X2, sK37(X0, X2)) & in(sK37(X0, X2), sK36(X0))) | ~ in(X2, sK36(X0))) & ~ (empty_set = sK36(X0)) & element(sK36(X0), powerset(powerset(X0)))) | ~ sP2(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f181, f183, f182])).
fof(f182, plain, ! [X0] : (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) => (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK36(X0))) | ~ in(X2, sK36(X0))) & ~ (empty_set = sK36(X0)) & element(sK36(X0), powerset(powerset(X0))))), introduced(choice_axiom, [])).
fof(f183, plain, ! [X2, X0] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK36(X0))) => (~ (sK37(X0, X2) = X2) & subset(X2, sK37(X0, X2)) & in(sK37(X0, X2), sK36(X0)))), introduced(choice_axiom, [])).
fof(f181, plain, ! [X0] : (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) | ~ sP2(X0)), inference(rectify, [], [f180])).
fof(f180, plain, ! [X0] : (? [X5] : (! [X6] : (? [X7] : (~ (X6 = X7) & subset(X6, X7) & in(X7, X5)) | ~ in(X6, X5)) & ~ (empty_set = X5) & element(X5, powerset(powerset(X0)))) | ~ sP2(X0)), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ! [X0] : (? [X5] : (! [X6] : (? [X7] : (~ (X6 = X7) & subset(X6, X7) & in(X7, X5)) | ~ in(X6, X5)) & ~ (empty_set = X5) & element(X5, powerset(powerset(X0)))) | ~ sP2(X0)), inference(usedef, [], [e117])).
fof(e117, plain, ! [X0] : (sP2(X0) <=> ? [X5] : (! [X6] : (? [X7] : (~ (X6 = X7) & subset(X6, X7) & in(X7, X5)) | ~ in(X6, X5)) & ~ (empty_set = X5) & element(X5, powerset(powerset(X0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f929, plain, ! [X5] : ((empty_set = sK36(set_union2(X5, singleton(X5)))) | in(sK13(sK36(set_union2(X5, singleton(X5)))), sK36(set_union2(X5, singleton(X5)))) | ~ in(set_union2(X5, singleton(X5)), omega) | sP0(X5) | ~ ordinal(X5) | ~ sP2(set_union2(X5, singleton(X5)))), inference(resolution, [], [f356, f331])).
fof(f331, plain, ! [X0] : (element(sK36(X0), powerset(powerset(X0))) | ~ sP2(X0)), inference(cnf_transformation, [], [f184])).
fof(f356, plain, ! [X8, X9] : (~ element(X9, powerset(powerset(set_union2(X8, singleton(X8))))) | (empty_set = X9) | in(sK13(X9), X9) | ~ in(set_union2(X8, singleton(X8)), omega) | sP0(X8) | ~ ordinal(X8)), inference(definition_unfolding, [], [f208, f339, f339])).
fof(f339, plain, ! [X0] : (succ(X0) = set_union2(X0, singleton(X0))), inference(cnf_transformation, [], [f51])).
fof(f51, plain, ! [X0] : (succ(X0) = set_union2(X0, singleton(X0))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', d1_ordinal1)).
fof(f208, plain, ! [X8, X9] : (in(sK13(X9), X9) | (empty_set = X9) | ~ element(X9, powerset(powerset(succ(X8)))) | ~ in(succ(X8), omega) | sP0(X8) | ~ ordinal(X8)), inference(cnf_transformation, [], [f137])).
fof(f137, plain, (((! [X2] : ((~ (sK11(X2) = X2) & subset(X2, sK11(X2)) & in(sK11(X2), sK10)) | ~ in(X2, sK10)) & ~ (empty_set = sK10) & element(sK10, powerset(powerset(sK9)))) & in(sK9, omega) & ordinal(sK9)) & ! [X4] : (! [X5] : ((! [X7] : ((sK12(X5) = X7) | ~ subset(sK12(X5), X7) | ~ in(X7, X5)) & in(sK12(X5), X5)) | (empty_set = X5) | ~ element(X5, powerset(powerset(X4)))) | ~ in(X4, omega) | (empty_set = X4) | sP1(X4) | ~ being_limit_ordinal(X4) | ~ ordinal(X4)) & ! [X8] : (! [X9] : ((! [X11] : ((sK13(X9) = X11) | ~ subset(sK13(X9), X11) | ~ in(X11, X9)) & in(sK13(X9), X9)) | (empty_set = X9) | ~ element(X9, powerset(powerset(succ(X8))))) | ~ in(succ(X8), omega) | sP0(X8) | ~ ordinal(X8)) & (! [X12] : ((! [X14] : ((sK14(X12) = X14) | ~ subset(sK14(X12), X14) | ~ in(X14, X12)) & in(sK14(X12), X12)) | (empty_set = X12) | ~ element(X12, powerset(powerset(empty_set)))) | ~ in(empty_set, omega))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11, sK12, sK13, sK14])], [f130, f136, f135, f134, f133, f132, f131])).
fof(f131, plain, (? [X0] : (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) & in(X0, omega) & ordinal(X0)) => (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(sK9)))) & in(sK9, omega) & ordinal(sK9))), introduced(choice_axiom, [])).
fof(f132, plain, (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(sK9)))) => (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK10)) | ~ in(X2, sK10)) & ~ (empty_set = sK10) & element(sK10, powerset(powerset(sK9))))), introduced(choice_axiom, [])).
fof(f133, plain, ! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK10)) => (~ (sK11(X2) = X2) & subset(X2, sK11(X2)) & in(sK11(X2), sK10))), introduced(choice_axiom, [])).
fof(f134, plain, ! [X5] : (? [X6] : (! [X7] : ((X6 = X7) | ~ subset(X6, X7) | ~ in(X7, X5)) & in(X6, X5)) => (! [X7] : ((sK12(X5) = X7) | ~ subset(sK12(X5), X7) | ~ in(X7, X5)) & in(sK12(X5), X5))), introduced(choice_axiom, [])).
fof(f135, plain, ! [X9] : (? [X10] : (! [X11] : ((X10 = X11) | ~ subset(X10, X11) | ~ in(X11, X9)) & in(X10, X9)) => (! [X11] : ((sK13(X9) = X11) | ~ subset(sK13(X9), X11) | ~ in(X11, X9)) & in(sK13(X9), X9))), introduced(choice_axiom, [])).
fof(f136, plain, ! [X12] : (? [X13] : (! [X14] : ((X13 = X14) | ~ subset(X13, X14) | ~ in(X14, X12)) & in(X13, X12)) => (! [X14] : ((sK14(X12) = X14) | ~ subset(sK14(X12), X14) | ~ in(X14, X12)) & in(sK14(X12), X12))), introduced(choice_axiom, [])).
fof(f130, plain, (? [X0] : (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) & in(X0, omega) & ordinal(X0)) & ! [X4] : (! [X5] : (? [X6] : (! [X7] : ((X6 = X7) | ~ subset(X6, X7) | ~ in(X7, X5)) & in(X6, X5)) | (empty_set = X5) | ~ element(X5, powerset(powerset(X4)))) | ~ in(X4, omega) | (empty_set = X4) | sP1(X4) | ~ being_limit_ordinal(X4) | ~ ordinal(X4)) & ! [X8] : (! [X9] : (? [X10] : (! [X11] : ((X10 = X11) | ~ subset(X10, X11) | ~ in(X11, X9)) & in(X10, X9)) | (empty_set = X9) | ~ element(X9, powerset(powerset(succ(X8))))) | ~ in(succ(X8), omega) | sP0(X8) | ~ ordinal(X8)) & (! [X12] : (? [X13] : (! [X14] : ((X13 = X14) | ~ subset(X13, X14) | ~ in(X14, X12)) & in(X13, X12)) | (empty_set = X12) | ~ element(X12, powerset(powerset(empty_set)))) | ~ in(empty_set, omega))), inference(rectify, [], [f116])).
fof(f116, plain, (? [X18] : (? [X19] : (! [X20] : (? [X21] : (~ (X20 = X21) & subset(X20, X21) & in(X21, X19)) | ~ in(X20, X19)) & ~ (empty_set = X19) & element(X19, powerset(powerset(X18)))) & in(X18, omega) & ordinal(X18)) & ! [X0] : (! [X5] : (? [X6] : (! [X7] : ((X6 = X7) | ~ subset(X6, X7) | ~ in(X7, X5)) & in(X6, X5)) | (empty_set = X5) | ~ element(X5, powerset(powerset(X0)))) | ~ in(X0, omega) | (empty_set = X0) | sP1(X0) | ~ being_limit_ordinal(X0) | ~ ordinal(X0)) & ! [X8] : (! [X12] : (? [X13] : (! [X14] : ((X13 = X14) | ~ subset(X13, X14) | ~ in(X14, X12)) & in(X13, X12)) | (empty_set = X12) | ~ element(X12, powerset(powerset(succ(X8))))) | ~ in(succ(X8), omega) | sP0(X8) | ~ ordinal(X8)) & (! [X15] : (? [X16] : (! [X17] : ((X16 = X17) | ~ subset(X16, X17) | ~ in(X17, X15)) & in(X16, X15)) | (empty_set = X15) | ~ element(X15, powerset(powerset(empty_set)))) | ~ in(empty_set, omega))), inference(definition_folding, [], [f82, e115, e114])).
fof(f114, plain, ! [X8] : ((? [X9] : (! [X10] : (? [X11] : (~ (X10 = X11) & subset(X10, X11) & in(X11, X9)) | ~ in(X10, X9)) & ~ (empty_set = X9) & element(X9, powerset(powerset(X8)))) & in(X8, omega)) | ~ sP0(X8)), inference(usedef, [], [e114])).
fof(e114, plain, ! [X8] : (sP0(X8) <=> (? [X9] : (! [X10] : (? [X11] : (~ (X10 = X11) & subset(X10, X11) & in(X11, X9)) | ~ in(X10, X9)) & ~ (empty_set = X9) & element(X9, powerset(powerset(X8)))) & in(X8, omega))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f115, plain, ! [X0] : (? [X1] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(X1)))) & in(X1, omega) & in(X1, X0) & ordinal(X1)) | ~ sP1(X0)), inference(usedef, [], [e115])).
fof(e115, plain, ! [X0] : (sP1(X0) <=> ? [X1] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(X1)))) & in(X1, omega) & in(X1, X0) & ordinal(X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f82, plain, (? [X18] : (? [X19] : (! [X20] : (? [X21] : (~ (X20 = X21) & subset(X20, X21) & in(X21, X19)) | ~ in(X20, X19)) & ~ (empty_set = X19) & element(X19, powerset(powerset(X18)))) & in(X18, omega) & ordinal(X18)) & ! [X0] : (! [X5] : (? [X6] : (! [X7] : ((X6 = X7) | ~ subset(X6, X7) | ~ in(X7, X5)) & in(X6, X5)) | (empty_set = X5) | ~ element(X5, powerset(powerset(X0)))) | ~ in(X0, omega) | (empty_set = X0) | ? [X1] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(X1)))) & in(X1, omega) & in(X1, X0) & ordinal(X1)) | ~ being_limit_ordinal(X0) | ~ ordinal(X0)) & ! [X8] : (! [X12] : (? [X13] : (! [X14] : ((X13 = X14) | ~ subset(X13, X14) | ~ in(X14, X12)) & in(X13, X12)) | (empty_set = X12) | ~ element(X12, powerset(powerset(succ(X8))))) | ~ in(succ(X8), omega) | (? [X9] : (! [X10] : (? [X11] : (~ (X10 = X11) & subset(X10, X11) & in(X11, X9)) | ~ in(X10, X9)) & ~ (empty_set = X9) & element(X9, powerset(powerset(X8)))) & in(X8, omega)) | ~ ordinal(X8)) & (! [X15] : (? [X16] : (! [X17] : ((X16 = X17) | ~ subset(X16, X17) | ~ in(X17, X15)) & in(X16, X15)) | (empty_set = X15) | ~ element(X15, powerset(powerset(empty_set)))) | ~ in(empty_set, omega))), inference(flattening, [], [f81])).
fof(f81, plain, (? [X18] : ((? [X19] : ((! [X20] : (? [X21] : (~ (X20 = X21) & (subset(X20, X21) & in(X21, X19))) | ~ in(X20, X19)) & ~ (empty_set = X19)) & element(X19, powerset(powerset(X18)))) & in(X18, omega)) & ordinal(X18)) & (! [X0] : ((((! [X5] : ((? [X6] : (! [X7] : ((X6 = X7) | (~ subset(X6, X7) | ~ in(X7, X5))) & in(X6, X5)) | (empty_set = X5)) | ~ element(X5, powerset(powerset(X0)))) | ~ in(X0, omega)) | (empty_set = X0)) | (? [X1] : (((? [X2] : ((! [X3] : (? [X4] : (~ (X3 = X4) & (subset(X3, X4) & in(X4, X2))) | ~ in(X3, X2)) & ~ (empty_set = X2)) & element(X2, powerset(powerset(X1)))) & in(X1, omega)) & in(X1, X0)) & ordinal(X1)) | ~ being_limit_ordinal(X0))) | ~ ordinal(X0)) & ! [X8] : (((! [X12] : ((? [X13] : (! [X14] : ((X13 = X14) | (~ subset(X13, X14) | ~ in(X14, X12))) & in(X13, X12)) | (empty_set = X12)) | ~ element(X12, powerset(powerset(succ(X8))))) | ~ in(succ(X8), omega)) | (? [X9] : ((! [X10] : (? [X11] : (~ (X10 = X11) & (subset(X10, X11) & in(X11, X9))) | ~ in(X10, X9)) & ~ (empty_set = X9)) & element(X9, powerset(powerset(X8)))) & in(X8, omega))) | ~ ordinal(X8)) & (! [X15] : ((? [X16] : (! [X17] : ((X16 = X17) | (~ subset(X16, X17) | ~ in(X17, X15))) & in(X16, X15)) | (empty_set = X15)) | ~ element(X15, powerset(powerset(empty_set)))) | ~ in(empty_set, omega)))), inference(ennf_transformation, [], [f67])).
fof(f67, plain, ~ ((! [X0] : (ordinal(X0) => ((! [X1] : (ordinal(X1) => (in(X1, X0) => (in(X1, omega) => ! [X2] : (element(X2, powerset(powerset(X1))) => ~ (! [X3] : ~ (! [X4] : ((subset(X3, X4) & in(X4, X2)) => (X3 = X4)) & in(X3, X2)) & ~ (empty_set = X2)))))) & being_limit_ordinal(X0)) => ((in(X0, omega) => ! [X5] : (element(X5, powerset(powerset(X0))) => ~ (! [X6] : ~ (! [X7] : ((subset(X6, X7) & in(X7, X5)) => (X6 = X7)) & in(X6, X5)) & ~ (empty_set = X5)))) | (empty_set = X0)))) & ! [X8] : (ordinal(X8) => ((in(X8, omega) => ! [X9] : (element(X9, powerset(powerset(X8))) => ~ (! [X10] : ~ (! [X11] : ((subset(X10, X11) & in(X11, X9)) => (X10 = X11)) & in(X10, X9)) & ~ (empty_set = X9)))) => (in(succ(X8), omega) => ! [X12] : (element(X12, powerset(powerset(succ(X8)))) => ~ (! [X13] : ~ (! [X14] : ((subset(X13, X14) & in(X14, X12)) => (X13 = X14)) & in(X13, X12)) & ~ (empty_set = X12)))))) & (in(empty_set, omega) => ! [X15] : (element(X15, powerset(powerset(empty_set))) => ~ (! [X16] : ~ (! [X17] : ((subset(X16, X17) & in(X17, X15)) => (X16 = X17)) & in(X16, X15)) & ~ (empty_set = X15))))) => ! [X18] : (ordinal(X18) => (in(X18, omega) => ! [X19] : (element(X19, powerset(powerset(X18))) => ~ (! [X20] : ~ (! [X21] : ((subset(X20, X21) & in(X21, X19)) => (X20 = X21)) & in(X20, X19)) & ~ (empty_set = X19)))))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ((! [X3] : (ordinal(X3) => ((! [X10] : (ordinal(X10) => (in(X10, X3) => (in(X10, omega) => ! [X11] : (element(X11, powerset(powerset(X10))) => ~ (! [X12] : ~ (! [X13] : ((subset(X12, X13) & in(X13, X11)) => (X12 = X13)) & in(X12, X11)) & ~ (empty_set = X11)))))) & being_limit_ordinal(X3)) => ((in(X3, omega) => ! [X14] : (element(X14, powerset(powerset(X3))) => ~ (! [X15] : ~ (! [X16] : ((subset(X15, X16) & in(X16, X14)) => (X15 = X16)) & in(X15, X14)) & ~ (empty_set = X14)))) | (empty_set = X3)))) & ! [X3] : (ordinal(X3) => ((in(X3, omega) => ! [X4] : (element(X4, powerset(powerset(X3))) => ~ (! [X5] : ~ (! [X6] : ((subset(X5, X6) & in(X6, X4)) => (X5 = X6)) & in(X5, X4)) & ~ (empty_set = X4)))) => (in(succ(X3), omega) => ! [X7] : (element(X7, powerset(powerset(succ(X3)))) => ~ (! [X8] : ~ (! [X9] : ((subset(X8, X9) & in(X9, X7)) => (X8 = X9)) & in(X8, X7)) & ~ (empty_set = X7)))))) & (in(empty_set, omega) => ! [X0] : (element(X0, powerset(powerset(empty_set))) => ~ (! [X1] : ~ (! [X2] : ((subset(X1, X2) & in(X2, X0)) => (X1 = X2)) & in(X1, X0)) & ~ (empty_set = X0))))) => ! [X3] : (ordinal(X3) => (in(X3, omega) => ! [X17] : (element(X17, powerset(powerset(X3))) => ~ (! [X18] : ~ (! [X19] : ((subset(X18, X19) & in(X19, X17)) => (X18 = X19)) & in(X18, X17)) & ~ (empty_set = X17)))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ((! [X3] : (ordinal(X3) => ((! [X10] : (ordinal(X10) => (in(X10, X3) => (in(X10, omega) => ! [X11] : (element(X11, powerset(powerset(X10))) => ~ (! [X12] : ~ (! [X13] : ((subset(X12, X13) & in(X13, X11)) => (X12 = X13)) & in(X12, X11)) & ~ (empty_set = X11)))))) & being_limit_ordinal(X3)) => ((in(X3, omega) => ! [X14] : (element(X14, powerset(powerset(X3))) => ~ (! [X15] : ~ (! [X16] : ((subset(X15, X16) & in(X16, X14)) => (X15 = X16)) & in(X15, X14)) & ~ (empty_set = X14)))) | (empty_set = X3)))) & ! [X3] : (ordinal(X3) => ((in(X3, omega) => ! [X4] : (element(X4, powerset(powerset(X3))) => ~ (! [X5] : ~ (! [X6] : ((subset(X5, X6) & in(X6, X4)) => (X5 = X6)) & in(X5, X4)) & ~ (empty_set = X4)))) => (in(succ(X3), omega) => ! [X7] : (element(X7, powerset(powerset(succ(X3)))) => ~ (! [X8] : ~ (! [X9] : ((subset(X8, X9) & in(X9, X7)) => (X8 = X9)) & in(X8, X7)) & ~ (empty_set = X7)))))) & (in(empty_set, omega) => ! [X0] : (element(X0, powerset(powerset(empty_set))) => ~ (! [X1] : ~ (! [X2] : ((subset(X1, X2) & in(X2, X0)) => (X1 = X2)) & in(X1, X0)) & ~ (empty_set = X0))))) => ! [X3] : (ordinal(X3) => (in(X3, omega) => ! [X17] : (element(X17, powerset(powerset(X3))) => ~ (! [X18] : ~ (! [X19] : ((subset(X18, X19) & in(X19, X17)) => (X18 = X19)) & in(X18, X17)) & ~ (empty_set = X17)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', s1_ordinal2__e18_27__finset_1)).
fof(f6167, plain, ! [X0] : (~ in(set_union2(X0, singleton(X0)), omega) | sP0(X0) | ~ ordinal(X0) | ~ sP2(set_union2(X0, singleton(X0))) | ~ in(sK13(sK36(set_union2(X0, singleton(X0)))), sK36(set_union2(X0, singleton(X0))))), inference(duplicate_literal_removal, [], [f6156])).
fof(f6156, plain, ! [X0] : (~ in(set_union2(X0, singleton(X0)), omega) | sP0(X0) | ~ ordinal(X0) | ~ sP2(set_union2(X0, singleton(X0))) | ~ in(sK13(sK36(set_union2(X0, singleton(X0)))), sK36(set_union2(X0, singleton(X0)))) | ~ sP2(set_union2(X0, singleton(X0))) | ~ in(sK13(sK36(set_union2(X0, singleton(X0)))), sK36(set_union2(X0, singleton(X0)))) | ~ sP2(set_union2(X0, singleton(X0)))), inference(resolution, [], [f1896, f333])).
fof(f333, plain, ! [X2, X0] : (in(sK37(X0, X2), sK36(X0)) | ~ in(X2, sK36(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f184])).
fof(f1896, plain, ! [X6, X7] : (~ in(sK37(X6, sK13(sK36(set_union2(X7, singleton(X7))))), sK36(set_union2(X7, singleton(X7)))) | ~ in(set_union2(X7, singleton(X7)), omega) | sP0(X7) | ~ ordinal(X7) | ~ sP2(set_union2(X7, singleton(X7))) | ~ in(sK13(sK36(set_union2(X7, singleton(X7)))), sK36(X6)) | ~ sP2(X6)), inference(subsumption_resolution, [], [f1890, f335])).
fof(f335, plain, ! [X2, X0] : (~ (sK37(X0, X2) = X2) | ~ in(X2, sK36(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f184])).
fof(f1890, plain, ! [X6, X7] : (~ in(sK37(X6, sK13(sK36(set_union2(X7, singleton(X7))))), sK36(set_union2(X7, singleton(X7)))) | (sK13(sK36(set_union2(X7, singleton(X7)))) = sK37(X6, sK13(sK36(set_union2(X7, singleton(X7)))))) | ~ in(set_union2(X7, singleton(X7)), omega) | sP0(X7) | ~ ordinal(X7) | ~ sP2(set_union2(X7, singleton(X7))) | ~ in(sK13(sK36(set_union2(X7, singleton(X7)))), sK36(X6)) | ~ sP2(X6)), inference(resolution, [], [f996, f334])).
fof(f334, plain, ! [X2, X0] : (subset(X2, sK37(X0, X2)) | ~ in(X2, sK36(X0)) | ~ sP2(X0)), inference(cnf_transformation, [], [f184])).
fof(f996, plain, ! [X10, X11] : (~ subset(sK13(sK36(set_union2(X10, singleton(X10)))), X11) | ~ in(X11, sK36(set_union2(X10, singleton(X10)))) | (sK13(sK36(set_union2(X10, singleton(X10)))) = X11) | ~ in(set_union2(X10, singleton(X10)), omega) | sP0(X10) | ~ ordinal(X10) | ~ sP2(set_union2(X10, singleton(X10)))), inference(subsumption_resolution, [], [f989, f332])).
fof(f989, plain, ! [X10, X11] : (~ subset(sK13(sK36(set_union2(X10, singleton(X10)))), X11) | ~ in(X11, sK36(set_union2(X10, singleton(X10)))) | (empty_set = sK36(set_union2(X10, singleton(X10)))) | (sK13(sK36(set_union2(X10, singleton(X10)))) = X11) | ~ in(set_union2(X10, singleton(X10)), omega) | sP0(X10) | ~ ordinal(X10) | ~ sP2(set_union2(X10, singleton(X10)))), inference(resolution, [], [f355, f331])).
fof(f355, plain, ! [X8, X11, X9] : (~ element(X9, powerset(powerset(set_union2(X8, singleton(X8))))) | ~ subset(sK13(X9), X11) | ~ in(X11, X9) | (empty_set = X9) | (sK13(X9) = X11) | ~ in(set_union2(X8, singleton(X8)), omega) | sP0(X8) | ~ ordinal(X8)), inference(definition_unfolding, [], [f209, f339, f339])).
fof(f209, plain, ! [X8, X11, X9] : ((sK13(X9) = X11) | ~ subset(sK13(X9), X11) | ~ in(X11, X9) | (empty_set = X9) | ~ element(X9, powerset(powerset(succ(X8)))) | ~ in(succ(X8), omega) | sP0(X8) | ~ ordinal(X8)), inference(cnf_transformation, [], [f137])).
fof(f5409, plain, (~ spl41_5 | ~ spl41_354 | ~ spl41_411 | ~ spl41_417), inference(avatar_contradiction_clause, [], [f5408])).
fof(f5408, plain, ($false | (~ spl41_5 | ~ spl41_354 | ~ spl41_411 | ~ spl41_417)), inference(subsumption_resolution, [], [f5407, f4883])).
fof(f4883, plain, (sP2(empty_set) | (~ spl41_5 | ~ spl41_354)), inference(backward_demodulation, [], [f390, f4326])).
fof(f4326, plain, ((empty_set = sK34) | ~ spl41_354), inference(avatar_component_clause, [], [f4324])).
fof(f4324, plain, (spl41_354 <=> (empty_set = sK34)), introduced(avatar_definition, [new_symbols(naming, [spl41_354])])).
fof(f5407, plain, (~ sP2(empty_set) | (~ spl41_411 | ~ spl41_417)), inference(subsumption_resolution, [], [f5406, f5052])).
fof(f5052, plain, (in(sK14(sK36(empty_set)), sK36(empty_set)) | ~ spl41_411), inference(avatar_component_clause, [], [f5050])).
fof(f5050, plain, (spl41_411 <=> in(sK14(sK36(empty_set)), sK36(empty_set))), introduced(avatar_definition, [new_symbols(naming, [spl41_411])])).
fof(f5406, plain, (~ in(sK14(sK36(empty_set)), sK36(empty_set)) | ~ sP2(empty_set) | ~ spl41_417), inference(duplicate_literal_removal, [], [f5405])).
fof(f5405, plain, (~ in(sK14(sK36(empty_set)), sK36(empty_set)) | ~ sP2(empty_set) | ~ in(sK14(sK36(empty_set)), sK36(empty_set)) | ~ sP2(empty_set) | ~ spl41_417), inference(resolution, [], [f5280, f333])).
fof(f5280, plain, (! [X2] : (~ in(sK37(X2, sK14(sK36(empty_set))), sK36(empty_set)) | ~ in(sK14(sK36(empty_set)), sK36(X2)) | ~ sP2(X2)) | ~ spl41_417), inference(subsumption_resolution, [], [f5267, f335])).
fof(f5267, plain, (! [X2] : ((sK14(sK36(empty_set)) = sK37(X2, sK14(sK36(empty_set)))) | ~ in(sK37(X2, sK14(sK36(empty_set))), sK36(empty_set)) | ~ in(sK14(sK36(empty_set)), sK36(X2)) | ~ sP2(X2)) | ~ spl41_417), inference(resolution, [], [f5174, f334])).
fof(f5174, plain, (! [X5] : (~ subset(sK14(sK36(empty_set)), X5) | (sK14(sK36(empty_set)) = X5) | ~ in(X5, sK36(empty_set))) | ~ spl41_417), inference(avatar_component_clause, [], [f5173])).
fof(f5173, plain, (spl41_417 <=> ! [X5] : ((sK14(sK36(empty_set)) = X5) | ~ subset(sK14(sK36(empty_set)), X5) | ~ in(X5, sK36(empty_set)))), introduced(avatar_definition, [new_symbols(naming, [spl41_417])])).
fof(f5175, plain, (spl41_410 | spl41_417 | ~ spl41_2 | ~ spl41_5 | ~ spl41_354), inference(avatar_split_clause, [], [f5171, f4324, f388, f376, f5173, f5046])).
fof(f5046, plain, (spl41_410 <=> (empty_set = sK36(empty_set))), introduced(avatar_definition, [new_symbols(naming, [spl41_410])])).
fof(f376, plain, (spl41_2 <=> ! [X12, X14] : ((sK14(X12) = X14) | ~ element(X12, powerset(powerset(empty_set))) | (empty_set = X12) | ~ in(X14, X12) | ~ subset(sK14(X12), X14))), introduced(avatar_definition, [new_symbols(naming, [spl41_2])])).
fof(f5171, plain, (! [X5] : ((sK14(sK36(empty_set)) = X5) | (empty_set = sK36(empty_set)) | ~ in(X5, sK36(empty_set)) | ~ subset(sK14(sK36(empty_set)), X5)) | (~ spl41_2 | ~ spl41_5 | ~ spl41_354)), inference(subsumption_resolution, [], [f5068, f4883])).
fof(f5068, plain, (! [X5] : ((sK14(sK36(empty_set)) = X5) | (empty_set = sK36(empty_set)) | ~ in(X5, sK36(empty_set)) | ~ subset(sK14(sK36(empty_set)), X5) | ~ sP2(empty_set)) | ~ spl41_2), inference(resolution, [], [f377, f331])).
fof(f377, plain, (! [X14, X12] : (~ element(X12, powerset(powerset(empty_set))) | (sK14(X12) = X14) | (empty_set = X12) | ~ in(X14, X12) | ~ subset(sK14(X12), X14)) | ~ spl41_2), inference(avatar_component_clause, [], [f376])).
fof(f5134, plain, (~ spl41_5 | ~ spl41_354 | ~ spl41_410), inference(avatar_contradiction_clause, [], [f5133])).
fof(f5133, plain, ($false | (~ spl41_5 | ~ spl41_354 | ~ spl41_410)), inference(subsumption_resolution, [], [f5132, f4883])).
fof(f5132, plain, (~ sP2(empty_set) | ~ spl41_410), inference(trivial_inequality_removal, [], [f5107])).
fof(f5107, plain, (~ (empty_set = empty_set) | ~ sP2(empty_set) | ~ spl41_410), inference(superposition, [], [f332, f5048])).
fof(f5048, plain, ((empty_set = sK36(empty_set)) | ~ spl41_410), inference(avatar_component_clause, [], [f5046])).
fof(f5053, plain, (spl41_410 | spl41_411 | ~ spl41_3 | ~ spl41_5 | ~ spl41_354), inference(avatar_split_clause, [], [f5044, f4324, f388, f380, f5050, f5046])).
fof(f380, plain, (spl41_3 <=> ! [X12] : (in(sK14(X12), X12) | ~ element(X12, powerset(powerset(empty_set))) | (empty_set = X12))), introduced(avatar_definition, [new_symbols(naming, [spl41_3])])).
fof(f5044, plain, (in(sK14(sK36(empty_set)), sK36(empty_set)) | (empty_set = sK36(empty_set)) | (~ spl41_3 | ~ spl41_5 | ~ spl41_354)), inference(subsumption_resolution, [], [f4999, f4883])).
fof(f4999, plain, (in(sK14(sK36(empty_set)), sK36(empty_set)) | (empty_set = sK36(empty_set)) | ~ sP2(empty_set) | ~ spl41_3), inference(resolution, [], [f381, f331])).
fof(f381, plain, (! [X12] : (~ element(X12, powerset(powerset(empty_set))) | in(sK14(X12), X12) | (empty_set = X12)) | ~ spl41_3), inference(avatar_component_clause, [], [f380])).
fof(f4882, plain, (~ spl41_29 | spl41_354 | ~ spl41_5 | ~ spl41_6 | ~ spl41_9 | spl41_99), inference(avatar_split_clause, [], [f4881, f1591, f406, f393, f388, f4324, f845])).
fof(f845, plain, (spl41_29 <=> being_limit_ordinal(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl41_29])])).
fof(f406, plain, (spl41_9 <=> ordinal(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl41_9])])).
fof(f1591, plain, (spl41_99 <=> sP1(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl41_99])])).
fof(f4881, plain, ((empty_set = sK34) | ~ being_limit_ordinal(sK34) | (~ spl41_5 | ~ spl41_6 | ~ spl41_9 | spl41_99)), inference(subsumption_resolution, [], [f4880, f390])).
fof(f4880, plain, ((empty_set = sK34) | ~ being_limit_ordinal(sK34) | ~ sP2(sK34) | (~ spl41_6 | ~ spl41_9 | spl41_99)), inference(subsumption_resolution, [], [f4879, f408])).
fof(f408, plain, (ordinal(sK34) | ~ spl41_9), inference(avatar_component_clause, [], [f406])).
fof(f4879, plain, ((empty_set = sK34) | ~ being_limit_ordinal(sK34) | ~ ordinal(sK34) | ~ sP2(sK34) | (~ spl41_6 | spl41_99)), inference(subsumption_resolution, [], [f4369, f1593])).
fof(f1593, plain, (~ sP1(sK34) | spl41_99), inference(avatar_component_clause, [], [f1591])).
fof(f4369, plain, ((empty_set = sK34) | sP1(sK34) | ~ being_limit_ordinal(sK34) | ~ ordinal(sK34) | ~ sP2(sK34) | ~ spl41_6), inference(resolution, [], [f395, f3967])).
fof(f3967, plain, ! [X0] : (~ in(X0, omega) | (empty_set = X0) | sP1(X0) | ~ being_limit_ordinal(X0) | ~ ordinal(X0) | ~ sP2(X0)), inference(subsumption_resolution, [], [f3966, f923])).
fof(f923, plain, ! [X6] : (in(sK12(sK36(X6)), sK36(X6)) | ~ in(X6, omega) | (empty_set = X6) | sP1(X6) | ~ being_limit_ordinal(X6) | ~ ordinal(X6) | ~ sP2(X6)), inference(subsumption_resolution, [], [f897, f332])).
fof(f897, plain, ! [X6] : ((empty_set = sK36(X6)) | in(sK12(sK36(X6)), sK36(X6)) | ~ in(X6, omega) | (empty_set = X6) | sP1(X6) | ~ being_limit_ordinal(X6) | ~ ordinal(X6) | ~ sP2(X6)), inference(resolution, [], [f210, f331])).
fof(f210, plain, ! [X4, X5] : (~ element(X5, powerset(powerset(X4))) | (empty_set = X5) | in(sK12(X5), X5) | ~ in(X4, omega) | (empty_set = X4) | sP1(X4) | ~ being_limit_ordinal(X4) | ~ ordinal(X4)), inference(cnf_transformation, [], [f137])).
fof(f3966, plain, ! [X0] : (~ in(X0, omega) | (empty_set = X0) | sP1(X0) | ~ being_limit_ordinal(X0) | ~ ordinal(X0) | ~ sP2(X0) | ~ in(sK12(sK36(X0)), sK36(X0))), inference(duplicate_literal_removal, [], [f3965])).
fof(f3965, plain, ! [X0] : (~ in(X0, omega) | (empty_set = X0) | sP1(X0) | ~ being_limit_ordinal(X0) | ~ ordinal(X0) | ~ sP2(X0) | ~ in(sK12(sK36(X0)), sK36(X0)) | ~ sP2(X0) | ~ in(sK12(sK36(X0)), sK36(X0)) | ~ sP2(X0)), inference(resolution, [], [f1770, f333])).
fof(f1770, plain, ! [X6, X7] : (~ in(sK37(X6, sK12(sK36(X7))), sK36(X7)) | ~ in(X7, omega) | (empty_set = X7) | sP1(X7) | ~ being_limit_ordinal(X7) | ~ ordinal(X7) | ~ sP2(X7) | ~ in(sK12(sK36(X7)), sK36(X6)) | ~ sP2(X6)), inference(subsumption_resolution, [], [f1766, f335])).
fof(f1766, plain, ! [X6, X7] : (~ in(sK37(X6, sK12(sK36(X7))), sK36(X7)) | (sK12(sK36(X7)) = sK37(X6, sK12(sK36(X7)))) | ~ in(X7, omega) | (empty_set = X7) | sP1(X7) | ~ being_limit_ordinal(X7) | ~ ordinal(X7) | ~ sP2(X7) | ~ in(sK12(sK36(X7)), sK36(X6)) | ~ sP2(X6)), inference(resolution, [], [f983, f334])).
fof(f983, plain, ! [X14, X13] : (~ subset(sK12(sK36(X13)), X14) | ~ in(X14, sK36(X13)) | (sK12(sK36(X13)) = X14) | ~ in(X13, omega) | (empty_set = X13) | sP1(X13) | ~ being_limit_ordinal(X13) | ~ ordinal(X13) | ~ sP2(X13)), inference(subsumption_resolution, [], [f973, f332])).
fof(f973, plain, ! [X14, X13] : (~ subset(sK12(sK36(X13)), X14) | ~ in(X14, sK36(X13)) | (empty_set = sK36(X13)) | (sK12(sK36(X13)) = X14) | ~ in(X13, omega) | (empty_set = X13) | sP1(X13) | ~ being_limit_ordinal(X13) | ~ ordinal(X13) | ~ sP2(X13)), inference(resolution, [], [f211, f331])).
fof(f211, plain, ! [X4, X7, X5] : (~ element(X5, powerset(powerset(X4))) | ~ subset(sK12(X5), X7) | ~ in(X7, X5) | (empty_set = X5) | (sK12(X5) = X7) | ~ in(X4, omega) | (empty_set = X4) | sP1(X4) | ~ being_limit_ordinal(X4) | ~ ordinal(X4)), inference(cnf_transformation, [], [f137])).
fof(f4878, plain, (spl41_28 | spl41_29 | ~ spl41_9), inference(avatar_split_clause, [], [f4364, f406, f845, f841])).
fof(f4364, plain, (being_limit_ordinal(sK34) | (sK34 = set_union2(sK40(sK34), singleton(sK40(sK34)))) | ~ spl41_9), inference(resolution, [], [f408, f369])).
fof(f369, plain, ! [X0] : (~ ordinal(X0) | being_limit_ordinal(X0) | (set_union2(sK40(X0), singleton(sK40(X0))) = X0)), inference(definition_unfolding, [], [f350, f339])).
fof(f350, plain, ! [X0] : ((succ(sK40(X0)) = X0) | being_limit_ordinal(X0) | ~ ordinal(X0)), inference(cnf_transformation, [], [f191])).
fof(f191, plain, ! [X0] : (((~ being_limit_ordinal(X0) | ! [X1] : (~ (succ(X1) = X0) | ~ ordinal(X1))) & (((succ(sK40(X0)) = X0) & ordinal(sK40(X0))) | being_limit_ordinal(X0))) | ~ ordinal(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK40])], [f110, f190])).
fof(f190, plain, ! [X0] : (? [X2] : ((succ(X2) = X0) & ordinal(X2)) => ((succ(sK40(X0)) = X0) & ordinal(sK40(X0)))), introduced(choice_axiom, [])).
fof(f110, plain, ! [X0] : (((~ being_limit_ordinal(X0) | ! [X1] : (~ (succ(X1) = X0) | ~ ordinal(X1))) & (? [X2] : ((succ(X2) = X0) & ordinal(X2)) | being_limit_ordinal(X0))) | ~ ordinal(X0)), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ! [X0] : (ordinal(X0) => (~ (being_limit_ordinal(X0) & ? [X1] : ((succ(X1) = X0) & ordinal(X1))) & ~ (! [X2] : (ordinal(X2) => ~ (succ(X2) = X0)) & ~ being_limit_ordinal(X0)))), inference(rectify, [], [f63])).
fof(f63, plain, ! [X0] : (ordinal(X0) => (~ (being_limit_ordinal(X0) & ? [X1] : ((succ(X1) = X0) & ordinal(X1))) & ~ (! [X1] : (ordinal(X1) => ~ (succ(X1) = X0)) & ~ being_limit_ordinal(X0)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', t42_ordinal1)).
fof(f4875, plain, (~ spl41_7 | ~ spl41_8 | ~ spl41_28 | ~ spl41_82), inference(avatar_contradiction_clause, [], [f4874])).
fof(f4874, plain, ($false | (~ spl41_7 | ~ spl41_8 | ~ spl41_28 | ~ spl41_82)), inference(subsumption_resolution, [], [f4871, f1395])).
fof(f1395, plain, (sP0(sK40(sK34)) | ~ spl41_82), inference(avatar_component_clause, [], [f1393])).
fof(f4871, plain, (~ sP0(sK40(sK34)) | (~ spl41_7 | ~ spl41_8 | ~ spl41_28)), inference(resolution, [], [f3126, f4412])).
fof(f4412, plain, (in(sK40(sK34), sK34) | ~ spl41_28), inference(superposition, [], [f367, f843])).
fof(f367, plain, ! [X0] : in(X0, set_union2(X0, singleton(X0))), inference(definition_unfolding, [], [f345, f339])).
fof(f345, plain, ! [X0] : in(X0, succ(X0)), inference(cnf_transformation, [], [f59])).
fof(f59, plain, ! [X0] : in(X0, succ(X0)), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', t10_ordinal1)).
fof(f3126, plain, (! [X0] : (~ in(X0, sK34) | ~ sP0(X0)) | (~ spl41_7 | ~ spl41_8)), inference(subsumption_resolution, [], [f3125, f874])).
fof(f874, plain, (! [X3] : (in(sK35(sK7(X3)), sK7(X3)) | ~ in(X3, sK34) | ~ sP0(X3)) | ~ spl41_8), inference(subsumption_resolution, [], [f873, f202])).
fof(f202, plain, ! [X0] : (~ (empty_set = sK7(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f129, plain, ! [X0] : (((! [X2] : ((~ (sK8(X0, X2) = X2) & subset(X2, sK8(X0, X2)) & in(sK8(X0, X2), sK7(X0))) | ~ in(X2, sK7(X0))) & ~ (empty_set = sK7(X0)) & element(sK7(X0), powerset(powerset(X0)))) & in(X0, omega)) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f126, f128, f127])).
fof(f127, plain, ! [X0] : (? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) => (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK7(X0))) | ~ in(X2, sK7(X0))) & ~ (empty_set = sK7(X0)) & element(sK7(X0), powerset(powerset(X0))))), introduced(choice_axiom, [])).
fof(f128, plain, ! [X2, X0] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, sK7(X0))) => (~ (sK8(X0, X2) = X2) & subset(X2, sK8(X0, X2)) & in(sK8(X0, X2), sK7(X0)))), introduced(choice_axiom, [])).
fof(f126, plain, ! [X0] : ((? [X1] : (! [X2] : (? [X3] : (~ (X2 = X3) & subset(X2, X3) & in(X3, X1)) | ~ in(X2, X1)) & ~ (empty_set = X1) & element(X1, powerset(powerset(X0)))) & in(X0, omega)) | ~ sP0(X0)), inference(rectify, [], [f125])).
fof(f125, plain, ! [X8] : ((? [X9] : (! [X10] : (? [X11] : (~ (X10 = X11) & subset(X10, X11) & in(X11, X9)) | ~ in(X10, X9)) & ~ (empty_set = X9) & element(X9, powerset(powerset(X8)))) & in(X8, omega)) | ~ sP0(X8)), inference(nnf_transformation, [], [f114])).
fof(f873, plain, (! [X3] : (~ in(X3, sK34) | in(sK35(sK7(X3)), sK7(X3)) | (empty_set = sK7(X3)) | ~ sP0(X3)) | ~ spl41_8), inference(subsumption_resolution, [], [f872, f200])).
fof(f200, plain, ! [X0] : (in(X0, omega) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f872, plain, (! [X3] : (~ in(X3, sK34) | ~ in(X3, omega) | in(sK35(sK7(X3)), sK7(X3)) | (empty_set = sK7(X3)) | ~ sP0(X3)) | ~ spl41_8), inference(subsumption_resolution, [], [f863, f598])).
fof(f598, plain, ! [X1] : (~ sP0(X1) | ordinal(X1)), inference(resolution, [], [f575, f309])).
fof(f309, plain, ! [X0] : (~ element(X0, omega) | ordinal(X0)), inference(cnf_transformation, [], [f100])).
fof(f100, plain, ! [X0] : ((natural(X0) & ordinal(X0) & epsilon_connected(X0) & epsilon_transitive(X0)) | ~ element(X0, omega)), inference(ennf_transformation, [], [f42])).
fof(f42, plain, ! [X0] : (element(X0, omega) => (natural(X0) & ordinal(X0) & epsilon_connected(X0) & epsilon_transitive(X0))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', cc3_arytm_3)).
fof(f575, plain, ! [X0] : (element(X0, omega) | ~ sP0(X0)), inference(resolution, [], [f347, f200])).
fof(f347, plain, ! [X0, X1] : (~ in(X0, X1) | element(X0, X1)), inference(cnf_transformation, [], [f107])).
fof(f107, plain, ! [X0, X1] : (element(X0, X1) | ~ in(X0, X1)), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0, X1] : (in(X0, X1) => element(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', t1_subset)).
fof(f863, plain, (! [X3] : (~ ordinal(X3) | ~ in(X3, sK34) | ~ in(X3, omega) | in(sK35(sK7(X3)), sK7(X3)) | (empty_set = sK7(X3)) | ~ sP0(X3)) | ~ spl41_8), inference(resolution, [], [f403, f201])).
fof(f201, plain, ! [X0] : (element(sK7(X0), powerset(powerset(X0))) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f403, plain, (! [X2, X1] : (~ element(X2, powerset(powerset(X1))) | ~ ordinal(X1) | ~ in(X1, sK34) | ~ in(X1, omega) | in(sK35(X2), X2) | (empty_set = X2)) | ~ spl41_8), inference(avatar_component_clause, [], [f402])).
fof(f402, plain, (spl41_8 <=> ! [X1, X2] : (in(sK35(X2), X2) | ~ ordinal(X1) | ~ in(X1, sK34) | ~ in(X1, omega) | ~ element(X2, powerset(powerset(X1))) | (empty_set = X2))), introduced(avatar_definition, [new_symbols(naming, [spl41_8])])).
fof(f3125, plain, (! [X0] : (~ in(X0, sK34) | ~ sP0(X0) | ~ in(sK35(sK7(X0)), sK7(X0))) | ~ spl41_7), inference(duplicate_literal_removal, [], [f3124])).
fof(f3124, plain, (! [X0] : (~ in(X0, sK34) | ~ sP0(X0) | ~ in(sK35(sK7(X0)), sK7(X0)) | ~ sP0(X0) | ~ in(sK35(sK7(X0)), sK7(X0)) | ~ sP0(X0)) | ~ spl41_7), inference(resolution, [], [f1601, f203])).
fof(f203, plain, ! [X2, X0] : (in(sK8(X0, X2), sK7(X0)) | ~ in(X2, sK7(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f1601, plain, (! [X4, X3] : (~ in(sK8(X4, sK35(sK7(X3))), sK7(X3)) | ~ in(X3, sK34) | ~ sP0(X3) | ~ in(sK35(sK7(X3)), sK7(X4)) | ~ sP0(X4)) | ~ spl41_7), inference(subsumption_resolution, [], [f1597, f205])).
fof(f205, plain, ! [X2, X0] : (~ (sK8(X0, X2) = X2) | ~ in(X2, sK7(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f1597, plain, (! [X4, X3] : ((sK35(sK7(X3)) = sK8(X4, sK35(sK7(X3)))) | ~ in(sK8(X4, sK35(sK7(X3))), sK7(X3)) | ~ in(X3, sK34) | ~ sP0(X3) | ~ in(sK35(sK7(X3)), sK7(X4)) | ~ sP0(X4)) | ~ spl41_7), inference(resolution, [], [f962, f204])).
fof(f204, plain, ! [X2, X0] : (subset(X2, sK8(X0, X2)) | ~ in(X2, sK7(X0)) | ~ sP0(X0)), inference(cnf_transformation, [], [f129])).
fof(f962, plain, (! [X6, X7] : (~ subset(sK35(sK7(X6)), X7) | (sK35(sK7(X6)) = X7) | ~ in(X7, sK7(X6)) | ~ in(X6, sK34) | ~ sP0(X6)) | ~ spl41_7), inference(subsumption_resolution, [], [f961, f202])).
fof(f961, plain, (! [X6, X7] : (~ in(X6, sK34) | (sK35(sK7(X6)) = X7) | (empty_set = sK7(X6)) | ~ in(X7, sK7(X6)) | ~ subset(sK35(sK7(X6)), X7) | ~ sP0(X6)) | ~ spl41_7), inference(subsumption_resolution, [], [f960, f200])).
fof(f960, plain, (! [X6, X7] : (~ in(X6, sK34) | ~ in(X6, omega) | (sK35(sK7(X6)) = X7) | (empty_set = sK7(X6)) | ~ in(X7, sK7(X6)) | ~ subset(sK35(sK7(X6)), X7) | ~ sP0(X6)) | ~ spl41_7), inference(subsumption_resolution, [], [f951, f598])).
fof(f951, plain, (! [X6, X7] : (~ ordinal(X6) | ~ in(X6, sK34) | ~ in(X6, omega) | (sK35(sK7(X6)) = X7) | (empty_set = sK7(X6)) | ~ in(X7, sK7(X6)) | ~ subset(sK35(sK7(X6)), X7) | ~ sP0(X6)) | ~ spl41_7), inference(resolution, [], [f399, f201])).
fof(f399, plain, (! [X4, X2, X1] : (~ element(X2, powerset(powerset(X1))) | ~ ordinal(X1) | ~ in(X1, sK34) | ~ in(X1, omega) | (sK35(X2) = X4) | (empty_set = X2) | ~ in(X4, X2) | ~ subset(sK35(X2), X4)) | ~ spl41_7), inference(avatar_component_clause, [], [f398])).
fof(f398, plain, (spl41_7 <=> ! [X1, X2, X4] : ((sK35(X2) = X4) | ~ ordinal(X1) | ~ in(X1, sK34) | ~ in(X1, omega) | ~ element(X2, powerset(powerset(X1))) | (empty_set = X2) | ~ in(X4, X2) | ~ subset(sK35(X2), X4))), introduced(avatar_definition, [new_symbols(naming, [spl41_7])])).
fof(f4335, plain, (~ spl41_99 | ~ spl41_7 | ~ spl41_8), inference(avatar_split_clause, [], [f3910, f402, f398, f1591])).
fof(f3910, plain, (~ sP1(sK34) | (~ spl41_7 | ~ spl41_8)), inference(duplicate_literal_removal, [], [f3909])).
fof(f3909, plain, (~ sP1(sK34) | ~ sP1(sK34) | (~ spl41_7 | ~ spl41_8)), inference(resolution, [], [f3173, f193])).
fof(f193, plain, ! [X0] : (in(sK4(X0), X0) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f124, plain, ! [X0] : (((! [X3] : ((~ (sK6(X0, X3) = X3) & subset(X3, sK6(X0, X3)) & in(sK6(X0, X3), sK5(X0))) | ~ in(X3, sK5(X0))) & ~ (empty_set = sK5(X0)) & element(sK5(X0), powerset(powerset(sK4(X0))))) & in(sK4(X0), omega) & in(sK4(X0), X0) & ordinal(sK4(X0))) | ~ sP1(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6])], [f120, f123, f122, f121])).
fof(f121, plain, ! [X0] : (? [X1] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(X1)))) & in(X1, omega) & in(X1, X0) & ordinal(X1)) => (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(sK4(X0))))) & in(sK4(X0), omega) & in(sK4(X0), X0) & ordinal(sK4(X0)))), introduced(choice_axiom, [])).
fof(f122, plain, ! [X0] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(sK4(X0))))) => (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, sK5(X0))) | ~ in(X3, sK5(X0))) & ~ (empty_set = sK5(X0)) & element(sK5(X0), powerset(powerset(sK4(X0)))))), introduced(choice_axiom, [])).
fof(f123, plain, ! [X3, X0] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, sK5(X0))) => (~ (sK6(X0, X3) = X3) & subset(X3, sK6(X0, X3)) & in(sK6(X0, X3), sK5(X0)))), introduced(choice_axiom, [])).
fof(f120, plain, ! [X0] : (? [X1] : (? [X2] : (! [X3] : (? [X4] : (~ (X3 = X4) & subset(X3, X4) & in(X4, X2)) | ~ in(X3, X2)) & ~ (empty_set = X2) & element(X2, powerset(powerset(X1)))) & in(X1, omega) & in(X1, X0) & ordinal(X1)) | ~ sP1(X0)), inference(nnf_transformation, [], [f115])).
fof(f3173, plain, (! [X0] : (~ in(sK4(X0), sK34) | ~ sP1(X0)) | (~ spl41_7 | ~ spl41_8)), inference(subsumption_resolution, [], [f3172, f871])).
fof(f871, plain, (! [X2] : (in(sK35(sK5(X2)), sK5(X2)) | ~ in(sK4(X2), sK34) | ~ sP1(X2)) | ~ spl41_8), inference(subsumption_resolution, [], [f870, f196])).
fof(f196, plain, ! [X0] : (~ (empty_set = sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f870, plain, (! [X2] : (~ in(sK4(X2), sK34) | in(sK35(sK5(X2)), sK5(X2)) | (empty_set = sK5(X2)) | ~ sP1(X2)) | ~ spl41_8), inference(subsumption_resolution, [], [f869, f194])).
fof(f194, plain, ! [X0] : (in(sK4(X0), omega) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f869, plain, (! [X2] : (~ in(sK4(X2), sK34) | ~ in(sK4(X2), omega) | in(sK35(sK5(X2)), sK5(X2)) | (empty_set = sK5(X2)) | ~ sP1(X2)) | ~ spl41_8), inference(subsumption_resolution, [], [f862, f192])).
fof(f192, plain, ! [X0] : (ordinal(sK4(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f862, plain, (! [X2] : (~ ordinal(sK4(X2)) | ~ in(sK4(X2), sK34) | ~ in(sK4(X2), omega) | in(sK35(sK5(X2)), sK5(X2)) | (empty_set = sK5(X2)) | ~ sP1(X2)) | ~ spl41_8), inference(resolution, [], [f403, f195])).
fof(f195, plain, ! [X0] : (element(sK5(X0), powerset(powerset(sK4(X0)))) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f3172, plain, (! [X0] : (~ in(sK4(X0), sK34) | ~ sP1(X0) | ~ in(sK35(sK5(X0)), sK5(X0))) | ~ spl41_7), inference(duplicate_literal_removal, [], [f3171])).
fof(f3171, plain, (! [X0] : (~ in(sK4(X0), sK34) | ~ sP1(X0) | ~ in(sK35(sK5(X0)), sK5(X0)) | ~ sP1(X0) | ~ in(sK35(sK5(X0)), sK5(X0)) | ~ sP1(X0)) | ~ spl41_7), inference(resolution, [], [f1612, f197])).
fof(f197, plain, ! [X0, X3] : (in(sK6(X0, X3), sK5(X0)) | ~ in(X3, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f1612, plain, (! [X2, X1] : (~ in(sK6(X2, sK35(sK5(X1))), sK5(X1)) | ~ in(sK4(X1), sK34) | ~ sP1(X1) | ~ in(sK35(sK5(X1)), sK5(X2)) | ~ sP1(X2)) | ~ spl41_7), inference(subsumption_resolution, [], [f1608, f199])).
fof(f199, plain, ! [X0, X3] : (~ (sK6(X0, X3) = X3) | ~ in(X3, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f1608, plain, (! [X2, X1] : ((sK35(sK5(X1)) = sK6(X2, sK35(sK5(X1)))) | ~ in(sK6(X2, sK35(sK5(X1))), sK5(X1)) | ~ in(sK4(X1), sK34) | ~ sP1(X1) | ~ in(sK35(sK5(X1)), sK5(X2)) | ~ sP1(X2)) | ~ spl41_7), inference(resolution, [], [f959, f198])).
fof(f198, plain, ! [X0, X3] : (subset(X3, sK6(X0, X3)) | ~ in(X3, sK5(X0)) | ~ sP1(X0)), inference(cnf_transformation, [], [f124])).
fof(f959, plain, (! [X4, X5] : (~ subset(sK35(sK5(X4)), X5) | (sK35(sK5(X4)) = X5) | ~ in(X5, sK5(X4)) | ~ in(sK4(X4), sK34) | ~ sP1(X4)) | ~ spl41_7), inference(subsumption_resolution, [], [f958, f196])).
fof(f958, plain, (! [X4, X5] : (~ in(sK4(X4), sK34) | (sK35(sK5(X4)) = X5) | (empty_set = sK5(X4)) | ~ in(X5, sK5(X4)) | ~ subset(sK35(sK5(X4)), X5) | ~ sP1(X4)) | ~ spl41_7), inference(subsumption_resolution, [], [f957, f194])).
fof(f957, plain, (! [X4, X5] : (~ in(sK4(X4), sK34) | ~ in(sK4(X4), omega) | (sK35(sK5(X4)) = X5) | (empty_set = sK5(X4)) | ~ in(X5, sK5(X4)) | ~ subset(sK35(sK5(X4)), X5) | ~ sP1(X4)) | ~ spl41_7), inference(subsumption_resolution, [], [f950, f192])).
fof(f950, plain, (! [X4, X5] : (~ ordinal(sK4(X4)) | ~ in(sK4(X4), sK34) | ~ in(sK4(X4), omega) | (sK35(sK5(X4)) = X5) | (empty_set = sK5(X4)) | ~ in(X5, sK5(X4)) | ~ subset(sK35(sK5(X4)), X5) | ~ sP1(X4)) | ~ spl41_7), inference(resolution, [], [f399, f195])).
fof(f4287, plain, (spl41_1 | ~ spl41_5 | ~ spl41_6 | ~ spl41_9 | ~ spl41_29 | spl41_99), inference(avatar_contradiction_clause, [], [f4286])).
fof(f4286, plain, ($false | (spl41_1 | ~ spl41_5 | ~ spl41_6 | ~ spl41_9 | ~ spl41_29 | spl41_99)), inference(subsumption_resolution, [], [f4165, f374])).
fof(f374, plain, (~ in(empty_set, omega) | spl41_1), inference(avatar_component_clause, [], [f372])).
fof(f372, plain, (spl41_1 <=> in(empty_set, omega)), introduced(avatar_definition, [new_symbols(naming, [spl41_1])])).
fof(f4165, plain, (in(empty_set, omega) | (~ spl41_5 | ~ spl41_6 | ~ spl41_9 | ~ spl41_29 | spl41_99)), inference(backward_demodulation, [], [f395, f4163])).
fof(f4163, plain, ((empty_set = sK34) | (~ spl41_5 | ~ spl41_6 | ~ spl41_9 | ~ spl41_29 | spl41_99)), inference(subsumption_resolution, [], [f4162, f390])).
fof(f4162, plain, ((empty_set = sK34) | ~ sP2(sK34) | (~ spl41_6 | ~ spl41_9 | ~ spl41_29 | spl41_99)), inference(subsumption_resolution, [], [f4161, f408])).
fof(f4161, plain, ((empty_set = sK34) | ~ ordinal(sK34) | ~ sP2(sK34) | (~ spl41_6 | ~ spl41_29 | spl41_99)), inference(subsumption_resolution, [], [f4160, f847])).
fof(f847, plain, (being_limit_ordinal(sK34) | ~ spl41_29), inference(avatar_component_clause, [], [f845])).
fof(f4160, plain, ((empty_set = sK34) | ~ being_limit_ordinal(sK34) | ~ ordinal(sK34) | ~ sP2(sK34) | (~ spl41_6 | spl41_99)), inference(subsumption_resolution, [], [f4157, f1593])).
fof(f4157, plain, ((empty_set = sK34) | sP1(sK34) | ~ being_limit_ordinal(sK34) | ~ ordinal(sK34) | ~ sP2(sK34) | ~ spl41_6), inference(resolution, [], [f3967, f395])).
fof(f3718, plain, (~ spl41_10 | ~ spl41_326 | ~ spl41_327), inference(avatar_contradiction_clause, [], [f3717])).
fof(f3717, plain, ($false | (~ spl41_10 | ~ spl41_326 | ~ spl41_327)), inference(subsumption_resolution, [], [f3716, f3715])).
fof(f3715, plain, (~ in(sK38(sK10), sK10) | ~ spl41_10), inference(subsumption_resolution, [], [f3714, f218])).
fof(f218, plain, ! [X2] : (~ (sK11(X2) = X2) | ~ in(X2, sK10)), inference(cnf_transformation, [], [f137])).
fof(f3714, plain, ((sK38(sK10) = sK11(sK38(sK10))) | ~ in(sK38(sK10), sK10) | ~ spl41_10), inference(subsumption_resolution, [], [f3478, f216])).
fof(f216, plain, ! [X2] : (in(sK11(X2), sK10) | ~ in(X2, sK10)), inference(cnf_transformation, [], [f137])).
fof(f3478, plain, (~ in(sK11(sK38(sK10)), sK10) | (sK38(sK10) = sK11(sK38(sK10))) | ~ in(sK38(sK10), sK10) | ~ spl41_10), inference(resolution, [], [f3449, f217])).
fof(f217, plain, ! [X2] : (subset(X2, sK11(X2)) | ~ in(X2, sK10)), inference(cnf_transformation, [], [f137])).
fof(f3449, plain, (! [X8] : (~ subset(sK38(sK10), X8) | ~ in(X8, sK10) | (sK38(sK10) = X8)) | ~ spl41_10), inference(subsumption_resolution, [], [f3448, f215])).
fof(f215, plain, ~ (empty_set = sK10), inference(cnf_transformation, [], [f137])).
fof(f3448, plain, (! [X8] : ((sK38(sK10) = X8) | (empty_set = sK10) | ~ in(X8, sK10) | ~ subset(sK38(sK10), X8)) | ~ spl41_10), inference(subsumption_resolution, [], [f3447, f213])).
fof(f213, plain, in(sK9, omega), inference(cnf_transformation, [], [f137])).
fof(f3447, plain, (! [X8] : (~ in(sK9, omega) | (sK38(sK10) = X8) | (empty_set = sK10) | ~ in(X8, sK10) | ~ subset(sK38(sK10), X8)) | ~ spl41_10), inference(subsumption_resolution, [], [f3436, f212])).
fof(f212, plain, ordinal(sK9), inference(cnf_transformation, [], [f137])).
fof(f3436, plain, (! [X8] : (~ ordinal(sK9) | ~ in(sK9, omega) | (sK38(sK10) = X8) | (empty_set = sK10) | ~ in(X8, sK10) | ~ subset(sK38(sK10), X8)) | ~ spl41_10), inference(resolution, [], [f412, f214])).
fof(f214, plain, element(sK10, powerset(powerset(sK9))), inference(cnf_transformation, [], [f137])).
fof(f412, plain, (! [X0, X3, X1] : (~ element(X1, powerset(powerset(X0))) | ~ ordinal(X0) | ~ in(X0, omega) | (sK38(X1) = X3) | (empty_set = X1) | ~ in(X3, X1) | ~ subset(sK38(X1), X3)) | ~ spl41_10), inference(avatar_component_clause, [], [f411])).
fof(f411, plain, (spl41_10 <=> ! [X1, X3, X0] : ((sK38(X1) = X3) | ~ ordinal(X0) | ~ in(X0, omega) | ~ element(X1, powerset(powerset(X0))) | (empty_set = X1) | ~ in(X3, X1) | ~ subset(sK38(X1), X3))), introduced(avatar_definition, [new_symbols(naming, [spl41_10])])).
fof(f3716, plain, (in(sK38(sK10), sK10) | (~ spl41_326 | ~ spl41_327)), inference(backward_demodulation, [], [f3489, f3486])).
fof(f3486, plain, ((sK38(sK10) = sK11(sK38(sK10))) | ~ spl41_326), inference(avatar_component_clause, [], [f3484])).
fof(f3484, plain, (spl41_326 <=> (sK38(sK10) = sK11(sK38(sK10)))), introduced(avatar_definition, [new_symbols(naming, [spl41_326])])).
fof(f3489, plain, (in(sK11(sK38(sK10)), sK10) | ~ spl41_327), inference(avatar_component_clause, [], [f3488])).
fof(f3488, plain, (spl41_327 <=> in(sK11(sK38(sK10)), sK10)), introduced(avatar_definition, [new_symbols(naming, [spl41_327])])).
fof(f3495, plain, (~ spl41_11 | spl41_327), inference(avatar_contradiction_clause, [], [f3494])).
fof(f3494, plain, ($false | (~ spl41_11 | spl41_327)), inference(subsumption_resolution, [], [f3493, f3425])).
fof(f3425, plain, (in(sK38(sK10), sK10) | ~ spl41_11), inference(subsumption_resolution, [], [f3424, f215])).
fof(f3424, plain, (in(sK38(sK10), sK10) | (empty_set = sK10) | ~ spl41_11), inference(subsumption_resolution, [], [f3423, f213])).
fof(f3423, plain, (~ in(sK9, omega) | in(sK38(sK10), sK10) | (empty_set = sK10) | ~ spl41_11), inference(subsumption_resolution, [], [f3412, f212])).
fof(f3412, plain, (~ ordinal(sK9) | ~ in(sK9, omega) | in(sK38(sK10), sK10) | (empty_set = sK10) | ~ spl41_11), inference(resolution, [], [f416, f214])).
fof(f416, plain, (! [X0, X1] : (~ element(X1, powerset(powerset(X0))) | ~ ordinal(X0) | ~ in(X0, omega) | in(sK38(X1), X1) | (empty_set = X1)) | ~ spl41_11), inference(avatar_component_clause, [], [f415])).
fof(f415, plain, (spl41_11 <=> ! [X1, X0] : (in(sK38(X1), X1) | ~ ordinal(X0) | ~ in(X0, omega) | ~ element(X1, powerset(powerset(X0))) | (empty_set = X1))), introduced(avatar_definition, [new_symbols(naming, [spl41_11])])).
fof(f3493, plain, (~ in(sK38(sK10), sK10) | spl41_327), inference(resolution, [], [f3490, f216])).
fof(f3490, plain, (~ in(sK11(sK38(sK10)), sK10) | spl41_327), inference(avatar_component_clause, [], [f3488])).
fof(f3491, plain, (spl41_326 | ~ spl41_327 | ~ spl41_10 | ~ spl41_11), inference(avatar_split_clause, [], [f3482, f415, f411, f3488, f3484])).
fof(f3482, plain, (~ in(sK11(sK38(sK10)), sK10) | (sK38(sK10) = sK11(sK38(sK10))) | (~ spl41_10 | ~ spl41_11)), inference(subsumption_resolution, [], [f3478, f3425])).
fof(f1417, plain, (spl41_81 | ~ spl41_9 | ~ spl41_28), inference(avatar_split_clause, [], [f1416, f841, f406, f1389])).
fof(f1416, plain, (ordinal(sK40(sK34)) | (~ spl41_9 | ~ spl41_28)), inference(subsumption_resolution, [], [f1382, f408])).
fof(f1382, plain, (~ ordinal(sK34) | ordinal(sK40(sK34)) | ~ spl41_28), inference(superposition, [], [f764, f843])).
fof(f764, plain, ! [X1] : (~ ordinal(set_union2(X1, singleton(X1))) | ordinal(X1)), inference(resolution, [], [f584, f306])).
fof(f306, plain, ! [X0, X1] : (~ element(X1, X0) | ordinal(X1) | ~ ordinal(X0)), inference(cnf_transformation, [], [f99])).
fof(f99, plain, ! [X0] : (! [X1] : ((ordinal(X1) & epsilon_connected(X1) & epsilon_transitive(X1)) | ~ element(X1, X0)) | ~ ordinal(X0)), inference(ennf_transformation, [], [f41])).
fof(f41, plain, ! [X0] : (ordinal(X0) => ! [X1] : (element(X1, X0) => (ordinal(X1) & epsilon_connected(X1) & epsilon_transitive(X1)))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', cc1_arytm_3)).
fof(f584, plain, ! [X0] : element(X0, set_union2(X0, singleton(X0))), inference(resolution, [], [f367, f347])).
fof(f417, plain, (spl41_4 | spl41_11), inference(avatar_split_clause, [], [f336, f415, f384])).
fof(f384, plain, (spl41_4 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl41_4])])).
fof(f336, plain, ! [X0, X1] : (in(sK38(X1), X1) | (empty_set = X1) | ~ element(X1, powerset(powerset(X0))) | ~ in(X0, omega) | ~ ordinal(X0) | sP3), inference(cnf_transformation, [], [f187])).
fof(f187, plain, (! [X0] : (! [X1] : ((! [X3] : ((sK38(X1) = X3) | ~ subset(sK38(X1), X3) | ~ in(X3, X1)) & in(sK38(X1), X1)) | (empty_set = X1) | ~ element(X1, powerset(powerset(X0)))) | ~ in(X0, omega) | ~ ordinal(X0)) | sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38])], [f185, f186])).
fof(f186, plain, ! [X1] : (? [X2] : (! [X3] : ((X2 = X3) | ~ subset(X2, X3) | ~ in(X3, X1)) & in(X2, X1)) => (! [X3] : ((sK38(X1) = X3) | ~ subset(sK38(X1), X3) | ~ in(X3, X1)) & in(sK38(X1), X1))), introduced(choice_axiom, [])).
fof(f185, plain, (! [X0] : (! [X1] : (? [X2] : (! [X3] : ((X2 = X3) | ~ subset(X2, X3) | ~ in(X3, X1)) & in(X2, X1)) | (empty_set = X1) | ~ element(X1, powerset(powerset(X0)))) | ~ in(X0, omega) | ~ ordinal(X0)) | sP3), inference(rectify, [], [f119])).
fof(f119, plain, (! [X8] : (! [X9] : (? [X10] : (! [X11] : ((X10 = X11) | ~ subset(X10, X11) | ~ in(X11, X9)) & in(X10, X9)) | (empty_set = X9) | ~ element(X9, powerset(powerset(X8)))) | ~ in(X8, omega) | ~ ordinal(X8)) | sP3), inference(definition_folding, [], [f104, e118, e117])).
fof(f118, plain, (? [X0] : (sP2(X0) & in(X0, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, X0) | ~ ordinal(X1)) & ordinal(X0)) | ~ sP3), inference(usedef, [], [e118])).
fof(e118, plain, (sP3 <=> ? [X0] : (sP2(X0) & in(X0, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, X0) | ~ ordinal(X1)) & ordinal(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f104, plain, (! [X8] : (! [X9] : (? [X10] : (! [X11] : ((X10 = X11) | ~ subset(X10, X11) | ~ in(X11, X9)) & in(X10, X9)) | (empty_set = X9) | ~ element(X9, powerset(powerset(X8)))) | ~ in(X8, omega) | ~ ordinal(X8)) | ? [X0] : (? [X5] : (! [X6] : (? [X7] : (~ (X6 = X7) & subset(X6, X7) & in(X7, X5)) | ~ in(X6, X5)) & ~ (empty_set = X5) & element(X5, powerset(powerset(X0)))) & in(X0, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, X0) | ~ ordinal(X1)) & ordinal(X0))), inference(flattening, [], [f103])).
fof(f103, plain, (! [X8] : ((! [X9] : ((? [X10] : (! [X11] : ((X10 = X11) | (~ subset(X10, X11) | ~ in(X11, X9))) & in(X10, X9)) | (empty_set = X9)) | ~ element(X9, powerset(powerset(X8)))) | ~ in(X8, omega)) | ~ ordinal(X8)) | ? [X0] : (((? [X5] : ((! [X6] : (? [X7] : (~ (X6 = X7) & (subset(X6, X7) & in(X7, X5))) | ~ in(X6, X5)) & ~ (empty_set = X5)) & element(X5, powerset(powerset(X0)))) & in(X0, omega)) & ! [X1] : (((! [X2] : ((? [X3] : (! [X4] : ((X3 = X4) | (~ subset(X3, X4) | ~ in(X4, X2))) & in(X3, X2)) | (empty_set = X2)) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega)) | ~ in(X1, X0)) | ~ ordinal(X1))) & ordinal(X0))), inference(ennf_transformation, [], [f69])).
fof(f69, plain, (! [X0] : (ordinal(X0) => (! [X1] : (ordinal(X1) => (in(X1, X0) => (in(X1, omega) => ! [X2] : (element(X2, powerset(powerset(X1))) => ~ (! [X3] : ~ (! [X4] : ((subset(X3, X4) & in(X4, X2)) => (X3 = X4)) & in(X3, X2)) & ~ (empty_set = X2)))))) => (in(X0, omega) => ! [X5] : (element(X5, powerset(powerset(X0))) => ~ (! [X6] : ~ (! [X7] : ((subset(X6, X7) & in(X7, X5)) => (X6 = X7)) & in(X6, X5)) & ~ (empty_set = X5)))))) => ! [X8] : (ordinal(X8) => (in(X8, omega) => ! [X9] : (element(X9, powerset(powerset(X8))) => ~ (! [X10] : ~ (! [X11] : ((subset(X10, X11) & in(X11, X9)) => (X10 = X11)) & in(X10, X9)) & ~ (empty_set = X9)))))), inference(rectify, [], [f49])).
fof(f49, plain, (! [X0] : (ordinal(X0) => (! [X1] : (ordinal(X1) => (in(X1, X0) => (in(X1, omega) => ! [X2] : (element(X2, powerset(powerset(X1))) => ~ (! [X3] : ~ (! [X4] : ((subset(X3, X4) & in(X4, X2)) => (X3 = X4)) & in(X3, X2)) & ~ (empty_set = X2)))))) => (in(X0, omega) => ! [X5] : (element(X5, powerset(powerset(X0))) => ~ (! [X6] : ~ (! [X7] : ((subset(X6, X7) & in(X7, X5)) => (X6 = X7)) & in(X6, X5)) & ~ (empty_set = X5)))))) => ! [X0] : (ordinal(X0) => (in(X0, omega) => ! [X8] : (element(X8, powerset(powerset(X0))) => ~ (! [X9] : ~ (! [X10] : ((subset(X9, X10) & in(X10, X8)) => (X9 = X10)) & in(X9, X8)) & ~ (empty_set = X8)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU301+1.p', s2_ordinal1__e18_27__finset_1__1)).
fof(f413, plain, (spl41_4 | spl41_10), inference(avatar_split_clause, [], [f337, f411, f384])).
fof(f337, plain, ! [X0, X3, X1] : ((sK38(X1) = X3) | ~ subset(sK38(X1), X3) | ~ in(X3, X1) | (empty_set = X1) | ~ element(X1, powerset(powerset(X0))) | ~ in(X0, omega) | ~ ordinal(X0) | sP3), inference(cnf_transformation, [], [f187])).
fof(f409, plain, (~ spl41_4 | spl41_9), inference(avatar_split_clause, [], [f326, f406, f384])).
fof(f326, plain, (ordinal(sK34) | ~ sP3), inference(cnf_transformation, [], [f179])).
fof(f179, plain, ((sP2(sK34) & in(sK34, omega) & ! [X1] : (! [X2] : ((! [X4] : ((sK35(X2) = X4) | ~ subset(sK35(X2), X4) | ~ in(X4, X2)) & in(sK35(X2), X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, sK34) | ~ ordinal(X1)) & ordinal(sK34)) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34, sK35])], [f176, f178, f177])).
fof(f177, plain, (? [X0] : (sP2(X0) & in(X0, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, X0) | ~ ordinal(X1)) & ordinal(X0)) => (sP2(sK34) & in(sK34, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, sK34) | ~ ordinal(X1)) & ordinal(sK34))), introduced(choice_axiom, [])).
fof(f178, plain, ! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) => (! [X4] : ((sK35(X2) = X4) | ~ subset(sK35(X2), X4) | ~ in(X4, X2)) & in(sK35(X2), X2))), introduced(choice_axiom, [])).
fof(f176, plain, (? [X0] : (sP2(X0) & in(X0, omega) & ! [X1] : (! [X2] : (? [X3] : (! [X4] : ((X3 = X4) | ~ subset(X3, X4) | ~ in(X4, X2)) & in(X3, X2)) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1)))) | ~ in(X1, omega) | ~ in(X1, X0) | ~ ordinal(X1)) & ordinal(X0)) | ~ sP3), inference(nnf_transformation, [], [f118])).
fof(f404, plain, (~ spl41_4 | spl41_8), inference(avatar_split_clause, [], [f327, f402, f384])).
fof(f327, plain, ! [X2, X1] : (in(sK35(X2), X2) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1))) | ~ in(X1, omega) | ~ in(X1, sK34) | ~ ordinal(X1) | ~ sP3), inference(cnf_transformation, [], [f179])).
fof(f400, plain, (~ spl41_4 | spl41_7), inference(avatar_split_clause, [], [f328, f398, f384])).
fof(f328, plain, ! [X4, X2, X1] : ((sK35(X2) = X4) | ~ subset(sK35(X2), X4) | ~ in(X4, X2) | (empty_set = X2) | ~ element(X2, powerset(powerset(X1))) | ~ in(X1, omega) | ~ in(X1, sK34) | ~ ordinal(X1) | ~ sP3), inference(cnf_transformation, [], [f179])).
fof(f396, plain, (~ spl41_4 | spl41_6), inference(avatar_split_clause, [], [f329, f393, f384])).
fof(f329, plain, (in(sK34, omega) | ~ sP3), inference(cnf_transformation, [], [f179])).
fof(f391, plain, (~ spl41_4 | spl41_5), inference(avatar_split_clause, [], [f330, f388, f384])).
fof(f330, plain, (sP2(sK34) | ~ sP3), inference(cnf_transformation, [], [f179])).
fof(f382, plain, (~ spl41_1 | spl41_3), inference(avatar_split_clause, [], [f206, f380, f372])).
fof(f206, plain, ! [X12] : (in(sK14(X12), X12) | (empty_set = X12) | ~ element(X12, powerset(powerset(empty_set))) | ~ in(empty_set, omega)), inference(cnf_transformation, [], [f137])).
fof(f378, plain, (~ spl41_1 | spl41_2), inference(avatar_split_clause, [], [f207, f376, f372])).
fof(f207, plain, ! [X14, X12] : ((sK14(X12) = X14) | ~ subset(sK14(X12), X14) | ~ in(X14, X12) | (empty_set = X12) | ~ element(X12, powerset(powerset(empty_set))) | ~ in(empty_set, omega)), inference(cnf_transformation, [], [f137])).