fof(f243, plain, $false, inference(avatar_sat_refutation, [], [f86, f90, f95, f140, f144, f146, f207, f210, f237, f242])).
fof(f242, plain, (spl10_5 | ~ spl10_2 | ~ spl10_3), inference(avatar_split_clause, [], [f241, f83, f79, f92])).
fof(f92, plain, (spl10_5 <=> relstr_set_smaller(sK7, sK9, sK8)), introduced(avatar_definition, [new_symbols(naming, [spl10_5])])).
fof(f79, plain, (spl10_2 <=> (sK8 = join_on_relstr(sK7, sK9))), introduced(avatar_definition, [new_symbols(naming, [spl10_2])])).
fof(f83, plain, (spl10_3 <=> ex_sup_of_relstr_set(sK7, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl10_3])])).
fof(f241, plain, (relstr_set_smaller(sK7, sK9, sK8) | (~ spl10_2 | ~ spl10_3)), inference(subsumption_resolution, [], [f240, f67])).
fof(f67, plain, rel_str(sK7), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (((((~ ex_sup_of_relstr_set(sK7, sK9) | ~ (sK8 = join_on_relstr(sK7, sK9))) & ! [X3] : (related(sK7, sK8, X3) | ~ relstr_set_smaller(sK7, sK9, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, sK9, sK8)) | sP0(sK8, sK7, sK9)) & element(sK8, the_carrier(sK7))) & rel_str(sK7) & antisymmetric_relstr(sK7)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8, sK9])], [f24, f45, f44, f43])).
fof(f43, plain, (? [X0] : (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(X0, X2) | ~ (join_on_relstr(X0, X2) = X1)) & ! [X3] : (related(X0, X1, X3) | ~ relstr_set_smaller(X0, X2, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X2, X1)) | sP0(X1, X0, X2)) & element(X1, the_carrier(X0))) & rel_str(X0) & antisymmetric_relstr(X0)) => (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(sK7, X2) | ~ (join_on_relstr(sK7, X2) = X1)) & ! [X3] : (related(sK7, X1, X3) | ~ relstr_set_smaller(sK7, X2, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, X2, X1)) | sP0(X1, sK7, X2)) & element(X1, the_carrier(sK7))) & rel_str(sK7) & antisymmetric_relstr(sK7))), introduced(choice_axiom, [])).
fof(f44, plain, (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(sK7, X2) | ~ (join_on_relstr(sK7, X2) = X1)) & ! [X3] : (related(sK7, X1, X3) | ~ relstr_set_smaller(sK7, X2, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, X2, X1)) | sP0(X1, sK7, X2)) & element(X1, the_carrier(sK7))) => (? [X2] : (((~ ex_sup_of_relstr_set(sK7, X2) | ~ (join_on_relstr(sK7, X2) = sK8)) & ! [X3] : (related(sK7, sK8, X3) | ~ relstr_set_smaller(sK7, X2, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, X2, sK8)) | sP0(sK8, sK7, X2)) & element(sK8, the_carrier(sK7)))), introduced(choice_axiom, [])).
fof(f45, plain, (? [X2] : (((~ ex_sup_of_relstr_set(sK7, X2) | ~ (join_on_relstr(sK7, X2) = sK8)) & ! [X3] : (related(sK7, sK8, X3) | ~ relstr_set_smaller(sK7, X2, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, X2, sK8)) | sP0(sK8, sK7, X2)) => (((~ ex_sup_of_relstr_set(sK7, sK9) | ~ (sK8 = join_on_relstr(sK7, sK9))) & ! [X3] : (related(sK7, sK8, X3) | ~ relstr_set_smaller(sK7, sK9, X3) | ~ element(X3, the_carrier(sK7))) & relstr_set_smaller(sK7, sK9, sK8)) | sP0(sK8, sK7, sK9))), introduced(choice_axiom, [])).
fof(f24, plain, ? [X0] : (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(X0, X2) | ~ (join_on_relstr(X0, X2) = X1)) & ! [X3] : (related(X0, X1, X3) | ~ relstr_set_smaller(X0, X2, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X2, X1)) | sP0(X1, X0, X2)) & element(X1, the_carrier(X0))) & rel_str(X0) & antisymmetric_relstr(X0)), inference(definition_folding, [], [f22, e23])).
fof(f23, plain, ! [X1, X0, X2] : (((? [X4] : (~ related(X0, X1, X4) & relstr_set_smaller(X0, X2, X4) & element(X4, the_carrier(X0))) | ~ relstr_set_smaller(X0, X2, X1)) & ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)) | ~ sP0(X1, X0, X2)), inference(usedef, [], [e23])).
fof(e23, plain, ! [X1, X0, X2] : (sP0(X1, X0, X2) <=> ((? [X4] : (~ related(X0, X1, X4) & relstr_set_smaller(X0, X2, X4) & element(X4, the_carrier(X0))) | ~ relstr_set_smaller(X0, X2, X1)) & ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f22, plain, ? [X0] : (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(X0, X2) | ~ (join_on_relstr(X0, X2) = X1)) & ! [X3] : (related(X0, X1, X3) | ~ relstr_set_smaller(X0, X2, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X2, X1)) | ((? [X4] : (~ related(X0, X1, X4) & relstr_set_smaller(X0, X2, X4) & element(X4, the_carrier(X0))) | ~ relstr_set_smaller(X0, X2, X1)) & ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1))) & element(X1, the_carrier(X0))) & rel_str(X0) & antisymmetric_relstr(X0)), inference(flattening, [], [f21])).
fof(f21, plain, ? [X0] : (? [X1] : (? [X2] : (((~ ex_sup_of_relstr_set(X0, X2) | ~ (join_on_relstr(X0, X2) = X1)) & (! [X3] : ((related(X0, X1, X3) | ~ relstr_set_smaller(X0, X2, X3)) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X2, X1))) | ((? [X4] : ((~ related(X0, X1, X4) & relstr_set_smaller(X0, X2, X4)) & element(X4, the_carrier(X0))) | ~ relstr_set_smaller(X0, X2, X1)) & (ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)))) & element(X1, the_carrier(X0))) & (rel_str(X0) & antisymmetric_relstr(X0))), inference(ennf_transformation, [], [f13])).
fof(f13, plain, ~ ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (element(X1, the_carrier(X0)) => ! [X2] : (((! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X3) => related(X0, X1, X3))) & relstr_set_smaller(X0, X2, X1)) => (ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1))) & ((ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)) => (! [X4] : (element(X4, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X4) => related(X0, X1, X4))) & relstr_set_smaller(X0, X2, X1)))))), inference(rectify, [], [f12])).
fof(f12, plain, ~ ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (element(X1, the_carrier(X0)) => ! [X2] : (((! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X3) => related(X0, X1, X3))) & relstr_set_smaller(X0, X2, X1)) => (ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1))) & ((ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)) => (! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X3) => related(X0, X1, X3))) & relstr_set_smaller(X0, X2, X1)))))), inference(negated_conjecture, [], [f11])).
fof(f11, plain, ~ ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (element(X1, the_carrier(X0)) => ! [X2] : (((! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X3) => related(X0, X1, X3))) & relstr_set_smaller(X0, X2, X1)) => (ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1))) & ((ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)) => (! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X2, X3) => related(X0, X1, X3))) & relstr_set_smaller(X0, X2, X1)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU359+1.p', t30_yellow_0)).
fof(f240, plain, (relstr_set_smaller(sK7, sK9, sK8) | ~ rel_str(sK7) | (~ spl10_2 | ~ spl10_3)), inference(subsumption_resolution, [], [f219, f84])).
fof(f84, plain, (ex_sup_of_relstr_set(sK7, sK9) | ~ spl10_3), inference(avatar_component_clause, [], [f83])).
fof(f219, plain, (relstr_set_smaller(sK7, sK9, sK8) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | ~ spl10_2), inference(superposition, [], [f117, f80])).
fof(f80, plain, ((sK8 = join_on_relstr(sK7, sK9)) | ~ spl10_2), inference(avatar_component_clause, [], [f79])).
fof(f117, plain, ! [X0, X1] : (relstr_set_smaller(X0, X1, join_on_relstr(X0, X1)) | ~ ex_sup_of_relstr_set(X0, X1) | ~ rel_str(X0)), inference(duplicate_literal_removal, [], [f116])).
fof(f116, plain, ! [X0, X1] : (~ ex_sup_of_relstr_set(X0, X1) | relstr_set_smaller(X0, X1, join_on_relstr(X0, X1)) | ~ rel_str(X0) | ~ rel_str(X0)), inference(resolution, [], [f73, f52])).
fof(f52, plain, ! [X0, X1] : (element(join_on_relstr(X0, X1), the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, ! [X0, X1] : (element(join_on_relstr(X0, X1), the_carrier(X0)) | ~ rel_str(X0)), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1] : (rel_str(X0) => element(join_on_relstr(X0, X1), the_carrier(X0))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU359+1.p', dt_k1_yellow_0)).
fof(f73, plain, ! [X0, X1] : (~ element(join_on_relstr(X0, X1), the_carrier(X0)) | ~ ex_sup_of_relstr_set(X0, X1) | relstr_set_smaller(X0, X1, join_on_relstr(X0, X1)) | ~ rel_str(X0)), inference(equality_resolution, [], [f47])).
fof(f47, plain, ! [X2, X0, X1] : (relstr_set_smaller(X0, X1, X2) | ~ (join_on_relstr(X0, X1) = X2) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (! [X1, X2] : ((((join_on_relstr(X0, X1) = X2) | (~ related(X0, X2, sK1(X0, X1, X2)) & relstr_set_smaller(X0, X1, sK1(X0, X1, X2)) & element(sK1(X0, X1, X2), the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2)) & ((! [X4] : (related(X0, X2, X4) | ~ relstr_set_smaller(X0, X1, X4) | ~ element(X4, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2)) | ~ (join_on_relstr(X0, X1) = X2))) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1])], [f27, f28])).
fof(f28, plain, ! [X2, X1, X0] : (? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) => (~ related(X0, X2, sK1(X0, X1, X2)) & relstr_set_smaller(X0, X1, sK1(X0, X1, X2)) & element(sK1(X0, X1, X2), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f27, plain, ! [X0] : (! [X1, X2] : ((((join_on_relstr(X0, X1) = X2) | ? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2)) & ((! [X4] : (related(X0, X2, X4) | ~ relstr_set_smaller(X0, X1, X4) | ~ element(X4, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2)) | ~ (join_on_relstr(X0, X1) = X2))) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(rectify, [], [f26])).
fof(f26, plain, ! [X0] : (! [X1, X2] : ((((join_on_relstr(X0, X1) = X2) | ? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2)) & ((! [X3] : (related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2)) | ~ (join_on_relstr(X0, X1) = X2))) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(flattening, [], [f25])).
fof(f25, plain, ! [X0] : (! [X1, X2] : ((((join_on_relstr(X0, X1) = X2) | (? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2))) & ((! [X3] : (related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2)) | ~ (join_on_relstr(X0, X1) = X2))) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ! [X0] : (! [X1, X2] : (((join_on_relstr(X0, X1) = X2) <=> (! [X3] : (related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2))) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(flattening, [], [f16])).
fof(f16, plain, ! [X0] : (! [X1, X2] : ((((join_on_relstr(X0, X1) = X2) <=> (! [X3] : ((related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3)) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2))) | ~ ex_sup_of_relstr_set(X0, X1)) | ~ element(X2, the_carrier(X0))) | ~ rel_str(X0)), inference(ennf_transformation, [], [f1])).
fof(f1, plain, ! [X0] : (rel_str(X0) => ! [X1, X2] : (element(X2, the_carrier(X0)) => (ex_sup_of_relstr_set(X0, X1) => ((join_on_relstr(X0, X1) = X2) <=> (! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X1, X3) => related(X0, X2, X3))) & relstr_set_smaller(X0, X1, X2)))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU359+1.p', d9_yellow_0)).
fof(f237, plain, (spl10_4 | ~ spl10_2 | ~ spl10_3), inference(avatar_split_clause, [], [f236, f83, f79, f88])).
fof(f88, plain, (spl10_4 <=> ! [X3] : (related(sK7, sK8, X3) | ~ element(X3, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X3))), introduced(avatar_definition, [new_symbols(naming, [spl10_4])])).
fof(f236, plain, (! [X0] : (related(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0)) | (~ spl10_2 | ~ spl10_3)), inference(subsumption_resolution, [], [f235, f67])).
fof(f235, plain, (! [X0] : (related(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ rel_str(sK7)) | (~ spl10_2 | ~ spl10_3)), inference(subsumption_resolution, [], [f217, f84])).
fof(f217, plain, (! [X0] : (related(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ rel_str(sK7)) | ~ spl10_2), inference(superposition, [], [f177, f80])).
fof(f177, plain, ! [X2, X0, X1] : (related(X0, join_on_relstr(X0, X1), X2) | ~ element(X2, the_carrier(X0)) | ~ ex_sup_of_relstr_set(X0, X1) | ~ relstr_set_smaller(X0, X1, X2) | ~ rel_str(X0)), inference(duplicate_literal_removal, [], [f176])).
fof(f176, plain, ! [X2, X0, X1] : (~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ ex_sup_of_relstr_set(X0, X1) | related(X0, join_on_relstr(X0, X1), X2) | ~ rel_str(X0) | ~ rel_str(X0)), inference(resolution, [], [f72, f52])).
fof(f72, plain, ! [X4, X0, X1] : (~ element(join_on_relstr(X0, X1), the_carrier(X0)) | ~ relstr_set_smaller(X0, X1, X4) | ~ element(X4, the_carrier(X0)) | ~ ex_sup_of_relstr_set(X0, X1) | related(X0, join_on_relstr(X0, X1), X4) | ~ rel_str(X0)), inference(equality_resolution, [], [f48])).
fof(f48, plain, ! [X4, X2, X0, X1] : (related(X0, X2, X4) | ~ relstr_set_smaller(X0, X1, X4) | ~ element(X4, the_carrier(X0)) | ~ (join_on_relstr(X0, X1) = X2) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f29])).
fof(f210, plain, (~ spl10_5 | spl10_2 | ~ spl10_3 | ~ spl10_4), inference(avatar_split_clause, [], [f209, f88, f83, f79, f92])).
fof(f209, plain, ((sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f208, f67])).
fof(f208, plain, ((sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ rel_str(sK7) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f201, f84])).
fof(f201, plain, ((sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f200, f68])).
fof(f68, plain, element(sK8, the_carrier(sK7)), inference(cnf_transformation, [], [f46])).
fof(f200, plain, (~ element(sK8, the_carrier(sK7)) | (sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | (~ spl10_3 | ~ spl10_4)), inference(duplicate_literal_removal, [], [f199])).
fof(f199, plain, (~ element(sK8, the_carrier(sK7)) | (sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | (sK8 = join_on_relstr(sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ element(sK8, the_carrier(sK7)) | ~ rel_str(sK7) | (~ spl10_3 | ~ spl10_4)), inference(resolution, [], [f198, f51])).
fof(f51, plain, ! [X2, X0, X1] : (~ related(X0, X2, sK1(X0, X1, X2)) | (join_on_relstr(X0, X1) = X2) | ~ relstr_set_smaller(X0, X1, X2) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f29])).
fof(f198, plain, (! [X0] : (related(sK7, sK8, sK1(sK7, sK9, X0)) | ~ element(X0, the_carrier(sK7)) | (join_on_relstr(sK7, sK9) = X0) | ~ relstr_set_smaller(sK7, sK9, X0)) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f197, f67])).
fof(f197, plain, (! [X0] : (~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | (join_on_relstr(sK7, sK9) = X0) | related(sK7, sK8, sK1(sK7, sK9, X0)) | ~ rel_str(sK7)) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f196, f84])).
fof(f196, plain, (! [X0] : (~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | (join_on_relstr(sK7, sK9) = X0) | related(sK7, sK8, sK1(sK7, sK9, X0)) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7)) | (~ spl10_3 | ~ spl10_4)), inference(duplicate_literal_removal, [], [f195])).
fof(f195, plain, (! [X0] : (~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | (join_on_relstr(sK7, sK9) = X0) | related(sK7, sK8, sK1(sK7, sK9, X0)) | (join_on_relstr(sK7, sK9) = X0) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7)) | (~ spl10_3 | ~ spl10_4)), inference(resolution, [], [f172, f49])).
fof(f49, plain, ! [X2, X0, X1] : (element(sK1(X0, X1, X2), the_carrier(X0)) | (join_on_relstr(X0, X1) = X2) | ~ relstr_set_smaller(X0, X1, X2) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f29])).
fof(f172, plain, (! [X0] : (~ element(sK1(sK7, sK9, X0), the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | (join_on_relstr(sK7, sK9) = X0) | related(sK7, sK8, sK1(sK7, sK9, X0))) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f171, f67])).
fof(f171, plain, (! [X0] : ((join_on_relstr(sK7, sK9) = X0) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ element(sK1(sK7, sK9, X0), the_carrier(sK7)) | related(sK7, sK8, sK1(sK7, sK9, X0))) | (~ spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f170, f84])).
fof(f170, plain, (! [X0] : ((join_on_relstr(sK7, sK9) = X0) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ ex_sup_of_relstr_set(sK7, sK9) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ element(sK1(sK7, sK9, X0), the_carrier(sK7)) | related(sK7, sK8, sK1(sK7, sK9, X0))) | ~ spl10_4), inference(resolution, [], [f50, f89])).
fof(f89, plain, (! [X3] : (~ relstr_set_smaller(sK7, sK9, X3) | ~ element(X3, the_carrier(sK7)) | related(sK7, sK8, X3)) | ~ spl10_4), inference(avatar_component_clause, [], [f88])).
fof(f50, plain, ! [X2, X0, X1] : (relstr_set_smaller(X0, X1, sK1(X0, X1, X2)) | (join_on_relstr(X0, X1) = X2) | ~ relstr_set_smaller(X0, X1, X2) | ~ ex_sup_of_relstr_set(X0, X1) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f29])).
fof(f207, plain, (~ spl10_1 | ~ spl10_5 | ~ spl10_4), inference(avatar_split_clause, [], [f182, f88, f92, f75])).
fof(f75, plain, (spl10_1 <=> sP0(sK8, sK7, sK9)), introduced(avatar_definition, [new_symbols(naming, [spl10_1])])).
fof(f182, plain, (~ relstr_set_smaller(sK7, sK9, sK8) | ~ sP0(sK8, sK7, sK9) | ~ spl10_4), inference(duplicate_literal_removal, [], [f181])).
fof(f181, plain, (~ relstr_set_smaller(sK7, sK9, sK8) | ~ sP0(sK8, sK7, sK9) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ sP0(sK8, sK7, sK9) | ~ spl10_4), inference(resolution, [], [f169, f65])).
fof(f65, plain, ! [X2, X0, X1] : (~ related(X1, X0, sK6(X0, X1, X2)) | ~ relstr_set_smaller(X1, X2, X0) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, ! [X0, X1, X2] : ((((~ related(X1, X0, sK6(X0, X1, X2)) & relstr_set_smaller(X1, X2, sK6(X0, X1, X2)) & element(sK6(X0, X1, X2), the_carrier(X1))) | ~ relstr_set_smaller(X1, X2, X0)) & ex_sup_of_relstr_set(X1, X2) & (join_on_relstr(X1, X2) = X0)) | ~ sP0(X0, X1, X2)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK6])], [f40, f41])).
fof(f41, plain, ! [X2, X1, X0] : (? [X3] : (~ related(X1, X0, X3) & relstr_set_smaller(X1, X2, X3) & element(X3, the_carrier(X1))) => (~ related(X1, X0, sK6(X0, X1, X2)) & relstr_set_smaller(X1, X2, sK6(X0, X1, X2)) & element(sK6(X0, X1, X2), the_carrier(X1)))), introduced(choice_axiom, [])).
fof(f40, plain, ! [X0, X1, X2] : (((? [X3] : (~ related(X1, X0, X3) & relstr_set_smaller(X1, X2, X3) & element(X3, the_carrier(X1))) | ~ relstr_set_smaller(X1, X2, X0)) & ex_sup_of_relstr_set(X1, X2) & (join_on_relstr(X1, X2) = X0)) | ~ sP0(X0, X1, X2)), inference(rectify, [], [f39])).
fof(f39, plain, ! [X1, X0, X2] : (((? [X4] : (~ related(X0, X1, X4) & relstr_set_smaller(X0, X2, X4) & element(X4, the_carrier(X0))) | ~ relstr_set_smaller(X0, X2, X1)) & ex_sup_of_relstr_set(X0, X2) & (join_on_relstr(X0, X2) = X1)) | ~ sP0(X1, X0, X2)), inference(nnf_transformation, [], [f23])).
fof(f169, plain, (! [X1] : (related(sK7, sK8, sK6(X1, sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, X1) | ~ sP0(X1, sK7, sK9)) | ~ spl10_4), inference(subsumption_resolution, [], [f156, f63])).
fof(f63, plain, ! [X2, X0, X1] : (element(sK6(X0, X1, X2), the_carrier(X1)) | ~ relstr_set_smaller(X1, X2, X0) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f42])).
fof(f156, plain, (! [X1] : (~ element(sK6(X1, sK7, sK9), the_carrier(sK7)) | related(sK7, sK8, sK6(X1, sK7, sK9)) | ~ relstr_set_smaller(sK7, sK9, X1) | ~ sP0(X1, sK7, sK9)) | ~ spl10_4), inference(resolution, [], [f89, f64])).
fof(f64, plain, ! [X2, X0, X1] : (relstr_set_smaller(X1, X2, sK6(X0, X1, X2)) | ~ relstr_set_smaller(X1, X2, X0) | ~ sP0(X0, X1, X2)), inference(cnf_transformation, [], [f42])).
fof(f146, plain, (~ spl10_1 | spl10_3), inference(avatar_contradiction_clause, [], [f145])).
fof(f145, plain, ($false | (~ spl10_1 | spl10_3)), inference(subsumption_resolution, [], [f142, f85])).
fof(f85, plain, (~ ex_sup_of_relstr_set(sK7, sK9) | spl10_3), inference(avatar_component_clause, [], [f83])).
fof(f142, plain, (ex_sup_of_relstr_set(sK7, sK9) | ~ spl10_1), inference(resolution, [], [f77, f62])).
fof(f62, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | ex_sup_of_relstr_set(X1, X2)), inference(cnf_transformation, [], [f42])).
fof(f77, plain, (sP0(sK8, sK7, sK9) | ~ spl10_1), inference(avatar_component_clause, [], [f75])).
fof(f144, plain, (~ spl10_1 | spl10_2), inference(avatar_contradiction_clause, [], [f143])).
fof(f143, plain, ($false | (~ spl10_1 | spl10_2)), inference(subsumption_resolution, [], [f141, f81])).
fof(f81, plain, (~ (sK8 = join_on_relstr(sK7, sK9)) | spl10_2), inference(avatar_component_clause, [], [f79])).
fof(f141, plain, ((sK8 = join_on_relstr(sK7, sK9)) | ~ spl10_1), inference(resolution, [], [f77, f61])).
fof(f61, plain, ! [X2, X0, X1] : (~ sP0(X0, X1, X2) | (join_on_relstr(X1, X2) = X0)), inference(cnf_transformation, [], [f42])).
fof(f140, plain, (spl10_3 | ~ spl10_4 | ~ spl10_5), inference(avatar_contradiction_clause, [], [f139])).
fof(f139, plain, ($false | (spl10_3 | ~ spl10_4 | ~ spl10_5)), inference(subsumption_resolution, [], [f138, f66])).
fof(f66, plain, antisymmetric_relstr(sK7), inference(cnf_transformation, [], [f46])).
fof(f138, plain, (~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4 | ~ spl10_5)), inference(subsumption_resolution, [], [f137, f67])).
fof(f137, plain, (~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4 | ~ spl10_5)), inference(subsumption_resolution, [], [f136, f85])).
fof(f136, plain, (ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4 | ~ spl10_5)), inference(subsumption_resolution, [], [f135, f68])).
fof(f135, plain, (~ element(sK8, the_carrier(sK7)) | ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4 | ~ spl10_5)), inference(subsumption_resolution, [], [f134, f94])).
fof(f94, plain, (relstr_set_smaller(sK7, sK9, sK8) | ~ spl10_5), inference(avatar_component_clause, [], [f92])).
fof(f134, plain, (~ relstr_set_smaller(sK7, sK9, sK8) | ~ element(sK8, the_carrier(sK7)) | ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4)), inference(duplicate_literal_removal, [], [f133])).
fof(f133, plain, (~ relstr_set_smaller(sK7, sK9, sK8) | ~ element(sK8, the_carrier(sK7)) | ex_sup_of_relstr_set(sK7, sK9) | ~ relstr_set_smaller(sK7, sK9, sK8) | ~ element(sK8, the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | (spl10_3 | ~ spl10_4)), inference(resolution, [], [f132, f60])).
fof(f60, plain, ! [X2, X0, X1] : (~ related(X0, X2, sK4(X0, X1, X2)) | ex_sup_of_relstr_set(X0, X1) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ! [X0] : (! [X1] : ((ex_sup_of_relstr_set(X0, X1) | ! [X2] : ((~ related(X0, X2, sK4(X0, X1, X2)) & relstr_set_smaller(X0, X1, sK4(X0, X1, X2)) & element(sK4(X0, X1, X2), the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & ((! [X5] : (related(X0, sK5(X0, X1), X5) | ~ relstr_set_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_set_smaller(X0, X1, sK5(X0, X1)) & element(sK5(X0, X1), the_carrier(X0))) | ~ ex_sup_of_relstr_set(X0, X1))) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5])], [f35, f37, f36])).
fof(f36, plain, ! [X2, X1, X0] : (? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) => (~ related(X0, X2, sK4(X0, X1, X2)) & relstr_set_smaller(X0, X1, sK4(X0, X1, X2)) & element(sK4(X0, X1, X2), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f37, plain, ! [X1, X0] : (? [X4] : (! [X5] : (related(X0, X4, X5) | ~ relstr_set_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_set_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) => (! [X5] : (related(X0, sK5(X0, X1), X5) | ~ relstr_set_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_set_smaller(X0, X1, sK5(X0, X1)) & element(sK5(X0, X1), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f35, plain, ! [X0] : (! [X1] : ((ex_sup_of_relstr_set(X0, X1) | ! [X2] : (? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & (? [X4] : (! [X5] : (related(X0, X4, X5) | ~ relstr_set_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_set_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ ex_sup_of_relstr_set(X0, X1))) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(rectify, [], [f34])).
fof(f34, plain, ! [X0] : (! [X1] : ((ex_sup_of_relstr_set(X0, X1) | ! [X2] : (? [X3] : (~ related(X0, X2, X3) & relstr_set_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & (? [X2] : (! [X3] : (related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2) & element(X2, the_carrier(X0))) | ~ ex_sup_of_relstr_set(X0, X1))) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (! [X1] : (ex_sup_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(flattening, [], [f19])).
fof(f19, plain, ! [X0] : (! [X1] : (ex_sup_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : ((related(X0, X2, X3) | ~ relstr_set_smaller(X0, X1, X3)) | ~ element(X3, the_carrier(X0))) & relstr_set_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) | (~ rel_str(X0) | ~ antisymmetric_relstr(X0))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (ex_sup_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (element(X3, the_carrier(X0)) => (relstr_set_smaller(X0, X1, X3) => related(X0, X2, X3))) & relstr_set_smaller(X0, X1, X2) & element(X2, the_carrier(X0))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU359+1.p', t15_yellow_0)).
fof(f132, plain, (! [X0] : (related(sK7, sK8, sK4(sK7, sK9, X0)) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7))) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f131, f66])).
fof(f131, plain, (! [X0] : (~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | related(sK7, sK8, sK4(sK7, sK9, X0)) | ~ antisymmetric_relstr(sK7)) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f130, f67])).
fof(f130, plain, (! [X0] : (~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | related(sK7, sK8, sK4(sK7, sK9, X0)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7)) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f129, f85])).
fof(f129, plain, (! [X0] : (~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | related(sK7, sK8, sK4(sK7, sK9, X0)) | ex_sup_of_relstr_set(sK7, sK9) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7)) | (spl10_3 | ~ spl10_4)), inference(duplicate_literal_removal, [], [f128])).
fof(f128, plain, (! [X0] : (~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | related(sK7, sK8, sK4(sK7, sK9, X0)) | ex_sup_of_relstr_set(sK7, sK9) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7)) | (spl10_3 | ~ spl10_4)), inference(resolution, [], [f124, f58])).
fof(f58, plain, ! [X2, X0, X1] : (element(sK4(X0, X1, X2), the_carrier(X0)) | ex_sup_of_relstr_set(X0, X1) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(cnf_transformation, [], [f38])).
fof(f124, plain, (! [X0] : (~ element(sK4(sK7, sK9, X0), the_carrier(sK7)) | ~ element(X0, the_carrier(sK7)) | ~ relstr_set_smaller(sK7, sK9, X0) | related(sK7, sK8, sK4(sK7, sK9, X0))) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f123, f66])).
fof(f123, plain, (! [X0] : (~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | ~ antisymmetric_relstr(sK7) | ~ element(sK4(sK7, sK9, X0), the_carrier(sK7)) | related(sK7, sK8, sK4(sK7, sK9, X0))) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f122, f67])).
fof(f122, plain, (! [X0] : (~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | ~ element(sK4(sK7, sK9, X0), the_carrier(sK7)) | related(sK7, sK8, sK4(sK7, sK9, X0))) | (spl10_3 | ~ spl10_4)), inference(subsumption_resolution, [], [f121, f85])).
fof(f121, plain, (! [X0] : (ex_sup_of_relstr_set(sK7, sK9) | ~ relstr_set_smaller(sK7, sK9, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | ~ element(sK4(sK7, sK9, X0), the_carrier(sK7)) | related(sK7, sK8, sK4(sK7, sK9, X0))) | ~ spl10_4), inference(resolution, [], [f59, f89])).
fof(f59, plain, ! [X2, X0, X1] : (relstr_set_smaller(X0, X1, sK4(X0, X1, X2)) | ex_sup_of_relstr_set(X0, X1) | ~ relstr_set_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(cnf_transformation, [], [f38])).
fof(f95, plain, (spl10_1 | spl10_5), inference(avatar_split_clause, [], [f69, f92, f75])).
fof(f69, plain, (relstr_set_smaller(sK7, sK9, sK8) | sP0(sK8, sK7, sK9)), inference(cnf_transformation, [], [f46])).
fof(f90, plain, (spl10_1 | spl10_4), inference(avatar_split_clause, [], [f70, f88, f75])).
fof(f70, plain, ! [X3] : (related(sK7, sK8, X3) | ~ relstr_set_smaller(sK7, sK9, X3) | ~ element(X3, the_carrier(sK7)) | sP0(sK8, sK7, sK9)), inference(cnf_transformation, [], [f46])).
fof(f86, plain, (spl10_1 | ~ spl10_2 | ~ spl10_3), inference(avatar_split_clause, [], [f71, f83, f79, f75])).
fof(f71, plain, (~ ex_sup_of_relstr_set(sK7, sK9) | ~ (sK8 = join_on_relstr(sK7, sK9)) | sP0(sK8, sK7, sK9)), inference(cnf_transformation, [], [f46])).