fof(f285, plain, $false, inference(avatar_sat_refutation, [], [f76, f80, f84, f88, f93, f98, f181, f205, f284])).
fof(f284, plain, (~ spl11_1 | ~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11), inference(avatar_contradiction_clause, [], [f283])).
fof(f283, plain, ($false | (~ spl11_1 | ~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11)), inference(subsumption_resolution, [], [f282, f61])).
fof(f61, plain, rel_str(sK7), inference(cnf_transformation, [], [f43])).
fof(f43, plain, (((! [X2] : ((~ related(sK7, sK9(X2), X2) & relstr_element_smaller(sK7, sK8, sK9(X2)) & element(sK9(X2), the_carrier(sK7))) | ~ relstr_element_smaller(sK7, sK8, X2) | ~ element(X2, the_carrier(sK7))) | ~ ex_inf_of_relstr_set(sK7, sK8)) & ((! [X5] : (related(sK7, X5, sK10) | ~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, sK8, sK10) & element(sK10, the_carrier(sK7))) | ex_inf_of_relstr_set(sK7, sK8))) & rel_str(sK7) & antisymmetric_relstr(sK7)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8, sK9, sK10])], [f38, f42, f41, f40, f39])).
fof(f39, plain, (? [X0] : (? [X1] : ((! [X2] : (? [X3] : (~ related(X0, X3, X2) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1)) & (? [X4] : (! [X5] : (related(X0, X5, X4) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ex_inf_of_relstr_set(X0, X1))) & rel_str(X0) & antisymmetric_relstr(X0)) => (? [X1] : ((! [X2] : (? [X3] : (~ related(sK7, X3, X2) & relstr_element_smaller(sK7, X1, X3) & element(X3, the_carrier(sK7))) | ~ relstr_element_smaller(sK7, X1, X2) | ~ element(X2, the_carrier(sK7))) | ~ ex_inf_of_relstr_set(sK7, X1)) & (? [X4] : (! [X5] : (related(sK7, X5, X4) | ~ relstr_element_smaller(sK7, X1, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, X1, X4) & element(X4, the_carrier(sK7))) | ex_inf_of_relstr_set(sK7, X1))) & rel_str(sK7) & antisymmetric_relstr(sK7))), introduced(choice_axiom, [])).
fof(f40, plain, (? [X1] : ((! [X2] : (? [X3] : (~ related(sK7, X3, X2) & relstr_element_smaller(sK7, X1, X3) & element(X3, the_carrier(sK7))) | ~ relstr_element_smaller(sK7, X1, X2) | ~ element(X2, the_carrier(sK7))) | ~ ex_inf_of_relstr_set(sK7, X1)) & (? [X4] : (! [X5] : (related(sK7, X5, X4) | ~ relstr_element_smaller(sK7, X1, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, X1, X4) & element(X4, the_carrier(sK7))) | ex_inf_of_relstr_set(sK7, X1))) => ((! [X2] : (? [X3] : (~ related(sK7, X3, X2) & relstr_element_smaller(sK7, sK8, X3) & element(X3, the_carrier(sK7))) | ~ relstr_element_smaller(sK7, sK8, X2) | ~ element(X2, the_carrier(sK7))) | ~ ex_inf_of_relstr_set(sK7, sK8)) & (? [X4] : (! [X5] : (related(sK7, X5, X4) | ~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, sK8, X4) & element(X4, the_carrier(sK7))) | ex_inf_of_relstr_set(sK7, sK8)))), introduced(choice_axiom, [])).
fof(f41, plain, ! [X2] : (? [X3] : (~ related(sK7, X3, X2) & relstr_element_smaller(sK7, sK8, X3) & element(X3, the_carrier(sK7))) => (~ related(sK7, sK9(X2), X2) & relstr_element_smaller(sK7, sK8, sK9(X2)) & element(sK9(X2), the_carrier(sK7)))), introduced(choice_axiom, [])).
fof(f42, plain, (? [X4] : (! [X5] : (related(sK7, X5, X4) | ~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, sK8, X4) & element(X4, the_carrier(sK7))) => (! [X5] : (related(sK7, X5, sK10) | ~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7))) & relstr_element_smaller(sK7, sK8, sK10) & element(sK10, the_carrier(sK7)))), introduced(choice_axiom, [])).
fof(f38, plain, ? [X0] : (? [X1] : ((! [X2] : (? [X3] : (~ related(X0, X3, X2) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1)) & (? [X4] : (! [X5] : (related(X0, X5, X4) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ex_inf_of_relstr_set(X0, X1))) & rel_str(X0) & antisymmetric_relstr(X0)), inference(rectify, [], [f37])).
fof(f37, plain, ? [X0] : (? [X1] : ((! [X2] : (? [X3] : (~ related(X0, X3, X2) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1)) & (? [X2] : (! [X3] : (related(X0, X3, X2) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))) | ex_inf_of_relstr_set(X0, X1))) & rel_str(X0) & antisymmetric_relstr(X0)), inference(nnf_transformation, [], [f18])).
fof(f18, plain, ? [X0] : (? [X1] : ~ (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (related(X0, X3, X2) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) & rel_str(X0) & antisymmetric_relstr(X0)), inference(flattening, [], [f17])).
fof(f17, plain, ? [X0] : (? [X1] : ~ (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : ((related(X0, X3, X2) | ~ relstr_element_smaller(X0, X1, X3)) | ~ element(X3, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) & (rel_str(X0) & antisymmetric_relstr(X0))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ~ ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (element(X3, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X3) => related(X0, X3, X2))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))))), inference(negated_conjecture, [], [f9])).
fof(f9, plain, ~ ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (element(X3, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X3) => related(X0, X3, X2))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU357+1.p', t16_yellow_0)).
fof(f282, plain, (~ rel_str(sK7) | (~ spl11_1 | ~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11)), inference(subsumption_resolution, [], [f281, f71])).
fof(f71, plain, (ex_inf_of_relstr_set(sK7, sK8) | ~ spl11_1), inference(avatar_component_clause, [], [f70])).
fof(f70, plain, (spl11_1 <=> ex_inf_of_relstr_set(sK7, sK8)), introduced(avatar_definition, [new_symbols(naming, [spl11_1])])).
fof(f281, plain, (~ ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7) | (~ spl11_1 | ~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11)), inference(resolution, [], [f275, f52])).
fof(f52, plain, ! [X0, X1] : (relstr_element_smaller(X0, X1, sK4(X0, X1)) | ~ ex_inf_of_relstr_set(X0, X1) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ! [X0] : (! [X1] : ((ex_inf_of_relstr_set(X0, X1) | ! [X2] : (~ sP0(X2, X0, X1) | (~ related(X0, sK3(X0, X1, X2), X2) & relstr_element_smaller(X0, X1, sK3(X0, X1, X2)) & element(sK3(X0, X1, X2), the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & ((sP0(sK4(X0, X1), X0, X1) & ! [X5] : (related(X0, X5, sK4(X0, X1)) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, sK4(X0, X1)) & element(sK4(X0, X1), the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1))) | ~ rel_str(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4])], [f29, f31, f30])).
fof(f30, plain, ! [X2, X1, X0] : (? [X3] : (~ related(X0, X3, X2) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) => (~ related(X0, sK3(X0, X1, X2), X2) & relstr_element_smaller(X0, X1, sK3(X0, X1, X2)) & element(sK3(X0, X1, X2), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X1, X0] : (? [X4] : (sP0(X4, X0, X1) & ! [X5] : (related(X0, X5, X4) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) => (sP0(sK4(X0, X1), X0, X1) & ! [X5] : (related(X0, X5, sK4(X0, X1)) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, sK4(X0, X1)) & element(sK4(X0, X1), the_carrier(X0)))), introduced(choice_axiom, [])).
fof(f29, plain, ! [X0] : (! [X1] : ((ex_inf_of_relstr_set(X0, X1) | ! [X2] : (~ sP0(X2, X0, X1) | ? [X3] : (~ related(X0, X3, X2) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & (? [X4] : (sP0(X4, X0, X1) & ! [X5] : (related(X0, X5, X4) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1))) | ~ rel_str(X0)), inference(rectify, [], [f28])).
fof(f28, plain, ! [X0] : (! [X1] : ((ex_inf_of_relstr_set(X0, X1) | ! [X2] : (~ sP0(X2, X0, X1) | ? [X5] : (~ related(X0, X5, X2) & relstr_element_smaller(X0, X1, X5) & element(X5, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)))) & (? [X2] : (sP0(X2, X0, X1) & ! [X5] : (related(X0, X5, X2) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))) | ~ ex_inf_of_relstr_set(X0, X1))) | ~ rel_str(X0)), inference(nnf_transformation, [], [f22])).
fof(f22, plain, ! [X0] : (! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (sP0(X2, X0, X1) & ! [X5] : (related(X0, X5, X2) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) | ~ rel_str(X0)), inference(definition_folding, [], [f16, e21])).
fof(f21, plain, ! [X2, X0, X1] : (sP0(X2, X0, X1) <=> ! [X3] : ((X2 = X3) | ? [X4] : (~ related(X0, X4, X3) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0)))), inference(usedef, [], [e21])).
fof(e21, plain, ! [X2, X0, X1] : (sP0(X2, X0, X1) <=> ! [X3] : ((X2 = X3) | ? [X4] : (~ related(X0, X4, X3) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f16, plain, ! [X0] : (! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : ((X2 = X3) | ? [X4] : (~ related(X0, X4, X3) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) & ! [X5] : (related(X0, X5, X2) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) | ~ rel_str(X0)), inference(flattening, [], [f15])).
fof(f15, plain, ! [X0] : (! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (((X2 = X3) | (? [X4] : ((~ related(X0, X4, X3) & relstr_element_smaller(X0, X1, X4)) & element(X4, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X3))) | ~ element(X3, the_carrier(X0))) & ! [X5] : ((related(X0, X5, X2) | ~ relstr_element_smaller(X0, X1, X5)) | ~ element(X5, the_carrier(X0))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0)))) | ~ rel_str(X0)), inference(ennf_transformation, [], [f12])).
fof(f12, plain, ! [X0] : (rel_str(X0) => ! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (element(X3, the_carrier(X0)) => ((! [X4] : (element(X4, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X4) => related(X0, X4, X3))) & relstr_element_smaller(X0, X1, X3)) => (X2 = X3))) & ! [X5] : (element(X5, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X5) => related(X0, X5, X2))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))))), inference(rectify, [], [f1])).
fof(f1, plain, ! [X0] : (rel_str(X0) => ! [X1] : (ex_inf_of_relstr_set(X0, X1) <=> ? [X2] : (! [X3] : (element(X3, the_carrier(X0)) => ((! [X4] : (element(X4, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X4) => related(X0, X4, X3))) & relstr_element_smaller(X0, X1, X3)) => (X2 = X3))) & ! [X3] : (element(X3, the_carrier(X0)) => (relstr_element_smaller(X0, X1, X3) => related(X0, X3, X2))) & relstr_element_smaller(X0, X1, X2) & element(X2, the_carrier(X0))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU357+1.p', d8_yellow_0)).
fof(f275, plain, (~ relstr_element_smaller(sK7, sK8, sK4(sK7, sK8)) | (~ spl11_1 | ~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11)), inference(subsumption_resolution, [], [f274, f71])).
fof(f274, plain, (~ relstr_element_smaller(sK7, sK8, sK4(sK7, sK8)) | ~ ex_inf_of_relstr_set(sK7, sK8) | (~ spl11_2 | ~ spl11_3 | ~ spl11_4 | ~ spl11_11)), inference(subsumption_resolution, [], [f273, f199])).
fof(f199, plain, (element(sK4(sK7, sK8), the_carrier(sK7)) | ~ spl11_11), inference(avatar_component_clause, [], [f198])).
fof(f198, plain, (spl11_11 <=> element(sK4(sK7, sK8), the_carrier(sK7))), introduced(avatar_definition, [new_symbols(naming, [spl11_11])])).
fof(f273, plain, (~ relstr_element_smaller(sK7, sK8, sK4(sK7, sK8)) | ~ element(sK4(sK7, sK8), the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, sK8) | (~ spl11_2 | ~ spl11_3 | ~ spl11_4)), inference(duplicate_literal_removal, [], [f272])).
fof(f272, plain, (~ relstr_element_smaller(sK7, sK8, sK4(sK7, sK8)) | ~ element(sK4(sK7, sK8), the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, sK8) | ~ element(sK4(sK7, sK8), the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, sK4(sK7, sK8)) | (~ spl11_2 | ~ spl11_3 | ~ spl11_4)), inference(resolution, [], [f210, f79])).
fof(f79, plain, (! [X2] : (relstr_element_smaller(sK7, sK8, sK9(X2)) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2)) | ~ spl11_3), inference(avatar_component_clause, [], [f78])).
fof(f78, plain, (spl11_3 <=> ! [X2] : (relstr_element_smaller(sK7, sK8, sK9(X2)) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2))), introduced(avatar_definition, [new_symbols(naming, [spl11_3])])).
fof(f210, plain, (! [X2] : (~ relstr_element_smaller(sK7, X2, sK9(sK4(sK7, X2))) | ~ relstr_element_smaller(sK7, sK8, sK4(sK7, X2)) | ~ element(sK4(sK7, X2), the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, X2)) | (~ spl11_2 | ~ spl11_4)), inference(subsumption_resolution, [], [f209, f83])).
fof(f83, plain, (! [X2] : (element(sK9(X2), the_carrier(sK7)) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2)) | ~ spl11_4), inference(avatar_component_clause, [], [f82])).
fof(f82, plain, (spl11_4 <=> ! [X2] : (element(sK9(X2), the_carrier(sK7)) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2))), introduced(avatar_definition, [new_symbols(naming, [spl11_4])])).
fof(f209, plain, (! [X2] : (~ element(sK4(sK7, X2), the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, sK4(sK7, X2)) | ~ relstr_element_smaller(sK7, X2, sK9(sK4(sK7, X2))) | ~ element(sK9(sK4(sK7, X2)), the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, X2)) | ~ spl11_2), inference(subsumption_resolution, [], [f207, f61])).
fof(f207, plain, (! [X2] : (~ element(sK4(sK7, X2), the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, sK4(sK7, X2)) | ~ relstr_element_smaller(sK7, X2, sK9(sK4(sK7, X2))) | ~ element(sK9(sK4(sK7, X2)), the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, X2) | ~ rel_str(sK7)) | ~ spl11_2), inference(resolution, [], [f75, f53])).
fof(f53, plain, ! [X0, X5, X1] : (related(X0, X5, sK4(X0, X1)) | ~ relstr_element_smaller(X0, X1, X5) | ~ element(X5, the_carrier(X0)) | ~ ex_inf_of_relstr_set(X0, X1) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f75, plain, (! [X2] : (~ related(sK7, sK9(X2), X2) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2)) | ~ spl11_2), inference(avatar_component_clause, [], [f74])).
fof(f74, plain, (spl11_2 <=> ! [X2] : (~ related(sK7, sK9(X2), X2) | ~ element(X2, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2))), introduced(avatar_definition, [new_symbols(naming, [spl11_2])])).
fof(f205, plain, (~ spl11_1 | spl11_11), inference(avatar_contradiction_clause, [], [f204])).
fof(f204, plain, ($false | (~ spl11_1 | spl11_11)), inference(subsumption_resolution, [], [f203, f61])).
fof(f203, plain, (~ rel_str(sK7) | (~ spl11_1 | spl11_11)), inference(subsumption_resolution, [], [f202, f71])).
fof(f202, plain, (~ ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7) | spl11_11), inference(resolution, [], [f200, f51])).
fof(f51, plain, ! [X0, X1] : (element(sK4(X0, X1), the_carrier(X0)) | ~ ex_inf_of_relstr_set(X0, X1) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f200, plain, (~ element(sK4(sK7, sK8), the_carrier(sK7)) | spl11_11), inference(avatar_component_clause, [], [f198])).
fof(f181, plain, (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7), inference(avatar_contradiction_clause, [], [f180])).
fof(f180, plain, ($false | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f176, f150])).
fof(f150, plain, (~ sP0(sK10, sK7, sK8) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f149, f61])).
fof(f149, plain, (~ sP0(sK10, sK7, sK8) | ~ rel_str(sK7) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f148, f72])).
fof(f72, plain, (~ ex_inf_of_relstr_set(sK7, sK8) | spl11_1), inference(avatar_component_clause, [], [f70])).
fof(f148, plain, (~ sP0(sK10, sK7, sK8) | ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f147, f92])).
fof(f92, plain, (relstr_element_smaller(sK7, sK8, sK10) | ~ spl11_6), inference(avatar_component_clause, [], [f90])).
fof(f90, plain, (spl11_6 <=> relstr_element_smaller(sK7, sK8, sK10)), introduced(avatar_definition, [new_symbols(naming, [spl11_6])])).
fof(f147, plain, (~ sP0(sK10, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7) | (spl11_1 | ~ spl11_5 | ~ spl11_7)), inference(subsumption_resolution, [], [f146, f97])).
fof(f97, plain, (element(sK10, the_carrier(sK7)) | ~ spl11_7), inference(avatar_component_clause, [], [f95])).
fof(f95, plain, (spl11_7 <=> element(sK10, the_carrier(sK7))), introduced(avatar_definition, [new_symbols(naming, [spl11_7])])).
fof(f146, plain, (~ element(sK10, the_carrier(sK7)) | ~ sP0(sK10, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7) | (spl11_1 | ~ spl11_5)), inference(duplicate_literal_removal, [], [f144])).
fof(f144, plain, (~ element(sK10, the_carrier(sK7)) | ~ sP0(sK10, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ~ sP0(sK10, sK7, sK8) | ex_inf_of_relstr_set(sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ~ element(sK10, the_carrier(sK7)) | ~ rel_str(sK7) | (spl11_1 | ~ spl11_5)), inference(resolution, [], [f143, f57])).
fof(f57, plain, ! [X2, X0, X1] : (~ related(X0, sK3(X0, X1, X2), X2) | ~ sP0(X2, X0, X1) | ex_inf_of_relstr_set(X0, X1) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f143, plain, (! [X0] : (related(sK7, sK3(sK7, sK8, X0), sK10) | ~ element(X0, the_carrier(sK7)) | ~ sP0(X0, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, X0)) | (spl11_1 | ~ spl11_5)), inference(subsumption_resolution, [], [f142, f61])).
fof(f142, plain, (! [X0] : (~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ sP0(X0, sK7, sK8) | related(sK7, sK3(sK7, sK8, X0), sK10) | ~ rel_str(sK7)) | (spl11_1 | ~ spl11_5)), inference(subsumption_resolution, [], [f141, f72])).
fof(f141, plain, (! [X0] : (~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ sP0(X0, sK7, sK8) | related(sK7, sK3(sK7, sK8, X0), sK10) | ex_inf_of_relstr_set(sK7, sK8) | ~ rel_str(sK7)) | (spl11_1 | ~ spl11_5)), inference(duplicate_literal_removal, [], [f140])).
fof(f140, plain, (! [X0] : (~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ sP0(X0, sK7, sK8) | related(sK7, sK3(sK7, sK8, X0), sK10) | ~ sP0(X0, sK7, sK8) | ex_inf_of_relstr_set(sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7)) | (spl11_1 | ~ spl11_5)), inference(resolution, [], [f131, f55])).
fof(f55, plain, ! [X2, X0, X1] : (element(sK3(X0, X1, X2), the_carrier(X0)) | ~ sP0(X2, X0, X1) | ex_inf_of_relstr_set(X0, X1) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f131, plain, (! [X0] : (~ element(sK3(sK7, sK8, X0), the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ sP0(X0, sK7, sK8) | related(sK7, sK3(sK7, sK8, X0), sK10)) | (spl11_1 | ~ spl11_5)), inference(subsumption_resolution, [], [f130, f61])).
fof(f130, plain, (! [X0] : (~ sP0(X0, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ element(sK3(sK7, sK8, X0), the_carrier(sK7)) | related(sK7, sK3(sK7, sK8, X0), sK10)) | (spl11_1 | ~ spl11_5)), inference(subsumption_resolution, [], [f129, f72])).
fof(f129, plain, (! [X0] : (~ sP0(X0, sK7, sK8) | ex_inf_of_relstr_set(sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, X0) | ~ element(X0, the_carrier(sK7)) | ~ rel_str(sK7) | ~ element(sK3(sK7, sK8, X0), the_carrier(sK7)) | related(sK7, sK3(sK7, sK8, X0), sK10)) | ~ spl11_5), inference(resolution, [], [f56, f87])).
fof(f87, plain, (! [X5] : (~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7)) | related(sK7, X5, sK10)) | ~ spl11_5), inference(avatar_component_clause, [], [f86])).
fof(f86, plain, (spl11_5 <=> ! [X5] : (related(sK7, X5, sK10) | ~ element(X5, the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X5))), introduced(avatar_definition, [new_symbols(naming, [spl11_5])])).
fof(f56, plain, ! [X2, X0, X1] : (relstr_element_smaller(X0, X1, sK3(X0, X1, X2)) | ~ sP0(X2, X0, X1) | ex_inf_of_relstr_set(X0, X1) | ~ relstr_element_smaller(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ rel_str(X0)), inference(cnf_transformation, [], [f32])).
fof(f176, plain, (sP0(sK10, sK7, sK8) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(trivial_inequality_removal, [], [f175])).
fof(f175, plain, (~ (sK10 = sK10) | sP0(sK10, sK7, sK8) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(superposition, [], [f50, f161])).
fof(f161, plain, ((sK10 = sK1(sK10, sK7, sK8)) | (spl11_1 | ~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(resolution, [], [f150, f126])).
fof(f126, plain, (! [X0] : (sP0(X0, sK7, sK8) | (sK10 = sK1(X0, sK7, sK8))) | (~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f125, f97])).
fof(f125, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | sP0(X0, sK7, sK8) | ~ element(sK10, the_carrier(sK7))) | (~ spl11_5 | ~ spl11_6 | ~ spl11_7)), inference(subsumption_resolution, [], [f124, f92])).
fof(f124, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | sP0(X0, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ~ element(sK10, the_carrier(sK7))) | (~ spl11_5 | ~ spl11_7)), inference(duplicate_literal_removal, [], [f123])).
fof(f123, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | sP0(X0, sK7, sK8) | sP0(X0, sK7, sK8) | ~ relstr_element_smaller(sK7, sK8, sK10) | ~ element(sK10, the_carrier(sK7))) | (~ spl11_5 | ~ spl11_7)), inference(resolution, [], [f120, f49])).
fof(f49, plain, ! [X4, X2, X0, X1] : (related(X1, X4, sK1(X0, X1, X2)) | sP0(X0, X1, X2) | ~ relstr_element_smaller(X1, X2, X4) | ~ element(X4, the_carrier(X1))), inference(cnf_transformation, [], [f27])).
fof(f27, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | (~ (sK1(X0, X1, X2) = X0) & ! [X4] : (related(X1, X4, sK1(X0, X1, X2)) | ~ relstr_element_smaller(X1, X2, X4) | ~ element(X4, the_carrier(X1))) & relstr_element_smaller(X1, X2, sK1(X0, X1, X2)) & element(sK1(X0, X1, X2), the_carrier(X1)))) & (! [X5] : ((X0 = X5) | (~ related(X1, sK2(X1, X2, X5), X5) & relstr_element_smaller(X1, X2, sK2(X1, X2, X5)) & element(sK2(X1, X2, X5), the_carrier(X1))) | ~ relstr_element_smaller(X1, X2, X5) | ~ element(X5, the_carrier(X1))) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK1, sK2])], [f24, f26, f25])).
fof(f25, plain, ! [X2, X1, X0] : (? [X3] : (~ (X0 = X3) & ! [X4] : (related(X1, X4, X3) | ~ relstr_element_smaller(X1, X2, X4) | ~ element(X4, the_carrier(X1))) & relstr_element_smaller(X1, X2, X3) & element(X3, the_carrier(X1))) => (~ (sK1(X0, X1, X2) = X0) & ! [X4] : (related(X1, X4, sK1(X0, X1, X2)) | ~ relstr_element_smaller(X1, X2, X4) | ~ element(X4, the_carrier(X1))) & relstr_element_smaller(X1, X2, sK1(X0, X1, X2)) & element(sK1(X0, X1, X2), the_carrier(X1)))), introduced(choice_axiom, [])).
fof(f26, plain, ! [X5, X2, X1] : (? [X6] : (~ related(X1, X6, X5) & relstr_element_smaller(X1, X2, X6) & element(X6, the_carrier(X1))) => (~ related(X1, sK2(X1, X2, X5), X5) & relstr_element_smaller(X1, X2, sK2(X1, X2, X5)) & element(sK2(X1, X2, X5), the_carrier(X1)))), introduced(choice_axiom, [])).
fof(f24, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ? [X3] : (~ (X0 = X3) & ! [X4] : (related(X1, X4, X3) | ~ relstr_element_smaller(X1, X2, X4) | ~ element(X4, the_carrier(X1))) & relstr_element_smaller(X1, X2, X3) & element(X3, the_carrier(X1)))) & (! [X5] : ((X0 = X5) | ? [X6] : (~ related(X1, X6, X5) & relstr_element_smaller(X1, X2, X6) & element(X6, the_carrier(X1))) | ~ relstr_element_smaller(X1, X2, X5) | ~ element(X5, the_carrier(X1))) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f23])).
fof(f23, plain, ! [X2, X0, X1] : ((sP0(X2, X0, X1) | ? [X3] : (~ (X2 = X3) & ! [X4] : (related(X0, X4, X3) | ~ relstr_element_smaller(X0, X1, X4) | ~ element(X4, the_carrier(X0))) & relstr_element_smaller(X0, X1, X3) & element(X3, the_carrier(X0)))) & (! [X3] : ((X2 = X3) | ? [X4] : (~ related(X0, X4, X3) & relstr_element_smaller(X0, X1, X4) & element(X4, the_carrier(X0))) | ~ relstr_element_smaller(X0, X1, X3) | ~ element(X3, the_carrier(X0))) | ~ sP0(X2, X0, X1))), inference(nnf_transformation, [], [f21])).
fof(f120, plain, (! [X0] : (~ related(sK7, sK10, sK1(X0, sK7, sK8)) | (sK10 = sK1(X0, sK7, sK8)) | sP0(X0, sK7, sK8)) | (~ spl11_5 | ~ spl11_7)), inference(subsumption_resolution, [], [f119, f47])).
fof(f47, plain, ! [X2, X0, X1] : (element(sK1(X0, X1, X2), the_carrier(X1)) | sP0(X0, X1, X2)), inference(cnf_transformation, [], [f27])).
fof(f119, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | ~ related(sK7, sK10, sK1(X0, sK7, sK8)) | ~ element(sK1(X0, sK7, sK8), the_carrier(sK7)) | sP0(X0, sK7, sK8)) | (~ spl11_5 | ~ spl11_7)), inference(subsumption_resolution, [], [f118, f60])).
fof(f60, plain, antisymmetric_relstr(sK7), inference(cnf_transformation, [], [f43])).
fof(f118, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | ~ related(sK7, sK10, sK1(X0, sK7, sK8)) | ~ element(sK1(X0, sK7, sK8), the_carrier(sK7)) | ~ antisymmetric_relstr(sK7) | sP0(X0, sK7, sK8)) | (~ spl11_5 | ~ spl11_7)), inference(subsumption_resolution, [], [f117, f61])).
fof(f117, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | ~ related(sK7, sK10, sK1(X0, sK7, sK8)) | ~ element(sK1(X0, sK7, sK8), the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | sP0(X0, sK7, sK8)) | (~ spl11_5 | ~ spl11_7)), inference(subsumption_resolution, [], [f111, f97])).
fof(f111, plain, (! [X0] : ((sK10 = sK1(X0, sK7, sK8)) | ~ related(sK7, sK10, sK1(X0, sK7, sK8)) | ~ element(sK1(X0, sK7, sK8), the_carrier(sK7)) | ~ element(sK10, the_carrier(sK7)) | ~ rel_str(sK7) | ~ antisymmetric_relstr(sK7) | sP0(X0, sK7, sK8)) | ~ spl11_5), inference(resolution, [], [f68, f103])).
fof(f103, plain, (! [X0] : (related(sK7, sK1(X0, sK7, sK8), sK10) | sP0(X0, sK7, sK8)) | ~ spl11_5), inference(subsumption_resolution, [], [f100, f47])).
fof(f100, plain, (! [X0] : (~ element(sK1(X0, sK7, sK8), the_carrier(sK7)) | related(sK7, sK1(X0, sK7, sK8), sK10) | sP0(X0, sK7, sK8)) | ~ spl11_5), inference(resolution, [], [f87, f48])).
fof(f48, plain, ! [X2, X0, X1] : (relstr_element_smaller(X1, X2, sK1(X0, X1, X2)) | sP0(X0, X1, X2)), inference(cnf_transformation, [], [f27])).
fof(f68, plain, ! [X2, X0, X1] : (~ related(X0, X2, X1) | (X1 = X2) | ~ related(X0, X1, X2) | ~ element(X2, the_carrier(X0)) | ~ element(X1, the_carrier(X0)) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(cnf_transformation, [], [f20])).
fof(f20, plain, ! [X0] : (! [X1] : (! [X2] : ((X1 = X2) | ~ related(X0, X2, X1) | ~ related(X0, X1, X2) | ~ element(X2, the_carrier(X0))) | ~ element(X1, the_carrier(X0))) | ~ rel_str(X0) | ~ antisymmetric_relstr(X0)), inference(flattening, [], [f19])).
fof(f19, plain, ! [X0] : (! [X1] : (! [X2] : (((X1 = X2) | (~ related(X0, X2, X1) | ~ related(X0, X1, X2))) | ~ element(X2, the_carrier(X0))) | ~ element(X1, the_carrier(X0))) | (~ rel_str(X0) | ~ antisymmetric_relstr(X0))), inference(ennf_transformation, [], [f11])).
fof(f11, plain, ! [X0] : ((rel_str(X0) & antisymmetric_relstr(X0)) => ! [X1] : (element(X1, the_carrier(X0)) => ! [X2] : (element(X2, the_carrier(X0)) => ((related(X0, X2, X1) & related(X0, X1, X2)) => (X1 = X2))))), file('/home/ubuntu/library/tptp/Problems/SEU/SEU357+1.p', t25_orders_2)).
fof(f50, plain, ! [X2, X0, X1] : (~ (sK1(X0, X1, X2) = X0) | sP0(X0, X1, X2)), inference(cnf_transformation, [], [f27])).
fof(f98, plain, (spl11_1 | spl11_7), inference(avatar_split_clause, [], [f62, f95, f70])).
fof(f62, plain, (element(sK10, the_carrier(sK7)) | ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).
fof(f93, plain, (spl11_1 | spl11_6), inference(avatar_split_clause, [], [f63, f90, f70])).
fof(f63, plain, (relstr_element_smaller(sK7, sK8, sK10) | ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).
fof(f88, plain, (spl11_1 | spl11_5), inference(avatar_split_clause, [], [f64, f86, f70])).
fof(f64, plain, ! [X5] : (related(sK7, X5, sK10) | ~ relstr_element_smaller(sK7, sK8, X5) | ~ element(X5, the_carrier(sK7)) | ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).
fof(f84, plain, (~ spl11_1 | spl11_4), inference(avatar_split_clause, [], [f65, f82, f70])).
fof(f65, plain, ! [X2] : (element(sK9(X2), the_carrier(sK7)) | ~ relstr_element_smaller(sK7, sK8, X2) | ~ element(X2, the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).
fof(f80, plain, (~ spl11_1 | spl11_3), inference(avatar_split_clause, [], [f66, f78, f70])).
fof(f66, plain, ! [X2] : (relstr_element_smaller(sK7, sK8, sK9(X2)) | ~ relstr_element_smaller(sK7, sK8, X2) | ~ element(X2, the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).
fof(f76, plain, (~ spl11_1 | spl11_2), inference(avatar_split_clause, [], [f67, f74, f70])).
fof(f67, plain, ! [X2] : (~ related(sK7, sK9(X2), X2) | ~ relstr_element_smaller(sK7, sK8, X2) | ~ element(X2, the_carrier(sK7)) | ~ ex_inf_of_relstr_set(sK7, sK8)), inference(cnf_transformation, [], [f43])).