fof(f645, plain, $false, inference(avatar_sat_refutation, [], [f75, f186, f516, f520, f527, f529, f531, f584, f600, f632])).
fof(f632, plain, (spl7_20 | ~ spl7_30), inference(avatar_contradiction_clause, [], [f631])).
fof(f631, plain, ($false | (spl7_20 | ~ spl7_30)), inference(subsumption_resolution, [], [f616, f185])).
fof(f185, plain, (~ icext(sK6, uri_ex_name) | spl7_20), inference(avatar_component_clause, [], [f183])).
fof(f183, plain, (spl7_20 <=> icext(sK6, uri_ex_name)), introduced(avatar_definition, [new_symbols(naming, [spl7_20])])).
fof(f616, plain, (icext(sK6, uri_ex_name) | ~ spl7_30), inference(resolution, [], [f565, f78])).
fof(f78, plain, icext(uri_ex_PersonAttribute, uri_ex_name), inference(resolution, [], [f33, f65])).
fof(f65, plain, iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute), inference(cnf_transformation, [], [f32])).
fof(f32, plain, (iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)) & iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute) & iext(uri_owl_hasValue, sK6, uri_foaf_Person) & iext(uri_owl_onProperty, sK6, uri_rdfs_domain) & iext(uri_rdf_type, sK6, uri_owl_Restriction) & iext(uri_rdf_rest, sK5, uri_rdf_nil) & iext(uri_rdf_first, sK5, sK6) & iext(uri_rdf_rest, sK4, sK5) & iext(uri_rdf_first, sK4, uri_owl_FunctionalProperty) & iext(uri_rdf_rest, sK3, sK4) & iext(uri_rdf_first, sK3, uri_owl_DatatypeProperty) & iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, sK3) & iext(uri_rdf_type, uri_foaf_Person, uri_owl_Class)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5, sK6])], [f11, f31])).
fof(f31, plain, (? [X0, X1, X2, X3] : (iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)) & iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute) & iext(uri_owl_hasValue, X3, uri_foaf_Person) & iext(uri_owl_onProperty, X3, uri_rdfs_domain) & iext(uri_rdf_type, X3, uri_owl_Restriction) & iext(uri_rdf_rest, X2, uri_rdf_nil) & iext(uri_rdf_first, X2, X3) & iext(uri_rdf_rest, X1, X2) & iext(uri_rdf_first, X1, uri_owl_FunctionalProperty) & iext(uri_rdf_rest, X0, X1) & iext(uri_rdf_first, X0, uri_owl_DatatypeProperty) & iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, X0) & iext(uri_rdf_type, uri_foaf_Person, uri_owl_Class)) => (iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)) & iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute) & iext(uri_owl_hasValue, sK6, uri_foaf_Person) & iext(uri_owl_onProperty, sK6, uri_rdfs_domain) & iext(uri_rdf_type, sK6, uri_owl_Restriction) & iext(uri_rdf_rest, sK5, uri_rdf_nil) & iext(uri_rdf_first, sK5, sK6) & iext(uri_rdf_rest, sK4, sK5) & iext(uri_rdf_first, sK4, uri_owl_FunctionalProperty) & iext(uri_rdf_rest, sK3, sK4) & iext(uri_rdf_first, sK3, uri_owl_DatatypeProperty) & iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, sK3) & iext(uri_rdf_type, uri_foaf_Person, uri_owl_Class))), introduced(choice_axiom, [])).
fof(f11, plain, ? [X0, X1, X2, X3] : (iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)) & iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute) & iext(uri_owl_hasValue, X3, uri_foaf_Person) & iext(uri_owl_onProperty, X3, uri_rdfs_domain) & iext(uri_rdf_type, X3, uri_owl_Restriction) & iext(uri_rdf_rest, X2, uri_rdf_nil) & iext(uri_rdf_first, X2, X3) & iext(uri_rdf_rest, X1, X2) & iext(uri_rdf_first, X1, uri_owl_FunctionalProperty) & iext(uri_rdf_rest, X0, X1) & iext(uri_rdf_first, X0, uri_owl_DatatypeProperty) & iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, X0) & iext(uri_rdf_type, uri_foaf_Person, uri_owl_Class)), inference(rectify, [], [f7])).
fof(f7, plain, ? [X12, X13, X14, X15] : (iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)) & iext(uri_rdf_type, uri_ex_name, uri_ex_PersonAttribute) & iext(uri_owl_hasValue, X15, uri_foaf_Person) & iext(uri_owl_onProperty, X15, uri_rdfs_domain) & iext(uri_rdf_type, X15, uri_owl_Restriction) & iext(uri_rdf_rest, X14, uri_rdf_nil) & iext(uri_rdf_first, X14, X15) & iext(uri_rdf_rest, X13, X14) & iext(uri_rdf_first, X13, uri_owl_FunctionalProperty) & iext(uri_rdf_rest, X12, X13) & iext(uri_rdf_first, X12, uri_owl_DatatypeProperty) & iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, X12) & iext(uri_rdf_type, uri_foaf_Person, uri_owl_Class)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', testcase_premise_fullish_012_Template_Class)).
fof(f33, plain, ! [X0, X1] : (~ iext(uri_rdf_type, X0, X1) | icext(X1, X0)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ! [X0, X1] : ((iext(uri_rdf_type, X0, X1) | ~ icext(X1, X0)) & (icext(X1, X0) | ~ iext(uri_rdf_type, X0, X1))), inference(nnf_transformation, [], [f1])).
fof(f1, plain, ! [X0, X1] : (iext(uri_rdf_type, X0, X1) <=> icext(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', rdfs_cext_def)).
fof(f565, plain, (! [X0] : (~ icext(uri_ex_PersonAttribute, X0) | icext(sK6, X0)) | ~ spl7_30), inference(resolution, [], [f519, f60])).
fof(f60, plain, iext(uri_rdf_first, sK5, sK6), inference(cnf_transformation, [], [f32])).
fof(f519, plain, (! [X23, X21] : (~ iext(uri_rdf_first, sK5, X21) | icext(X21, X23) | ~ icext(uri_ex_PersonAttribute, X23)) | ~ spl7_30), inference(avatar_component_clause, [], [f518])).
fof(f518, plain, (spl7_30 <=> ! [X21, X23] : (~ iext(uri_rdf_first, sK5, X21) | icext(X21, X23) | ~ icext(uri_ex_PersonAttribute, X23))), introduced(avatar_definition, [new_symbols(naming, [spl7_30])])).
fof(f600, plain, (~ spl7_19 | spl7_2), inference(avatar_split_clause, [], [f599, f72, f179])).
fof(f179, plain, (spl7_19 <=> icext(uri_foaf_Person, uri_ex_alice)), introduced(avatar_definition, [new_symbols(naming, [spl7_19])])).
fof(f72, plain, (spl7_2 <=> iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person)), introduced(avatar_definition, [new_symbols(naming, [spl7_2])])).
fof(f599, plain, (~ icext(uri_foaf_Person, uri_ex_alice) | spl7_2), inference(resolution, [], [f74, f34])).
fof(f34, plain, ! [X0, X1] : (iext(uri_rdf_type, X0, X1) | ~ icext(X1, X0)), inference(cnf_transformation, [], [f22])).
fof(f74, plain, (~ iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person) | spl7_2), inference(avatar_component_clause, [], [f72])).
fof(f584, plain, (spl7_1 | ~ spl7_29), inference(avatar_contradiction_clause, [], [f583])).
fof(f583, plain, ($false | (spl7_1 | ~ spl7_29)), inference(subsumption_resolution, [], [f568, f79])).
fof(f79, plain, (~ icext(uri_owl_FunctionalProperty, uri_ex_name) | spl7_1), inference(resolution, [], [f34, f70])).
fof(f70, plain, (~ iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty) | spl7_1), inference(avatar_component_clause, [], [f68])).
fof(f68, plain, (spl7_1 <=> iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty)), introduced(avatar_definition, [new_symbols(naming, [spl7_1])])).
fof(f568, plain, (icext(uri_owl_FunctionalProperty, uri_ex_name) | ~ spl7_29), inference(resolution, [], [f537, f78])).
fof(f537, plain, (! [X0] : (~ icext(uri_ex_PersonAttribute, X0) | icext(uri_owl_FunctionalProperty, X0)) | ~ spl7_29), inference(resolution, [], [f515, f58])).
fof(f58, plain, iext(uri_rdf_first, sK4, uri_owl_FunctionalProperty), inference(cnf_transformation, [], [f32])).
fof(f515, plain, (! [X19, X18] : (~ iext(uri_rdf_first, sK4, X18) | icext(X18, X19) | ~ icext(uri_ex_PersonAttribute, X19)) | ~ spl7_29), inference(avatar_component_clause, [], [f514])).
fof(f514, plain, (spl7_29 <=> ! [X18, X19] : (~ iext(uri_rdf_first, sK4, X18) | icext(X18, X19) | ~ icext(uri_ex_PersonAttribute, X19))), introduced(avatar_definition, [new_symbols(naming, [spl7_29])])).
fof(f531, plain, ~ spl7_24, inference(avatar_contradiction_clause, [], [f530])).
fof(f530, plain, ($false | ~ spl7_24), inference(resolution, [], [f495, f56])).
fof(f56, plain, iext(uri_rdf_first, sK3, uri_owl_DatatypeProperty), inference(cnf_transformation, [], [f32])).
fof(f495, plain, (! [X0] : ~ iext(uri_rdf_first, sK3, X0) | ~ spl7_24), inference(avatar_component_clause, [], [f494])).
fof(f494, plain, (spl7_24 <=> ! [X0] : ~ iext(uri_rdf_first, sK3, X0)), introduced(avatar_definition, [new_symbols(naming, [spl7_24])])).
fof(f529, plain, ~ spl7_23, inference(avatar_contradiction_clause, [], [f528])).
fof(f528, plain, ($false | ~ spl7_23), inference(resolution, [], [f492, f60])).
fof(f492, plain, (! [X1] : ~ iext(uri_rdf_first, sK5, X1) | ~ spl7_23), inference(avatar_component_clause, [], [f491])).
fof(f491, plain, (spl7_23 <=> ! [X1] : ~ iext(uri_rdf_first, sK5, X1)), introduced(avatar_definition, [new_symbols(naming, [spl7_23])])).
fof(f527, plain, ~ spl7_22, inference(avatar_contradiction_clause, [], [f526])).
fof(f526, plain, ($false | ~ spl7_22), inference(resolution, [], [f489, f58])).
fof(f489, plain, (! [X2] : ~ iext(uri_rdf_first, sK4, X2) | ~ spl7_22), inference(avatar_component_clause, [], [f488])).
fof(f488, plain, (spl7_22 <=> ! [X2] : ~ iext(uri_rdf_first, sK4, X2)), introduced(avatar_definition, [new_symbols(naming, [spl7_22])])).
fof(f520, plain, (spl7_22 | spl7_30 | spl7_24), inference(avatar_split_clause, [], [f479, f494, f518, f488])).
fof(f479, plain, ! [X23, X21, X22, X20] : (~ iext(uri_rdf_first, sK3, X20) | ~ iext(uri_rdf_first, sK5, X21) | ~ iext(uri_rdf_first, sK4, X22) | ~ icext(uri_ex_PersonAttribute, X23) | icext(X21, X23)), inference(resolution, [], [f416, f44])).
fof(f44, plain, ! [X2, X0, X5, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ icext(X3, X5) | icext(X0, X5)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0, X1, X2, X3] : ((sP0(X0, X1, X2, X3) | ((~ icext(X0, sK2(X0, X1, X2, X3)) | ~ icext(X1, sK2(X0, X1, X2, X3)) | ~ icext(X2, sK2(X0, X1, X2, X3)) | ~ icext(X3, sK2(X0, X1, X2, X3))) & ((icext(X0, sK2(X0, X1, X2, X3)) & icext(X1, sK2(X0, X1, X2, X3)) & icext(X2, sK2(X0, X1, X2, X3))) | icext(X3, sK2(X0, X1, X2, X3)))) | ~ ic(X0) | ~ ic(X1) | ~ ic(X2) | ~ ic(X3)) & ((! [X5] : ((icext(X3, X5) | ~ icext(X0, X5) | ~ icext(X1, X5) | ~ icext(X2, X5)) & ((icext(X0, X5) & icext(X1, X5) & icext(X2, X5)) | ~ icext(X3, X5))) & ic(X0) & ic(X1) & ic(X2) & ic(X3)) | ~ sP0(X0, X1, X2, X3))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2])], [f27, f28])).
fof(f28, plain, ! [X3, X2, X1, X0] : (? [X4] : ((~ icext(X0, X4) | ~ icext(X1, X4) | ~ icext(X2, X4) | ~ icext(X3, X4)) & ((icext(X0, X4) & icext(X1, X4) & icext(X2, X4)) | icext(X3, X4))) => ((~ icext(X0, sK2(X0, X1, X2, X3)) | ~ icext(X1, sK2(X0, X1, X2, X3)) | ~ icext(X2, sK2(X0, X1, X2, X3)) | ~ icext(X3, sK2(X0, X1, X2, X3))) & ((icext(X0, sK2(X0, X1, X2, X3)) & icext(X1, sK2(X0, X1, X2, X3)) & icext(X2, sK2(X0, X1, X2, X3))) | icext(X3, sK2(X0, X1, X2, X3))))), introduced(choice_axiom, [])).
fof(f27, plain, ! [X0, X1, X2, X3] : ((sP0(X0, X1, X2, X3) | ? [X4] : ((~ icext(X0, X4) | ~ icext(X1, X4) | ~ icext(X2, X4) | ~ icext(X3, X4)) & ((icext(X0, X4) & icext(X1, X4) & icext(X2, X4)) | icext(X3, X4))) | ~ ic(X0) | ~ ic(X1) | ~ ic(X2) | ~ ic(X3)) & ((! [X5] : ((icext(X3, X5) | ~ icext(X0, X5) | ~ icext(X1, X5) | ~ icext(X2, X5)) & ((icext(X0, X5) & icext(X1, X5) & icext(X2, X5)) | ~ icext(X3, X5))) & ic(X0) & ic(X1) & ic(X2) & ic(X3)) | ~ sP0(X0, X1, X2, X3))), inference(rectify, [], [f26])).
fof(f26, plain, ! [X6, X4, X2, X0] : ((sP0(X6, X4, X2, X0) | ? [X7] : ((~ icext(X6, X7) | ~ icext(X4, X7) | ~ icext(X2, X7) | ~ icext(X0, X7)) & ((icext(X6, X7) & icext(X4, X7) & icext(X2, X7)) | icext(X0, X7))) | ~ ic(X6) | ~ ic(X4) | ~ ic(X2) | ~ ic(X0)) & ((! [X7] : ((icext(X0, X7) | ~ icext(X6, X7) | ~ icext(X4, X7) | ~ icext(X2, X7)) & ((icext(X6, X7) & icext(X4, X7) & icext(X2, X7)) | ~ icext(X0, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0)) | ~ sP0(X6, X4, X2, X0))), inference(flattening, [], [f25])).
fof(f25, plain, ! [X6, X4, X2, X0] : ((sP0(X6, X4, X2, X0) | (? [X7] : (((~ icext(X6, X7) | ~ icext(X4, X7) | ~ icext(X2, X7)) | ~ icext(X0, X7)) & ((icext(X6, X7) & icext(X4, X7) & icext(X2, X7)) | icext(X0, X7))) | ~ ic(X6) | ~ ic(X4) | ~ ic(X2) | ~ ic(X0))) & ((! [X7] : ((icext(X0, X7) | (~ icext(X6, X7) | ~ icext(X4, X7) | ~ icext(X2, X7))) & ((icext(X6, X7) & icext(X4, X7) & icext(X2, X7)) | ~ icext(X0, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0)) | ~ sP0(X6, X4, X2, X0))), inference(nnf_transformation, [], [f19])).
fof(f19, plain, ! [X6, X4, X2, X0] : (sP0(X6, X4, X2, X0) <=> (! [X7] : (icext(X0, X7) <=> (icext(X6, X7) & icext(X4, X7) & icext(X2, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0))), inference(usedef, [], [e19])).
fof(e19, plain, ! [X6, X4, X2, X0] : (sP0(X6, X4, X2, X0) <=> (! [X7] : (icext(X0, X7) <=> (icext(X6, X7) & icext(X4, X7) & icext(X2, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f416, plain, ! [X2, X0, X1] : (sP0(X2, X0, X1, uri_ex_PersonAttribute) | ~ iext(uri_rdf_first, sK3, X1) | ~ iext(uri_rdf_first, sK5, X2) | ~ iext(uri_rdf_first, sK4, X0)), inference(subsumption_resolution, [], [f415, f57])).
fof(f57, plain, iext(uri_rdf_rest, sK3, sK4), inference(cnf_transformation, [], [f32])).
fof(f415, plain, ! [X2, X0, X1] : (~ iext(uri_rdf_first, sK4, X0) | ~ iext(uri_rdf_rest, sK3, sK4) | ~ iext(uri_rdf_first, sK3, X1) | ~ iext(uri_rdf_first, sK5, X2) | sP0(X2, X0, X1, uri_ex_PersonAttribute)), inference(resolution, [], [f217, f55])).
fof(f55, plain, iext(uri_owl_intersectionOf, uri_ex_PersonAttribute, sK3), inference(cnf_transformation, [], [f32])).
fof(f217, plain, ! [X6, X8, X7, X5, X9] : (~ iext(uri_owl_intersectionOf, X9, X7) | ~ iext(uri_rdf_first, sK4, X6) | ~ iext(uri_rdf_rest, X7, sK4) | ~ iext(uri_rdf_first, X7, X8) | ~ iext(uri_rdf_first, sK5, X5) | sP0(X5, X6, X8, X9)), inference(resolution, [], [f214, f36])).
fof(f36, plain, ! [X4, X2, X0, X3, X1] : (~ sP1(X0, X1, X2, X3, X4) | ~ iext(uri_owl_intersectionOf, X0, X4) | sP0(X3, X2, X1, X0)), inference(cnf_transformation, [], [f24])).
fof(f24, plain, ! [X0, X1, X2, X3, X4] : (((iext(uri_owl_intersectionOf, X0, X4) | ~ sP0(X3, X2, X1, X0)) & (sP0(X3, X2, X1, X0) | ~ iext(uri_owl_intersectionOf, X0, X4))) | ~ sP1(X0, X1, X2, X3, X4)), inference(rectify, [], [f23])).
fof(f23, plain, ! [X0, X2, X4, X6, X1] : (((iext(uri_owl_intersectionOf, X0, X1) | ~ sP0(X6, X4, X2, X0)) & (sP0(X6, X4, X2, X0) | ~ iext(uri_owl_intersectionOf, X0, X1))) | ~ sP1(X0, X2, X4, X6, X1)), inference(nnf_transformation, [], [f20])).
fof(f20, plain, ! [X0, X2, X4, X6, X1] : ((iext(uri_owl_intersectionOf, X0, X1) <=> sP0(X6, X4, X2, X0)) | ~ sP1(X0, X2, X4, X6, X1)), inference(usedef, [], [e20])).
fof(e20, plain, ! [X0, X2, X4, X6, X1] : (sP1(X0, X2, X4, X6, X1) <=> (iext(uri_owl_intersectionOf, X0, X1) <=> sP0(X6, X4, X2, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f214, plain, ! [X4, X2, X0, X3, X1] : (sP1(X1, X2, X3, X0, X4) | ~ iext(uri_rdf_first, sK5, X0) | ~ iext(uri_rdf_first, sK4, X3) | ~ iext(uri_rdf_rest, X4, sK4) | ~ iext(uri_rdf_first, X4, X2)), inference(resolution, [], [f191, f59])).
fof(f59, plain, iext(uri_rdf_rest, sK4, sK5), inference(cnf_transformation, [], [f32])).
fof(f191, plain, ! [X4, X2, X0, X5, X3, X1] : (~ iext(uri_rdf_rest, X5, sK5) | ~ iext(uri_rdf_first, sK5, X3) | sP1(X0, X1, X2, X3, X4) | ~ iext(uri_rdf_first, X5, X2) | ~ iext(uri_rdf_rest, X4, X5) | ~ iext(uri_rdf_first, X4, X1)), inference(resolution, [], [f50, f61])).
fof(f61, plain, iext(uri_rdf_rest, sK5, uri_rdf_nil), inference(cnf_transformation, [], [f32])).
fof(f50, plain, ! [X6, X4, X2, X0, X5, X3, X1] : (~ iext(uri_rdf_rest, X5, uri_rdf_nil) | sP1(X0, X2, X4, X6, X1) | ~ iext(uri_rdf_first, X5, X6) | ~ iext(uri_rdf_rest, X3, X5) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(cnf_transformation, [], [f21])).
fof(f21, plain, ! [X0, X1, X2, X3, X4, X5, X6] : (sP1(X0, X2, X4, X6, X1) | ~ iext(uri_rdf_rest, X5, uri_rdf_nil) | ~ iext(uri_rdf_first, X5, X6) | ~ iext(uri_rdf_rest, X3, X5) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(definition_folding, [], [f15, e20, e19])).
fof(f15, plain, ! [X0, X1, X2, X3, X4, X5, X6] : ((iext(uri_owl_intersectionOf, X0, X1) <=> (! [X7] : (icext(X0, X7) <=> (icext(X6, X7) & icext(X4, X7) & icext(X2, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0))) | ~ iext(uri_rdf_rest, X5, uri_rdf_nil) | ~ iext(uri_rdf_first, X5, X6) | ~ iext(uri_rdf_rest, X3, X5) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(flattening, [], [f14])).
fof(f14, plain, ! [X0, X1, X2, X3, X4, X5, X6] : ((iext(uri_owl_intersectionOf, X0, X1) <=> (! [X7] : (icext(X0, X7) <=> (icext(X6, X7) & icext(X4, X7) & icext(X2, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0))) | (~ iext(uri_rdf_rest, X5, uri_rdf_nil) | ~ iext(uri_rdf_first, X5, X6) | ~ iext(uri_rdf_rest, X3, X5) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1, X2, X3, X4, X5, X6] : ((iext(uri_rdf_rest, X5, uri_rdf_nil) & iext(uri_rdf_first, X5, X6) & iext(uri_rdf_rest, X3, X5) & iext(uri_rdf_first, X3, X4) & iext(uri_rdf_rest, X1, X3) & iext(uri_rdf_first, X1, X2)) => (iext(uri_owl_intersectionOf, X0, X1) <=> (! [X7] : (icext(X0, X7) <=> (icext(X6, X7) & icext(X4, X7) & icext(X2, X7))) & ic(X6) & ic(X4) & ic(X2) & ic(X0)))), inference(rectify, [], [f3])).
fof(f3, plain, ! [X4, X5, X6, X7, X8, X9, X10] : ((iext(uri_rdf_rest, X9, uri_rdf_nil) & iext(uri_rdf_first, X9, X10) & iext(uri_rdf_rest, X7, X9) & iext(uri_rdf_first, X7, X8) & iext(uri_rdf_rest, X5, X7) & iext(uri_rdf_first, X5, X6)) => (iext(uri_owl_intersectionOf, X4, X5) <=> (! [X0] : (icext(X4, X0) <=> (icext(X10, X0) & icext(X8, X0) & icext(X6, X0))) & ic(X10) & ic(X8) & ic(X6) & ic(X4)))), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', owl_bool_intersectionof_class_003)).
fof(f516, plain, (spl7_29 | spl7_23 | spl7_24), inference(avatar_split_clause, [], [f478, f494, f491, f514])).
fof(f478, plain, ! [X19, X17, X18, X16] : (~ iext(uri_rdf_first, sK3, X16) | ~ iext(uri_rdf_first, sK5, X17) | ~ iext(uri_rdf_first, sK4, X18) | ~ icext(uri_ex_PersonAttribute, X19) | icext(X18, X19)), inference(resolution, [], [f416, f43])).
fof(f43, plain, ! [X2, X0, X5, X3, X1] : (~ sP0(X0, X1, X2, X3) | ~ icext(X3, X5) | icext(X1, X5)), inference(cnf_transformation, [], [f29])).
fof(f186, plain, (spl7_19 | ~ spl7_20), inference(avatar_split_clause, [], [f101, f183, f179])).
fof(f101, plain, (~ icext(sK6, uri_ex_name) | icext(uri_foaf_Person, uri_ex_alice)), inference(resolution, [], [f84, f66])).
fof(f66, plain, iext(uri_ex_name, uri_ex_alice, literal_plain(dat_str_alice)), inference(cnf_transformation, [], [f32])).
fof(f84, plain, ! [X2, X0, X1] : (~ iext(X0, X1, X2) | ~ icext(sK6, X0) | icext(uri_foaf_Person, X1)), inference(resolution, [], [f83, f35])).
fof(f35, plain, ! [X2, X0, X3, X1] : (~ iext(uri_rdfs_domain, X0, X1) | ~ iext(X0, X2, X3) | icext(X1, X2)), inference(cnf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X1, X2, X3] : (icext(X1, X2) | ~ iext(X0, X2, X3) | ~ iext(uri_rdfs_domain, X0, X1)), inference(flattening, [], [f12])).
fof(f12, plain, ! [X0, X1, X2, X3] : (icext(X1, X2) | (~ iext(X0, X2, X3) | ~ iext(uri_rdfs_domain, X0, X1))), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1, X2, X3] : ((iext(X0, X2, X3) & iext(uri_rdfs_domain, X0, X1)) => icext(X1, X2)), inference(rectify, [], [f2])).
fof(f2, plain, ! [X2, X1, X0, X3] : ((iext(X2, X0, X3) & iext(uri_rdfs_domain, X2, X1)) => icext(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', rdfs_domain_main)).
fof(f83, plain, ! [X0] : (iext(uri_rdfs_domain, X0, uri_foaf_Person) | ~ icext(sK6, X0)), inference(resolution, [], [f81, f64])).
fof(f64, plain, iext(uri_owl_hasValue, sK6, uri_foaf_Person), inference(cnf_transformation, [], [f32])).
fof(f81, plain, ! [X0, X1] : (~ iext(uri_owl_hasValue, sK6, X1) | iext(uri_rdfs_domain, X0, X1) | ~ icext(sK6, X0)), inference(resolution, [], [f51, f63])).
fof(f63, plain, iext(uri_owl_onProperty, sK6, uri_rdfs_domain), inference(cnf_transformation, [], [f32])).
fof(f51, plain, ! [X2, X0, X3, X1] : (~ iext(uri_owl_onProperty, X0, X1) | ~ icext(X0, X3) | iext(X1, X3, X2) | ~ iext(uri_owl_hasValue, X0, X2)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0, X1, X2] : (! [X3] : ((icext(X0, X3) | ~ iext(X1, X3, X2)) & (iext(X1, X3, X2) | ~ icext(X0, X3))) | ~ iext(uri_owl_onProperty, X0, X1) | ~ iext(uri_owl_hasValue, X0, X2)), inference(nnf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1, X2] : (! [X3] : (icext(X0, X3) <=> iext(X1, X3, X2)) | ~ iext(uri_owl_onProperty, X0, X1) | ~ iext(uri_owl_hasValue, X0, X2)), inference(flattening, [], [f16])).
fof(f16, plain, ! [X0, X1, X2] : (! [X3] : (icext(X0, X3) <=> iext(X1, X3, X2)) | (~ iext(uri_owl_onProperty, X0, X1) | ~ iext(uri_owl_hasValue, X0, X2))), inference(ennf_transformation, [], [f10])).
fof(f10, plain, ! [X0, X1, X2] : ((iext(uri_owl_onProperty, X0, X1) & iext(uri_owl_hasValue, X0, X2)) => ! [X3] : (icext(X0, X3) <=> iext(X1, X3, X2))), inference(rectify, [], [f4])).
fof(f4, plain, ! [X4, X2, X11] : ((iext(uri_owl_onProperty, X4, X2) & iext(uri_owl_hasValue, X4, X11)) => ! [X0] : (icext(X4, X0) <=> iext(X2, X0, X11))), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', owl_restrict_hasvalue)).
fof(f75, plain, (~ spl7_1 | ~ spl7_2), inference(avatar_split_clause, [], [f53, f72, f68])).
fof(f53, plain, (~ iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person) | ~ iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty)), inference(cnf_transformation, [], [f18])).
fof(f18, plain, (~ iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person) | ~ iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty)), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ~ (iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person) & iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty)), inference(negated_conjecture, [], [f5])).
fof(f5, plain, ~ (iext(uri_rdf_type, uri_ex_alice, uri_foaf_Person) & iext(uri_rdf_type, uri_ex_name, uri_owl_FunctionalProperty)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB012+2.p', testcase_conclusion_fullish_012_Template_Class)).