fof(f200, plain, $false, inference(subsumption_resolution, [], [f199, f56])).
fof(f56, plain, ~ icext(uri_ex_Falcon, uri_ex_harry), inference(resolution, [], [f53, f41])).
fof(f41, plain, iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species), inference(cnf_transformation, [], [f24])).
fof(f24, plain, (iext(uri_rdf_rest, sK5, uri_rdf_nil) & iext(uri_rdf_first, sK5, uri_ex_Falcon) & iext(uri_rdf_rest, sK4, sK5) & iext(uri_rdf_first, sK4, uri_ex_Eagle) & iext(uri_owl_unionOf, sK3, sK4) & iext(uri_rdf_type, uri_ex_harry, sK3) & iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species) & iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK3, sK4, sK5])], [f8, f23])).
fof(f23, plain, (? [X0, X1, X2] : (iext(uri_rdf_rest, X2, uri_rdf_nil) & iext(uri_rdf_first, X2, uri_ex_Falcon) & iext(uri_rdf_rest, X1, X2) & iext(uri_rdf_first, X1, uri_ex_Eagle) & iext(uri_owl_unionOf, X0, X1) & iext(uri_rdf_type, uri_ex_harry, X0) & iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species) & iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species)) => (iext(uri_rdf_rest, sK5, uri_rdf_nil) & iext(uri_rdf_first, sK5, uri_ex_Falcon) & iext(uri_rdf_rest, sK4, sK5) & iext(uri_rdf_first, sK4, uri_ex_Eagle) & iext(uri_owl_unionOf, sK3, sK4) & iext(uri_rdf_type, uri_ex_harry, sK3) & iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species) & iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species))), introduced(choice_axiom, [])).
fof(f8, plain, ? [X0, X1, X2] : (iext(uri_rdf_rest, X2, uri_rdf_nil) & iext(uri_rdf_first, X2, uri_ex_Falcon) & iext(uri_rdf_rest, X1, X2) & iext(uri_rdf_first, X1, uri_ex_Eagle) & iext(uri_owl_unionOf, X0, X1) & iext(uri_rdf_type, uri_ex_harry, X0) & iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species) & iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species)), inference(rectify, [], [f5])).
fof(f5, plain, ? [X8, X9, X10] : (iext(uri_rdf_rest, X10, uri_rdf_nil) & iext(uri_rdf_first, X10, uri_ex_Falcon) & iext(uri_rdf_rest, X9, X10) & iext(uri_rdf_first, X9, uri_ex_Eagle) & iext(uri_owl_unionOf, X8, X9) & iext(uri_rdf_type, uri_ex_harry, X8) & iext(uri_rdf_type, uri_ex_Falcon, uri_ex_Species) & iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB014+2.p', testcase_premise_fullish_014_Harry_belongs_to_some_Species)).
fof(f53, plain, ! [X0] : (~ iext(uri_rdf_type, X0, uri_ex_Species) | ~ icext(X0, uri_ex_harry)), inference(resolution, [], [f39, f26])).
fof(f26, plain, ! [X0, X1] : (iext(uri_rdf_type, X0, X1) | ~ icext(X1, X0)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, ! [X0, X1] : ((iext(uri_rdf_type, X0, X1) | ~ icext(X1, X0)) & (icext(X1, X0) | ~ iext(uri_rdf_type, X0, X1))), inference(nnf_transformation, [], [f1])).
fof(f1, plain, ! [X0, X1] : (iext(uri_rdf_type, X0, X1) <=> icext(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB014+2.p', rdfs_cext_def)).
fof(f39, plain, ! [X0] : (~ iext(uri_rdf_type, uri_ex_harry, X0) | ~ iext(uri_rdf_type, X0, uri_ex_Species)), inference(cnf_transformation, [], [f11])).
fof(f11, plain, ! [X0] : (~ iext(uri_rdf_type, X0, uri_ex_Species) | ~ iext(uri_rdf_type, uri_ex_harry, X0)), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ? [X0] : (iext(uri_rdf_type, X0, uri_ex_Species) & iext(uri_rdf_type, uri_ex_harry, X0)), inference(rectify, [], [f4])).
fof(f4, plain, ~ ? [X7] : (iext(uri_rdf_type, X7, uri_ex_Species) & iext(uri_rdf_type, uri_ex_harry, X7)), inference(negated_conjecture, [], [f3])).
fof(f3, plain, ~ ? [X7] : (iext(uri_rdf_type, X7, uri_ex_Species) & iext(uri_rdf_type, uri_ex_harry, X7)), file('/home/ubuntu/library/tptp/Problems/SWB/SWB014+2.p', testcase_conclusion_fullish_014_Harry_belongs_to_some_Species)).
fof(f199, plain, icext(uri_ex_Falcon, uri_ex_harry), inference(subsumption_resolution, [], [f191, f55])).
fof(f55, plain, ~ icext(uri_ex_Eagle, uri_ex_harry), inference(resolution, [], [f53, f40])).
fof(f40, plain, iext(uri_rdf_type, uri_ex_Eagle, uri_ex_Species), inference(cnf_transformation, [], [f24])).
fof(f191, plain, (icext(uri_ex_Eagle, uri_ex_harry) | icext(uri_ex_Falcon, uri_ex_harry)), inference(resolution, [], [f155, f50])).
fof(f50, plain, icext(sK3, uri_ex_harry), inference(resolution, [], [f25, f42])).
fof(f42, plain, iext(uri_rdf_type, uri_ex_harry, sK3), inference(cnf_transformation, [], [f24])).
fof(f25, plain, ! [X0, X1] : (~ iext(uri_rdf_type, X0, X1) | icext(X1, X0)), inference(cnf_transformation, [], [f15])).
fof(f155, plain, ! [X0] : (~ icext(sK3, X0) | icext(uri_ex_Eagle, X0) | icext(uri_ex_Falcon, X0)), inference(resolution, [], [f151, f32])).
fof(f32, plain, ! [X4, X2, X0, X1] : (~ sP0(X0, X1, X2) | icext(X1, X4) | ~ icext(X2, X4) | icext(X0, X4)), inference(cnf_transformation, [], [f22])).
fof(f22, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | (((~ icext(X0, sK2(X0, X1, X2)) & ~ icext(X1, sK2(X0, X1, X2))) | ~ icext(X2, sK2(X0, X1, X2))) & (icext(X0, sK2(X0, X1, X2)) | icext(X1, sK2(X0, X1, X2)) | icext(X2, sK2(X0, X1, X2)))) | ~ ic(X0) | ~ ic(X1) | ~ ic(X2)) & ((! [X4] : ((icext(X2, X4) | (~ icext(X0, X4) & ~ icext(X1, X4))) & (icext(X0, X4) | icext(X1, X4) | ~ icext(X2, X4))) & ic(X0) & ic(X1) & ic(X2)) | ~ sP0(X0, X1, X2))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK2])], [f20, f21])).
fof(f21, plain, ! [X2, X1, X0] : (? [X3] : (((~ icext(X0, X3) & ~ icext(X1, X3)) | ~ icext(X2, X3)) & (icext(X0, X3) | icext(X1, X3) | icext(X2, X3))) => (((~ icext(X0, sK2(X0, X1, X2)) & ~ icext(X1, sK2(X0, X1, X2))) | ~ icext(X2, sK2(X0, X1, X2))) & (icext(X0, sK2(X0, X1, X2)) | icext(X1, sK2(X0, X1, X2)) | icext(X2, sK2(X0, X1, X2))))), introduced(choice_axiom, [])).
fof(f20, plain, ! [X0, X1, X2] : ((sP0(X0, X1, X2) | ? [X3] : (((~ icext(X0, X3) & ~ icext(X1, X3)) | ~ icext(X2, X3)) & (icext(X0, X3) | icext(X1, X3) | icext(X2, X3))) | ~ ic(X0) | ~ ic(X1) | ~ ic(X2)) & ((! [X4] : ((icext(X2, X4) | (~ icext(X0, X4) & ~ icext(X1, X4))) & (icext(X0, X4) | icext(X1, X4) | ~ icext(X2, X4))) & ic(X0) & ic(X1) & ic(X2)) | ~ sP0(X0, X1, X2))), inference(rectify, [], [f19])).
fof(f19, plain, ! [X4, X2, X0] : ((sP0(X4, X2, X0) | ? [X5] : (((~ icext(X4, X5) & ~ icext(X2, X5)) | ~ icext(X0, X5)) & (icext(X4, X5) | icext(X2, X5) | icext(X0, X5))) | ~ ic(X4) | ~ ic(X2) | ~ ic(X0)) & ((! [X5] : ((icext(X0, X5) | (~ icext(X4, X5) & ~ icext(X2, X5))) & (icext(X4, X5) | icext(X2, X5) | ~ icext(X0, X5))) & ic(X4) & ic(X2) & ic(X0)) | ~ sP0(X4, X2, X0))), inference(flattening, [], [f18])).
fof(f18, plain, ! [X4, X2, X0] : ((sP0(X4, X2, X0) | (? [X5] : (((~ icext(X4, X5) & ~ icext(X2, X5)) | ~ icext(X0, X5)) & ((icext(X4, X5) | icext(X2, X5)) | icext(X0, X5))) | ~ ic(X4) | ~ ic(X2) | ~ ic(X0))) & ((! [X5] : ((icext(X0, X5) | (~ icext(X4, X5) & ~ icext(X2, X5))) & ((icext(X4, X5) | icext(X2, X5)) | ~ icext(X0, X5))) & ic(X4) & ic(X2) & ic(X0)) | ~ sP0(X4, X2, X0))), inference(nnf_transformation, [], [f12])).
fof(f12, plain, ! [X4, X2, X0] : (sP0(X4, X2, X0) <=> (! [X5] : (icext(X0, X5) <=> (icext(X4, X5) | icext(X2, X5))) & ic(X4) & ic(X2) & ic(X0))), inference(usedef, [], [e12])).
fof(e12, plain, ! [X4, X2, X0] : (sP0(X4, X2, X0) <=> (! [X5] : (icext(X0, X5) <=> (icext(X4, X5) | icext(X2, X5))) & ic(X4) & ic(X2) & ic(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f151, plain, sP0(uri_ex_Falcon, uri_ex_Eagle, sK3), inference(resolution, [], [f150, f44])).
fof(f44, plain, iext(uri_rdf_first, sK4, uri_ex_Eagle), inference(cnf_transformation, [], [f24])).
fof(f150, plain, ! [X0] : (~ iext(uri_rdf_first, sK4, X0) | sP0(uri_ex_Falcon, X0, sK3)), inference(resolution, [], [f149, f46])).
fof(f46, plain, iext(uri_rdf_first, sK5, uri_ex_Falcon), inference(cnf_transformation, [], [f24])).
fof(f149, plain, ! [X0, X1] : (~ iext(uri_rdf_first, sK5, X1) | ~ iext(uri_rdf_first, sK4, X0) | sP0(X1, X0, sK3)), inference(subsumption_resolution, [], [f148, f45])).
fof(f45, plain, iext(uri_rdf_rest, sK4, sK5), inference(cnf_transformation, [], [f24])).
fof(f148, plain, ! [X0, X1] : (~ iext(uri_rdf_rest, sK4, sK5) | ~ iext(uri_rdf_first, sK4, X0) | ~ iext(uri_rdf_first, sK5, X1) | sP0(X1, X0, sK3)), inference(resolution, [], [f62, f43])).
fof(f43, plain, iext(uri_owl_unionOf, sK3, sK4), inference(cnf_transformation, [], [f24])).
fof(f62, plain, ! [X6, X4, X7, X5] : (~ iext(uri_owl_unionOf, X7, X5) | ~ iext(uri_rdf_rest, X5, sK5) | ~ iext(uri_rdf_first, X5, X6) | ~ iext(uri_rdf_first, sK5, X4) | sP0(X4, X6, X7)), inference(resolution, [], [f58, f27])).
fof(f27, plain, ! [X2, X0, X3, X1] : (~ sP1(X0, X1, X2, X3) | ~ iext(uri_owl_unionOf, X0, X3) | sP0(X2, X1, X0)), inference(cnf_transformation, [], [f17])).
fof(f17, plain, ! [X0, X1, X2, X3] : (((iext(uri_owl_unionOf, X0, X3) | ~ sP0(X2, X1, X0)) & (sP0(X2, X1, X0) | ~ iext(uri_owl_unionOf, X0, X3))) | ~ sP1(X0, X1, X2, X3)), inference(rectify, [], [f16])).
fof(f16, plain, ! [X0, X2, X4, X1] : (((iext(uri_owl_unionOf, X0, X1) | ~ sP0(X4, X2, X0)) & (sP0(X4, X2, X0) | ~ iext(uri_owl_unionOf, X0, X1))) | ~ sP1(X0, X2, X4, X1)), inference(nnf_transformation, [], [f13])).
fof(f13, plain, ! [X0, X2, X4, X1] : ((iext(uri_owl_unionOf, X0, X1) <=> sP0(X4, X2, X0)) | ~ sP1(X0, X2, X4, X1)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X0, X2, X4, X1] : (sP1(X0, X2, X4, X1) <=> (iext(uri_owl_unionOf, X0, X1) <=> sP0(X4, X2, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f58, plain, ! [X2, X0, X3, X1] : (sP1(X0, X1, X2, X3) | ~ iext(uri_rdf_first, sK5, X2) | ~ iext(uri_rdf_rest, X3, sK5) | ~ iext(uri_rdf_first, X3, X1)), inference(resolution, [], [f38, f47])).
fof(f47, plain, iext(uri_rdf_rest, sK5, uri_rdf_nil), inference(cnf_transformation, [], [f24])).
fof(f38, plain, ! [X4, X2, X0, X3, X1] : (~ iext(uri_rdf_rest, X3, uri_rdf_nil) | sP1(X0, X2, X4, X1) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(cnf_transformation, [], [f14])).
fof(f14, plain, ! [X0, X1, X2, X3, X4] : (sP1(X0, X2, X4, X1) | ~ iext(uri_rdf_rest, X3, uri_rdf_nil) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(definition_folding, [], [f10, e13, e12])).
fof(f10, plain, ! [X0, X1, X2, X3, X4] : ((iext(uri_owl_unionOf, X0, X1) <=> (! [X5] : (icext(X0, X5) <=> (icext(X4, X5) | icext(X2, X5))) & ic(X4) & ic(X2) & ic(X0))) | ~ iext(uri_rdf_rest, X3, uri_rdf_nil) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2)), inference(flattening, [], [f9])).
fof(f9, plain, ! [X0, X1, X2, X3, X4] : ((iext(uri_owl_unionOf, X0, X1) <=> (! [X5] : (icext(X0, X5) <=> (icext(X4, X5) | icext(X2, X5))) & ic(X4) & ic(X2) & ic(X0))) | (~ iext(uri_rdf_rest, X3, uri_rdf_nil) | ~ iext(uri_rdf_first, X3, X4) | ~ iext(uri_rdf_rest, X1, X3) | ~ iext(uri_rdf_first, X1, X2))), inference(ennf_transformation, [], [f6])).
fof(f6, plain, ! [X0, X1, X2, X3, X4] : ((iext(uri_rdf_rest, X3, uri_rdf_nil) & iext(uri_rdf_first, X3, X4) & iext(uri_rdf_rest, X1, X3) & iext(uri_rdf_first, X1, X2)) => (iext(uri_owl_unionOf, X0, X1) <=> (! [X5] : (icext(X0, X5) <=> (icext(X4, X5) | icext(X2, X5))) & ic(X4) & ic(X2) & ic(X0)))), inference(rectify, [], [f2])).
fof(f2, plain, ! [X2, X3, X4, X5, X6] : ((iext(uri_rdf_rest, X5, uri_rdf_nil) & iext(uri_rdf_first, X5, X6) & iext(uri_rdf_rest, X3, X5) & iext(uri_rdf_first, X3, X4)) => (iext(uri_owl_unionOf, X2, X3) <=> (! [X0] : (icext(X2, X0) <=> (icext(X6, X0) | icext(X4, X0))) & ic(X6) & ic(X4) & ic(X2)))), file('/home/ubuntu/library/tptp/Problems/SWB/SWB014+2.p', owl_bool_unionof_class_002)).