fof(f14823, plain, $false, inference(avatar_sat_refutation, [], [f615, f621, f14822])).
fof(f14822, plain, ~ spl61_1, inference(avatar_contradiction_clause, [], [f14821])).
fof(f14821, plain, ($false | ~ spl61_1), inference(resolution, [], [f13373, f624])).
fof(f624, plain, (sP6(sK57, sK58, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f623, f568])).
fof(f568, plain, (sK57 = sK59), inference(cnf_transformation, [], [f356])).
fof(f356, plain, ((((((~ neq(sK60, nil) & neq(sK58, nil)) | sP6(sK59, sK60, sK57, sK58)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60)) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57, sK58, sK59, sK60])], [f234, f355, f354, f353, f352])).
fof(f352, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X2, X3, X0, X1)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X2, X3, sK57, X1)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK57))), introduced(choice_axiom, [])).
fof(f353, plain, (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X2, X3, sK57, X1)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X2, X3, sK57, sK58)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK58))), introduced(choice_axiom, [])).
fof(f354, plain, (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X2, X3, sK57, sK58)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(sK59, X3, sK57, sK58)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) & ssList(sK59))), introduced(choice_axiom, [])).
fof(f355, plain, (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(sK59, X3, sK57, sK58)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) => (((~ neq(sK60, nil) & neq(sK58, nil)) | sP6(sK59, sK60, sK57, sK58)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60))), introduced(choice_axiom, [])).
fof(f234, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X2, X3, X0, X1)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f223, e233])).
fof(f233, plain, ! [X2, X3, X0, X1] : ((? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X0) | ~ (app(app(X8, cons(X7, nil)), X9) = X1) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X1, nil)) | ~ sP6(X2, X3, X0, X1)), inference(usedef, [], [e233])).
fof(e233, plain, ! [X2, X3, X0, X1] : (sP6(X2, X3, X0, X1) <=> (? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X0) | ~ (app(app(X8, cons(X7, nil)), X9) = X1) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X1, nil))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f223, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X0) | ~ (app(app(X8, cons(X7, nil)), X9) = X1) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : (? [X6] : (((app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3)) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X0) | ~ (app(app(X8, cons(X7, nil)), X9) = X1) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3)) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (ssList(X6) => (~ (app(X5, X6) = X2) | ~ (app(app(X5, cons(X4, nil)), X6) = X3))))) | ? [X7] : (? [X8] : (? [X9] : ((app(X8, X9) = X0) & (app(app(X8, cons(X7, nil)), X9) = X1) & ssList(X9)) & ssList(X8)) & ssItem(X7)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X7] : (ssItem(X7) => ! [X8] : (ssList(X8) => ! [X9] : (ssList(X9) => (~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3))))) | ? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X7] : (ssItem(X7) => ! [X8] : (ssList(X8) => ! [X9] : (ssList(X9) => (~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3))))) | ? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC011+1.p', co1)).
fof(f623, plain, (sP6(sK59, sK58, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f610, f567])).
fof(f567, plain, (sK58 = sK60), inference(cnf_transformation, [], [f356])).
fof(f610, plain, (sP6(sK59, sK60, sK57, sK58) | ~ spl61_1), inference(avatar_component_clause, [], [f608])).
fof(f608, plain, (spl61_1 <=> sP6(sK59, sK60, sK57, sK58)), introduced(avatar_definition, [new_symbols(naming, [spl61_1])])).
fof(f13373, plain, (! [X0, X1] : ~ sP6(X0, X1, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f13372, f5027])).
fof(f5027, plain, ((sK57 = app(sK55(sK57, sK58), sK56(sK57, sK58))) | ~ spl61_1), inference(resolution, [], [f624, f562])).
fof(f562, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(sK55(X0, X1), sK56(X0, X1)) = X0)), inference(cnf_transformation, [], [f351])).
fof(f351, plain, ! [X0, X1, X2, X3] : ((((((app(sK55(X0, X1), sK56(X0, X1)) = X0) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X1) & ssList(sK56(X0, X1))) & ssList(sK55(X0, X1))) & ssItem(sK54(X0, X1))) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X3, nil)) | ~ sP6(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55, sK56])], [f347, f350, f349, f348])).
fof(f348, plain, ! [X1, X0] : (? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) => (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(sK54(X0, X1), nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(sK54(X0, X1)))), introduced(choice_axiom, [])).
fof(f349, plain, ! [X1, X0] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(sK54(X0, X1), nil)), X6) = X1) & ssList(X6)) & ssList(X5)) => (? [X6] : ((app(sK55(X0, X1), X6) = X0) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), X6) = X1) & ssList(X6)) & ssList(sK55(X0, X1)))), introduced(choice_axiom, [])).
fof(f350, plain, ! [X1, X0] : (? [X6] : ((app(sK55(X0, X1), X6) = X0) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), X6) = X1) & ssList(X6)) => ((app(sK55(X0, X1), sK56(X0, X1)) = X0) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X1) & ssList(sK56(X0, X1)))), introduced(choice_axiom, [])).
fof(f347, plain, ! [X0, X1, X2, X3] : ((? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X3, nil)) | ~ sP6(X0, X1, X2, X3)), inference(rectify, [], [f346])).
fof(f346, plain, ! [X2, X3, X0, X1] : ((? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(X8, X9) = X0) | ~ (app(app(X8, cons(X7, nil)), X9) = X1) | ~ ssList(X9)) | ~ ssList(X8)) | ~ ssItem(X7)) & neq(X1, nil)) | ~ sP6(X2, X3, X0, X1)), inference(nnf_transformation, [], [f233])).
fof(f13372, plain, (! [X0, X1] : ~ sP6(X0, X1, app(sK55(sK57, sK58), sK56(sK57, sK58)), sK58) | ~ spl61_1), inference(subsumption_resolution, [], [f13371, f5028])).
fof(f5028, plain, (ssList(sK56(sK57, sK58)) | ~ spl61_1), inference(resolution, [], [f624, f560])).
fof(f560, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssList(sK56(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f13371, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK57, sK58), sK56(sK57, sK58)), sK58) | ~ ssList(sK56(sK57, sK58))) | ~ spl61_1), inference(subsumption_resolution, [], [f13370, f991])).
fof(f991, plain, (ssItem(sK54(sK57, sK58)) | ~ spl61_1), inference(resolution, [], [f558, f624])).
fof(f558, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssItem(sK54(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f13370, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK57, sK58), sK56(sK57, sK58)), sK58) | ~ ssItem(sK54(sK57, sK58)) | ~ ssList(sK56(sK57, sK58))) | ~ spl61_1), inference(subsumption_resolution, [], [f13343, f5029])).
fof(f5029, plain, (ssList(sK55(sK57, sK58)) | ~ spl61_1), inference(resolution, [], [f624, f559])).
fof(f559, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssList(sK55(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f13343, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK57, sK58), sK56(sK57, sK58)), sK58) | ~ ssList(sK55(sK57, sK58)) | ~ ssItem(sK54(sK57, sK58)) | ~ ssList(sK56(sK57, sK58))) | ~ spl61_1), inference(superposition, [], [f599, f5026])).
fof(f5026, plain, ((sK58 = app(app(sK55(sK57, sK58), cons(sK54(sK57, sK58), nil)), sK56(sK57, sK58))) | ~ spl61_1), inference(resolution, [], [f624, f561])).
fof(f561, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X1)), inference(cnf_transformation, [], [f351])).
fof(f599, plain, ! [X0, X8, X7, X1, X9] : (~ sP6(X0, X1, app(X8, X9), app(app(X8, cons(X7, nil)), X9)) | ~ ssList(X8) | ~ ssItem(X7) | ~ ssList(X9)), inference(equality_resolution, [], [f598])).
fof(f598, plain, ! [X0, X8, X7, X3, X1, X9] : (~ (app(app(X8, cons(X7, nil)), X9) = X3) | ~ ssList(X9) | ~ ssList(X8) | ~ ssItem(X7) | ~ sP6(X0, X1, app(X8, X9), X3)), inference(equality_resolution, [], [f557])).
fof(f557, plain, ! [X2, X0, X8, X7, X3, X1, X9] : (~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3) | ~ ssList(X9) | ~ ssList(X8) | ~ ssItem(X7) | ~ sP6(X0, X1, X2, X3)), inference(cnf_transformation, [], [f351])).
fof(f621, plain, spl61_2, inference(avatar_contradiction_clause, [], [f620])).
fof(f620, plain, ($false | spl61_2), inference(subsumption_resolution, [], [f619, f616])).
fof(f616, plain, neq(sK58, nil), inference(subsumption_resolution, [], [f569, f556])).
fof(f556, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | neq(X3, nil)), inference(cnf_transformation, [], [f351])).
fof(f569, plain, (neq(sK58, nil) | sP6(sK59, sK60, sK57, sK58)), inference(cnf_transformation, [], [f356])).
fof(f619, plain, (~ neq(sK58, nil) | spl61_2), inference(forward_demodulation, [], [f614, f567])).
fof(f614, plain, (~ neq(sK60, nil) | spl61_2), inference(avatar_component_clause, [], [f612])).
fof(f612, plain, (spl61_2 <=> neq(sK60, nil)), introduced(avatar_definition, [new_symbols(naming, [spl61_2])])).
fof(f615, plain, (spl61_1 | ~ spl61_2), inference(avatar_split_clause, [], [f570, f612, f608])).
fof(f570, plain, (~ neq(sK60, nil) | sP6(sK59, sK60, sK57, sK58)), inference(cnf_transformation, [], [f356])).