fof(f12774, plain, $false, inference(avatar_sat_refutation, [], [f616, f622, f12773])).
fof(f12773, plain, ~ spl61_1, inference(avatar_contradiction_clause, [], [f12772])).
fof(f12772, plain, ($false | ~ spl61_1), inference(resolution, [], [f12629, f625])).
fof(f625, plain, (sP6(sK58, sK57, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f624, f568])).
fof(f568, plain, (sK58 = sK60), inference(cnf_transformation, [], [f356])).
fof(f356, plain, ((((((~ neq(sK60, nil) & neq(sK58, nil)) | sP6(sK60, sK59, sK57, sK58)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60)) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57, sK58, sK59, sK60])], [f234, f355, f354, f353, f352])).
fof(f352, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X0, X1)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, sK57, X1)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK57))), introduced(choice_axiom, [])).
fof(f353, plain, (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, sK57, X1)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X3, X2, sK57, sK58)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK58))), introduced(choice_axiom, [])).
fof(f354, plain, (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X3, X2, sK57, sK58)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X3, sK59, sK57, sK58)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) & ssList(sK59))), introduced(choice_axiom, [])).
fof(f355, plain, (? [X3] : (((~ neq(X3, nil) & neq(sK58, nil)) | sP6(X3, sK59, sK57, sK58)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) => (((~ neq(sK60, nil) & neq(sK58, nil)) | sP6(sK60, sK59, sK57, sK58)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60))), introduced(choice_axiom, [])).
fof(f234, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X0, X1)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f223, e233])).
fof(f233, plain, ! [X3, X2, X0, X1] : ((? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X3, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X0) | ~ (app(app(X9, cons(X8, nil)), X10) = X1) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X1, nil)) | ~ sP6(X3, X2, X0, X1)), inference(usedef, [], [e233])).
fof(e233, plain, ! [X3, X2, X0, X1] : (sP6(X3, X2, X0, X1) <=> (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X3, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X0) | ~ (app(app(X9, cons(X8, nil)), X10) = X1) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X1, nil))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f223, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X3, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X0) | ~ (app(app(X9, cons(X8, nil)), X10) = X1) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~ geq(X7, X4) | ~ memberP(X3, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3)) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X0) | ~ (app(app(X9, cons(X8, nil)), X10) = X1) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3)) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (ssList(X6) => (? [X7] : (geq(X7, X4) & memberP(X3, X7) & ~ (X4 = X7) & ssItem(X7)) | ~ (app(X5, X6) = X2) | ~ (app(app(X5, cons(X4, nil)), X6) = X3))))) | ? [X8] : (? [X9] : (? [X10] : ((app(X9, X10) = X0) & (app(app(X9, cons(X8, nil)), X10) = X1) & ssList(X10)) & ssList(X9)) & ssItem(X8)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X7] : (ssItem(X7) => ! [X8] : (ssList(X8) => ! [X9] : (ssList(X9) => (? [X10] : (geq(X10, X7) & memberP(X3, X10) & ~ (X7 = X10) & ssItem(X10)) | ~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3))))) | ? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X7] : (ssItem(X7) => ! [X8] : (ssList(X8) => ! [X9] : (ssList(X9) => (? [X10] : (geq(X10, X7) & memberP(X3, X10) & ~ (X7 = X10) & ssItem(X10)) | ~ (app(X8, X9) = X2) | ~ (app(app(X8, cons(X7, nil)), X9) = X3))))) | ? [X4] : (? [X5] : (? [X6] : ((app(X5, X6) = X0) & (app(app(X5, cons(X4, nil)), X6) = X1) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC010+1.p', co1)).
fof(f624, plain, (sP6(sK60, sK57, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f611, f569])).
fof(f569, plain, (sK57 = sK59), inference(cnf_transformation, [], [f356])).
fof(f611, plain, (sP6(sK60, sK59, sK57, sK58) | ~ spl61_1), inference(avatar_component_clause, [], [f609])).
fof(f609, plain, (spl61_1 <=> sP6(sK60, sK59, sK57, sK58)), introduced(avatar_definition, [new_symbols(naming, [spl61_1])])).
fof(f12629, plain, (! [X0, X1] : ~ sP6(X0, X1, sK57, sK58) | ~ spl61_1), inference(forward_demodulation, [], [f12628, f5775])).
fof(f5775, plain, ((sK57 = app(sK55(sK58, sK57), sK56(sK58, sK57))) | ~ spl61_1), inference(resolution, [], [f625, f562])).
fof(f562, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(sK55(X0, X1), sK56(X0, X1)) = X1)), inference(cnf_transformation, [], [f351])).
fof(f351, plain, ! [X0, X1, X2, X3] : (((((! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(sK55(X0, X1), sK56(X0, X1)) = X1) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X0) & ssList(sK56(X0, X1))) & ssList(sK55(X0, X1))) & ssItem(sK54(X0, X1))) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X2) | ~ (app(app(X9, cons(X8, nil)), X10) = X3) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X3, nil)) | ~ sP6(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55, sK56])], [f347, f350, f349, f348])).
fof(f348, plain, ! [X1, X0] : (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X0, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X1) & (app(app(X5, cons(X4, nil)), X6) = X0) & ssList(X6)) & ssList(X5)) & ssItem(X4)) => (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(X5, X6) = X1) & (app(app(X5, cons(sK54(X0, X1), nil)), X6) = X0) & ssList(X6)) & ssList(X5)) & ssItem(sK54(X0, X1)))), introduced(choice_axiom, [])).
fof(f349, plain, ! [X1, X0] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(X5, X6) = X1) & (app(app(X5, cons(sK54(X0, X1), nil)), X6) = X0) & ssList(X6)) & ssList(X5)) => (? [X6] : (! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(sK55(X0, X1), X6) = X1) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), X6) = X0) & ssList(X6)) & ssList(sK55(X0, X1)))), introduced(choice_axiom, [])).
fof(f350, plain, ! [X1, X0] : (? [X6] : (! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(sK55(X0, X1), X6) = X1) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), X6) = X0) & ssList(X6)) => (! [X7] : (~ geq(X7, sK54(X0, X1)) | ~ memberP(X0, X7) | (sK54(X0, X1) = X7) | ~ ssItem(X7)) & (app(sK55(X0, X1), sK56(X0, X1)) = X1) & (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X0) & ssList(sK56(X0, X1)))), introduced(choice_axiom, [])).
fof(f347, plain, ! [X0, X1, X2, X3] : ((? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X0, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X1) & (app(app(X5, cons(X4, nil)), X6) = X0) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X2) | ~ (app(app(X9, cons(X8, nil)), X10) = X3) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X3, nil)) | ~ sP6(X0, X1, X2, X3)), inference(rectify, [], [f346])).
fof(f346, plain, ! [X3, X2, X0, X1] : ((? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ geq(X7, X4) | ~ memberP(X3, X7) | (X4 = X7) | ~ ssItem(X7)) & (app(X5, X6) = X2) & (app(app(X5, cons(X4, nil)), X6) = X3) & ssList(X6)) & ssList(X5)) & ssItem(X4)) & ! [X8] : (! [X9] : (! [X10] : (~ (app(X9, X10) = X0) | ~ (app(app(X9, cons(X8, nil)), X10) = X1) | ~ ssList(X10)) | ~ ssList(X9)) | ~ ssItem(X8)) & neq(X1, nil)) | ~ sP6(X3, X2, X0, X1)), inference(nnf_transformation, [], [f233])).
fof(f12628, plain, (! [X0, X1] : ~ sP6(X0, X1, app(sK55(sK58, sK57), sK56(sK58, sK57)), sK58) | ~ spl61_1), inference(subsumption_resolution, [], [f12627, f1009])).
fof(f1009, plain, (ssList(sK56(sK58, sK57)) | ~ spl61_1), inference(resolution, [], [f560, f625])).
fof(f560, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssList(sK56(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f12627, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK58, sK57), sK56(sK58, sK57)), sK58) | ~ ssList(sK56(sK58, sK57))) | ~ spl61_1), inference(subsumption_resolution, [], [f12626, f992])).
fof(f992, plain, (ssItem(sK54(sK58, sK57)) | ~ spl61_1), inference(resolution, [], [f558, f625])).
fof(f558, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssItem(sK54(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f12626, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK58, sK57), sK56(sK58, sK57)), sK58) | ~ ssItem(sK54(sK58, sK57)) | ~ ssList(sK56(sK58, sK57))) | ~ spl61_1), inference(subsumption_resolution, [], [f12599, f993])).
fof(f993, plain, (ssList(sK55(sK58, sK57)) | ~ spl61_1), inference(resolution, [], [f559, f625])).
fof(f559, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssList(sK55(X0, X1))), inference(cnf_transformation, [], [f351])).
fof(f12599, plain, (! [X0, X1] : (~ sP6(X0, X1, app(sK55(sK58, sK57), sK56(sK58, sK57)), sK58) | ~ ssList(sK55(sK58, sK57)) | ~ ssItem(sK54(sK58, sK57)) | ~ ssList(sK56(sK58, sK57))) | ~ spl61_1), inference(superposition, [], [f600, f5774])).
fof(f5774, plain, ((sK58 = app(app(sK55(sK58, sK57), cons(sK54(sK58, sK57), nil)), sK56(sK58, sK57))) | ~ spl61_1), inference(resolution, [], [f625, f561])).
fof(f561, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(app(sK55(X0, X1), cons(sK54(X0, X1), nil)), sK56(X0, X1)) = X0)), inference(cnf_transformation, [], [f351])).
fof(f600, plain, ! [X0, X10, X8, X1, X9] : (~ sP6(X0, X1, app(X9, X10), app(app(X9, cons(X8, nil)), X10)) | ~ ssList(X9) | ~ ssItem(X8) | ~ ssList(X10)), inference(equality_resolution, [], [f599])).
fof(f599, plain, ! [X0, X10, X8, X3, X1, X9] : (~ (app(app(X9, cons(X8, nil)), X10) = X3) | ~ ssList(X10) | ~ ssList(X9) | ~ ssItem(X8) | ~ sP6(X0, X1, app(X9, X10), X3)), inference(equality_resolution, [], [f557])).
fof(f557, plain, ! [X2, X0, X10, X8, X3, X1, X9] : (~ (app(X9, X10) = X2) | ~ (app(app(X9, cons(X8, nil)), X10) = X3) | ~ ssList(X10) | ~ ssList(X9) | ~ ssItem(X8) | ~ sP6(X0, X1, X2, X3)), inference(cnf_transformation, [], [f351])).
fof(f622, plain, spl61_2, inference(avatar_contradiction_clause, [], [f621])).
fof(f621, plain, ($false | spl61_2), inference(subsumption_resolution, [], [f620, f617])).
fof(f617, plain, neq(sK58, nil), inference(subsumption_resolution, [], [f570, f556])).
fof(f556, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | neq(X3, nil)), inference(cnf_transformation, [], [f351])).
fof(f570, plain, (neq(sK58, nil) | sP6(sK60, sK59, sK57, sK58)), inference(cnf_transformation, [], [f356])).
fof(f620, plain, (~ neq(sK58, nil) | spl61_2), inference(forward_demodulation, [], [f615, f568])).
fof(f615, plain, (~ neq(sK60, nil) | spl61_2), inference(avatar_component_clause, [], [f613])).
fof(f613, plain, (spl61_2 <=> neq(sK60, nil)), introduced(avatar_definition, [new_symbols(naming, [spl61_2])])).
fof(f616, plain, (spl61_1 | ~ spl61_2), inference(avatar_split_clause, [], [f571, f613, f609])).
fof(f571, plain, (~ neq(sK60, nil) | sP6(sK60, sK59, sK57, sK58)), inference(cnf_transformation, [], [f356])).