fof(f1540, plain, $false, inference(avatar_sat_refutation, [], [f613, f627, f1485, f1539])).
fof(f1539, plain, (~ spl61_1 | spl61_23), inference(avatar_contradiction_clause, [], [f1538])).
fof(f1538, plain, ($false | (~ spl61_1 | spl61_23)), inference(subsumption_resolution, [], [f1537, f1150])).
fof(f1150, plain, (sP6(sK58, sK57) | ~ spl61_1), inference(backward_demodulation, [], [f628, f566])).
fof(f566, plain, (sK58 = sK60), inference(cnf_transformation, [], [f355])).
fof(f355, plain, (((((((nil = sK59) & (nil = sK60)) | sP6(sK60, sK59)) & ~ strictorderedP(sK57) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60)) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57, sK58, sK59, sK60])], [f233, f354, f353, f352, f351])).
fof(f351, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(X0) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(sK57) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK57))), introduced(choice_axiom, [])).
fof(f352, plain, (? [X1] : (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(sK57) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(sK57) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK58))), introduced(choice_axiom, [])).
fof(f353, plain, (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(sK57) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : ((((nil = sK59) & (nil = X3)) | sP6(X3, sK59)) & ~ strictorderedP(sK57) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) & ssList(sK59))), introduced(choice_axiom, [])).
fof(f354, plain, (? [X3] : ((((nil = sK59) & (nil = X3)) | sP6(X3, sK59)) & ~ strictorderedP(sK57) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) => ((((nil = sK59) & (nil = sK60)) | sP6(sK60, sK59)) & ~ strictorderedP(sK57) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60))), introduced(choice_axiom, [])).
fof(f233, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | sP6(X3, X2)) & ~ strictorderedP(X0) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f222, e232])).
fof(f232, plain, ! [X3, X2] : (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ lt(X7, X4) | ~ memberP(X6, X7) | ~ ssItem(X7)) & ! [X8] : (~ lt(X4, X8) | ~ memberP(X5, X8) | ~ ssItem(X8)) & (app(app(X5, X2), X6) = X3) & (cons(X4, nil) = X2) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ sP6(X3, X2)), inference(usedef, [], [e232])).
fof(e232, plain, ! [X3, X2] : (sP6(X3, X2) <=> ? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ lt(X7, X4) | ~ memberP(X6, X7) | ~ ssItem(X7)) & ! [X8] : (~ lt(X4, X8) | ~ memberP(X5, X8) | ~ ssItem(X8)) & (app(app(X5, X2), X6) = X3) & (cons(X4, nil) = X2) & ssList(X6)) & ssList(X5)) & ssItem(X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((nil = X2) & (nil = X3)) | ? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ lt(X7, X4) | ~ memberP(X6, X7) | ~ ssItem(X7)) & ! [X8] : (~ lt(X4, X8) | ~ memberP(X5, X8) | ~ ssItem(X8)) & (app(app(X5, X2), X6) = X3) & (cons(X4, nil) = X2) & ssList(X6)) & ssList(X5)) & ssItem(X4))) & ~ strictorderedP(X0) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (((~ (nil = X2) | ~ (nil = X3)) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (? [X7] : (lt(X7, X4) & memberP(X6, X7) & ssItem(X7)) | ? [X8] : (lt(X4, X8) & memberP(X5, X8) & ssItem(X8)) | ~ (app(app(X5, X2), X6) = X3) | ~ (cons(X4, nil) = X2) | ~ ssList(X6))))) | strictorderedP(X0) | ~ (X0 = X2) | ~ (X1 = X3) | ~ ssList(X3))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (((~ (nil = X2) | ~ (nil = X3)) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (? [X8] : (lt(X8, X4) & memberP(X6, X8) & ssItem(X8)) | ? [X7] : (lt(X4, X7) & memberP(X5, X7) & ssItem(X7)) | ~ (app(app(X5, X2), X6) = X3) | ~ (cons(X4, nil) = X2) | ~ ssList(X6))))) | strictorderedP(X0) | ~ (X0 = X2) | ~ (X1 = X3) | ~ ssList(X3))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (((~ (nil = X2) | ~ (nil = X3)) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (? [X8] : (lt(X8, X4) & memberP(X6, X8) & ssItem(X8)) | ? [X7] : (lt(X4, X7) & memberP(X5, X7) & ssItem(X7)) | ~ (app(app(X5, X2), X6) = X3) | ~ (cons(X4, nil) = X2) | ~ ssList(X6))))) | strictorderedP(X0) | ~ (X0 = X2) | ~ (X1 = X3) | ~ ssList(X3))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC288+1.p', co1)).
fof(f628, plain, (sP6(sK60, sK57) | ~ spl61_1), inference(forward_demodulation, [], [f608, f567])).
fof(f567, plain, (sK57 = sK59), inference(cnf_transformation, [], [f355])).
fof(f608, plain, (sP6(sK60, sK59) | ~ spl61_1), inference(avatar_component_clause, [], [f606])).
fof(f606, plain, (spl61_1 <=> sP6(sK60, sK59)), introduced(avatar_definition, [new_symbols(naming, [spl61_1])])).
fof(f1537, plain, (~ sP6(sK58, sK57) | spl61_23), inference(resolution, [], [f1466, f555])).
fof(f555, plain, ! [X0, X1] : (ssItem(sK54(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f350])).
fof(f350, plain, ! [X0, X1] : ((((! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(sK56(X0, X1), X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(sK55(X0, X1), X6) | ~ ssItem(X6)) & (app(app(sK55(X0, X1), X1), sK56(X0, X1)) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(sK56(X0, X1))) & ssList(sK55(X0, X1))) & ssItem(sK54(X0, X1))) | ~ sP6(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55, sK56])], [f346, f349, f348, f347])).
fof(f347, plain, ! [X1, X0] : (? [X2] : (? [X3] : (? [X4] : (! [X5] : (~ lt(X5, X2) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(X2, X6) | ~ memberP(X3, X6) | ~ ssItem(X6)) & (app(app(X3, X1), X4) = X0) & (cons(X2, nil) = X1) & ssList(X4)) & ssList(X3)) & ssItem(X2)) => (? [X3] : (? [X4] : (! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(X3, X6) | ~ ssItem(X6)) & (app(app(X3, X1), X4) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(X4)) & ssList(X3)) & ssItem(sK54(X0, X1)))), introduced(choice_axiom, [])).
fof(f348, plain, ! [X1, X0] : (? [X3] : (? [X4] : (! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(X3, X6) | ~ ssItem(X6)) & (app(app(X3, X1), X4) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(X4)) & ssList(X3)) => (? [X4] : (! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(sK55(X0, X1), X6) | ~ ssItem(X6)) & (app(app(sK55(X0, X1), X1), X4) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(X4)) & ssList(sK55(X0, X1)))), introduced(choice_axiom, [])).
fof(f349, plain, ! [X1, X0] : (? [X4] : (! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(sK55(X0, X1), X6) | ~ ssItem(X6)) & (app(app(sK55(X0, X1), X1), X4) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(X4)) => (! [X5] : (~ lt(X5, sK54(X0, X1)) | ~ memberP(sK56(X0, X1), X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(sK54(X0, X1), X6) | ~ memberP(sK55(X0, X1), X6) | ~ ssItem(X6)) & (app(app(sK55(X0, X1), X1), sK56(X0, X1)) = X0) & (cons(sK54(X0, X1), nil) = X1) & ssList(sK56(X0, X1)))), introduced(choice_axiom, [])).
fof(f346, plain, ! [X0, X1] : (? [X2] : (? [X3] : (? [X4] : (! [X5] : (~ lt(X5, X2) | ~ memberP(X4, X5) | ~ ssItem(X5)) & ! [X6] : (~ lt(X2, X6) | ~ memberP(X3, X6) | ~ ssItem(X6)) & (app(app(X3, X1), X4) = X0) & (cons(X2, nil) = X1) & ssList(X4)) & ssList(X3)) & ssItem(X2)) | ~ sP6(X0, X1)), inference(rectify, [], [f345])).
fof(f345, plain, ! [X3, X2] : (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~ lt(X7, X4) | ~ memberP(X6, X7) | ~ ssItem(X7)) & ! [X8] : (~ lt(X4, X8) | ~ memberP(X5, X8) | ~ ssItem(X8)) & (app(app(X5, X2), X6) = X3) & (cons(X4, nil) = X2) & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~ sP6(X3, X2)), inference(nnf_transformation, [], [f232])).
fof(f1466, plain, (~ ssItem(sK54(sK58, sK57)) | spl61_23), inference(avatar_component_clause, [], [f1464])).
fof(f1464, plain, (spl61_23 <=> ssItem(sK54(sK58, sK57))), introduced(avatar_definition, [new_symbols(naming, [spl61_23])])).
fof(f1485, plain, (~ spl61_23 | ~ spl61_1), inference(avatar_split_clause, [], [f1484, f606, f1464])).
fof(f1484, plain, (~ ssItem(sK54(sK58, sK57)) | ~ spl61_1), inference(subsumption_resolution, [], [f1445, f568])).
fof(f568, plain, ~ strictorderedP(sK57), inference(cnf_transformation, [], [f355])).
fof(f1445, plain, (strictorderedP(sK57) | ~ ssItem(sK54(sK58, sK57)) | ~ spl61_1), inference(superposition, [], [f517, f1297])).
fof(f1297, plain, ((sK57 = cons(sK54(sK58, sK57), nil)) | ~ spl61_1), inference(resolution, [], [f1150, f558])).
fof(f558, plain, ! [X0, X1] : (~ sP6(X0, X1) | (cons(sK54(X0, X1), nil) = X1)), inference(cnf_transformation, [], [f350])).
fof(f517, plain, ! [X0] : (strictorderedP(cons(X0, nil)) | ~ ssItem(X0)), inference(cnf_transformation, [], [f183])).
fof(f183, plain, ! [X0] : (strictorderedP(cons(X0, nil)) | ~ ssItem(X0)), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ! [X0] : (ssItem(X0) => strictorderedP(cons(X0, nil))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC288+1.p', ax68)).
fof(f627, plain, ~ spl61_2, inference(avatar_contradiction_clause, [], [f626])).
fof(f626, plain, ($false | ~ spl61_2), inference(subsumption_resolution, [], [f625, f518])).
fof(f518, plain, strictorderedP(nil), inference(cnf_transformation, [], [f69])).
fof(f69, plain, strictorderedP(nil), file('/home/ubuntu/library/tptp/Problems/SWC/SWC288+1.p', ax69)).
fof(f625, plain, (~ strictorderedP(nil) | ~ spl61_2), inference(backward_demodulation, [], [f568, f623])).
fof(f623, plain, ((nil = sK57) | ~ spl61_2), inference(backward_demodulation, [], [f567, f612])).
fof(f612, plain, ((nil = sK59) | ~ spl61_2), inference(avatar_component_clause, [], [f610])).
fof(f610, plain, (spl61_2 <=> (nil = sK59)), introduced(avatar_definition, [new_symbols(naming, [spl61_2])])).
fof(f613, plain, (spl61_1 | spl61_2), inference(avatar_split_clause, [], [f570, f610, f606])).
fof(f570, plain, ((nil = sK59) | sP6(sK60, sK59)), inference(cnf_transformation, [], [f355])).