fof(f7240, plain, $false, inference(avatar_sat_refutation, [], [f613, f619, f7239])).
fof(f7239, plain, ~ spl60_1, inference(avatar_contradiction_clause, [], [f7238])).
fof(f7238, plain, ($false | ~ spl60_1), inference(resolution, [], [f6063, f622])).
fof(f622, plain, (sP6(sK57, sK56, sK57, sK56) | ~ spl60_1), inference(forward_demodulation, [], [f621, f565])).
fof(f565, plain, (sK57 = sK59), inference(cnf_transformation, [], [f355])).
fof(f355, plain, ((((((~ neq(sK59, nil) & neq(sK57, nil)) | sP6(sK59, sK58, sK57, sK56)) & (sK56 = sK58) & (sK57 = sK59) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)) & ssList(sK56)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK56, sK57, sK58, sK59])], [f234, f354, f353, f352, f351])).
fof(f351, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X1, X0)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X1, sK56)) & (sK56 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK56))), introduced(choice_axiom, [])).
fof(f352, plain, (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X1, sK56)) & (sK56 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK57, nil)) | sP6(X3, X2, sK57, sK56)) & (sK56 = X2) & (sK57 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK57))), introduced(choice_axiom, [])).
fof(f353, plain, (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(sK57, nil)) | sP6(X3, X2, sK57, sK56)) & (sK56 = X2) & (sK57 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : (((~ neq(X3, nil) & neq(sK57, nil)) | sP6(X3, sK58, sK57, sK56)) & (sK56 = sK58) & (sK57 = X3) & ssList(X3)) & ssList(sK58))), introduced(choice_axiom, [])).
fof(f354, plain, (? [X3] : (((~ neq(X3, nil) & neq(sK57, nil)) | sP6(X3, sK58, sK57, sK56)) & (sK56 = sK58) & (sK57 = X3) & ssList(X3)) => (((~ neq(sK59, nil) & neq(sK57, nil)) | sP6(sK59, sK58, sK57, sK56)) & (sK56 = sK58) & (sK57 = sK59) & ssList(sK59))), introduced(choice_axiom, [])).
fof(f234, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | sP6(X3, X2, X1, X0)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f223, e233])).
fof(f233, plain, ! [X3, X2, X1, X0] : ((? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X3) & (app(cons(X4, nil), X5) = X2) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X1) | ~ (app(cons(X6, nil), X7) = X0) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X1, nil)) | ~ sP6(X3, X2, X1, X0)), inference(usedef, [], [e233])).
fof(e233, plain, ! [X3, X2, X1, X0] : (sP6(X3, X2, X1, X0) <=> (? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X3) & (app(cons(X4, nil), X5) = X2) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X1) | ~ (app(cons(X6, nil), X7) = X0) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X1, nil))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f223, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X3) & (app(cons(X4, nil), X5) = X2) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X1) | ~ (app(cons(X6, nil), X7) = X0) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((~ neq(X3, nil) & neq(X1, nil)) | (? [X4] : (? [X5] : (((app(X5, cons(X4, nil)) = X3) & (app(cons(X4, nil), X5) = X2)) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X1) | ~ (app(cons(X6, nil), X7) = X0) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X1, nil))) & (X0 = X2) & (X1 = X3)) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => (~ (app(X5, cons(X4, nil)) = X3) | ~ (app(cons(X4, nil), X5) = X2)))) | ? [X6] : (? [X7] : ((app(X7, cons(X6, nil)) = X1) & (app(cons(X6, nil), X7) = X0) & ssList(X7)) & ssItem(X6)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X6] : (ssItem(X6) => ! [X7] : (ssList(X7) => (~ (app(X7, cons(X6, nil)) = X3) | ~ (app(cons(X6, nil), X7) = X2)))) | ? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X1) & (app(cons(X4, nil), X5) = X0) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((neq(X3, nil) | ~ neq(X1, nil)) & (! [X6] : (ssItem(X6) => ! [X7] : (ssList(X7) => (~ (app(X7, cons(X6, nil)) = X3) | ~ (app(cons(X6, nil), X7) = X2)))) | ? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X1) & (app(cons(X4, nil), X5) = X0) & ssList(X5)) & ssItem(X4)) | ~ neq(X1, nil))) | ~ (X0 = X2) | ~ (X1 = X3)))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC317+1.p', co1)).
fof(f621, plain, (sP6(sK59, sK56, sK57, sK56) | ~ spl60_1), inference(forward_demodulation, [], [f608, f566])).
fof(f566, plain, (sK56 = sK58), inference(cnf_transformation, [], [f355])).
fof(f608, plain, (sP6(sK59, sK58, sK57, sK56) | ~ spl60_1), inference(avatar_component_clause, [], [f606])).
fof(f606, plain, (spl60_1 <=> sP6(sK59, sK58, sK57, sK56)), introduced(avatar_definition, [new_symbols(naming, [spl60_1])])).
fof(f6063, plain, (! [X0, X1] : ~ sP6(X0, X1, sK57, sK56) | ~ spl60_1), inference(forward_demodulation, [], [f6062, f4289])).
fof(f4289, plain, ((sK56 = app(cons(sK54(sK57, sK56), nil), sK55(sK57, sK56))) | ~ spl60_1), inference(resolution, [], [f622, f559])).
fof(f559, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(cons(sK54(X0, X1), nil), sK55(X0, X1)) = X1)), inference(cnf_transformation, [], [f350])).
fof(f350, plain, ! [X0, X1, X2, X3] : (((((app(sK55(X0, X1), cons(sK54(X0, X1), nil)) = X0) & (app(cons(sK54(X0, X1), nil), sK55(X0, X1)) = X1) & ssList(sK55(X0, X1))) & ssItem(sK54(X0, X1))) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X2) | ~ (app(cons(X6, nil), X7) = X3) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X2, nil)) | ~ sP6(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55])], [f347, f349, f348])).
fof(f348, plain, ! [X1, X0] : (? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X0) & (app(cons(X4, nil), X5) = X1) & ssList(X5)) & ssItem(X4)) => (? [X5] : ((app(X5, cons(sK54(X0, X1), nil)) = X0) & (app(cons(sK54(X0, X1), nil), X5) = X1) & ssList(X5)) & ssItem(sK54(X0, X1)))), introduced(choice_axiom, [])).
fof(f349, plain, ! [X1, X0] : (? [X5] : ((app(X5, cons(sK54(X0, X1), nil)) = X0) & (app(cons(sK54(X0, X1), nil), X5) = X1) & ssList(X5)) => ((app(sK55(X0, X1), cons(sK54(X0, X1), nil)) = X0) & (app(cons(sK54(X0, X1), nil), sK55(X0, X1)) = X1) & ssList(sK55(X0, X1)))), introduced(choice_axiom, [])).
fof(f347, plain, ! [X0, X1, X2, X3] : ((? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X0) & (app(cons(X4, nil), X5) = X1) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X2) | ~ (app(cons(X6, nil), X7) = X3) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X2, nil)) | ~ sP6(X0, X1, X2, X3)), inference(rectify, [], [f346])).
fof(f346, plain, ! [X3, X2, X1, X0] : ((? [X4] : (? [X5] : ((app(X5, cons(X4, nil)) = X3) & (app(cons(X4, nil), X5) = X2) & ssList(X5)) & ssItem(X4)) & ! [X6] : (! [X7] : (~ (app(X7, cons(X6, nil)) = X1) | ~ (app(cons(X6, nil), X7) = X0) | ~ ssList(X7)) | ~ ssItem(X6)) & neq(X1, nil)) | ~ sP6(X3, X2, X1, X0)), inference(nnf_transformation, [], [f233])).
fof(f6062, plain, (! [X0, X1] : ~ sP6(X0, X1, sK57, app(cons(sK54(sK57, sK56), nil), sK55(sK57, sK56))) | ~ spl60_1), inference(subsumption_resolution, [], [f6061, f990])).
fof(f990, plain, (ssList(sK55(sK57, sK56)) | ~ spl60_1), inference(resolution, [], [f558, f622])).
fof(f558, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssList(sK55(X0, X1))), inference(cnf_transformation, [], [f350])).
fof(f6061, plain, (! [X0, X1] : (~ sP6(X0, X1, sK57, app(cons(sK54(sK57, sK56), nil), sK55(sK57, sK56))) | ~ ssList(sK55(sK57, sK56))) | ~ spl60_1), inference(subsumption_resolution, [], [f6034, f989])).
fof(f989, plain, (ssItem(sK54(sK57, sK56)) | ~ spl60_1), inference(resolution, [], [f557, f622])).
fof(f557, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | ssItem(sK54(X0, X1))), inference(cnf_transformation, [], [f350])).
fof(f6034, plain, (! [X0, X1] : (~ sP6(X0, X1, sK57, app(cons(sK54(sK57, sK56), nil), sK55(sK57, sK56))) | ~ ssItem(sK54(sK57, sK56)) | ~ ssList(sK55(sK57, sK56))) | ~ spl60_1), inference(superposition, [], [f597, f4288])).
fof(f4288, plain, ((sK57 = app(sK55(sK57, sK56), cons(sK54(sK57, sK56), nil))) | ~ spl60_1), inference(resolution, [], [f622, f560])).
fof(f560, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | (app(sK55(X0, X1), cons(sK54(X0, X1), nil)) = X0)), inference(cnf_transformation, [], [f350])).
fof(f597, plain, ! [X6, X0, X7, X1] : (~ sP6(X0, X1, app(X7, cons(X6, nil)), app(cons(X6, nil), X7)) | ~ ssItem(X6) | ~ ssList(X7)), inference(equality_resolution, [], [f596])).
fof(f596, plain, ! [X6, X0, X7, X3, X1] : (~ (app(cons(X6, nil), X7) = X3) | ~ ssList(X7) | ~ ssItem(X6) | ~ sP6(X0, X1, app(X7, cons(X6, nil)), X3)), inference(equality_resolution, [], [f556])).
fof(f556, plain, ! [X6, X2, X0, X7, X3, X1] : (~ (app(X7, cons(X6, nil)) = X2) | ~ (app(cons(X6, nil), X7) = X3) | ~ ssList(X7) | ~ ssItem(X6) | ~ sP6(X0, X1, X2, X3)), inference(cnf_transformation, [], [f350])).
fof(f619, plain, spl60_2, inference(avatar_contradiction_clause, [], [f618])).
fof(f618, plain, ($false | spl60_2), inference(subsumption_resolution, [], [f617, f614])).
fof(f614, plain, neq(sK57, nil), inference(subsumption_resolution, [], [f567, f555])).
fof(f555, plain, ! [X2, X0, X3, X1] : (~ sP6(X0, X1, X2, X3) | neq(X2, nil)), inference(cnf_transformation, [], [f350])).
fof(f567, plain, (neq(sK57, nil) | sP6(sK59, sK58, sK57, sK56)), inference(cnf_transformation, [], [f355])).
fof(f617, plain, (~ neq(sK57, nil) | spl60_2), inference(forward_demodulation, [], [f612, f565])).
fof(f612, plain, (~ neq(sK59, nil) | spl60_2), inference(avatar_component_clause, [], [f610])).
fof(f610, plain, (spl60_2 <=> neq(sK59, nil)), introduced(avatar_definition, [new_symbols(naming, [spl60_2])])).
fof(f613, plain, (spl60_1 | ~ spl60_2), inference(avatar_split_clause, [], [f568, f610, f606])).
fof(f568, plain, (~ neq(sK59, nil) | sP6(sK59, sK58, sK57, sK56)), inference(cnf_transformation, [], [f355])).