fof(f336068, plain, $false, inference(avatar_sat_refutation, [], [f623, f628, f637, f650, f1574, f336067])).
fof(f336067, plain, ~ spl62_1, inference(avatar_contradiction_clause, [], [f336066])).
fof(f336066, plain, ($false | ~ spl62_1), inference(subsumption_resolution, [], [f336065, f53418])).
fof(f53418, plain, (ssList(sK56(sK57, sK61)) | ~ spl62_1), inference(subsumption_resolution, [], [f53417, f639])).
fof(f639, plain, equalelemsP(sK57), inference(backward_demodulation, [], [f571, f568])).
fof(f568, plain, (sK57 = sK59), inference(cnf_transformation, [], [f358])).
fof(f358, plain, ((((((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = sK59) | (nil = sK60)) & (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = sK61) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (sK60 = app(sK59, sK61)) & ssList(sK61)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60)) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57, sK58, sK59, sK60, sK61])], [f352, f357, f356, f355, f354, f353])).
fof(f353, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = X0)) | sP6(X0, X1)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = sK57)) | sP6(sK57, X1)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK57))), introduced(choice_axiom, [])).
fof(f354, plain, (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = sK57)) | sP6(sK57, X1)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (sK57 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK58))), introduced(choice_axiom, [])).
fof(f355, plain, (? [X2] : (? [X3] : (((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (sK57 = X2) & (sK58 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : (((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = sK59) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (app(sK59, X4) = X3) & ssList(X4)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) & ssList(sK59))), introduced(choice_axiom, [])).
fof(f356, plain, (? [X3] : (((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = sK59) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (app(sK59, X4) = X3) & ssList(X4)) & (sK57 = sK59) & (sK58 = X3) & ssList(X3)) => (((~ (nil = sK58) & (nil = sK57)) | sP6(sK57, sK58)) & (~ (nil = sK59) | (nil = sK60)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (app(sK59, X4) = sK60) & ssList(X4)) & (sK57 = sK59) & (sK58 = sK60) & ssList(sK60))), introduced(choice_axiom, [])).
fof(f357, plain, (? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (app(sK59, X4) = sK60) & ssList(X4)) => (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = sK61) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(sK59) & (sK60 = app(sK59, sK61)) & ssList(sK61))), introduced(choice_axiom, [])).
fof(f352, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = X0)) | sP6(X0, X1)) & (~ (nil = X2) | (nil = X3)) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (~ (app(X7, cons(X5, nil)) = X2) | ~ ssList(X7)) | ~ (app(cons(X5, nil), X6) = X4) | ~ ssList(X6)) | ~ ssItem(X5)) & equalelemsP(X2) & (app(X2, X4) = X3) & ssList(X4)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(rectify, [], [f234])).
fof(f234, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = X0)) | sP6(X0, X1)) & (~ (nil = X2) | (nil = X3)) & ? [X8] : (! [X9] : (! [X10] : (! [X11] : (~ (app(X11, cons(X9, nil)) = X2) | ~ ssList(X11)) | ~ (app(cons(X9, nil), X10) = X8) | ~ ssList(X10)) | ~ ssItem(X9)) & equalelemsP(X2) & (app(X2, X8) = X3) & ssList(X8)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f223, e233])).
fof(f233, plain, ! [X0, X1] : (! [X4] : (~ equalelemsP(X0) | ? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X0) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ (app(X0, X4) = X1) | ~ ssList(X4)) | ~ sP6(X0, X1)), inference(usedef, [], [e233])).
fof(e233, plain, ! [X0, X1] : (sP6(X0, X1) <=> ! [X4] : (~ equalelemsP(X0) | ? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X0) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ (app(X0, X4) = X1) | ~ ssList(X4))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f223, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~ (nil = X1) & (nil = X0)) | ! [X4] : (~ equalelemsP(X0) | ? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X0) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ (app(X0, X4) = X1) | ~ ssList(X4))) & (~ (nil = X2) | (nil = X3)) & ? [X8] : (! [X9] : (! [X10] : (! [X11] : (~ (app(X11, cons(X9, nil)) = X2) | ~ ssList(X11)) | ~ (app(cons(X9, nil), X10) = X8) | ~ ssList(X10)) | ~ ssItem(X9)) & equalelemsP(X2) & (app(X2, X8) = X3) & ssList(X8)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((~ (nil = X1) & (nil = X0)) | ! [X4] : (~ equalelemsP(X0) | ? [X5] : (? [X6] : ((? [X7] : ((app(X7, cons(X5, nil)) = X0) & ssList(X7)) & (app(cons(X5, nil), X6) = X4)) & ssList(X6)) & ssItem(X5)) | ~ (app(X0, X4) = X1) | ~ ssList(X4))) & (~ (nil = X2) | (nil = X3)) & ? [X8] : ((! [X9] : (! [X10] : (! [X11] : (~ (app(X11, cons(X9, nil)) = X2) | ~ ssList(X11)) | ~ (app(cons(X9, nil), X10) = X8) | ~ ssList(X10)) | ~ ssItem(X9)) & equalelemsP(X2) & (app(X2, X8) = X3)) & ssList(X8)) & (X0 = X2) & (X1 = X3)) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => ((((nil = X1) | ~ (nil = X0)) & ? [X4] : (equalelemsP(X0) & ! [X5] : (ssItem(X5) => ! [X6] : (ssList(X6) => (! [X7] : (ssList(X7) => ~ (app(X7, cons(X5, nil)) = X0)) | ~ (app(cons(X5, nil), X6) = X4)))) & (app(X0, X4) = X1) & ssList(X4))) | ((nil = X2) & ~ (nil = X3)) | ! [X8] : (ssList(X8) => (? [X9] : (? [X10] : (? [X11] : ((app(X11, cons(X9, nil)) = X2) & ssList(X11)) & (app(cons(X9, nil), X10) = X8) & ssList(X10)) & ssItem(X9)) | ~ equalelemsP(X2) | ~ (app(X2, X8) = X3))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => ((((nil = X1) | ~ (nil = X0)) & ? [X8] : (equalelemsP(X0) & ! [X9] : (ssItem(X9) => ! [X10] : (ssList(X10) => (! [X11] : (ssList(X11) => ~ (app(X11, cons(X9, nil)) = X0)) | ~ (app(cons(X9, nil), X10) = X8)))) & (app(X0, X8) = X1) & ssList(X8))) | ((nil = X2) & ~ (nil = X3)) | ! [X4] : (ssList(X4) => (? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X2) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ equalelemsP(X2) | ~ (app(X2, X4) = X3))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => ((((nil = X1) | ~ (nil = X0)) & ? [X8] : (equalelemsP(X0) & ! [X9] : (ssItem(X9) => ! [X10] : (ssList(X10) => (! [X11] : (ssList(X11) => ~ (app(X11, cons(X9, nil)) = X0)) | ~ (app(cons(X9, nil), X10) = X8)))) & (app(X0, X8) = X1) & ssList(X8))) | ((nil = X2) & ~ (nil = X3)) | ! [X4] : (ssList(X4) => (? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X2) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ equalelemsP(X2) | ~ (app(X2, X4) = X3))) | ~ (X0 = X2) | ~ (X1 = X3)))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC327+1.p', co1)).
fof(f571, plain, equalelemsP(sK59), inference(cnf_transformation, [], [f358])).
fof(f53417, plain, (ssList(sK56(sK57, sK61)) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f53416, f569])).
fof(f569, plain, ssList(sK61), inference(cnf_transformation, [], [f358])).
fof(f53416, plain, (ssList(sK56(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f53363, f618])).
fof(f618, plain, (sP6(sK57, sK58) | ~ spl62_1), inference(avatar_component_clause, [], [f616])).
fof(f616, plain, (spl62_1 <=> sP6(sK57, sK58)), introduced(avatar_definition, [new_symbols(naming, [spl62_1])])).
fof(f53363, plain, (~ sP6(sK57, sK58) | ssList(sK56(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57)), inference(superposition, [], [f604, f641])).
fof(f641, plain, (sK58 = app(sK57, sK61)), inference(backward_demodulation, [], [f638, f567])).
fof(f567, plain, (sK58 = sK60), inference(cnf_transformation, [], [f358])).
fof(f638, plain, (sK60 = app(sK57, sK61)), inference(backward_demodulation, [], [f570, f568])).
fof(f570, plain, (sK60 = app(sK59, sK61)), inference(cnf_transformation, [], [f358])).
fof(f604, plain, ! [X2, X0] : (~ sP6(X0, app(X0, X2)) | ssList(sK56(X0, X2)) | ~ ssList(X2) | ~ equalelemsP(X0)), inference(equality_resolution, [], [f561])).
fof(f561, plain, ! [X2, X0, X1] : (~ equalelemsP(X0) | ssList(sK56(X0, X2)) | ~ (app(X0, X2) = X1) | ~ ssList(X2) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f351])).
fof(f351, plain, ! [X0, X1] : (! [X2] : (~ equalelemsP(X0) | ((((app(sK56(X0, X2), cons(sK54(X0, X2), nil)) = X0) & ssList(sK56(X0, X2))) & (app(cons(sK54(X0, X2), nil), sK55(X0, X2)) = X2) & ssList(sK55(X0, X2))) & ssItem(sK54(X0, X2))) | ~ (app(X0, X2) = X1) | ~ ssList(X2)) | ~ sP6(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54, sK55, sK56])], [f347, f350, f349, f348])).
fof(f348, plain, ! [X2, X0] : (? [X3] : (? [X4] : (? [X5] : ((app(X5, cons(X3, nil)) = X0) & ssList(X5)) & (app(cons(X3, nil), X4) = X2) & ssList(X4)) & ssItem(X3)) => (? [X4] : (? [X5] : ((app(X5, cons(sK54(X0, X2), nil)) = X0) & ssList(X5)) & (app(cons(sK54(X0, X2), nil), X4) = X2) & ssList(X4)) & ssItem(sK54(X0, X2)))), introduced(choice_axiom, [])).
fof(f349, plain, ! [X2, X0] : (? [X4] : (? [X5] : ((app(X5, cons(sK54(X0, X2), nil)) = X0) & ssList(X5)) & (app(cons(sK54(X0, X2), nil), X4) = X2) & ssList(X4)) => (? [X5] : ((app(X5, cons(sK54(X0, X2), nil)) = X0) & ssList(X5)) & (app(cons(sK54(X0, X2), nil), sK55(X0, X2)) = X2) & ssList(sK55(X0, X2)))), introduced(choice_axiom, [])).
fof(f350, plain, ! [X2, X0] : (? [X5] : ((app(X5, cons(sK54(X0, X2), nil)) = X0) & ssList(X5)) => ((app(sK56(X0, X2), cons(sK54(X0, X2), nil)) = X0) & ssList(sK56(X0, X2)))), introduced(choice_axiom, [])).
fof(f347, plain, ! [X0, X1] : (! [X2] : (~ equalelemsP(X0) | ? [X3] : (? [X4] : (? [X5] : ((app(X5, cons(X3, nil)) = X0) & ssList(X5)) & (app(cons(X3, nil), X4) = X2) & ssList(X4)) & ssItem(X3)) | ~ (app(X0, X2) = X1) | ~ ssList(X2)) | ~ sP6(X0, X1)), inference(rectify, [], [f346])).
fof(f346, plain, ! [X0, X1] : (! [X4] : (~ equalelemsP(X0) | ? [X5] : (? [X6] : (? [X7] : ((app(X7, cons(X5, nil)) = X0) & ssList(X7)) & (app(cons(X5, nil), X6) = X4) & ssList(X6)) & ssItem(X5)) | ~ (app(X0, X4) = X1) | ~ ssList(X4)) | ~ sP6(X0, X1)), inference(nnf_transformation, [], [f233])).
fof(f336065, plain, (~ ssList(sK56(sK57, sK61)) | ~ spl62_1), inference(trivial_inequality_removal, [], [f336064])).
fof(f336064, plain, (~ (sK57 = sK57) | ~ ssList(sK56(sK57, sK61)) | ~ spl62_1), inference(superposition, [], [f68820, f60978])).
fof(f60978, plain, ((sK57 = app(sK56(sK57, sK61), cons(sK54(sK57, sK61), nil))) | ~ spl62_1), inference(subsumption_resolution, [], [f60977, f639])).
fof(f60977, plain, ((sK57 = app(sK56(sK57, sK61), cons(sK54(sK57, sK61), nil))) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60976, f569])).
fof(f60976, plain, ((sK57 = app(sK56(sK57, sK61), cons(sK54(sK57, sK61), nil))) | ~ ssList(sK61) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60949, f618])).
fof(f60949, plain, (~ sP6(sK57, sK58) | (sK57 = app(sK56(sK57, sK61), cons(sK54(sK57, sK61), nil))) | ~ ssList(sK61) | ~ equalelemsP(sK57)), inference(superposition, [], [f603, f641])).
fof(f603, plain, ! [X2, X0] : (~ sP6(X0, app(X0, X2)) | (app(sK56(X0, X2), cons(sK54(X0, X2), nil)) = X0) | ~ ssList(X2) | ~ equalelemsP(X0)), inference(equality_resolution, [], [f562])).
fof(f562, plain, ! [X2, X0, X1] : (~ equalelemsP(X0) | (app(sK56(X0, X2), cons(sK54(X0, X2), nil)) = X0) | ~ (app(X0, X2) = X1) | ~ ssList(X2) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f351])).
fof(f68820, plain, (! [X0] : (~ (sK57 = app(X0, cons(sK54(sK57, sK61), nil))) | ~ ssList(X0)) | ~ spl62_1), inference(subsumption_resolution, [], [f68819, f60993])).
fof(f60993, plain, (ssItem(sK54(sK57, sK61)) | ~ spl62_1), inference(subsumption_resolution, [], [f60992, f639])).
fof(f60992, plain, (ssItem(sK54(sK57, sK61)) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60991, f569])).
fof(f60991, plain, (ssItem(sK54(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60953, f618])).
fof(f60953, plain, (~ sP6(sK57, sK58) | ssItem(sK54(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57)), inference(superposition, [], [f607, f641])).
fof(f607, plain, ! [X2, X0] : (~ sP6(X0, app(X0, X2)) | ssItem(sK54(X0, X2)) | ~ ssList(X2) | ~ equalelemsP(X0)), inference(equality_resolution, [], [f558])).
fof(f558, plain, ! [X2, X0, X1] : (~ equalelemsP(X0) | ssItem(sK54(X0, X2)) | ~ (app(X0, X2) = X1) | ~ ssList(X2) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f351])).
fof(f68819, plain, (! [X0] : (~ ssList(X0) | ~ (sK57 = app(X0, cons(sK54(sK57, sK61), nil))) | ~ ssItem(sK54(sK57, sK61))) | ~ spl62_1), inference(subsumption_resolution, [], [f68818, f60989])).
fof(f60989, plain, (ssList(sK55(sK57, sK61)) | ~ spl62_1), inference(subsumption_resolution, [], [f60988, f639])).
fof(f60988, plain, (ssList(sK55(sK57, sK61)) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60987, f569])).
fof(f60987, plain, (ssList(sK55(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60952, f618])).
fof(f60952, plain, (~ sP6(sK57, sK58) | ssList(sK55(sK57, sK61)) | ~ ssList(sK61) | ~ equalelemsP(sK57)), inference(superposition, [], [f606, f641])).
fof(f606, plain, ! [X2, X0] : (~ sP6(X0, app(X0, X2)) | ssList(sK55(X0, X2)) | ~ ssList(X2) | ~ equalelemsP(X0)), inference(equality_resolution, [], [f559])).
fof(f559, plain, ! [X2, X0, X1] : (~ equalelemsP(X0) | ssList(sK55(X0, X2)) | ~ (app(X0, X2) = X1) | ~ ssList(X2) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f351])).
fof(f68818, plain, (! [X0] : (~ ssList(X0) | ~ (sK57 = app(X0, cons(sK54(sK57, sK61), nil))) | ~ ssList(sK55(sK57, sK61)) | ~ ssItem(sK54(sK57, sK61))) | ~ spl62_1), inference(trivial_inequality_removal, [], [f68774])).
fof(f68774, plain, (! [X0] : (~ (sK61 = sK61) | ~ ssList(X0) | ~ (sK57 = app(X0, cons(sK54(sK57, sK61), nil))) | ~ ssList(sK55(sK57, sK61)) | ~ ssItem(sK54(sK57, sK61))) | ~ spl62_1), inference(superposition, [], [f640, f60983])).
fof(f60983, plain, ((sK61 = app(cons(sK54(sK57, sK61), nil), sK55(sK57, sK61))) | ~ spl62_1), inference(subsumption_resolution, [], [f60982, f639])).
fof(f60982, plain, ((sK61 = app(cons(sK54(sK57, sK61), nil), sK55(sK57, sK61))) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60981, f569])).
fof(f60981, plain, ((sK61 = app(cons(sK54(sK57, sK61), nil), sK55(sK57, sK61))) | ~ ssList(sK61) | ~ equalelemsP(sK57) | ~ spl62_1), inference(subsumption_resolution, [], [f60951, f618])).
fof(f60951, plain, (~ sP6(sK57, sK58) | (sK61 = app(cons(sK54(sK57, sK61), nil), sK55(sK57, sK61))) | ~ ssList(sK61) | ~ equalelemsP(sK57)), inference(superposition, [], [f605, f641])).
fof(f605, plain, ! [X2, X0] : (~ sP6(X0, app(X0, X2)) | (app(cons(sK54(X0, X2), nil), sK55(X0, X2)) = X2) | ~ ssList(X2) | ~ equalelemsP(X0)), inference(equality_resolution, [], [f560])).
fof(f560, plain, ! [X2, X0, X1] : (~ equalelemsP(X0) | (app(cons(sK54(X0, X2), nil), sK55(X0, X2)) = X2) | ~ (app(X0, X2) = X1) | ~ ssList(X2) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f351])).
fof(f640, plain, ! [X6, X7, X5] : (~ (app(cons(X5, nil), X6) = sK61) | ~ ssList(X7) | ~ (app(X7, cons(X5, nil)) = sK57) | ~ ssList(X6) | ~ ssItem(X5)), inference(backward_demodulation, [], [f572, f568])).
fof(f572, plain, ! [X6, X7, X5] : (~ (app(X7, cons(X5, nil)) = sK59) | ~ ssList(X7) | ~ (app(cons(X5, nil), X6) = sK61) | ~ ssList(X6) | ~ ssItem(X5)), inference(cnf_transformation, [], [f358])).
fof(f1574, plain, (spl62_2 | ~ spl62_4), inference(avatar_split_clause, [], [f1532, f630, f620])).
fof(f620, plain, (spl62_2 <=> (nil = sK58)), introduced(avatar_definition, [new_symbols(naming, [spl62_2])])).
fof(f630, plain, (spl62_4 <=> (nil = sK60)), introduced(avatar_definition, [new_symbols(naming, [spl62_4])])).
fof(f1532, plain, ((nil = sK58) | ~ spl62_4), inference(forward_demodulation, [], [f567, f632])).
fof(f632, plain, ((nil = sK60) | ~ spl62_4), inference(avatar_component_clause, [], [f630])).
fof(f650, plain, (~ spl62_3 | spl62_5), inference(avatar_contradiction_clause, [], [f649])).
fof(f649, plain, ($false | (~ spl62_3 | spl62_5)), inference(subsumption_resolution, [], [f648, f636])).
fof(f636, plain, (~ (nil = sK59) | spl62_5), inference(avatar_component_clause, [], [f634])).
fof(f634, plain, (spl62_5 <=> (nil = sK59)), introduced(avatar_definition, [new_symbols(naming, [spl62_5])])).
fof(f648, plain, ((nil = sK59) | ~ spl62_3), inference(backward_demodulation, [], [f568, f627])).
fof(f627, plain, ((nil = sK57) | ~ spl62_3), inference(avatar_component_clause, [], [f625])).
fof(f625, plain, (spl62_3 <=> (nil = sK57)), introduced(avatar_definition, [new_symbols(naming, [spl62_3])])).
fof(f637, plain, (spl62_4 | ~ spl62_5), inference(avatar_split_clause, [], [f573, f634, f630])).
fof(f573, plain, (~ (nil = sK59) | (nil = sK60)), inference(cnf_transformation, [], [f358])).
fof(f628, plain, (spl62_1 | spl62_3), inference(avatar_split_clause, [], [f574, f625, f616])).
fof(f574, plain, ((nil = sK57) | sP6(sK57, sK58)), inference(cnf_transformation, [], [f358])).
fof(f623, plain, (spl62_1 | ~ spl62_2), inference(avatar_split_clause, [], [f575, f620, f616])).
fof(f575, plain, (~ (nil = sK58) | sP6(sK57, sK58)), inference(cnf_transformation, [], [f358])).