fof(f411050, plain, $false, inference(avatar_sat_refutation, [], [f632, f637, f642, f647, f651, f2250, f262555, f262691, f263008, f263182, f411015])).
fof(f411015, plain, (spl68_10159 | ~ spl68_1 | ~ spl68_10156 | ~ spl68_10157 | ~ spl68_10201), inference(avatar_split_clause, [], [f411014, f263090, f262400, f262396, f625, f262410])).
fof(f262410, plain, (spl68_10159 <=> ! [X3, X2] : ~ sP7(X2, X3, sK62, sK61)), introduced(avatar_definition, [new_symbols(naming, [spl68_10159])])).
fof(f625, plain, (spl68_1 <=> sP7(sK64, sK63, sK62, sK61)), introduced(avatar_definition, [new_symbols(naming, [spl68_1])])).
fof(f262396, plain, (spl68_10156 <=> ssItem(sK59(sK61, sK62))), introduced(avatar_definition, [new_symbols(naming, [spl68_10156])])).
fof(f262400, plain, (spl68_10157 <=> ssItem(sK58(sK61, sK62))), introduced(avatar_definition, [new_symbols(naming, [spl68_10157])])).
fof(f263090, plain, (spl68_10201 <=> ssList(sK60(sK61, sK62))), introduced(avatar_definition, [new_symbols(naming, [spl68_10201])])).
fof(f411014, plain, (! [X2, X3] : ~ sP7(X2, X3, sK62, sK61) | (~ spl68_1 | ~ spl68_10156 | ~ spl68_10157 | ~ spl68_10201)), inference(forward_demodulation, [], [f411013, f93571])).
fof(f93571, plain, ((sK61 = app(app(cons(sK59(sK61, sK62), nil), cons(sK58(sK61, sK62), nil)), sK60(sK61, sK62))) | ~ spl68_1), inference(resolution, [], [f13964, f576])).
fof(f576, plain, ! [X0, X1] : (~ sP6(X0, X1) | (app(app(cons(sK59(X0, X1), nil), cons(sK58(X0, X1), nil)), sK60(X0, X1)) = X0)), inference(cnf_transformation, [], [f358])).
fof(f358, plain, ! [X0, X1] : (((((app(app(cons(sK59(X0, X1), nil), cons(sK58(X0, X1), nil)), sK60(X0, X1)) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(sK59(X0, X1), nil)), sK60(X0, X1)) = X1) & ssList(sK60(X0, X1))) & ssItem(sK59(X0, X1))) & ssItem(sK58(X0, X1))) | ~ sP6(X0, X1)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK58, sK59, sK60])], [f354, f357, f356, f355])).
fof(f355, plain, ! [X1, X0] : (? [X2] : (? [X3] : (? [X4] : ((app(app(cons(X3, nil), cons(X2, nil)), X4) = X0) & (app(app(cons(X2, nil), cons(X3, nil)), X4) = X1) & ssList(X4)) & ssItem(X3)) & ssItem(X2)) => (? [X3] : (? [X4] : ((app(app(cons(X3, nil), cons(sK58(X0, X1), nil)), X4) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(X3, nil)), X4) = X1) & ssList(X4)) & ssItem(X3)) & ssItem(sK58(X0, X1)))), introduced(choice_axiom, [])).
fof(f356, plain, ! [X1, X0] : (? [X3] : (? [X4] : ((app(app(cons(X3, nil), cons(sK58(X0, X1), nil)), X4) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(X3, nil)), X4) = X1) & ssList(X4)) & ssItem(X3)) => (? [X4] : ((app(app(cons(sK59(X0, X1), nil), cons(sK58(X0, X1), nil)), X4) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(sK59(X0, X1), nil)), X4) = X1) & ssList(X4)) & ssItem(sK59(X0, X1)))), introduced(choice_axiom, [])).
fof(f357, plain, ! [X1, X0] : (? [X4] : ((app(app(cons(sK59(X0, X1), nil), cons(sK58(X0, X1), nil)), X4) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(sK59(X0, X1), nil)), X4) = X1) & ssList(X4)) => ((app(app(cons(sK59(X0, X1), nil), cons(sK58(X0, X1), nil)), sK60(X0, X1)) = X0) & (app(app(cons(sK58(X0, X1), nil), cons(sK59(X0, X1), nil)), sK60(X0, X1)) = X1) & ssList(sK60(X0, X1)))), introduced(choice_axiom, [])).
fof(f354, plain, ! [X0, X1] : (? [X2] : (? [X3] : (? [X4] : ((app(app(cons(X3, nil), cons(X2, nil)), X4) = X0) & (app(app(cons(X2, nil), cons(X3, nil)), X4) = X1) & ssList(X4)) & ssItem(X3)) & ssItem(X2)) | ~ sP6(X0, X1)), inference(rectify, [], [f353])).
fof(f353, plain, ! [X2, X3] : (? [X10] : (? [X11] : (? [X12] : ((app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) & (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3) & ssList(X12)) & ssItem(X11)) & ssItem(X10)) | ~ sP6(X2, X3)), inference(nnf_transformation, [], [f233])).
fof(f233, plain, ! [X2, X3] : (? [X10] : (? [X11] : (? [X12] : ((app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) & (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3) & ssList(X12)) & ssItem(X11)) & ssItem(X10)) | ~ sP6(X2, X3)), inference(usedef, [], [e233])).
fof(e233, plain, ! [X2, X3] : (sP6(X2, X3) <=> ? [X10] : (? [X11] : (? [X12] : ((app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) & (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3) & ssList(X12)) & ssItem(X11)) & ssItem(X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f13964, plain, (sP6(sK61, sK62) | ~ spl68_1), inference(resolution, [], [f6388, f571])).
fof(f571, plain, ! [X2, X0, X3, X1] : (~ sP7(X0, X1, X2, X3) | sP6(X1, X0)), inference(cnf_transformation, [], [f352])).
fof(f352, plain, ! [X0, X1, X2, X3] : ((sP6(X1, X0) & ((((app(app(cons(sK55(X2), nil), cons(sK56(X2), nil)), sK57(X2)) = X2) & ssList(sK57(X2))) & ssItem(sK56(X2))) & ssItem(sK55(X2))) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X8, nil), cons(X7, nil)), X9) = X3) | ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X2) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | ~ sP7(X0, X1, X2, X3)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK55, sK56, sK57])], [f348, f351, f350, f349])).
fof(f349, plain, ! [X2] : (? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X2) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) => (? [X5] : (? [X6] : ((app(app(cons(sK55(X2), nil), cons(X5, nil)), X6) = X2) & ssList(X6)) & ssItem(X5)) & ssItem(sK55(X2)))), introduced(choice_axiom, [])).
fof(f350, plain, ! [X2] : (? [X5] : (? [X6] : ((app(app(cons(sK55(X2), nil), cons(X5, nil)), X6) = X2) & ssList(X6)) & ssItem(X5)) => (? [X6] : ((app(app(cons(sK55(X2), nil), cons(sK56(X2), nil)), X6) = X2) & ssList(X6)) & ssItem(sK56(X2)))), introduced(choice_axiom, [])).
fof(f351, plain, ! [X2] : (? [X6] : ((app(app(cons(sK55(X2), nil), cons(sK56(X2), nil)), X6) = X2) & ssList(X6)) => ((app(app(cons(sK55(X2), nil), cons(sK56(X2), nil)), sK57(X2)) = X2) & ssList(sK57(X2)))), introduced(choice_axiom, [])).
fof(f348, plain, ! [X0, X1, X2, X3] : ((sP6(X1, X0) & ? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X2) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X8, nil), cons(X7, nil)), X9) = X3) | ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X2) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | ~ sP7(X0, X1, X2, X3)), inference(rectify, [], [f347])).
fof(f347, plain, ! [X3, X2, X1, X0] : ((sP6(X2, X3) & ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X1) & ssList(X15)) & ssItem(X14)) & ssItem(X13)) & ! [X16] : (! [X17] : (! [X18] : (~ (app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) | ~ (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) | ~ ssList(X18)) | ~ ssItem(X17)) | ~ ssItem(X16))) | ~ sP7(X3, X2, X1, X0)), inference(nnf_transformation, [], [f234])).
fof(f234, plain, ! [X3, X2, X1, X0] : ((sP6(X2, X3) & ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X1) & ssList(X15)) & ssItem(X14)) & ssItem(X13)) & ! [X16] : (! [X17] : (! [X18] : (~ (app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) | ~ (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) | ~ ssList(X18)) | ~ ssItem(X17)) | ~ ssItem(X16))) | ~ sP7(X3, X2, X1, X0)), inference(usedef, [], [e234])).
fof(e234, plain, ! [X3, X2, X1, X0] : (sP7(X3, X2, X1, X0) <=> (sP6(X2, X3) & ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X1) & ssList(X15)) & ssItem(X14)) & ssItem(X13)) & ! [X16] : (! [X17] : (! [X18] : (~ (app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) | ~ (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) | ~ ssList(X18)) | ~ ssItem(X17)) | ~ ssItem(X16)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f6388, plain, (sP7(sK62, sK61, sK62, sK61) | ~ spl68_1), inference(backward_demodulation, [], [f6387, f581])).
fof(f581, plain, (sK62 = sK64), inference(cnf_transformation, [], [f366])).
fof(f366, plain, ((((((((((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), sK67)) & ssList(sK67)) & ssItem(sK66)) & ssItem(sK65)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK64) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(sK64, sK63, sK62, sK61)) & (sK61 = sK63) & (sK62 = sK64) & ssList(sK64)) & ssList(sK63)) & ssList(sK62)) & ssList(sK61)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK61, sK62, sK63, sK64, sK65, sK66, sK67])], [f235, f365, f364, f363, f362, f361, f360, f359])).
fof(f359, plain, (? [X0] : (? [X1] : (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, X1, X0)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, X1, sK61)) & (sK61 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK61))), introduced(choice_axiom, [])).
fof(f360, plain, (? [X1] : (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, X1, sK61)) & (sK61 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, sK62, sK61)) & (sK61 = X2) & (sK62 = X3) & ssList(X3)) & ssList(X2)) & ssList(sK62))), introduced(choice_axiom, [])).
fof(f361, plain, (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, sK62, sK61)) & (sK61 = X2) & (sK62 = X3) & ssList(X3)) & ssList(X2)) => (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, sK63, sK62, sK61)) & (sK61 = sK63) & (sK62 = X3) & ssList(X3)) & ssList(sK63))), introduced(choice_axiom, [])).
fof(f362, plain, (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, sK63, sK62, sK61)) & (sK61 = sK63) & (sK62 = X3) & ssList(X3)) => (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK64) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(sK64, sK63, sK62, sK61)) & (sK61 = sK63) & (sK62 = sK64) & ssList(sK64))), introduced(choice_axiom, [])).
fof(f363, plain, (? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = sK62) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) => (? [X5] : (? [X6] : ((sK62 = app(app(cons(sK65, nil), cons(X5, nil)), X6)) & ssList(X6)) & ssItem(X5)) & ssItem(sK65))), introduced(choice_axiom, [])).
fof(f364, plain, (? [X5] : (? [X6] : ((sK62 = app(app(cons(sK65, nil), cons(X5, nil)), X6)) & ssList(X6)) & ssItem(X5)) => (? [X6] : ((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), X6)) & ssList(X6)) & ssItem(sK66))), introduced(choice_axiom, [])).
fof(f365, plain, (? [X6] : ((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), X6)) & ssList(X6)) => ((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), sK67)) & ssList(sK67))), introduced(choice_axiom, [])).
fof(f235, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | sP7(X3, X2, X1, X0)) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(definition_folding, [], [f223, e234, e233])).
fof(f223, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | (? [X10] : (? [X11] : (? [X12] : ((app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) & (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3) & ssList(X12)) & ssItem(X11)) & ssItem(X10)) & ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X1) & ssList(X15)) & ssItem(X14)) & ssItem(X13)) & ! [X16] : (! [X17] : (! [X18] : (~ (app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) | ~ (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) | ~ ssList(X18)) | ~ ssItem(X17)) | ~ ssItem(X16)))) & (X0 = X2) & (X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)) & ! [X7] : (! [X8] : (! [X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) | ~ ssList(X9)) | ~ ssItem(X8)) | ~ ssItem(X7))) | (? [X10] : (? [X11] : (? [X12] : (((app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) & (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3)) & ssList(X12)) & ssItem(X11)) & ssItem(X10)) & ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X1) & ssList(X15)) & ssItem(X14)) & ssItem(X13)) & ! [X16] : (! [X17] : (! [X18] : (~ (app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) | ~ (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) | ~ ssList(X18)) | ~ ssItem(X17)) | ~ ssItem(X16)))) & (X0 = X2) & (X1 = X3)) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)), inference(ennf_transformation, [], [f98])).
fof(f98, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((! [X4] : (ssItem(X4) => ! [X5] : (ssItem(X5) => ! [X6] : (ssList(X6) => ~ (app(app(cons(X4, nil), cons(X5, nil)), X6) = X1)))) | ? [X7] : (? [X8] : (? [X9] : ((app(app(cons(X7, nil), cons(X8, nil)), X9) = X3) & ssList(X9)) & ssItem(X8)) & ssItem(X7))) & (! [X10] : (ssItem(X10) => ! [X11] : (ssItem(X11) => ! [X12] : (ssList(X12) => (~ (app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) | ~ (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3))))) | ! [X13] : (ssItem(X13) => ! [X14] : (ssItem(X14) => ! [X15] : (ssList(X15) => ~ (app(app(cons(X13, nil), cons(X14, nil)), X15) = X1)))) | ? [X16] : (? [X17] : (? [X18] : ((app(app(cons(X17, nil), cons(X16, nil)), X18) = X0) & (app(app(cons(X16, nil), cons(X17, nil)), X18) = X1) & ssList(X18)) & ssItem(X17)) & ssItem(X16)))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(rectify, [], [f97])).
fof(f97, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((! [X7] : (ssItem(X7) => ! [X8] : (ssItem(X8) => ! [X9] : (ssList(X9) => ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X1)))) | ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X3) & ssList(X15)) & ssItem(X14)) & ssItem(X13))) & (! [X10] : (ssItem(X10) => ! [X11] : (ssItem(X11) => ! [X12] : (ssList(X12) => (~ (app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) | ~ (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3))))) | ! [X7] : (ssItem(X7) => ! [X8] : (ssItem(X8) => ! [X9] : (ssList(X9) => ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X1)))) | ? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X5, nil), cons(X4, nil)), X6) = X0) & (app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)))) | ~ (X0 = X2) | ~ (X1 = X3)))))), inference(negated_conjecture, [], [f96])).
fof(f96, plain, ~ ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((! [X7] : (ssItem(X7) => ! [X8] : (ssItem(X8) => ! [X9] : (ssList(X9) => ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X1)))) | ? [X13] : (? [X14] : (? [X15] : ((app(app(cons(X13, nil), cons(X14, nil)), X15) = X3) & ssList(X15)) & ssItem(X14)) & ssItem(X13))) & (! [X10] : (ssItem(X10) => ! [X11] : (ssItem(X11) => ! [X12] : (ssList(X12) => (~ (app(app(cons(X11, nil), cons(X10, nil)), X12) = X2) | ~ (app(app(cons(X10, nil), cons(X11, nil)), X12) = X3))))) | ! [X7] : (ssItem(X7) => ! [X8] : (ssItem(X8) => ! [X9] : (ssList(X9) => ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X1)))) | ? [X4] : (? [X5] : (? [X6] : ((app(app(cons(X5, nil), cons(X4, nil)), X6) = X0) & (app(app(cons(X4, nil), cons(X5, nil)), X6) = X1) & ssList(X6)) & ssItem(X5)) & ssItem(X4)))) | ~ (X0 = X2) | ~ (X1 = X3)))))), file('/home/ubuntu/library/tptp/Problems/SWC/SWC413+1.p', co1)).
fof(f6387, plain, (sP7(sK64, sK61, sK62, sK61) | ~ spl68_1), inference(forward_demodulation, [], [f627, f582])).
fof(f582, plain, (sK61 = sK63), inference(cnf_transformation, [], [f366])).
fof(f627, plain, (sP7(sK64, sK63, sK62, sK61) | ~ spl68_1), inference(avatar_component_clause, [], [f625])).
fof(f411013, plain, (! [X2, X3] : ~ sP7(X2, X3, sK62, app(app(cons(sK59(sK61, sK62), nil), cons(sK58(sK61, sK62), nil)), sK60(sK61, sK62))) | (~ spl68_1 | ~ spl68_10156 | ~ spl68_10157 | ~ spl68_10201)), inference(subsumption_resolution, [], [f411012, f263091])).
fof(f263091, plain, (ssList(sK60(sK61, sK62)) | ~ spl68_10201), inference(avatar_component_clause, [], [f263090])).
fof(f411012, plain, (! [X2, X3] : (~ sP7(X2, X3, sK62, app(app(cons(sK59(sK61, sK62), nil), cons(sK58(sK61, sK62), nil)), sK60(sK61, sK62))) | ~ ssList(sK60(sK61, sK62))) | (~ spl68_1 | ~ spl68_10156 | ~ spl68_10157)), inference(subsumption_resolution, [], [f411011, f262401])).
fof(f262401, plain, (ssItem(sK58(sK61, sK62)) | ~ spl68_10157), inference(avatar_component_clause, [], [f262400])).
fof(f411011, plain, (! [X2, X3] : (~ sP7(X2, X3, sK62, app(app(cons(sK59(sK61, sK62), nil), cons(sK58(sK61, sK62), nil)), sK60(sK61, sK62))) | ~ ssItem(sK58(sK61, sK62)) | ~ ssList(sK60(sK61, sK62))) | (~ spl68_1 | ~ spl68_10156)), inference(subsumption_resolution, [], [f410951, f262397])).
fof(f262397, plain, (ssItem(sK59(sK61, sK62)) | ~ spl68_10156), inference(avatar_component_clause, [], [f262396])).
fof(f410951, plain, (! [X2, X3] : (~ sP7(X2, X3, sK62, app(app(cons(sK59(sK61, sK62), nil), cons(sK58(sK61, sK62), nil)), sK60(sK61, sK62))) | ~ ssItem(sK59(sK61, sK62)) | ~ ssItem(sK58(sK61, sK62)) | ~ ssList(sK60(sK61, sK62))) | ~ spl68_1), inference(superposition, [], [f616, f93570])).
fof(f93570, plain, ((sK62 = app(app(cons(sK58(sK61, sK62), nil), cons(sK59(sK61, sK62), nil)), sK60(sK61, sK62))) | ~ spl68_1), inference(resolution, [], [f13964, f575])).
fof(f575, plain, ! [X0, X1] : (~ sP6(X0, X1) | (app(app(cons(sK58(X0, X1), nil), cons(sK59(X0, X1), nil)), sK60(X0, X1)) = X1)), inference(cnf_transformation, [], [f358])).
fof(f616, plain, ! [X0, X8, X7, X1, X9] : (~ sP7(X0, X1, app(app(cons(X7, nil), cons(X8, nil)), X9), app(app(cons(X8, nil), cons(X7, nil)), X9)) | ~ ssItem(X8) | ~ ssItem(X7) | ~ ssList(X9)), inference(equality_resolution, [], [f615])).
fof(f615, plain, ! [X2, X0, X8, X7, X1, X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X2) | ~ ssList(X9) | ~ ssItem(X8) | ~ ssItem(X7) | ~ sP7(X0, X1, X2, app(app(cons(X8, nil), cons(X7, nil)), X9))), inference(equality_resolution, [], [f566])).
fof(f566, plain, ! [X2, X0, X8, X7, X3, X1, X9] : (~ (app(app(cons(X8, nil), cons(X7, nil)), X9) = X3) | ~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = X2) | ~ ssList(X9) | ~ ssItem(X8) | ~ ssItem(X7) | ~ sP7(X0, X1, X2, X3)), inference(cnf_transformation, [], [f352])).
fof(f263182, plain, (~ spl68_1 | spl68_10201), inference(avatar_contradiction_clause, [], [f263181])).
fof(f263181, plain, ($false | (~ spl68_1 | spl68_10201)), inference(subsumption_resolution, [], [f263180, f13964])).
fof(f263180, plain, (~ sP6(sK61, sK62) | spl68_10201), inference(resolution, [], [f263092, f574])).
fof(f574, plain, ! [X0, X1] : (ssList(sK60(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f358])).
fof(f263092, plain, (~ ssList(sK60(sK61, sK62)) | spl68_10201), inference(avatar_component_clause, [], [f263090])).
fof(f263008, plain, (~ spl68_1 | ~ spl68_10159), inference(avatar_contradiction_clause, [], [f263007])).
fof(f263007, plain, ($false | (~ spl68_1 | ~ spl68_10159)), inference(resolution, [], [f262411, f6388])).
fof(f262411, plain, (! [X2, X3] : ~ sP7(X2, X3, sK62, sK61) | ~ spl68_10159), inference(avatar_component_clause, [], [f262410])).
fof(f262691, plain, (~ spl68_1 | spl68_10156), inference(avatar_contradiction_clause, [], [f262690])).
fof(f262690, plain, ($false | (~ spl68_1 | spl68_10156)), inference(subsumption_resolution, [], [f262689, f13964])).
fof(f262689, plain, (~ sP6(sK61, sK62) | spl68_10156), inference(resolution, [], [f262398, f573])).
fof(f573, plain, ! [X0, X1] : (ssItem(sK59(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f358])).
fof(f262398, plain, (~ ssItem(sK59(sK61, sK62)) | spl68_10156), inference(avatar_component_clause, [], [f262396])).
fof(f262555, plain, (~ spl68_1 | spl68_10157), inference(avatar_contradiction_clause, [], [f262554])).
fof(f262554, plain, ($false | (~ spl68_1 | spl68_10157)), inference(subsumption_resolution, [], [f262553, f13964])).
fof(f262553, plain, (~ sP6(sK61, sK62) | spl68_10157), inference(resolution, [], [f262402, f572])).
fof(f572, plain, ! [X0, X1] : (ssItem(sK58(X0, X1)) | ~ sP6(X0, X1)), inference(cnf_transformation, [], [f358])).
fof(f262402, plain, (~ ssItem(sK58(sK61, sK62)) | spl68_10157), inference(avatar_component_clause, [], [f262400])).
fof(f2250, plain, (~ spl68_4 | ~ spl68_2 | ~ spl68_3 | ~ spl68_5 | ~ spl68_6), inference(avatar_split_clause, [], [f2249, f649, f644, f634, f629, f639])).
fof(f639, plain, (spl68_4 <=> ssItem(sK66)), introduced(avatar_definition, [new_symbols(naming, [spl68_4])])).
fof(f629, plain, (spl68_2 <=> (sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), sK67))), introduced(avatar_definition, [new_symbols(naming, [spl68_2])])).
fof(f634, plain, (spl68_3 <=> ssList(sK67)), introduced(avatar_definition, [new_symbols(naming, [spl68_3])])).
fof(f644, plain, (spl68_5 <=> ssItem(sK65)), introduced(avatar_definition, [new_symbols(naming, [spl68_5])])).
fof(f649, plain, (spl68_6 <=> ! [X9, X7, X8] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK64) | ~ ssItem(X7) | ~ ssItem(X8) | ~ ssList(X9))), introduced(avatar_definition, [new_symbols(naming, [spl68_6])])).
fof(f2249, plain, (~ ssItem(sK66) | (~ spl68_2 | ~ spl68_3 | ~ spl68_5 | ~ spl68_6)), inference(subsumption_resolution, [], [f2245, f636])).
fof(f636, plain, (ssList(sK67) | ~ spl68_3), inference(avatar_component_clause, [], [f634])).
fof(f2245, plain, (~ ssItem(sK66) | ~ ssList(sK67) | (~ spl68_2 | ~ spl68_5 | ~ spl68_6)), inference(subsumption_resolution, [], [f2244, f646])).
fof(f646, plain, (ssItem(sK65) | ~ spl68_5), inference(avatar_component_clause, [], [f644])).
fof(f2244, plain, (~ ssItem(sK65) | ~ ssItem(sK66) | ~ ssList(sK67) | (~ spl68_2 | ~ spl68_6)), inference(trivial_inequality_removal, [], [f2243])).
fof(f2243, plain, (~ (sK62 = sK62) | ~ ssItem(sK65) | ~ ssItem(sK66) | ~ ssList(sK67) | (~ spl68_2 | ~ spl68_6)), inference(superposition, [], [f654, f631])).
fof(f631, plain, ((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), sK67)) | ~ spl68_2), inference(avatar_component_clause, [], [f629])).
fof(f654, plain, (! [X8, X7, X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK62) | ~ ssItem(X7) | ~ ssItem(X8) | ~ ssList(X9)) | ~ spl68_6), inference(forward_demodulation, [], [f650, f581])).
fof(f650, plain, (! [X8, X7, X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK64) | ~ ssItem(X7) | ~ ssItem(X8) | ~ ssList(X9)) | ~ spl68_6), inference(avatar_component_clause, [], [f649])).
fof(f651, plain, (spl68_1 | spl68_6), inference(avatar_split_clause, [], [f583, f649, f625])).
fof(f583, plain, ! [X8, X7, X9] : (~ (app(app(cons(X7, nil), cons(X8, nil)), X9) = sK64) | ~ ssList(X9) | ~ ssItem(X8) | ~ ssItem(X7) | sP7(sK64, sK63, sK62, sK61)), inference(cnf_transformation, [], [f366])).
fof(f647, plain, (spl68_1 | spl68_5), inference(avatar_split_clause, [], [f584, f644, f625])).
fof(f584, plain, (ssItem(sK65) | sP7(sK64, sK63, sK62, sK61)), inference(cnf_transformation, [], [f366])).
fof(f642, plain, (spl68_1 | spl68_4), inference(avatar_split_clause, [], [f585, f639, f625])).
fof(f585, plain, (ssItem(sK66) | sP7(sK64, sK63, sK62, sK61)), inference(cnf_transformation, [], [f366])).
fof(f637, plain, (spl68_1 | spl68_3), inference(avatar_split_clause, [], [f586, f634, f625])).
fof(f586, plain, (ssList(sK67) | sP7(sK64, sK63, sK62, sK61)), inference(cnf_transformation, [], [f366])).
fof(f632, plain, (spl68_1 | spl68_2), inference(avatar_split_clause, [], [f587, f629, f625])).
fof(f587, plain, ((sK62 = app(app(cons(sK65, nil), cons(sK66, nil)), sK67)) | sP7(sK64, sK63, sK62, sK61)), inference(cnf_transformation, [], [f366])).