fof(f417936, plain, $false, inference(avatar_sat_refutation, [], [f532, f537, f542, f551, f556, f561, f570, f575, f580, f585, f590, f675, f680, f685, f690, f691, f692, f693, f694, f695, f696, f697, f698, f699, f700, f701, f702, f703, f704, f705, f706, f707, f708, f709, f1869, f1876, f1879, f1950, f1997, f415394, f415808, f416337, f416779, f417219, f417891])).
fof(f417891, plain, (spl39_2 | ~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16), inference(avatar_contradiction_clause, [], [f417890])).
fof(f417890, plain, ($false | (spl39_2 | ~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f417889, f417665])).
fof(f417665, plain, (~ (s_best7_init = a_select2(s_center7_init, sK34)) | (spl39_2 | ~ spl39_15)), inference(forward_demodulation, [], [f531, f593])).
fof(f593, plain, ((s_best7_init = s_worst7_init) | ~ spl39_15), inference(avatar_component_clause, [], [f592])).
fof(f592, plain, (spl39_15 <=> (s_best7_init = s_worst7_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_15])])).
fof(f531, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK34)) | spl39_2), inference(avatar_component_clause, [], [f529])).
fof(f529, plain, (spl39_2 <=> (s_worst7_init = a_select2(s_center7_init, sK34))), introduced(avatar_definition, [new_symbols(naming, [spl39_2])])).
fof(f417889, plain, ((s_best7_init = a_select2(s_center7_init, sK34)) | (~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f417789, f536])).
fof(f536, plain, (leq(sK34, n2) | ~ spl39_3), inference(avatar_component_clause, [], [f534])).
fof(f534, plain, (spl39_3 <=> leq(sK34, n2)), introduced(avatar_definition, [new_symbols(naming, [spl39_3])])).
fof(f417789, plain, (~ leq(sK34, n2) | (s_best7_init = a_select2(s_center7_init, sK34)) | (~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f541, f730])).
fof(f730, plain, (! [X2] : (~ leq(n0, X2) | ~ leq(X2, n2) | (s_best7_init = a_select2(s_center7_init, X2))) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f726, f728])).
fof(f728, plain, ((s_best7_init = s_sworst7_init) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f597, f593])).
fof(f597, plain, ((s_sworst7_init = s_worst7_init) | ~ spl39_16), inference(avatar_component_clause, [], [f596])).
fof(f596, plain, (spl39_16 <=> (s_sworst7_init = s_worst7_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_16])])).
fof(f726, plain, (! [X2] : ((s_sworst7_init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) | ~ spl39_16), inference(backward_demodulation, [], [f485, f597])).
fof(f485, plain, ! [X2] : ((s_worst7_init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)), inference(definition_unfolding, [], [f393, f376])).
fof(f376, plain, (init = s_worst7_init), inference(cnf_transformation, [], [f251])).
fof(f251, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_try7_init, sK38)) & leq(sK38, minus(n3, n1)) & leq(n0, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) & ! [X4] : (! [X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5)) | ~ leq(X4, n2) | ~ leq(n0, X4)) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38])], [f249, f250])).
fof(f250, plain, (? [X0] : (~ (init = a_select2(s_try7_init, X0)) & leq(X0, minus(n3, n1)) & leq(n0, X0)) => (~ (init = a_select2(s_try7_init, sK38)) & leq(sK38, minus(n3, n1)) & leq(n0, sK38))), introduced(choice_axiom, [])).
fof(f249, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_try7_init, X0)) & leq(X0, minus(n3, n1)) & leq(n0, X0)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) & ! [X4] : (! [X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5)) | ~ leq(X4, n2) | ~ leq(n0, X4)) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(rectify, [], [f188])).
fof(f188, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & leq(X5, minus(n3, n1)) & leq(n0, X5)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | ~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(definition_folding, [], [f165, e187, e186, e185])).
fof(f185, plain, (? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ sP4), inference(usedef, [], [e185])).
fof(e185, plain, (sP4 <=> ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f186, plain, (? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ~ sP5), inference(usedef, [], [e186])).
fof(e186, plain, (sP5 <=> ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f187, plain, (? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ~ sP6), inference(usedef, [], [e187])).
fof(e187, plain, (sP6 <=> ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f165, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & leq(X5, minus(n3, n1)) & leq(n0, X5)) | ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | ~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(flattening, [], [f164])).
fof(f164, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & (leq(X5, minus(n3, n1)) & leq(n0, X5))) | ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & (leq(X6, n2) & leq(n0, X6))) | ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & (leq(X7, n3) & leq(n0, X7))) | ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & (leq(X9, n3) & leq(n0, X9))) & (leq(X8, n2) & leq(n0, X8))) | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & ((((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | (~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0))) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | (~ leq(X1, n2) | ~ leq(n0, X1))) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | (~ leq(X2, n3) | ~ leq(n0, X2))) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | (~ leq(X4, n3) | ~ leq(n0, X4))) | (~ leq(X3, n2) | ~ leq(n0, X3))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(ennf_transformation, [], [f122])).
fof(f122, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X0] : ((leq(X0, minus(n3, n1)) & leq(n0, X0)) => (init = a_select2(s_try7_init, X0))) & ! [X1] : ((leq(X1, n2) & leq(n0, X1)) => (init = a_select2(s_center7_init, X1))) & ! [X2] : ((leq(X2, n3) & leq(n0, X2)) => (init = a_select2(s_values7_init, X2))) & ! [X3] : ((leq(X3, n2) & leq(n0, X3)) => ! [X4] : ((leq(X4, n3) & leq(n0, X4)) => (init = a_select3(simplex7_init, X4, X3)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X5] : ((leq(X5, minus(n3, n1)) & leq(n0, X5)) => (init = a_select2(s_try7_init, X5))) & ! [X6] : ((leq(X6, n2) & leq(n0, X6)) => (init = a_select2(s_center7_init, X6))) & ! [X7] : ((leq(X7, n3) & leq(n0, X7)) => (init = a_select2(s_values7_init, X7))) & ! [X8] : ((leq(X8, n2) & leq(n0, X8)) => ! [X9] : ((leq(X9, n3) & leq(n0, X9)) => (init = a_select3(simplex7_init, X9, X8)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = a_select2(s_try7_init, n2)) & (init = a_select2(s_try7_init, n1)) & (init = a_select2(s_try7_init, n0)) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(rectify, [], [f54])).
fof(f54, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X20] : ((leq(X20, minus(n3, n1)) & leq(n0, X20)) => (init = a_select2(s_try7_init, X20))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X10] : ((leq(X10, minus(n3, n1)) & leq(n0, X10)) => (init = a_select2(s_try7_init, X10))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => (init = a_select2(s_center7_init, X4))) & ! [X28] : ((leq(X28, n3) & leq(n0, X28)) => (init = a_select2(s_values7_init, X28))) & ! [X21] : ((leq(X21, n2) & leq(n0, X21)) => ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select3(simplex7_init, X27, X21)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = a_select2(s_try7_init, n2)) & (init = a_select2(s_try7_init, n1)) & (init = a_select2(s_try7_init, n0)) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(negated_conjecture, [], [f53])).
fof(f53, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X20] : ((leq(X20, minus(n3, n1)) & leq(n0, X20)) => (init = a_select2(s_try7_init, X20))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv8, minus(n330, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv8) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X10] : ((leq(X10, minus(n3, n1)) & leq(n0, X10)) => (init = a_select2(s_try7_init, X10))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => (init = a_select2(s_center7_init, X4))) & ! [X28] : ((leq(X28, n3) & leq(n0, X28)) => (init = a_select2(s_values7_init, X28))) & ! [X21] : ((leq(X21, n2) & leq(n0, X21)) => ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select3(simplex7_init, X27, X21)))) & leq(pv20, minus(n330, n1)) & leq(pv19, minus(n410, n1)) & leq(pv7, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv20) & leq(n0, pv19) & leq(n0, pv7) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = a_select2(s_try7_init, n2)) & (init = a_select2(s_try7_init, n1)) & (init = a_select2(s_try7_init, n0)) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gauss_init_0009)).
fof(f393, plain, ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)), inference(cnf_transformation, [], [f251])).
fof(f541, plain, (leq(n0, sK34) | ~ spl39_4), inference(avatar_component_clause, [], [f539])).
fof(f539, plain, (spl39_4 <=> leq(n0, sK34)), introduced(avatar_definition, [new_symbols(naming, [spl39_4])])).
fof(f417219, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_911), inference(avatar_contradiction_clause, [], [f417218])).
fof(f417218, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_911)), inference(subsumption_resolution, [], [f415804, f417217])).
fof(f417217, plain, (~ (s_best7_init = a_select3(simplex7_init, n1, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_911)), inference(forward_demodulation, [], [f1894, f414814])).
fof(f414814, plain, ((n1 = sK37) | ~ spl39_911), inference(avatar_component_clause, [], [f414812])).
fof(f414812, plain, (spl39_911 <=> (n1 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_911])])).
fof(f1894, plain, (~ (s_best7_init = a_select3(simplex7_init, sK37, sK36)) | (spl39_10 | ~ spl39_15)), inference(forward_demodulation, [], [f569, f593])).
fof(f569, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK37, sK36)) | spl39_10), inference(avatar_component_clause, [], [f567])).
fof(f567, plain, (spl39_10 <=> (s_worst7_init = a_select3(simplex7_init, sK37, sK36))), introduced(avatar_definition, [new_symbols(naming, [spl39_10])])).
fof(f415804, plain, ((s_best7_init = a_select3(simplex7_init, n1, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f415701, f584])).
fof(f584, plain, (leq(sK36, n2) | ~ spl39_13), inference(avatar_component_clause, [], [f582])).
fof(f582, plain, (spl39_13 <=> leq(sK36, n2)), introduced(avatar_definition, [new_symbols(naming, [spl39_13])])).
fof(f415701, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n1, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f589, f2284])).
fof(f2284, plain, (! [X3] : (~ leq(n0, X3) | ~ leq(X3, n2) | (s_best7_init = a_select3(simplex7_init, n1, X3))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2259, f755])).
fof(f755, plain, leq(n1, n3), inference(resolution, [], [f257, f430])).
fof(f430, plain, gt(n3, n1), inference(cnf_transformation, [], [f81])).
fof(f81, plain, gt(n3, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_3_1)).
fof(f257, plain, ! [X0, X1] : (~ gt(X1, X0) | leq(X0, X1)), inference(cnf_transformation, [], [f128])).
fof(f128, plain, ! [X0, X1] : (leq(X0, X1) | ~ gt(X1, X0)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (gt(X1, X0) => leq(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', leq_gt1)).
fof(f2259, plain, (! [X3] : (~ leq(n1, n3) | (s_best7_init = a_select3(simplex7_init, n1, X3)) | ~ leq(X3, n2) | ~ leq(n0, X3)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f732, f739])).
fof(f739, plain, leq(n0, n1), inference(resolution, [], [f257, f422])).
fof(f422, plain, gt(n1, n0), inference(cnf_transformation, [], [f73])).
fof(f73, plain, gt(n1, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_1_0)).
fof(f732, plain, (! [X4, X5] : (~ leq(n0, X5) | ~ leq(X5, n3) | (s_best7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X4, n2) | ~ leq(n0, X4)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f724, f728])).
fof(f724, plain, (! [X4, X5] : ((s_sworst7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)) | ~ spl39_16), inference(backward_demodulation, [], [f487, f597])).
fof(f487, plain, ! [X4, X5] : ((s_worst7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)), inference(definition_unfolding, [], [f391, f376])).
fof(f391, plain, ! [X4, X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)), inference(cnf_transformation, [], [f251])).
fof(f589, plain, (leq(n0, sK36) | ~ spl39_14), inference(avatar_component_clause, [], [f587])).
fof(f587, plain, (spl39_14 <=> leq(n0, sK36)), introduced(avatar_definition, [new_symbols(naming, [spl39_14])])).
fof(f416779, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_910), inference(avatar_contradiction_clause, [], [f416778])).
fof(f416778, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_910)), inference(subsumption_resolution, [], [f415809, f416777])).
fof(f416777, plain, (~ (s_best7_init = a_select3(simplex7_init, n0, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_910)), inference(forward_demodulation, [], [f1894, f414810])).
fof(f414810, plain, ((n0 = sK37) | ~ spl39_910), inference(avatar_component_clause, [], [f414808])).
fof(f414808, plain, (spl39_910 <=> (n0 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_910])])).
fof(f415809, plain, ((s_best7_init = a_select3(simplex7_init, n0, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f415704, f584])).
fof(f415704, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n0, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f589, f2290])).
fof(f2290, plain, (! [X41] : (~ leq(n0, X41) | ~ leq(X41, n2) | (s_best7_init = a_select3(simplex7_init, n0, X41))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2281, f741])).
fof(f741, plain, leq(n0, n3), inference(resolution, [], [f257, f424])).
fof(f424, plain, gt(n3, n0), inference(cnf_transformation, [], [f75])).
fof(f75, plain, gt(n3, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_3_0)).
fof(f2281, plain, (! [X41] : (~ leq(n0, n3) | (s_best7_init = a_select3(simplex7_init, n0, X41)) | ~ leq(X41, n2) | ~ leq(n0, X41)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f732, f255])).
fof(f255, plain, ! [X0] : leq(X0, X0), inference(cnf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : leq(X0, X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', reflexivity_leq)).
fof(f416337, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_912), inference(avatar_contradiction_clause, [], [f416336])).
fof(f416336, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_912)), inference(subsumption_resolution, [], [f416335, f415805])).
fof(f415805, plain, ((s_best7_init = a_select3(simplex7_init, n2, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f415702, f584])).
fof(f415702, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n2, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f589, f2285])).
fof(f2285, plain, (! [X4] : (~ leq(n0, X4) | ~ leq(X4, n2) | (s_best7_init = a_select3(simplex7_init, n2, X4))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2260, f760])).
fof(f760, plain, leq(n2, n3), inference(resolution, [], [f257, f435])).
fof(f435, plain, gt(n3, n2), inference(cnf_transformation, [], [f86])).
fof(f86, plain, gt(n3, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_3_2)).
fof(f2260, plain, (! [X4] : (~ leq(n2, n3) | (s_best7_init = a_select3(simplex7_init, n2, X4)) | ~ leq(X4, n2) | ~ leq(n0, X4)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f732, f740])).
fof(f740, plain, leq(n0, n2), inference(resolution, [], [f257, f423])).
fof(f423, plain, gt(n2, n0), inference(cnf_transformation, [], [f74])).
fof(f74, plain, gt(n2, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_2_0)).
fof(f416335, plain, (~ (s_best7_init = a_select3(simplex7_init, n2, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_912)), inference(forward_demodulation, [], [f1894, f414818])).
fof(f414818, plain, ((n2 = sK37) | ~ spl39_912), inference(avatar_component_clause, [], [f414816])).
fof(f414816, plain, (spl39_912 <=> (n2 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_912])])).
fof(f415808, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_139), inference(avatar_contradiction_clause, [], [f415807])).
fof(f415807, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_139)), inference(subsumption_resolution, [], [f415806, f415409])).
fof(f415409, plain, (~ (s_best7_init = a_select3(simplex7_init, n3, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_139)), inference(backward_demodulation, [], [f1894, f1938])).
fof(f1938, plain, ((n3 = sK37) | ~ spl39_139), inference(avatar_component_clause, [], [f1936])).
fof(f1936, plain, (spl39_139 <=> (n3 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_139])])).
fof(f415806, plain, ((s_best7_init = a_select3(simplex7_init, n3, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f415703, f584])).
fof(f415703, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n3, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f589, f2286])).
fof(f2286, plain, (! [X5] : (~ leq(n0, X5) | ~ leq(X5, n2) | (s_best7_init = a_select3(simplex7_init, n3, X5))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2261, f255])).
fof(f2261, plain, (! [X5] : (~ leq(n3, n3) | (s_best7_init = a_select3(simplex7_init, n3, X5)) | ~ leq(X5, n2) | ~ leq(n0, X5)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f732, f741])).
fof(f415394, plain, (spl39_139 | spl39_910 | spl39_911 | spl39_912 | ~ spl39_11 | ~ spl39_12), inference(avatar_split_clause, [], [f414842, f577, f572, f414816, f414812, f414808, f1936])).
fof(f572, plain, (spl39_11 <=> leq(sK37, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_11])])).
fof(f577, plain, (spl39_12 <=> leq(n0, sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_12])])).
fof(f414842, plain, ((n2 = sK37) | (n1 = sK37) | (n0 = sK37) | (n3 = sK37) | (~ spl39_11 | ~ spl39_12)), inference(subsumption_resolution, [], [f414733, f574])).
fof(f574, plain, (leq(sK37, n3) | ~ spl39_11), inference(avatar_component_clause, [], [f572])).
fof(f414733, plain, ((n2 = sK37) | (n1 = sK37) | (n0 = sK37) | ~ leq(sK37, n3) | (n3 = sK37) | ~ spl39_12), inference(resolution, [], [f579, f445])).
fof(f445, plain, ! [X0] : (~ leq(n0, X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | (n3 = X0)), inference(cnf_transformation, [], [f177])).
fof(f177, plain, ! [X0] : ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | ~ leq(n0, X0)), inference(flattening, [], [f176])).
fof(f176, plain, ! [X0] : (((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0)) | (~ leq(X0, n3) | ~ leq(n0, X0))), inference(ennf_transformation, [], [f96])).
fof(f96, plain, ! [X0] : ((leq(X0, n3) & leq(n0, X0)) => ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', finite_domain_3)).
fof(f579, plain, (leq(n0, sK37) | ~ spl39_12), inference(avatar_component_clause, [], [f577])).
fof(f1997, plain, (spl39_6 | ~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16), inference(avatar_contradiction_clause, [], [f1996])).
fof(f1996, plain, ($false | (spl39_6 | ~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1995, f1882])).
fof(f1882, plain, (~ (s_best7_init = a_select2(s_values7_init, sK35)) | (spl39_6 | ~ spl39_15)), inference(forward_demodulation, [], [f550, f593])).
fof(f550, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK35)) | spl39_6), inference(avatar_component_clause, [], [f548])).
fof(f548, plain, (spl39_6 <=> (s_worst7_init = a_select2(s_values7_init, sK35))), introduced(avatar_definition, [new_symbols(naming, [spl39_6])])).
fof(f1995, plain, ((s_best7_init = a_select2(s_values7_init, sK35)) | (~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1990, f555])).
fof(f555, plain, (leq(sK35, n3) | ~ spl39_7), inference(avatar_component_clause, [], [f553])).
fof(f553, plain, (spl39_7 <=> leq(sK35, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_7])])).
fof(f1990, plain, (~ leq(sK35, n3) | (s_best7_init = a_select2(s_values7_init, sK35)) | (~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f560, f731])).
fof(f731, plain, (! [X3] : (~ leq(n0, X3) | ~ leq(X3, n3) | (s_best7_init = a_select2(s_values7_init, X3))) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f725, f728])).
fof(f725, plain, (! [X3] : ((s_sworst7_init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) | ~ spl39_16), inference(backward_demodulation, [], [f486, f597])).
fof(f486, plain, ! [X3] : ((s_worst7_init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)), inference(definition_unfolding, [], [f392, f376])).
fof(f392, plain, ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)), inference(cnf_transformation, [], [f251])).
fof(f560, plain, (leq(n0, sK35) | ~ spl39_8), inference(avatar_component_clause, [], [f558])).
fof(f558, plain, (spl39_8 <=> leq(n0, sK35)), introduced(avatar_definition, [new_symbols(naming, [spl39_8])])).
fof(f1950, plain, (~ spl39_15 | ~ spl39_16 | spl39_32 | ~ spl39_36 | ~ spl39_37), inference(avatar_contradiction_clause, [], [f1949])).
fof(f1949, plain, ($false | (~ spl39_15 | ~ spl39_16 | spl39_32 | ~ spl39_36 | ~ spl39_37)), inference(subsumption_resolution, [], [f1948, f1940])).
fof(f1940, plain, (leq(sK38, n2) | ~ spl39_36), inference(forward_demodulation, [], [f679, f791])).
fof(f791, plain, (n2 = minus(n3, n1)), inference(superposition, [], [f467, f722])).
fof(f722, plain, (n3 = plus(n2, n1)), inference(forward_demodulation, [], [f721, f717])).
fof(f717, plain, (n2 = plus(n1, n1)), inference(forward_demodulation, [], [f494, f493])).
fof(f493, plain, (n1 = plus(n0, n1)), inference(definition_unfolding, [], [f448, f332])).
fof(f332, plain, ! [X0] : (succ(X0) = plus(X0, n1)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (succ(X0) = plus(X0, n1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', succ_plus_1_r)).
fof(f448, plain, (n1 = succ(n0)), inference(cnf_transformation, [], [f99])).
fof(f99, plain, (n1 = succ(n0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', successor_1)).
fof(f494, plain, (n2 = plus(plus(n0, n1), n1)), inference(definition_unfolding, [], [f449, f332, f332])).
fof(f449, plain, (n2 = succ(succ(n0))), inference(cnf_transformation, [], [f100])).
fof(f100, plain, (n2 = succ(succ(n0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', successor_2)).
fof(f721, plain, (n3 = plus(plus(n1, n1), n1)), inference(forward_demodulation, [], [f720, f493])).
fof(f720, plain, (n3 = plus(plus(n1, plus(n0, n1)), n1)), inference(forward_demodulation, [], [f495, f458])).
fof(f458, plain, ! [X0] : (plus(X0, n1) = plus(n1, X0)), inference(definition_unfolding, [], [f333, f332])).
fof(f333, plain, ! [X0] : (succ(X0) = plus(n1, X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (succ(X0) = plus(n1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', succ_plus_1_l)).
fof(f495, plain, (n3 = plus(plus(plus(n0, n1), n1), n1)), inference(definition_unfolding, [], [f450, f332, f332, f332])).
fof(f450, plain, (n3 = succ(succ(succ(n0)))), inference(cnf_transformation, [], [f101])).
fof(f101, plain, (n3 = succ(succ(succ(n0)))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', successor_3)).
fof(f467, plain, ! [X0] : (minus(plus(X0, n1), n1) = X0), inference(definition_unfolding, [], [f343, f342, f332])).
fof(f342, plain, ! [X0] : (minus(X0, n1) = pred(X0)), inference(cnf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (minus(X0, n1) = pred(X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', pred_minus_1)).
fof(f343, plain, ! [X0] : (pred(succ(X0)) = X0), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (pred(succ(X0)) = X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', pred_succ)).
fof(f679, plain, (leq(sK38, minus(n3, n1)) | ~ spl39_36), inference(avatar_component_clause, [], [f677])).
fof(f677, plain, (spl39_36 <=> leq(sK38, minus(n3, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_36])])).
fof(f1948, plain, (~ leq(sK38, n2) | (~ spl39_15 | ~ spl39_16 | spl39_32 | ~ spl39_37)), inference(subsumption_resolution, [], [f1942, f1941])).
fof(f1941, plain, (~ (s_best7_init = a_select2(s_try7_init, sK38)) | (~ spl39_15 | spl39_32)), inference(forward_demodulation, [], [f662, f593])).
fof(f662, plain, (~ (s_worst7_init = a_select2(s_try7_init, sK38)) | spl39_32), inference(avatar_component_clause, [], [f660])).
fof(f660, plain, (spl39_32 <=> (s_worst7_init = a_select2(s_try7_init, sK38))), introduced(avatar_definition, [new_symbols(naming, [spl39_32])])).
fof(f1942, plain, ((s_best7_init = a_select2(s_try7_init, sK38)) | ~ leq(sK38, n2) | (~ spl39_15 | ~ spl39_16 | ~ spl39_37)), inference(resolution, [], [f684, f794])).
fof(f794, plain, (! [X1] : (~ leq(n0, X1) | (s_best7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, n2)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f729, f791])).
fof(f729, plain, (! [X1] : ((s_best7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f727, f728])).
fof(f727, plain, (! [X1] : ((s_sworst7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) | ~ spl39_16), inference(backward_demodulation, [], [f484, f597])).
fof(f484, plain, ! [X1] : ((s_worst7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)), inference(definition_unfolding, [], [f394, f376])).
fof(f394, plain, ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)), inference(cnf_transformation, [], [f251])).
fof(f684, plain, (leq(n0, sK38) | ~ spl39_37), inference(avatar_component_clause, [], [f682])).
fof(f682, plain, (spl39_37 <=> leq(n0, sK38)), introduced(avatar_definition, [new_symbols(naming, [spl39_37])])).
fof(f1879, plain, (~ spl39_15 | ~ spl39_16 | spl39_19), inference(avatar_contradiction_clause, [], [f1878])).
fof(f1878, plain, ($false | (~ spl39_15 | ~ spl39_16 | spl39_19)), inference(subsumption_resolution, [], [f1877, f593])).
fof(f1877, plain, (~ (s_best7_init = s_worst7_init) | (~ spl39_15 | ~ spl39_16 | spl39_19)), inference(forward_demodulation, [], [f610, f1870])).
fof(f1870, plain, ((s_best7_init = a_select2(s_try7_init, n2)) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1843, f255])).
fof(f1843, plain, ((s_best7_init = a_select2(s_try7_init, n2)) | ~ leq(n2, n2) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f794, f740])).
fof(f610, plain, (~ (s_worst7_init = a_select2(s_try7_init, n2)) | spl39_19), inference(avatar_component_clause, [], [f608])).
fof(f608, plain, (spl39_19 <=> (s_worst7_init = a_select2(s_try7_init, n2))), introduced(avatar_definition, [new_symbols(naming, [spl39_19])])).
fof(f1876, plain, (~ spl39_15 | ~ spl39_16 | spl39_17), inference(avatar_contradiction_clause, [], [f1875])).
fof(f1875, plain, ($false | (~ spl39_15 | ~ spl39_16 | spl39_17)), inference(subsumption_resolution, [], [f1874, f593])).
fof(f1874, plain, (~ (s_best7_init = s_worst7_init) | (~ spl39_15 | ~ spl39_16 | spl39_17)), inference(forward_demodulation, [], [f602, f1871])).
fof(f1871, plain, ((s_best7_init = a_select2(s_try7_init, n0)) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1864, f740])).
fof(f1864, plain, ((s_best7_init = a_select2(s_try7_init, n0)) | ~ leq(n0, n2) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f794, f255])).
fof(f602, plain, (~ (s_worst7_init = a_select2(s_try7_init, n0)) | spl39_17), inference(avatar_component_clause, [], [f600])).
fof(f600, plain, (spl39_17 <=> (s_worst7_init = a_select2(s_try7_init, n0))), introduced(avatar_definition, [new_symbols(naming, [spl39_17])])).
fof(f1869, plain, (~ spl39_15 | ~ spl39_16 | spl39_18), inference(avatar_contradiction_clause, [], [f1868])).
fof(f1868, plain, ($false | (~ spl39_15 | ~ spl39_16 | spl39_18)), inference(subsumption_resolution, [], [f1867, f754])).
fof(f754, plain, leq(n1, n2), inference(resolution, [], [f257, f429])).
fof(f429, plain, gt(n2, n1), inference(cnf_transformation, [], [f80])).
fof(f80, plain, gt(n2, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV024+1.p', gt_2_1)).
fof(f1867, plain, (~ leq(n1, n2) | (~ spl39_15 | ~ spl39_16 | spl39_18)), inference(subsumption_resolution, [], [f1842, f733])).
fof(f733, plain, (~ (s_best7_init = a_select2(s_try7_init, n1)) | (~ spl39_15 | ~ spl39_16 | spl39_18)), inference(backward_demodulation, [], [f723, f728])).
fof(f723, plain, (~ (s_sworst7_init = a_select2(s_try7_init, n1)) | (~ spl39_16 | spl39_18)), inference(backward_demodulation, [], [f606, f597])).
fof(f606, plain, (~ (s_worst7_init = a_select2(s_try7_init, n1)) | spl39_18), inference(avatar_component_clause, [], [f604])).
fof(f604, plain, (spl39_18 <=> (s_worst7_init = a_select2(s_try7_init, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_18])])).
fof(f1842, plain, ((s_best7_init = a_select2(s_try7_init, n1)) | ~ leq(n1, n2) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f794, f739])).
fof(f709, plain, spl39_15, inference(avatar_split_clause, [], [f489, f592])).
fof(f489, plain, (s_best7_init = s_worst7_init), inference(definition_unfolding, [], [f374, f376])).
fof(f374, plain, (init = s_best7_init), inference(cnf_transformation, [], [f251])).
fof(f708, plain, spl39_16, inference(avatar_split_clause, [], [f488, f596])).
fof(f488, plain, (s_sworst7_init = s_worst7_init), inference(definition_unfolding, [], [f375, f376])).
fof(f375, plain, (init = s_sworst7_init), inference(cnf_transformation, [], [f251])).
fof(f707, plain, spl39_20, inference(avatar_split_clause, [], [f377, f612])).
fof(f612, plain, (spl39_20 <=> leq(n0, s_best7)), introduced(avatar_definition, [new_symbols(naming, [spl39_20])])).
fof(f377, plain, leq(n0, s_best7), inference(cnf_transformation, [], [f251])).
fof(f706, plain, spl39_21, inference(avatar_split_clause, [], [f378, f616])).
fof(f616, plain, (spl39_21 <=> leq(n0, s_sworst7)), introduced(avatar_definition, [new_symbols(naming, [spl39_21])])).
fof(f378, plain, leq(n0, s_sworst7), inference(cnf_transformation, [], [f251])).
fof(f705, plain, spl39_22, inference(avatar_split_clause, [], [f379, f620])).
fof(f620, plain, (spl39_22 <=> leq(n0, s_worst7)), introduced(avatar_definition, [new_symbols(naming, [spl39_22])])).
fof(f379, plain, leq(n0, s_worst7), inference(cnf_transformation, [], [f251])).
fof(f704, plain, spl39_23, inference(avatar_split_clause, [], [f380, f624])).
fof(f624, plain, (spl39_23 <=> leq(n0, pv7)), introduced(avatar_definition, [new_symbols(naming, [spl39_23])])).
fof(f380, plain, leq(n0, pv7), inference(cnf_transformation, [], [f251])).
fof(f703, plain, spl39_24, inference(avatar_split_clause, [], [f382, f628])).
fof(f628, plain, (spl39_24 <=> leq(n0, pv19)), introduced(avatar_definition, [new_symbols(naming, [spl39_24])])).
fof(f382, plain, leq(n0, pv19), inference(cnf_transformation, [], [f251])).
fof(f702, plain, spl39_25, inference(avatar_split_clause, [], [f383, f632])).
fof(f632, plain, (spl39_25 <=> leq(n0, pv20)), introduced(avatar_definition, [new_symbols(naming, [spl39_25])])).
fof(f383, plain, leq(n0, pv20), inference(cnf_transformation, [], [f251])).
fof(f701, plain, spl39_26, inference(avatar_split_clause, [], [f384, f636])).
fof(f636, plain, (spl39_26 <=> leq(s_best7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_26])])).
fof(f384, plain, leq(s_best7, n3), inference(cnf_transformation, [], [f251])).
fof(f700, plain, spl39_27, inference(avatar_split_clause, [], [f385, f640])).
fof(f640, plain, (spl39_27 <=> leq(s_sworst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_27])])).
fof(f385, plain, leq(s_sworst7, n3), inference(cnf_transformation, [], [f251])).
fof(f699, plain, spl39_28, inference(avatar_split_clause, [], [f386, f644])).
fof(f644, plain, (spl39_28 <=> leq(s_worst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_28])])).
fof(f386, plain, leq(s_worst7, n3), inference(cnf_transformation, [], [f251])).
fof(f698, plain, spl39_29, inference(avatar_split_clause, [], [f387, f648])).
fof(f648, plain, (spl39_29 <=> leq(pv7, minus(n410, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_29])])).
fof(f387, plain, leq(pv7, minus(n410, n1)), inference(cnf_transformation, [], [f251])).
fof(f697, plain, spl39_30, inference(avatar_split_clause, [], [f389, f652])).
fof(f652, plain, (spl39_30 <=> leq(pv19, minus(n410, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_30])])).
fof(f389, plain, leq(pv19, minus(n410, n1)), inference(cnf_transformation, [], [f251])).
fof(f696, plain, spl39_31, inference(avatar_split_clause, [], [f390, f656])).
fof(f656, plain, (spl39_31 <=> leq(pv20, minus(n330, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_31])])).
fof(f390, plain, leq(pv20, minus(n330, n1)), inference(cnf_transformation, [], [f251])).
fof(f695, plain, (~ spl39_38 | spl39_33), inference(avatar_split_clause, [], [f483, f664, f687])).
fof(f687, plain, (spl39_38 <=> gt(loopcounter, n1)), introduced(avatar_definition, [new_symbols(naming, [spl39_38])])).
fof(f664, plain, (spl39_33 <=> (s_worst7_init = pvar1400_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_33])])).
fof(f483, plain, ((s_worst7_init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f395, f376])).
fof(f395, plain, ((init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f251])).
fof(f694, plain, (~ spl39_38 | spl39_34), inference(avatar_split_clause, [], [f482, f668, f687])).
fof(f668, plain, (spl39_34 <=> (s_worst7_init = pvar1401_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_34])])).
fof(f482, plain, ((s_worst7_init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f396, f376])).
fof(f396, plain, ((init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f251])).
fof(f693, plain, (~ spl39_38 | spl39_35), inference(avatar_split_clause, [], [f481, f672, f687])).
fof(f672, plain, (spl39_35 <=> (s_worst7_init = pvar1402_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_35])])).
fof(f481, plain, ((s_worst7_init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f397, f376])).
fof(f397, plain, ((init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f251])).
fof(f692, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | spl39_37 | spl39_38), inference(avatar_split_clause, [], [f500, f687, f682, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f563, plain, (spl39_9 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl39_9])])).
fof(f544, plain, (spl39_5 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl39_5])])).
fof(f525, plain, (spl39_1 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl39_1])])).
fof(f500, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f499])).
fof(f499, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f480])).
fof(f480, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f398, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f398, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f691, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | spl39_36 | spl39_38), inference(avatar_split_clause, [], [f502, f687, f677, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f502, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f501])).
fof(f501, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f479])).
fof(f479, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f399, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f399, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f690, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | ~ spl39_32 | spl39_38), inference(avatar_split_clause, [], [f504, f687, f660, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f504, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f503])).
fof(f503, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f478])).
fof(f478, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f400, f376, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f400, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f685, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | spl39_37 | ~ spl39_33 | ~ spl39_34 | ~ spl39_35), inference(avatar_split_clause, [], [f506, f672, f668, f664, f682, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f506, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f505])).
fof(f505, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f477])).
fof(f477, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f401, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f401, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f680, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | spl39_36 | ~ spl39_33 | ~ spl39_34 | ~ spl39_35), inference(avatar_split_clause, [], [f508, f672, f668, f664, f677, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f508, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f507])).
fof(f507, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f476])).
fof(f476, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f402, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f402, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f675, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28 | ~ spl39_29 | ~ spl39_30 | ~ spl39_31 | spl39_9 | spl39_5 | spl39_1 | ~ spl39_32 | ~ spl39_33 | ~ spl39_34 | ~ spl39_35), inference(avatar_split_clause, [], [f510, f672, f668, f664, f660, f525, f544, f563, f656, f652, f648, f644, f640, f636, f632, f628, f624, f620, f616, f612, f608, f604, f600, f596, f592])).
fof(f510, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f509])).
fof(f509, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f475])).
fof(f475, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = a_select2(s_try7_init, n2)) | ~ (s_worst7_init = a_select2(s_try7_init, n1)) | ~ (s_worst7_init = a_select2(s_try7_init, n0)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f403, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376, f376])).
fof(f403, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv20, minus(n330, n1)) | ~ leq(pv19, minus(n410, n1)) | ~ leq(pv7, minus(n410, n1)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv20) | ~ leq(n0, pv19) | ~ leq(n0, pv7) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = a_select2(s_try7_init, n2)) | ~ (init = a_select2(s_try7_init, n1)) | ~ (init = a_select2(s_try7_init, n0)) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f251])).
fof(f590, plain, (~ spl39_9 | spl39_14), inference(avatar_split_clause, [], [f368, f587, f563])).
fof(f368, plain, (leq(n0, sK36) | ~ sP4), inference(cnf_transformation, [], [f248])).
fof(f248, plain, (((~ (init = a_select3(simplex7_init, sK37, sK36)) & leq(sK37, n3) & leq(n0, sK37)) & leq(sK36, n2) & leq(n0, sK36)) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f245, f247, f246])).
fof(f246, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK36)) & leq(X1, n3) & leq(n0, X1)) & leq(sK36, n2) & leq(n0, sK36))), introduced(choice_axiom, [])).
fof(f247, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK36)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK37, sK36)) & leq(sK37, n3) & leq(n0, sK37))), introduced(choice_axiom, [])).
fof(f245, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP4), inference(rectify, [], [f244])).
fof(f244, plain, (? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ sP4), inference(nnf_transformation, [], [f185])).
fof(f585, plain, (~ spl39_9 | spl39_13), inference(avatar_split_clause, [], [f369, f582, f563])).
fof(f369, plain, (leq(sK36, n2) | ~ sP4), inference(cnf_transformation, [], [f248])).
fof(f580, plain, (~ spl39_9 | spl39_12), inference(avatar_split_clause, [], [f370, f577, f563])).
fof(f370, plain, (leq(n0, sK37) | ~ sP4), inference(cnf_transformation, [], [f248])).
fof(f575, plain, (~ spl39_9 | spl39_11), inference(avatar_split_clause, [], [f371, f572, f563])).
fof(f371, plain, (leq(sK37, n3) | ~ sP4), inference(cnf_transformation, [], [f248])).
fof(f570, plain, (~ spl39_9 | ~ spl39_10), inference(avatar_split_clause, [], [f474, f567, f563])).
fof(f474, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK37, sK36)) | ~ sP4), inference(definition_unfolding, [], [f372, f376])).
fof(f372, plain, (~ (init = a_select3(simplex7_init, sK37, sK36)) | ~ sP4), inference(cnf_transformation, [], [f248])).
fof(f561, plain, (~ spl39_5 | spl39_8), inference(avatar_split_clause, [], [f365, f558, f544])).
fof(f365, plain, (leq(n0, sK35) | ~ sP5), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((~ (init = a_select2(s_values7_init, sK35)) & leq(sK35, n3) & leq(n0, sK35)) | ~ sP5), inference(skolemisation, [status(esa), new_symbols(skolem, [sK35])], [f241, f242])).
fof(f242, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK35)) & leq(sK35, n3) & leq(n0, sK35))), introduced(choice_axiom, [])).
fof(f241, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP5), inference(rectify, [], [f240])).
fof(f240, plain, (? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ~ sP5), inference(nnf_transformation, [], [f186])).
fof(f556, plain, (~ spl39_5 | spl39_7), inference(avatar_split_clause, [], [f366, f553, f544])).
fof(f366, plain, (leq(sK35, n3) | ~ sP5), inference(cnf_transformation, [], [f243])).
fof(f551, plain, (~ spl39_5 | ~ spl39_6), inference(avatar_split_clause, [], [f473, f548, f544])).
fof(f473, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK35)) | ~ sP5), inference(definition_unfolding, [], [f367, f376])).
fof(f367, plain, (~ (init = a_select2(s_values7_init, sK35)) | ~ sP5), inference(cnf_transformation, [], [f243])).
fof(f542, plain, (~ spl39_1 | spl39_4), inference(avatar_split_clause, [], [f362, f539, f525])).
fof(f362, plain, (leq(n0, sK34) | ~ sP6), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ((~ (init = a_select2(s_center7_init, sK34)) & leq(sK34, n2) & leq(n0, sK34)) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34])], [f237, f238])).
fof(f238, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK34)) & leq(sK34, n2) & leq(n0, sK34))), introduced(choice_axiom, [])).
fof(f237, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | ~ sP6), inference(rectify, [], [f236])).
fof(f236, plain, (? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ~ sP6), inference(nnf_transformation, [], [f187])).
fof(f537, plain, (~ spl39_1 | spl39_3), inference(avatar_split_clause, [], [f363, f534, f525])).
fof(f363, plain, (leq(sK34, n2) | ~ sP6), inference(cnf_transformation, [], [f239])).
fof(f532, plain, (~ spl39_1 | ~ spl39_2), inference(avatar_split_clause, [], [f472, f529, f525])).
fof(f472, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK34)) | ~ sP6), inference(definition_unfolding, [], [f364, f376])).
fof(f364, plain, (~ (init = a_select2(s_center7_init, sK34)) | ~ sP6), inference(cnf_transformation, [], [f239])).