fof(f248840, plain, $false, inference(avatar_sat_refutation, [], [f703, f708, f713, f718, f719, f720, f739, f744, f749, f758, f763, f768, f773, f778, f803, f808, f813, f814, f815, f816, f827, f832, f837, f842, f847, f852, f857, f862, f883, f888, f893, f894, f895, f896, f902, f907, f912, f917, f922, f927, f932, f937, f942, f947, f952, f957, f962, f967, f972, f977, f986, f991, f996, f997, f998, f999, f1001, f1002, f1003, f1004, f1005, f1006, f1007, f1008, f1009, f1010, f1011, f8376, f8475, f9024, f9604, f10232, f235913, f235978, f236118, f236255, f236357, f236613, f236848, f237061, f237542, f238731, f238942, f239368, f239378, f239405, f239419, f239666, f240484, f240730, f240962, f241457, f242090, f242139, f242579, f242792, f243191, f243393, f243611, f243874, f244121, f245265, f245717, f246166, f246798, f247102, f247566, f247813, f248045, f248198, f248342, f248472, f248507, f248640, f248767, f248836])).
fof(f248836, plain, (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_185), inference(avatar_contradiction_clause, [], [f248835])).
fof(f248835, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_185)), inference(subsumption_resolution, [], [f248834, f27611])).
fof(f27611, plain, ((s_sworst7_init = a_select3(simplex7_init, n0, n1)) | ~ spl58_32), inference(subsumption_resolution, [], [f27568, f1039])).
fof(f1039, plain, leq(n1, n2), inference(resolution, [], [f285, f483])).
fof(f483, plain, gt(n2, n1), inference(cnf_transformation, [], [f69])).
fof(f69, plain, gt(n2, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_2_1)).
fof(f285, plain, ! [X0, X1] : (~ gt(X1, X0) | leq(X0, X1)), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ! [X0, X1] : (leq(X0, X1) | ~ gt(X1, X0)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (gt(X1, X0) => leq(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', leq_gt1)).
fof(f27568, plain, (~ leq(n1, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, n1)) | ~ spl58_32), inference(resolution, [], [f2090, f1035])).
fof(f1035, plain, leq(n0, n1), inference(resolution, [], [f285, f478])).
fof(f478, plain, gt(n1, n0), inference(cnf_transformation, [], [f64])).
fof(f64, plain, gt(n1, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_1_0)).
fof(f2090, plain, (! [X35] : (~ leq(n0, X35) | ~ leq(X35, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, X35))) | ~ spl58_32), inference(subsumption_resolution, [], [f2080, f1040])).
fof(f1040, plain, leq(n0, n3), inference(resolution, [], [f285, f480])).
fof(f480, plain, gt(n3, n0), inference(cnf_transformation, [], [f66])).
fof(f66, plain, gt(n3, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_3_0)).
fof(f2080, plain, (! [X35] : (~ leq(n0, n3) | (s_sworst7_init = a_select3(simplex7_init, n0, X35)) | ~ leq(X35, n2) | ~ leq(n0, X35)) | ~ spl58_32), inference(resolution, [], [f1025, f283])).
fof(f283, plain, ! [X0] : leq(X0, X0), inference(cnf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : leq(X0, X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', reflexivity_leq)).
fof(f1025, plain, (! [X4, X3] : (~ leq(n0, X4) | ~ leq(X4, n3) | (s_sworst7_init = a_select3(simplex7_init, X4, X3)) | ~ leq(X3, n2) | ~ leq(n0, X3)) | ~ spl58_32), inference(backward_demodulation, [], [f563, f781])).
fof(f781, plain, ((s_sworst7_init = s_worst7_init) | ~ spl58_32), inference(avatar_component_clause, [], [f780])).
fof(f780, plain, (spl58_32 <=> (s_sworst7_init = s_worst7_init)), introduced(avatar_definition, [new_symbols(naming, [spl58_32])])).
fof(f563, plain, ! [X4, X3] : ((s_worst7_init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4) | ~ leq(X3, n2) | ~ leq(n0, X3)), inference(definition_unfolding, [], [f456, f447])).
fof(f447, plain, (init = s_worst7_init), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_center7_init, sK57)) & leq(sK57, n2) & leq(n0, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init)) & gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK57])], [f277, f278])).
fof(f278, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK57)) & leq(sK57, n2) & leq(n0, sK57))), introduced(choice_axiom, [])).
fof(f277, plain, ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init)) & gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)), inference(rectify, [], [f181])).
fof(f181, plain, ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X4] : (~ (init = a_select2(s_center7_init, X4)) & leq(X4, n2) & leq(n0, X4)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init)) & gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_center7_init, X0)) | ~ leq(X0, n2) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_values7_init, X1)) | ~ leq(X1, n3) | ~ leq(n0, X1)) & ! [X2] : (! [X3] : ((init = a_select3(simplex7_init, X3, X2)) | ~ leq(X3, n3) | ~ leq(n0, X3)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)), inference(definition_folding, [], [f150, e180, e179, e178, e177, e176, e175, e174, e173, e172, e171, e170])).
fof(f170, plain, (? [X18] : (? [X19] : (~ (init = a_select3(simplex7_init, X19, X18)) & leq(X19, n3) & leq(n0, X19)) & leq(X18, n2) & leq(n0, X18)) | ~ sP4), inference(usedef, [], [e170])).
fof(e170, plain, (sP4 <=> ? [X18] : (? [X19] : (~ (init = a_select3(simplex7_init, X19, X18)) & leq(X19, n3) & leq(n0, X19)) & leq(X18, n2) & leq(n0, X18))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f171, plain, (? [X17] : (~ (init = a_select2(s_values7_init, X17)) & leq(X17, n3) & leq(n0, X17)) | ~ sP5), inference(usedef, [], [e171])).
fof(e171, plain, (sP5 <=> ? [X17] : (~ (init = a_select2(s_values7_init, X17)) & leq(X17, n3) & leq(n0, X17))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f172, plain, (? [X14] : (? [X15] : (~ (init = a_select3(simplex7_init, X15, X14)) & leq(X15, n3) & leq(n0, X15)) & leq(X14, n2) & leq(n0, X14)) | ~ sP6), inference(usedef, [], [e172])).
fof(e172, plain, (sP6 <=> ? [X14] : (? [X15] : (~ (init = a_select3(simplex7_init, X15, X14)) & leq(X15, n3) & leq(n0, X15)) & leq(X14, n2) & leq(n0, X14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f173, plain, (? [X13] : (~ (init = a_select2(s_values7_init, X13)) & leq(X13, n3) & leq(n0, X13)) | ~ sP7), inference(usedef, [], [e173])).
fof(e173, plain, (sP7 <=> ? [X13] : (~ (init = a_select2(s_values7_init, X13)) & leq(X13, n3) & leq(n0, X13))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f174, plain, (((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X16] : (~ (init = a_select2(s_center7_init, X16)) & leq(X16, n2) & leq(n0, X16)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ sP8), inference(usedef, [], [e174])).
fof(e174, plain, (sP8 <=> ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X16] : (~ (init = a_select2(s_center7_init, X16)) & leq(X16, n2) & leq(n0, X16)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f175, plain, (? [X10] : (? [X11] : (~ (init = a_select3(simplex7_init, X11, X10)) & leq(X11, n3) & leq(n0, X11)) & leq(X10, n2) & leq(n0, X10)) | ~ sP9), inference(usedef, [], [e175])).
fof(e175, plain, (sP9 <=> ? [X10] : (? [X11] : (~ (init = a_select3(simplex7_init, X11, X10)) & leq(X11, n3) & leq(n0, X11)) & leq(X10, n2) & leq(n0, X10))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f176, plain, (? [X9] : (~ (init = a_select2(s_values7_init, X9)) & leq(X9, n3) & leq(n0, X9)) | ~ sP10), inference(usedef, [], [e176])).
fof(e176, plain, (sP10 <=> ? [X9] : (~ (init = a_select2(s_values7_init, X9)) & leq(X9, n3) & leq(n0, X9))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f177, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X12] : (~ (init = a_select2(s_center7_init, X12)) & leq(X12, n2) & leq(n0, X12)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ sP11), inference(usedef, [], [e177])).
fof(e177, plain, (sP11 <=> ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X12] : (~ (init = a_select2(s_center7_init, X12)) & leq(X12, n2) & leq(n0, X12)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f178, plain, (? [X6] : (? [X7] : (~ (init = a_select3(simplex7_init, X7, X6)) & leq(X7, n3) & leq(n0, X7)) & leq(X6, n2) & leq(n0, X6)) | ~ sP12), inference(usedef, [], [e178])).
fof(e178, plain, (sP12 <=> ? [X6] : (? [X7] : (~ (init = a_select3(simplex7_init, X7, X6)) & leq(X7, n3) & leq(n0, X7)) & leq(X6, n2) & leq(n0, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f179, plain, (? [X5] : (~ (init = a_select2(s_values7_init, X5)) & leq(X5, n3) & leq(n0, X5)) | ~ sP13), inference(usedef, [], [e179])).
fof(e179, plain, (sP13 <=> ? [X5] : (~ (init = a_select2(s_values7_init, X5)) & leq(X5, n3) & leq(n0, X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f180, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X8] : (~ (init = a_select2(s_center7_init, X8)) & leq(X8, n2) & leq(n0, X8)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ sP14), inference(usedef, [], [e180])).
fof(e180, plain, (sP14 <=> ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X8] : (~ (init = a_select2(s_center7_init, X8)) & leq(X8, n2) & leq(n0, X8)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f150, plain, ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X4] : (~ (init = a_select2(s_center7_init, X4)) & leq(X4, n2) & leq(n0, X4)) | ? [X5] : (~ (init = a_select2(s_values7_init, X5)) & leq(X5, n3) & leq(n0, X5)) | ? [X6] : (? [X7] : (~ (init = a_select3(simplex7_init, X7, X6)) & leq(X7, n3) & leq(n0, X7)) & leq(X6, n2) & leq(n0, X6)) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init)) & gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X8] : (~ (init = a_select2(s_center7_init, X8)) & leq(X8, n2) & leq(n0, X8)) | ? [X9] : (~ (init = a_select2(s_values7_init, X9)) & leq(X9, n3) & leq(n0, X9)) | ? [X10] : (? [X11] : (~ (init = a_select3(simplex7_init, X11, X10)) & leq(X11, n3) & leq(n0, X11)) & leq(X10, n2) & leq(n0, X10)) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X12] : (~ (init = a_select2(s_center7_init, X12)) & leq(X12, n2) & leq(n0, X12)) | ? [X13] : (~ (init = a_select2(s_values7_init, X13)) & leq(X13, n3) & leq(n0, X13)) | ? [X14] : (? [X15] : (~ (init = a_select3(simplex7_init, X15, X14)) & leq(X15, n3) & leq(n0, X15)) & leq(X14, n2) & leq(n0, X14)) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X16] : (~ (init = a_select2(s_center7_init, X16)) & leq(X16, n2) & leq(n0, X16)) | ? [X17] : (~ (init = a_select2(s_values7_init, X17)) & leq(X17, n3) & leq(n0, X17)) | ? [X18] : (? [X19] : (~ (init = a_select3(simplex7_init, X19, X18)) & leq(X19, n3) & leq(n0, X19)) & leq(X18, n2) & leq(n0, X18)) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_center7_init, X0)) | ~ leq(X0, n2) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_values7_init, X1)) | ~ leq(X1, n3) | ~ leq(n0, X1)) & ! [X2] : (! [X3] : ((init = a_select3(simplex7_init, X3, X2)) | ~ leq(X3, n3) | ~ leq(n0, X3)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)), inference(flattening, [], [f149])).
fof(f149, plain, ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X4] : (~ (init = a_select2(s_center7_init, X4)) & (leq(X4, n2) & leq(n0, X4))) | ? [X5] : (~ (init = a_select2(s_values7_init, X5)) & (leq(X5, n3) & leq(n0, X5))) | ? [X6] : (? [X7] : (~ (init = a_select3(simplex7_init, X7, X6)) & (leq(X7, n3) & leq(n0, X7))) & (leq(X6, n2) & leq(n0, X6))) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init)) & gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X8] : (~ (init = a_select2(s_center7_init, X8)) & (leq(X8, n2) & leq(n0, X8))) | ? [X9] : (~ (init = a_select2(s_values7_init, X9)) & (leq(X9, n3) & leq(n0, X9))) | ? [X10] : (? [X11] : (~ (init = a_select3(simplex7_init, X11, X10)) & (leq(X11, n3) & leq(n0, X11))) & (leq(X10, n2) & leq(n0, X10))) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X12] : (~ (init = a_select2(s_center7_init, X12)) & (leq(X12, n2) & leq(n0, X12))) | ? [X13] : (~ (init = a_select2(s_values7_init, X13)) & (leq(X13, n3) & leq(n0, X13))) | ? [X14] : (? [X15] : (~ (init = a_select3(simplex7_init, X15, X14)) & (leq(X15, n3) & leq(n0, X15))) & (leq(X14, n2) & leq(n0, X14))) | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X16] : (~ (init = a_select2(s_center7_init, X16)) & (leq(X16, n2) & leq(n0, X16))) | ? [X17] : (~ (init = a_select2(s_values7_init, X17)) & (leq(X17, n3) & leq(n0, X17))) | ? [X18] : (? [X19] : (~ (init = a_select3(simplex7_init, X19, X18)) & (leq(X19, n3) & leq(n0, X19))) & (leq(X18, n2) & leq(n0, X18))) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)) & ((((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_center7_init, X0)) | (~ leq(X0, n2) | ~ leq(n0, X0))) & ! [X1] : ((init = a_select2(s_values7_init, X1)) | (~ leq(X1, n3) | ~ leq(n0, X1))) & ! [X2] : (! [X3] : ((init = a_select3(simplex7_init, X3, X2)) | (~ leq(X3, n3) | ~ leq(n0, X3))) | (~ leq(X2, n2) | ~ leq(n0, X2))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X0] : ((leq(X0, n2) & leq(n0, X0)) => (init = a_select2(s_center7_init, X0))) & ! [X1] : ((leq(X1, n3) & leq(n0, X1)) => (init = a_select2(s_values7_init, X1))) & ! [X2] : ((leq(X2, n2) & leq(n0, X2)) => ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select3(simplex7_init, X3, X2)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)) => ((gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => (init = a_select2(s_center7_init, X4))) & ! [X5] : ((leq(X5, n3) & leq(n0, X5)) => (init = a_select2(s_values7_init, X5))) & ! [X6] : ((leq(X6, n2) & leq(n0, X6)) => ! [X7] : ((leq(X7, n3) & leq(n0, X7)) => (init = a_select3(simplex7_init, X7, X6)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = init))) & (~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X8] : ((leq(X8, n2) & leq(n0, X8)) => (init = a_select2(s_center7_init, X8))) & ! [X9] : ((leq(X9, n3) & leq(n0, X9)) => (init = a_select2(s_values7_init, X9))) & ! [X10] : ((leq(X10, n2) & leq(n0, X10)) => ! [X11] : ((leq(X11, n3) & leq(n0, X11)) => (init = a_select3(simplex7_init, X11, X10)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X12] : ((leq(X12, n2) & leq(n0, X12)) => (init = a_select2(s_center7_init, X12))) & ! [X13] : ((leq(X13, n3) & leq(n0, X13)) => (init = a_select2(s_values7_init, X13))) & ! [X14] : ((leq(X14, n2) & leq(n0, X14)) => ! [X15] : ((leq(X15, n3) & leq(n0, X15)) => (init = a_select3(simplex7_init, X15, X14)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X16] : ((leq(X16, n2) & leq(n0, X16)) => (init = a_select2(s_center7_init, X16))) & ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select2(s_values7_init, X17))) & ! [X18] : ((leq(X18, n2) & leq(n0, X18)) => ! [X19] : ((leq(X19, n3) & leq(n0, X19)) => (init = a_select3(simplex7_init, X19, X18)))) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_sworst7)) & (init = s_sworst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_worst7)) & (init = s_worst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_best7)) & (s_best7_init = init) & (init = init))), inference(rectify, [], [f54])).
fof(f54, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)) => ((gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X34] : ((leq(X34, n2) & leq(n0, X34)) => (init = a_select2(s_center7_init, X34))) & ! [X33] : ((leq(X33, n3) & leq(n0, X33)) => (init = a_select2(s_values7_init, X33))) & ! [X31] : ((leq(X31, n2) & leq(n0, X31)) => ! [X32] : ((leq(X32, n3) & leq(n0, X32)) => (init = a_select3(simplex7_init, X32, X31)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = init))) & (~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X30] : ((leq(X30, n2) & leq(n0, X30)) => (init = a_select2(s_center7_init, X30))) & ! [X29] : ((leq(X29, n3) & leq(n0, X29)) => (init = a_select2(s_values7_init, X29))) & ! [X18] : ((leq(X18, n2) & leq(n0, X18)) => ! [X14] : ((leq(X14, n3) & leq(n0, X14)) => (init = a_select3(simplex7_init, X14, X18)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X5] : ((leq(X5, n2) & leq(n0, X5)) => (init = a_select2(s_center7_init, X5))) & ! [X15] : ((leq(X15, n3) & leq(n0, X15)) => (init = a_select2(s_values7_init, X15))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => ! [X10] : ((leq(X10, n3) & leq(n0, X10)) => (init = a_select3(simplex7_init, X10, X4)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X28] : ((leq(X28, n2) & leq(n0, X28)) => (init = a_select2(s_center7_init, X28))) & ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select2(s_values7_init, X27))) & ! [X20] : ((leq(X20, n2) & leq(n0, X20)) => ! [X21] : ((leq(X21, n3) & leq(n0, X21)) => (init = a_select3(simplex7_init, X21, X20)))) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_sworst7)) & (init = s_sworst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_worst7)) & (init = s_worst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_best7)) & (s_best7_init = init) & (init = init))), inference(negated_conjecture, [], [f53])).
fof(f53, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n2, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init)) => ((gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X34] : ((leq(X34, n2) & leq(n0, X34)) => (init = a_select2(s_center7_init, X34))) & ! [X33] : ((leq(X33, n3) & leq(n0, X33)) => (init = a_select2(s_values7_init, X33))) & ! [X31] : ((leq(X31, n2) & leq(n0, X31)) => ! [X32] : ((leq(X32, n3) & leq(n0, X32)) => (init = a_select3(simplex7_init, X32, X31)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_sworst7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = init))) & (~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X30] : ((leq(X30, n2) & leq(n0, X30)) => (init = a_select2(s_center7_init, X30))) & ! [X29] : ((leq(X29, n3) & leq(n0, X29)) => (init = a_select2(s_values7_init, X29))) & ! [X18] : ((leq(X18, n2) & leq(n0, X18)) => ! [X14] : ((leq(X14, n3) & leq(n0, X14)) => (init = a_select3(simplex7_init, X14, X18)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7)) => ((leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X5] : ((leq(X5, n2) & leq(n0, X5)) => (init = a_select2(s_center7_init, X5))) & ! [X15] : ((leq(X15, n3) & leq(n0, X15)) => (init = a_select2(s_values7_init, X15))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => ! [X10] : ((leq(X10, n3) & leq(n0, X10)) => (init = a_select3(simplex7_init, X10, X4)))) & leq(pv1388, n3) & leq(s_worst7, n3) & leq(s_best7, n3) & leq(n0, pv1388) & leq(n0, s_worst7) & leq(n0, s_best7) & (init = s_worst7_init) & (s_best7_init = init) & (init = init))) & (~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X28] : ((leq(X28, n2) & leq(n0, X28)) => (init = a_select2(s_center7_init, X28))) & ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select2(s_values7_init, X27))) & ! [X20] : ((leq(X20, n2) & leq(n0, X20)) => ! [X21] : ((leq(X21, n3) & leq(n0, X21)) => (init = a_select3(simplex7_init, X21, X20)))) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (s_best7_init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_sworst7)) & (init = s_sworst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_worst7)) & (init = s_worst7_init) & (init = init))) & (init = a_select2(s_values7_init, pv1388)) & (init = a_select2(s_values7_init, s_best7)) & (s_best7_init = init) & (init = init))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gauss_init_0057)).
fof(f456, plain, ! [X4, X3] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4) | ~ leq(X3, n2) | ~ leq(n0, X3)), inference(cnf_transformation, [], [f279])).
fof(f248834, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_185)), inference(backward_demodulation, [], [f248643, f9569])).
fof(f9569, plain, ((n0 = sK45) | ~ spl58_170), inference(avatar_component_clause, [], [f9567])).
fof(f9567, plain, (spl58_170 <=> (n0 = sK45)), introduced(avatar_definition, [new_symbols(naming, [spl58_170])])).
fof(f248643, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_185)), inference(backward_demodulation, [], [f9159, f10206])).
fof(f10206, plain, ((n1 = sK44) | ~ spl58_185), inference(avatar_component_clause, [], [f10204])).
fof(f10204, plain, (spl58_185 <=> (n1 = sK44)), introduced(avatar_definition, [new_symbols(naming, [spl58_185])])).
fof(f9159, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, sK44)) | (spl58_27 | ~ spl58_32)), inference(forward_demodulation, [], [f757, f781])).
fof(f757, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK45, sK44)) | spl58_27), inference(avatar_component_clause, [], [f755])).
fof(f755, plain, (spl58_27 <=> (s_worst7_init = a_select3(simplex7_init, sK45, sK44))), introduced(avatar_definition, [new_symbols(naming, [spl58_27])])).
fof(f248767, plain, (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_185), inference(avatar_contradiction_clause, [], [f248766])).
fof(f248766, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_185)), inference(subsumption_resolution, [], [f248725, f235909])).
fof(f235909, plain, ((s_sworst7_init = a_select3(simplex7_init, n3, n1)) | ~ spl58_32), inference(subsumption_resolution, [], [f235815, f1039])).
fof(f235815, plain, (~ leq(n1, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, n1)) | ~ spl58_32), inference(resolution, [], [f2086, f1035])).
fof(f2086, plain, (! [X5] : (~ leq(n0, X5) | ~ leq(X5, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, X5))) | ~ spl58_32), inference(subsumption_resolution, [], [f2066, f283])).
fof(f2066, plain, (! [X5] : (~ leq(n3, n3) | (s_sworst7_init = a_select3(simplex7_init, n3, X5)) | ~ leq(X5, n2) | ~ leq(n0, X5)) | ~ spl58_32), inference(resolution, [], [f1025, f1040])).
fof(f248725, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_185)), inference(backward_demodulation, [], [f248643, f9505])).
fof(f9505, plain, ((n3 = sK45) | ~ spl58_167), inference(avatar_component_clause, [], [f9503])).
fof(f9503, plain, (spl58_167 <=> (n3 = sK45)), introduced(avatar_definition, [new_symbols(naming, [spl58_167])])).
fof(f248640, plain, (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_184), inference(avatar_contradiction_clause, [], [f248639])).
fof(f248639, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_184)), inference(subsumption_resolution, [], [f248566, f27613])).
fof(f27613, plain, ((s_sworst7_init = a_select3(simplex7_init, n0, n0)) | ~ spl58_32), inference(subsumption_resolution, [], [f27587, f1038])).
fof(f1038, plain, leq(n0, n2), inference(resolution, [], [f285, f479])).
fof(f479, plain, gt(n2, n0), inference(cnf_transformation, [], [f65])).
fof(f65, plain, gt(n2, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_2_0)).
fof(f27587, plain, (~ leq(n0, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, n0)) | ~ spl58_32), inference(resolution, [], [f2090, f283])).
fof(f248566, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, n0)) | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_184)), inference(backward_demodulation, [], [f248508, f9569])).
fof(f248508, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, n0)) | (spl58_27 | ~ spl58_32 | ~ spl58_184)), inference(forward_demodulation, [], [f9159, f10202])).
fof(f10202, plain, ((n0 = sK44) | ~ spl58_184), inference(avatar_component_clause, [], [f10200])).
fof(f10200, plain, (spl58_184 <=> (n0 = sK44)), introduced(avatar_definition, [new_symbols(naming, [spl58_184])])).
fof(f248507, plain, (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_184), inference(avatar_contradiction_clause, [], [f248506])).
fof(f248506, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_184)), inference(subsumption_resolution, [], [f248503, f235699])).
fof(f235699, plain, ((s_sworst7_init = a_select3(simplex7_init, n1, n0)) | ~ spl58_32), inference(subsumption_resolution, [], [f235640, f1038])).
fof(f235640, plain, (~ leq(n0, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, n0)) | ~ spl58_32), inference(resolution, [], [f2084, f283])).
fof(f2084, plain, (! [X3] : (~ leq(n0, X3) | ~ leq(X3, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, X3))) | ~ spl58_32), inference(subsumption_resolution, [], [f2064, f1042])).
fof(f1042, plain, leq(n1, n3), inference(resolution, [], [f285, f484])).
fof(f484, plain, gt(n3, n1), inference(cnf_transformation, [], [f70])).
fof(f70, plain, gt(n3, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_3_1)).
fof(f2064, plain, (! [X3] : (~ leq(n1, n3) | (s_sworst7_init = a_select3(simplex7_init, n1, X3)) | ~ leq(X3, n2) | ~ leq(n0, X3)) | ~ spl58_32), inference(resolution, [], [f1025, f1035])).
fof(f248503, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n0)) | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_184)), inference(backward_demodulation, [], [f248473, f10202])).
fof(f248473, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, sK44)) | (spl58_27 | ~ spl58_32 | ~ spl58_171)), inference(forward_demodulation, [], [f9159, f9573])).
fof(f9573, plain, ((n1 = sK45) | ~ spl58_171), inference(avatar_component_clause, [], [f9571])).
fof(f9571, plain, (spl58_171 <=> (n1 = sK45)), introduced(avatar_definition, [new_symbols(naming, [spl58_171])])).
fof(f248472, plain, (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_185), inference(avatar_contradiction_clause, [], [f248471])).
fof(f248471, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_185)), inference(subsumption_resolution, [], [f248430, f235697])).
fof(f235697, plain, ((s_sworst7_init = a_select3(simplex7_init, n1, n1)) | ~ spl58_32), inference(subsumption_resolution, [], [f235603, f1039])).
fof(f235603, plain, (~ leq(n1, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, n1)) | ~ spl58_32), inference(resolution, [], [f2084, f1035])).
fof(f248430, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_185)), inference(backward_demodulation, [], [f248345, f9573])).
fof(f248345, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_185)), inference(backward_demodulation, [], [f9159, f10206])).
fof(f248342, plain, (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_181), inference(avatar_contradiction_clause, [], [f248341])).
fof(f248341, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_181)), inference(subsumption_resolution, [], [f248262, f27612])).
fof(f27612, plain, ((s_sworst7_init = a_select3(simplex7_init, n0, n2)) | ~ spl58_32), inference(subsumption_resolution, [], [f27575, f283])).
fof(f27575, plain, (~ leq(n2, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, n2)) | ~ spl58_32), inference(resolution, [], [f2090, f1038])).
fof(f248262, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_170 | ~ spl58_181)), inference(backward_demodulation, [], [f248069, f9569])).
fof(f248069, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_181)), inference(backward_demodulation, [], [f9159, f10138])).
fof(f10138, plain, ((n2 = sK44) | ~ spl58_181), inference(avatar_component_clause, [], [f10136])).
fof(f10136, plain, (spl58_181 <=> (n2 = sK44)), introduced(avatar_definition, [new_symbols(naming, [spl58_181])])).
fof(f248198, plain, (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_181), inference(avatar_contradiction_clause, [], [f248197])).
fof(f248197, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_181)), inference(subsumption_resolution, [], [f248154, f235698])).
fof(f235698, plain, ((s_sworst7_init = a_select3(simplex7_init, n1, n2)) | ~ spl58_32), inference(subsumption_resolution, [], [f235620, f283])).
fof(f235620, plain, (~ leq(n2, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, n2)) | ~ spl58_32), inference(resolution, [], [f2084, f1038])).
fof(f248154, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_171 | ~ spl58_181)), inference(backward_demodulation, [], [f248069, f9573])).
fof(f248045, plain, (~ spl58_5 | ~ spl58_8 | ~ spl58_32 | spl58_72), inference(avatar_contradiction_clause, [], [f248044])).
fof(f248044, plain, ($false | (~ spl58_5 | ~ spl58_8 | ~ spl58_32 | spl58_72)), inference(subsumption_resolution, [], [f1831, f248043])).
fof(f248043, plain, (~ (s_sworst7_init = a_select2(s_values7_init, s_best7)) | (~ spl58_32 | spl58_72)), inference(forward_demodulation, [], [f981, f781])).
fof(f981, plain, (~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | spl58_72), inference(avatar_component_clause, [], [f979])).
fof(f979, plain, (spl58_72 <=> (s_worst7_init = a_select2(s_values7_init, s_best7))), introduced(avatar_definition, [new_symbols(naming, [spl58_72])])).
fof(f1831, plain, ((s_sworst7_init = a_select2(s_values7_init, s_best7)) | (~ spl58_5 | ~ spl58_8 | ~ spl58_32)), inference(subsumption_resolution, [], [f1804, f669])).
fof(f669, plain, (leq(s_best7, n3) | ~ spl58_8), inference(avatar_component_clause, [], [f668])).
fof(f668, plain, (spl58_8 <=> leq(s_best7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_8])])).
fof(f1804, plain, (~ leq(s_best7, n3) | (s_sworst7_init = a_select2(s_values7_init, s_best7)) | (~ spl58_5 | ~ spl58_32)), inference(resolution, [], [f1026, f657])).
fof(f657, plain, (leq(n0, s_best7) | ~ spl58_5), inference(avatar_component_clause, [], [f656])).
fof(f656, plain, (spl58_5 <=> leq(n0, s_best7)), introduced(avatar_definition, [new_symbols(naming, [spl58_5])])).
fof(f1026, plain, (! [X2] : (~ leq(n0, X2) | ~ leq(X2, n3) | (s_sworst7_init = a_select2(s_values7_init, X2))) | ~ spl58_32), inference(backward_demodulation, [], [f562, f781])).
fof(f562, plain, ! [X2] : ((s_worst7_init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)), inference(definition_unfolding, [], [f457, f447])).
fof(f457, plain, ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)), inference(cnf_transformation, [], [f279])).
fof(f247813, plain, (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_733), inference(avatar_contradiction_clause, [], [f247812])).
fof(f247812, plain, ($false | (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_733)), inference(subsumption_resolution, [], [f241454, f247811])).
fof(f247811, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, sK52)) | (~ spl58_32 | spl58_59 | ~ spl58_733)), inference(forward_demodulation, [], [f247309, f240460])).
fof(f240460, plain, ((n2 = sK53) | ~ spl58_733), inference(avatar_component_clause, [], [f240458])).
fof(f240458, plain, (spl58_733 <=> (n2 = sK53)), introduced(avatar_definition, [new_symbols(naming, [spl58_733])])).
fof(f247309, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK53, sK52)) | (~ spl58_32 | spl58_59)), inference(forward_demodulation, [], [f916, f781])).
fof(f916, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK53, sK52)) | spl58_59), inference(avatar_component_clause, [], [f914])).
fof(f914, plain, (spl58_59 <=> (s_worst7_init = a_select3(simplex7_init, sK53, sK52))), introduced(avatar_definition, [new_symbols(naming, [spl58_59])])).
fof(f241454, plain, ((s_sworst7_init = a_select3(simplex7_init, n2, sK52)) | (~ spl58_32 | ~ spl58_62 | ~ spl58_63)), inference(subsumption_resolution, [], [f241356, f931])).
fof(f931, plain, (leq(sK52, n2) | ~ spl58_62), inference(avatar_component_clause, [], [f929])).
fof(f929, plain, (spl58_62 <=> leq(sK52, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_62])])).
fof(f241356, plain, (~ leq(sK52, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, sK52)) | (~ spl58_32 | ~ spl58_63)), inference(resolution, [], [f936, f2085])).
fof(f2085, plain, (! [X4] : (~ leq(n0, X4) | ~ leq(X4, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, X4))) | ~ spl58_32), inference(subsumption_resolution, [], [f2065, f1043])).
fof(f1043, plain, leq(n2, n3), inference(resolution, [], [f285, f487])).
fof(f487, plain, gt(n3, n2), inference(cnf_transformation, [], [f73])).
fof(f73, plain, gt(n3, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', gt_3_2)).
fof(f2065, plain, (! [X4] : (~ leq(n2, n3) | (s_sworst7_init = a_select3(simplex7_init, n2, X4)) | ~ leq(X4, n2) | ~ leq(n0, X4)) | ~ spl58_32), inference(resolution, [], [f1025, f1038])).
fof(f936, plain, (leq(n0, sK52) | ~ spl58_63), inference(avatar_component_clause, [], [f934])).
fof(f934, plain, (spl58_63 <=> leq(n0, sK52)), introduced(avatar_definition, [new_symbols(naming, [spl58_63])])).
fof(f247566, plain, (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_731), inference(avatar_contradiction_clause, [], [f247565])).
fof(f247565, plain, ($false | (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_731)), inference(subsumption_resolution, [], [f241458, f247564])).
fof(f247564, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, sK52)) | (~ spl58_32 | spl58_59 | ~ spl58_731)), inference(forward_demodulation, [], [f247309, f240452])).
fof(f240452, plain, ((n0 = sK53) | ~ spl58_731), inference(avatar_component_clause, [], [f240450])).
fof(f240450, plain, (spl58_731 <=> (n0 = sK53)), introduced(avatar_definition, [new_symbols(naming, [spl58_731])])).
fof(f241458, plain, ((s_sworst7_init = a_select3(simplex7_init, n0, sK52)) | (~ spl58_32 | ~ spl58_62 | ~ spl58_63)), inference(subsumption_resolution, [], [f241358, f931])).
fof(f241358, plain, (~ leq(sK52, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, sK52)) | (~ spl58_32 | ~ spl58_63)), inference(resolution, [], [f936, f2090])).
fof(f247102, plain, (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_732), inference(avatar_contradiction_clause, [], [f247101])).
fof(f247101, plain, ($false | (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_732)), inference(subsumption_resolution, [], [f241453, f247100])).
fof(f247100, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, sK52)) | (~ spl58_32 | spl58_59 | ~ spl58_732)), inference(forward_demodulation, [], [f247099, f781])).
fof(f247099, plain, (~ (s_worst7_init = a_select3(simplex7_init, n1, sK52)) | (spl58_59 | ~ spl58_732)), inference(forward_demodulation, [], [f916, f240456])).
fof(f240456, plain, ((n1 = sK53) | ~ spl58_732), inference(avatar_component_clause, [], [f240454])).
fof(f240454, plain, (spl58_732 <=> (n1 = sK53)), introduced(avatar_definition, [new_symbols(naming, [spl58_732])])).
fof(f241453, plain, ((s_sworst7_init = a_select3(simplex7_init, n1, sK52)) | (~ spl58_32 | ~ spl58_62 | ~ spl58_63)), inference(subsumption_resolution, [], [f241355, f931])).
fof(f241355, plain, (~ leq(sK52, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, sK52)) | (~ spl58_32 | ~ spl58_63)), inference(resolution, [], [f936, f2084])).
fof(f246798, plain, (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_856), inference(avatar_contradiction_clause, [], [f246797])).
fof(f246797, plain, ($false | (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_856)), inference(subsumption_resolution, [], [f246796, f245267])).
fof(f245267, plain, ((s_sworst7_init = a_select3(simplex7_init, n0, sK55)) | (~ spl58_32 | ~ spl58_70 | ~ spl58_71)), inference(subsumption_resolution, [], [f245163, f971])).
fof(f971, plain, (leq(sK55, n2) | ~ spl58_70), inference(avatar_component_clause, [], [f969])).
fof(f969, plain, (spl58_70 <=> leq(sK55, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_70])])).
fof(f245163, plain, (~ leq(sK55, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, sK55)) | (~ spl58_32 | ~ spl58_71)), inference(resolution, [], [f976, f2090])).
fof(f976, plain, (leq(n0, sK55) | ~ spl58_71), inference(avatar_component_clause, [], [f974])).
fof(f974, plain, (spl58_71 <=> leq(n0, sK55)), introduced(avatar_definition, [new_symbols(naming, [spl58_71])])).
fof(f246796, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, sK55)) | (~ spl58_32 | spl58_67 | ~ spl58_856)), inference(forward_demodulation, [], [f243937, f244089])).
fof(f244089, plain, ((n0 = sK56) | ~ spl58_856), inference(avatar_component_clause, [], [f244087])).
fof(f244087, plain, (spl58_856 <=> (n0 = sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_856])])).
fof(f243937, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK56, sK55)) | (~ spl58_32 | spl58_67)), inference(forward_demodulation, [], [f956, f781])).
fof(f956, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK56, sK55)) | spl58_67), inference(avatar_component_clause, [], [f954])).
fof(f954, plain, (spl58_67 <=> (s_worst7_init = a_select3(simplex7_init, sK56, sK55))), introduced(avatar_definition, [new_symbols(naming, [spl58_67])])).
fof(f246166, plain, (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_857), inference(avatar_contradiction_clause, [], [f246165])).
fof(f246165, plain, ($false | (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_857)), inference(subsumption_resolution, [], [f246164, f245262])).
fof(f245262, plain, ((s_sworst7_init = a_select3(simplex7_init, n1, sK55)) | (~ spl58_32 | ~ spl58_70 | ~ spl58_71)), inference(subsumption_resolution, [], [f245160, f971])).
fof(f245160, plain, (~ leq(sK55, n2) | (s_sworst7_init = a_select3(simplex7_init, n1, sK55)) | (~ spl58_32 | ~ spl58_71)), inference(resolution, [], [f976, f2084])).
fof(f246164, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, sK55)) | (~ spl58_32 | spl58_67 | ~ spl58_857)), inference(forward_demodulation, [], [f243937, f244093])).
fof(f244093, plain, ((n1 = sK56) | ~ spl58_857), inference(avatar_component_clause, [], [f244091])).
fof(f244091, plain, (spl58_857 <=> (n1 = sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_857])])).
fof(f245717, plain, (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_854), inference(avatar_contradiction_clause, [], [f245716])).
fof(f245716, plain, ($false | (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_854)), inference(subsumption_resolution, [], [f245715, f245266])).
fof(f245266, plain, ((s_sworst7_init = a_select3(simplex7_init, n3, sK55)) | (~ spl58_32 | ~ spl58_70 | ~ spl58_71)), inference(subsumption_resolution, [], [f245162, f971])).
fof(f245162, plain, (~ leq(sK55, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, sK55)) | (~ spl58_32 | ~ spl58_71)), inference(resolution, [], [f976, f2086])).
fof(f245715, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, sK55)) | (~ spl58_32 | spl58_67 | ~ spl58_854)), inference(forward_demodulation, [], [f243937, f243971])).
fof(f243971, plain, ((n3 = sK56) | ~ spl58_854), inference(avatar_component_clause, [], [f243969])).
fof(f243969, plain, (spl58_854 <=> (n3 = sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_854])])).
fof(f245265, plain, (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_858), inference(avatar_contradiction_clause, [], [f245264])).
fof(f245264, plain, ($false | (~ spl58_32 | spl58_67 | ~ spl58_70 | ~ spl58_71 | ~ spl58_858)), inference(subsumption_resolution, [], [f245263, f244886])).
fof(f244886, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, sK55)) | (~ spl58_32 | spl58_67 | ~ spl58_858)), inference(backward_demodulation, [], [f243937, f244097])).
fof(f244097, plain, ((n2 = sK56) | ~ spl58_858), inference(avatar_component_clause, [], [f244095])).
fof(f244095, plain, (spl58_858 <=> (n2 = sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_858])])).
fof(f245263, plain, ((s_sworst7_init = a_select3(simplex7_init, n2, sK55)) | (~ spl58_32 | ~ spl58_70 | ~ spl58_71)), inference(subsumption_resolution, [], [f245161, f971])).
fof(f245161, plain, (~ leq(sK55, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, sK55)) | (~ spl58_32 | ~ spl58_71)), inference(resolution, [], [f976, f2085])).
fof(f244121, plain, (spl58_854 | spl58_856 | spl58_857 | spl58_858 | ~ spl58_68 | ~ spl58_69), inference(avatar_split_clause, [], [f244120, f964, f959, f244095, f244091, f244087, f243969])).
fof(f959, plain, (spl58_68 <=> leq(sK56, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_68])])).
fof(f964, plain, (spl58_69 <=> leq(n0, sK56)), introduced(avatar_definition, [new_symbols(naming, [spl58_69])])).
fof(f244120, plain, ((n2 = sK56) | (n1 = sK56) | (n0 = sK56) | (n3 = sK56) | (~ spl58_68 | ~ spl58_69)), inference(subsumption_resolution, [], [f244018, f961])).
fof(f961, plain, (leq(sK56, n3) | ~ spl58_68), inference(avatar_component_clause, [], [f959])).
fof(f244018, plain, ((n2 = sK56) | (n1 = sK56) | (n0 = sK56) | ~ leq(sK56, n3) | (n3 = sK56) | ~ spl58_69), inference(resolution, [], [f966, f495])).
fof(f495, plain, ! [X0] : (~ leq(n0, X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | (n3 = X0)), inference(cnf_transformation, [], [f162])).
fof(f162, plain, ! [X0] : ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | ~ leq(n0, X0)), inference(flattening, [], [f161])).
fof(f161, plain, ! [X0] : (((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0)) | (~ leq(X0, n3) | ~ leq(n0, X0))), inference(ennf_transformation, [], [f81])).
fof(f81, plain, ! [X0] : ((leq(X0, n3) & leq(n0, X0)) => ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', finite_domain_3)).
fof(f966, plain, (leq(n0, sK56) | ~ spl58_69), inference(avatar_component_clause, [], [f964])).
fof(f243874, plain, (~ spl58_32 | spl58_64 | ~ spl58_65 | ~ spl58_66), inference(avatar_contradiction_clause, [], [f243873])).
fof(f243873, plain, ($false | (~ spl58_32 | spl58_64 | ~ spl58_65 | ~ spl58_66)), inference(subsumption_resolution, [], [f243872, f243600])).
fof(f243600, plain, (~ (s_sworst7_init = a_select2(s_values7_init, sK54)) | (~ spl58_32 | spl58_64)), inference(forward_demodulation, [], [f941, f781])).
fof(f941, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK54)) | spl58_64), inference(avatar_component_clause, [], [f939])).
fof(f939, plain, (spl58_64 <=> (s_worst7_init = a_select2(s_values7_init, sK54))), introduced(avatar_definition, [new_symbols(naming, [spl58_64])])).
fof(f243872, plain, ((s_sworst7_init = a_select2(s_values7_init, sK54)) | (~ spl58_32 | ~ spl58_65 | ~ spl58_66)), inference(subsumption_resolution, [], [f243766, f946])).
fof(f946, plain, (leq(sK54, n3) | ~ spl58_65), inference(avatar_component_clause, [], [f944])).
fof(f944, plain, (spl58_65 <=> leq(sK54, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_65])])).
fof(f243766, plain, (~ leq(sK54, n3) | (s_sworst7_init = a_select2(s_values7_init, sK54)) | (~ spl58_32 | ~ spl58_66)), inference(resolution, [], [f951, f1026])).
fof(f951, plain, (leq(n0, sK54) | ~ spl58_66), inference(avatar_component_clause, [], [f949])).
fof(f949, plain, (spl58_66 <=> leq(n0, sK54)), introduced(avatar_definition, [new_symbols(naming, [spl58_66])])).
fof(f243611, plain, (~ spl58_689 | ~ spl58_32 | spl58_44 | ~ spl58_650), inference(avatar_split_clause, [], [f243608, f237516, f839, f780, f239375])).
fof(f239375, plain, (spl58_689 <=> (s_sworst7_init = a_select3(simplex7_init, n2, sK48))), introduced(avatar_definition, [new_symbols(naming, [spl58_689])])).
fof(f839, plain, (spl58_44 <=> (s_worst7_init = a_select3(simplex7_init, sK49, sK48))), introduced(avatar_definition, [new_symbols(naming, [spl58_44])])).
fof(f237516, plain, (spl58_650 <=> (n2 = sK49)), introduced(avatar_definition, [new_symbols(naming, [spl58_650])])).
fof(f243608, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, sK48)) | (~ spl58_32 | spl58_44 | ~ spl58_650)), inference(forward_demodulation, [], [f237103, f237518])).
fof(f237518, plain, ((n2 = sK49) | ~ spl58_650), inference(avatar_component_clause, [], [f237516])).
fof(f237103, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK49, sK48)) | (~ spl58_32 | spl58_44)), inference(forward_demodulation, [], [f841, f781])).
fof(f841, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK49, sK48)) | spl58_44), inference(avatar_component_clause, [], [f839])).
fof(f243393, plain, (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_669), inference(avatar_contradiction_clause, [], [f243392])).
fof(f243392, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_669)), inference(subsumption_resolution, [], [f243314, f235910])).
fof(f235910, plain, ((s_sworst7_init = a_select3(simplex7_init, n3, n2)) | ~ spl58_32), inference(subsumption_resolution, [], [f235832, f283])).
fof(f235832, plain, (~ leq(n2, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, n2)) | ~ spl58_32), inference(resolution, [], [f2086, f1038])).
fof(f243314, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n2)) | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_669)), inference(backward_demodulation, [], [f242801, f237392])).
fof(f237392, plain, ((n3 = sK49) | ~ spl58_646), inference(avatar_component_clause, [], [f237390])).
fof(f237390, plain, (spl58_646 <=> (n3 = sK49)), introduced(avatar_definition, [new_symbols(naming, [spl58_646])])).
fof(f242801, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK49, n2)) | (~ spl58_32 | spl58_44 | ~ spl58_669)), inference(backward_demodulation, [], [f237103, f237833])).
fof(f237833, plain, ((n2 = sK48) | ~ spl58_669), inference(avatar_component_clause, [], [f237831])).
fof(f237831, plain, (spl58_669 <=> (n2 = sK48)), introduced(avatar_definition, [new_symbols(naming, [spl58_669])])).
fof(f243191, plain, (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_669), inference(avatar_contradiction_clause, [], [f243190])).
fof(f243190, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_669)), inference(subsumption_resolution, [], [f243111, f235698])).
fof(f243111, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n2)) | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_669)), inference(backward_demodulation, [], [f242801, f237514])).
fof(f237514, plain, ((n1 = sK49) | ~ spl58_649), inference(avatar_component_clause, [], [f237512])).
fof(f237512, plain, (spl58_649 <=> (n1 = sK49)), introduced(avatar_definition, [new_symbols(naming, [spl58_649])])).
fof(f242792, plain, (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_672), inference(avatar_contradiction_clause, [], [f242791])).
fof(f242791, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_672)), inference(subsumption_resolution, [], [f242709, f235909])).
fof(f242709, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n1)) | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_672)), inference(backward_demodulation, [], [f242187, f237392])).
fof(f242187, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK49, n1)) | (~ spl58_32 | spl58_44 | ~ spl58_672)), inference(backward_demodulation, [], [f237103, f238219])).
fof(f238219, plain, ((n1 = sK48) | ~ spl58_672), inference(avatar_component_clause, [], [f238217])).
fof(f238217, plain, (spl58_672 <=> (n1 = sK48)), introduced(avatar_definition, [new_symbols(naming, [spl58_672])])).
fof(f242579, plain, (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_672), inference(avatar_contradiction_clause, [], [f242578])).
fof(f242578, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_672)), inference(subsumption_resolution, [], [f242499, f235697])).
fof(f242499, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n1)) | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_672)), inference(backward_demodulation, [], [f242187, f237514])).
fof(f242139, plain, (spl58_2 | ~ spl58_6 | ~ spl58_9 | ~ spl58_32), inference(avatar_contradiction_clause, [], [f242138])).
fof(f242138, plain, ($false | (spl58_2 | ~ spl58_6 | ~ spl58_9 | ~ spl58_32)), inference(subsumption_resolution, [], [f1833, f242137])).
fof(f242137, plain, (~ (s_sworst7_init = a_select2(s_values7_init, s_worst7)) | (spl58_2 | ~ spl58_32)), inference(forward_demodulation, [], [f646, f781])).
fof(f646, plain, (~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | spl58_2), inference(avatar_component_clause, [], [f644])).
fof(f644, plain, (spl58_2 <=> (s_worst7_init = a_select2(s_values7_init, s_worst7))), introduced(avatar_definition, [new_symbols(naming, [spl58_2])])).
fof(f1833, plain, ((s_sworst7_init = a_select2(s_values7_init, s_worst7)) | (~ spl58_6 | ~ spl58_9 | ~ spl58_32)), inference(subsumption_resolution, [], [f1806, f673])).
fof(f673, plain, (leq(s_worst7, n3) | ~ spl58_9), inference(avatar_component_clause, [], [f672])).
fof(f672, plain, (spl58_9 <=> leq(s_worst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_9])])).
fof(f1806, plain, (~ leq(s_worst7, n3) | (s_sworst7_init = a_select2(s_values7_init, s_worst7)) | (~ spl58_6 | ~ spl58_32)), inference(resolution, [], [f1026, f661])).
fof(f661, plain, (leq(n0, s_worst7) | ~ spl58_6), inference(avatar_component_clause, [], [f660])).
fof(f660, plain, (spl58_6 <=> leq(n0, s_worst7)), introduced(avatar_definition, [new_symbols(naming, [spl58_6])])).
fof(f242090, plain, (~ spl58_32 | spl58_37 | ~ spl58_38 | ~ spl58_39), inference(avatar_contradiction_clause, [], [f242089])).
fof(f242089, plain, ($false | (~ spl58_32 | spl58_37 | ~ spl58_38 | ~ spl58_39)), inference(subsumption_resolution, [], [f242088, f241695])).
fof(f241695, plain, (~ (s_sworst7_init = a_select2(s_center7_init, sK46)) | (~ spl58_32 | spl58_37)), inference(forward_demodulation, [], [f802, f781])).
fof(f802, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK46)) | spl58_37), inference(avatar_component_clause, [], [f800])).
fof(f800, plain, (spl58_37 <=> (s_worst7_init = a_select2(s_center7_init, sK46))), introduced(avatar_definition, [new_symbols(naming, [spl58_37])])).
fof(f242088, plain, ((s_sworst7_init = a_select2(s_center7_init, sK46)) | (~ spl58_32 | ~ spl58_38 | ~ spl58_39)), inference(subsumption_resolution, [], [f241984, f807])).
fof(f807, plain, (leq(sK46, n2) | ~ spl58_38), inference(avatar_component_clause, [], [f805])).
fof(f805, plain, (spl58_38 <=> leq(sK46, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_38])])).
fof(f241984, plain, (~ leq(sK46, n2) | (s_sworst7_init = a_select2(s_center7_init, sK46)) | (~ spl58_32 | ~ spl58_39)), inference(resolution, [], [f812, f1027])).
fof(f1027, plain, (! [X1] : (~ leq(n0, X1) | ~ leq(X1, n2) | (s_sworst7_init = a_select2(s_center7_init, X1))) | ~ spl58_32), inference(backward_demodulation, [], [f561, f781])).
fof(f561, plain, ! [X1] : ((s_worst7_init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)), inference(definition_unfolding, [], [f458, f447])).
fof(f458, plain, ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)), inference(cnf_transformation, [], [f279])).
fof(f812, plain, (leq(n0, sK46) | ~ spl58_39), inference(avatar_component_clause, [], [f810])).
fof(f810, plain, (spl58_39 <=> leq(n0, sK46)), introduced(avatar_definition, [new_symbols(naming, [spl58_39])])).
fof(f241457, plain, (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_729), inference(avatar_contradiction_clause, [], [f241456])).
fof(f241456, plain, ($false | (~ spl58_32 | spl58_59 | ~ spl58_62 | ~ spl58_63 | ~ spl58_729)), inference(subsumption_resolution, [], [f241455, f241182])).
fof(f241182, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, sK52)) | (~ spl58_32 | spl58_59 | ~ spl58_729)), inference(backward_demodulation, [], [f239888, f240334])).
fof(f240334, plain, ((n3 = sK53) | ~ spl58_729), inference(avatar_component_clause, [], [f240332])).
fof(f240332, plain, (spl58_729 <=> (n3 = sK53)), introduced(avatar_definition, [new_symbols(naming, [spl58_729])])).
fof(f239888, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK53, sK52)) | (~ spl58_32 | spl58_59)), inference(forward_demodulation, [], [f916, f781])).
fof(f241455, plain, ((s_sworst7_init = a_select3(simplex7_init, n3, sK52)) | (~ spl58_32 | ~ spl58_62 | ~ spl58_63)), inference(subsumption_resolution, [], [f241357, f931])).
fof(f241357, plain, (~ leq(sK52, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, sK52)) | (~ spl58_32 | ~ spl58_63)), inference(resolution, [], [f936, f2086])).
fof(f240962, plain, (~ spl58_32 | spl58_56 | ~ spl58_57 | ~ spl58_58), inference(avatar_contradiction_clause, [], [f240961])).
fof(f240961, plain, ($false | (~ spl58_32 | spl58_56 | ~ spl58_57 | ~ spl58_58)), inference(subsumption_resolution, [], [f240960, f240772])).
fof(f240772, plain, (~ (s_sworst7_init = a_select2(s_values7_init, sK51)) | (~ spl58_32 | spl58_56)), inference(forward_demodulation, [], [f901, f781])).
fof(f901, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK51)) | spl58_56), inference(avatar_component_clause, [], [f899])).
fof(f899, plain, (spl58_56 <=> (s_worst7_init = a_select2(s_values7_init, sK51))), introduced(avatar_definition, [new_symbols(naming, [spl58_56])])).
fof(f240960, plain, ((s_sworst7_init = a_select2(s_values7_init, sK51)) | (~ spl58_32 | ~ spl58_57 | ~ spl58_58)), inference(subsumption_resolution, [], [f240857, f906])).
fof(f906, plain, (leq(sK51, n3) | ~ spl58_57), inference(avatar_component_clause, [], [f904])).
fof(f904, plain, (spl58_57 <=> leq(sK51, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_57])])).
fof(f240857, plain, (~ leq(sK51, n3) | (s_sworst7_init = a_select2(s_values7_init, sK51)) | (~ spl58_32 | ~ spl58_58)), inference(resolution, [], [f911, f1026])).
fof(f911, plain, (leq(n0, sK51) | ~ spl58_58), inference(avatar_component_clause, [], [f909])).
fof(f909, plain, (spl58_58 <=> leq(n0, sK51)), introduced(avatar_definition, [new_symbols(naming, [spl58_58])])).
fof(f240730, plain, (~ spl58_32 | spl58_53 | ~ spl58_54 | ~ spl58_55), inference(avatar_contradiction_clause, [], [f240729])).
fof(f240729, plain, ($false | (~ spl58_32 | spl58_53 | ~ spl58_54 | ~ spl58_55)), inference(subsumption_resolution, [], [f240728, f240549])).
fof(f240549, plain, (~ (s_sworst7_init = a_select2(s_center7_init, sK50)) | (~ spl58_32 | spl58_53)), inference(forward_demodulation, [], [f882, f781])).
fof(f882, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK50)) | spl58_53), inference(avatar_component_clause, [], [f880])).
fof(f880, plain, (spl58_53 <=> (s_worst7_init = a_select2(s_center7_init, sK50))), introduced(avatar_definition, [new_symbols(naming, [spl58_53])])).
fof(f240728, plain, ((s_sworst7_init = a_select2(s_center7_init, sK50)) | (~ spl58_32 | ~ spl58_54 | ~ spl58_55)), inference(subsumption_resolution, [], [f240628, f887])).
fof(f887, plain, (leq(sK50, n2) | ~ spl58_54), inference(avatar_component_clause, [], [f885])).
fof(f885, plain, (spl58_54 <=> leq(sK50, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_54])])).
fof(f240628, plain, (~ leq(sK50, n2) | (s_sworst7_init = a_select2(s_center7_init, sK50)) | (~ spl58_32 | ~ spl58_55)), inference(resolution, [], [f892, f1027])).
fof(f892, plain, (leq(n0, sK50) | ~ spl58_55), inference(avatar_component_clause, [], [f890])).
fof(f890, plain, (spl58_55 <=> leq(n0, sK50)), introduced(avatar_definition, [new_symbols(naming, [spl58_55])])).
fof(f240484, plain, (spl58_729 | spl58_731 | spl58_732 | spl58_733 | ~ spl58_60 | ~ spl58_61), inference(avatar_split_clause, [], [f240483, f924, f919, f240458, f240454, f240450, f240332])).
fof(f919, plain, (spl58_60 <=> leq(sK53, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_60])])).
fof(f924, plain, (spl58_61 <=> leq(n0, sK53)), introduced(avatar_definition, [new_symbols(naming, [spl58_61])])).
fof(f240483, plain, ((n2 = sK53) | (n1 = sK53) | (n0 = sK53) | (n3 = sK53) | (~ spl58_60 | ~ spl58_61)), inference(subsumption_resolution, [], [f240381, f921])).
fof(f921, plain, (leq(sK53, n3) | ~ spl58_60), inference(avatar_component_clause, [], [f919])).
fof(f240381, plain, ((n2 = sK53) | (n1 = sK53) | (n0 = sK53) | ~ leq(sK53, n3) | (n3 = sK53) | ~ spl58_61), inference(resolution, [], [f926, f495])).
fof(f926, plain, (leq(n0, sK53) | ~ spl58_61), inference(avatar_component_clause, [], [f924])).
fof(f239666, plain, (~ spl58_42 | ~ spl58_32 | spl58_41 | ~ spl58_43), inference(avatar_split_clause, [], [f239665, f834, f824, f780, f829])).
fof(f829, plain, (spl58_42 <=> leq(sK47, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_42])])).
fof(f824, plain, (spl58_41 <=> (s_worst7_init = a_select2(s_values7_init, sK47))), introduced(avatar_definition, [new_symbols(naming, [spl58_41])])).
fof(f834, plain, (spl58_43 <=> leq(n0, sK47)), introduced(avatar_definition, [new_symbols(naming, [spl58_43])])).
fof(f239665, plain, (~ leq(sK47, n3) | (~ spl58_32 | spl58_41 | ~ spl58_43)), inference(subsumption_resolution, [], [f239497, f239423])).
fof(f239423, plain, (~ (s_sworst7_init = a_select2(s_values7_init, sK47)) | (~ spl58_32 | spl58_41)), inference(forward_demodulation, [], [f826, f781])).
fof(f826, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK47)) | spl58_41), inference(avatar_component_clause, [], [f824])).
fof(f239497, plain, (~ leq(sK47, n3) | (s_sworst7_init = a_select2(s_values7_init, sK47)) | (~ spl58_32 | ~ spl58_43)), inference(resolution, [], [f836, f1026])).
fof(f836, plain, (leq(n0, sK47) | ~ spl58_43), inference(avatar_component_clause, [], [f834])).
fof(f239419, plain, (~ spl58_687 | ~ spl58_32 | spl58_44 | ~ spl58_648), inference(avatar_split_clause, [], [f239418, f237508, f839, f780, f239365])).
fof(f239365, plain, (spl58_687 <=> (s_sworst7_init = a_select3(simplex7_init, n0, sK48))), introduced(avatar_definition, [new_symbols(naming, [spl58_687])])).
fof(f237508, plain, (spl58_648 <=> (n0 = sK49)), introduced(avatar_definition, [new_symbols(naming, [spl58_648])])).
fof(f239418, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n0, sK48)) | (~ spl58_32 | spl58_44 | ~ spl58_648)), inference(forward_demodulation, [], [f237103, f237510])).
fof(f237510, plain, ((n0 = sK49) | ~ spl58_648), inference(avatar_component_clause, [], [f237508])).
fof(f239405, plain, (spl58_669 | ~ spl58_47 | spl58_672 | ~ spl58_48 | spl58_671), inference(avatar_split_clause, [], [f239404, f238213, f859, f238217, f854, f237831])).
fof(f854, plain, (spl58_47 <=> leq(sK48, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_47])])).
fof(f859, plain, (spl58_48 <=> leq(n0, sK48)), introduced(avatar_definition, [new_symbols(naming, [spl58_48])])).
fof(f238213, plain, (spl58_671 <=> (n0 = sK48)), introduced(avatar_definition, [new_symbols(naming, [spl58_671])])).
fof(f239404, plain, ((n1 = sK48) | ~ leq(sK48, n2) | (n2 = sK48) | (~ spl58_48 | spl58_671)), inference(subsumption_resolution, [], [f238142, f238214])).
fof(f238214, plain, (~ (n0 = sK48) | spl58_671), inference(avatar_component_clause, [], [f238213])).
fof(f238142, plain, ((n1 = sK48) | (n0 = sK48) | ~ leq(sK48, n2) | (n2 = sK48) | ~ spl58_48), inference(resolution, [], [f861, f494])).
fof(f494, plain, ! [X0] : (~ leq(n0, X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n2) | (n2 = X0)), inference(cnf_transformation, [], [f160])).
fof(f160, plain, ! [X0] : ((n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n2) | ~ leq(n0, X0)), inference(flattening, [], [f159])).
fof(f159, plain, ! [X0] : (((n2 = X0) | (n1 = X0) | (n0 = X0)) | (~ leq(X0, n2) | ~ leq(n0, X0))), inference(ennf_transformation, [], [f80])).
fof(f80, plain, ! [X0] : ((leq(X0, n2) & leq(n0, X0)) => ((n2 = X0) | (n1 = X0) | (n0 = X0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', finite_domain_2)).
fof(f861, plain, (leq(n0, sK48) | ~ spl58_48), inference(avatar_component_clause, [], [f859])).
fof(f239378, plain, (spl58_689 | ~ spl58_47 | ~ spl58_32 | ~ spl58_48), inference(avatar_split_clause, [], [f238150, f859, f780, f854, f239375])).
fof(f238150, plain, (~ leq(sK48, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, sK48)) | (~ spl58_32 | ~ spl58_48)), inference(resolution, [], [f861, f2085])).
fof(f239368, plain, (spl58_687 | ~ spl58_47 | ~ spl58_32 | ~ spl58_48), inference(avatar_split_clause, [], [f238152, f859, f780, f854, f239365])).
fof(f238152, plain, (~ leq(sK48, n2) | (s_sworst7_init = a_select3(simplex7_init, n0, sK48)) | (~ spl58_32 | ~ spl58_48)), inference(resolution, [], [f861, f2090])).
fof(f238942, plain, (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_671), inference(avatar_contradiction_clause, [], [f238941])).
fof(f238941, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_671)), inference(subsumption_resolution, [], [f238858, f235699])).
fof(f238858, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n1, n0)) | (~ spl58_32 | spl58_44 | ~ spl58_649 | ~ spl58_671)), inference(backward_demodulation, [], [f238296, f237514])).
fof(f238296, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK49, n0)) | (~ spl58_32 | spl58_44 | ~ spl58_671)), inference(backward_demodulation, [], [f237103, f238215])).
fof(f238215, plain, ((n0 = sK48) | ~ spl58_671), inference(avatar_component_clause, [], [f238213])).
fof(f238731, plain, (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_671), inference(avatar_contradiction_clause, [], [f238730])).
fof(f238730, plain, ($false | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_671)), inference(subsumption_resolution, [], [f238708, f235911])).
fof(f235911, plain, ((s_sworst7_init = a_select3(simplex7_init, n3, n0)) | ~ spl58_32), inference(subsumption_resolution, [], [f235852, f1038])).
fof(f235852, plain, (~ leq(n0, n2) | (s_sworst7_init = a_select3(simplex7_init, n3, n0)) | ~ spl58_32), inference(resolution, [], [f2086, f283])).
fof(f238708, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n0)) | (~ spl58_32 | spl58_44 | ~ spl58_646 | ~ spl58_671)), inference(backward_demodulation, [], [f238296, f237392])).
fof(f237542, plain, (spl58_646 | spl58_648 | spl58_649 | spl58_650 | ~ spl58_45 | ~ spl58_46), inference(avatar_split_clause, [], [f237541, f849, f844, f237516, f237512, f237508, f237390])).
fof(f844, plain, (spl58_45 <=> leq(sK49, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_45])])).
fof(f849, plain, (spl58_46 <=> leq(n0, sK49)), introduced(avatar_definition, [new_symbols(naming, [spl58_46])])).
fof(f237541, plain, ((n2 = sK49) | (n1 = sK49) | (n0 = sK49) | (n3 = sK49) | (~ spl58_45 | ~ spl58_46)), inference(subsumption_resolution, [], [f237439, f846])).
fof(f846, plain, (leq(sK49, n3) | ~ spl58_45), inference(avatar_component_clause, [], [f844])).
fof(f237439, plain, ((n2 = sK49) | (n1 = sK49) | (n0 = sK49) | ~ leq(sK49, n3) | (n3 = sK49) | ~ spl58_46), inference(resolution, [], [f851, f495])).
fof(f851, plain, (leq(n0, sK49) | ~ spl58_46), inference(avatar_component_clause, [], [f849])).
fof(f237061, plain, (spl58_13 | ~ spl58_17 | ~ spl58_18 | ~ spl58_32), inference(avatar_contradiction_clause, [], [f237060])).
fof(f237060, plain, ($false | (spl58_13 | ~ spl58_17 | ~ spl58_18 | ~ spl58_32)), inference(subsumption_resolution, [], [f237059, f236882])).
fof(f236882, plain, (~ (s_sworst7_init = a_select2(s_center7_init, sK42)) | (spl58_13 | ~ spl58_32)), inference(forward_demodulation, [], [f690, f781])).
fof(f690, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK42)) | spl58_13), inference(avatar_component_clause, [], [f688])).
fof(f688, plain, (spl58_13 <=> (s_worst7_init = a_select2(s_center7_init, sK42))), introduced(avatar_definition, [new_symbols(naming, [spl58_13])])).
fof(f237059, plain, ((s_sworst7_init = a_select2(s_center7_init, sK42)) | (~ spl58_17 | ~ spl58_18 | ~ spl58_32)), inference(subsumption_resolution, [], [f236959, f707])).
fof(f707, plain, (leq(sK42, n2) | ~ spl58_17), inference(avatar_component_clause, [], [f705])).
fof(f705, plain, (spl58_17 <=> leq(sK42, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_17])])).
fof(f236959, plain, (~ leq(sK42, n2) | (s_sworst7_init = a_select2(s_center7_init, sK42)) | (~ spl58_18 | ~ spl58_32)), inference(resolution, [], [f712, f1027])).
fof(f712, plain, (leq(n0, sK42) | ~ spl58_18), inference(avatar_component_clause, [], [f710])).
fof(f710, plain, (spl58_18 <=> leq(n0, sK42)), introduced(avatar_definition, [new_symbols(naming, [spl58_18])])).
fof(f236848, plain, (~ spl58_32 | spl58_33 | ~ spl58_49 | ~ spl58_50), inference(avatar_contradiction_clause, [], [f236847])).
fof(f236847, plain, ($false | (~ spl58_32 | spl58_33 | ~ spl58_49 | ~ spl58_50)), inference(subsumption_resolution, [], [f1832, f236846])).
fof(f236846, plain, (~ (s_sworst7_init = a_select2(s_values7_init, s_sworst7)) | (~ spl58_32 | spl58_33)), inference(forward_demodulation, [], [f786, f781])).
fof(f786, plain, (~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | spl58_33), inference(avatar_component_clause, [], [f784])).
fof(f784, plain, (spl58_33 <=> (s_worst7_init = a_select2(s_values7_init, s_sworst7))), introduced(avatar_definition, [new_symbols(naming, [spl58_33])])).
fof(f1832, plain, ((s_sworst7_init = a_select2(s_values7_init, s_sworst7)) | (~ spl58_32 | ~ spl58_49 | ~ spl58_50)), inference(subsumption_resolution, [], [f1805, f869])).
fof(f869, plain, (leq(s_sworst7, n3) | ~ spl58_50), inference(avatar_component_clause, [], [f868])).
fof(f868, plain, (spl58_50 <=> leq(s_sworst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_50])])).
fof(f1805, plain, (~ leq(s_sworst7, n3) | (s_sworst7_init = a_select2(s_values7_init, s_sworst7)) | (~ spl58_32 | ~ spl58_49)), inference(resolution, [], [f1026, f865])).
fof(f865, plain, (leq(n0, s_sworst7) | ~ spl58_49), inference(avatar_component_clause, [], [f864])).
fof(f864, plain, (spl58_49 <=> leq(n0, s_sworst7)), introduced(avatar_definition, [new_symbols(naming, [spl58_49])])).
fof(f236613, plain, (~ spl58_74 | ~ spl58_32 | spl58_73 | ~ spl58_75), inference(avatar_split_clause, [], [f236612, f993, f983, f780, f988])).
fof(f988, plain, (spl58_74 <=> leq(sK57, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_74])])).
fof(f983, plain, (spl58_73 <=> (s_worst7_init = a_select2(s_center7_init, sK57))), introduced(avatar_definition, [new_symbols(naming, [spl58_73])])).
fof(f993, plain, (spl58_75 <=> leq(n0, sK57)), introduced(avatar_definition, [new_symbols(naming, [spl58_75])])).
fof(f236612, plain, (~ leq(sK57, n2) | (~ spl58_32 | spl58_73 | ~ spl58_75)), inference(subsumption_resolution, [], [f236436, f236358])).
fof(f236358, plain, (~ (s_sworst7_init = a_select2(s_center7_init, sK57)) | (~ spl58_32 | spl58_73)), inference(forward_demodulation, [], [f985, f781])).
fof(f985, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK57)) | spl58_73), inference(avatar_component_clause, [], [f983])).
fof(f236436, plain, (~ leq(sK57, n2) | (s_sworst7_init = a_select2(s_center7_init, sK57)) | (~ spl58_32 | ~ spl58_75)), inference(resolution, [], [f995, f1027])).
fof(f995, plain, (leq(n0, sK57) | ~ spl58_75), inference(avatar_component_clause, [], [f993])).
fof(f236357, plain, (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_185), inference(avatar_contradiction_clause, [], [f236356])).
fof(f236356, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_185)), inference(subsumption_resolution, [], [f236313, f235803])).
fof(f235803, plain, ((s_sworst7_init = a_select3(simplex7_init, n2, n1)) | ~ spl58_32), inference(subsumption_resolution, [], [f235709, f1039])).
fof(f235709, plain, (~ leq(n1, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, n1)) | ~ spl58_32), inference(resolution, [], [f2085, f1035])).
fof(f236313, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, n1)) | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_185)), inference(backward_demodulation, [], [f236121, f10206])).
fof(f236121, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, sK44)) | (spl58_27 | ~ spl58_32 | ~ spl58_172)), inference(forward_demodulation, [], [f9159, f9577])).
fof(f9577, plain, ((n2 = sK45) | ~ spl58_172), inference(avatar_component_clause, [], [f9575])).
fof(f9575, plain, (spl58_172 <=> (n2 = sK45)), introduced(avatar_definition, [new_symbols(naming, [spl58_172])])).
fof(f236255, plain, (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_184), inference(avatar_contradiction_clause, [], [f236254])).
fof(f236254, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_184)), inference(subsumption_resolution, [], [f236175, f235805])).
fof(f235805, plain, ((s_sworst7_init = a_select3(simplex7_init, n2, n0)) | ~ spl58_32), inference(subsumption_resolution, [], [f235746, f1038])).
fof(f235746, plain, (~ leq(n0, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, n0)) | ~ spl58_32), inference(resolution, [], [f2085, f283])).
fof(f236175, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, n0)) | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_184)), inference(backward_demodulation, [], [f236121, f10202])).
fof(f236118, plain, (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_181), inference(avatar_contradiction_clause, [], [f236117])).
fof(f236117, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_181)), inference(subsumption_resolution, [], [f236076, f235804])).
fof(f235804, plain, ((s_sworst7_init = a_select3(simplex7_init, n2, n2)) | ~ spl58_32), inference(subsumption_resolution, [], [f235726, f283])).
fof(f235726, plain, (~ leq(n2, n2) | (s_sworst7_init = a_select3(simplex7_init, n2, n2)) | ~ spl58_32), inference(resolution, [], [f2085, f1038])).
fof(f236076, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n2, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_172 | ~ spl58_181)), inference(backward_demodulation, [], [f236021, f9577])).
fof(f236021, plain, (~ (s_sworst7_init = a_select3(simplex7_init, sK45, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_181)), inference(forward_demodulation, [], [f9159, f10138])).
fof(f235978, plain, (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_181), inference(avatar_contradiction_clause, [], [f235977])).
fof(f235977, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_181)), inference(subsumption_resolution, [], [f235924, f235910])).
fof(f235924, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n2)) | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_181)), inference(backward_demodulation, [], [f10033, f10138])).
fof(f10033, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, sK44)) | (spl58_27 | ~ spl58_32 | ~ spl58_167)), inference(backward_demodulation, [], [f9159, f9505])).
fof(f235913, plain, (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_184), inference(avatar_contradiction_clause, [], [f235912])).
fof(f235912, plain, ($false | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_184)), inference(subsumption_resolution, [], [f235911, f10255])).
fof(f10255, plain, (~ (s_sworst7_init = a_select3(simplex7_init, n3, n0)) | (spl58_27 | ~ spl58_32 | ~ spl58_167 | ~ spl58_184)), inference(backward_demodulation, [], [f10033, f10202])).
fof(f10232, plain, (spl58_181 | spl58_184 | spl58_185 | ~ spl58_30 | ~ spl58_31), inference(avatar_split_clause, [], [f10231, f775, f770, f10204, f10200, f10136])).
fof(f770, plain, (spl58_30 <=> leq(sK44, n2)), introduced(avatar_definition, [new_symbols(naming, [spl58_30])])).
fof(f775, plain, (spl58_31 <=> leq(n0, sK44)), introduced(avatar_definition, [new_symbols(naming, [spl58_31])])).
fof(f10231, plain, ((n1 = sK44) | (n0 = sK44) | (n2 = sK44) | (~ spl58_30 | ~ spl58_31)), inference(subsumption_resolution, [], [f10182, f772])).
fof(f772, plain, (leq(sK44, n2) | ~ spl58_30), inference(avatar_component_clause, [], [f770])).
fof(f10182, plain, ((n1 = sK44) | (n0 = sK44) | ~ leq(sK44, n2) | (n2 = sK44) | ~ spl58_31), inference(resolution, [], [f777, f494])).
fof(f777, plain, (leq(n0, sK44) | ~ spl58_31), inference(avatar_component_clause, [], [f775])).
fof(f9604, plain, (spl58_167 | spl58_170 | spl58_171 | spl58_172 | ~ spl58_28 | ~ spl58_29), inference(avatar_split_clause, [], [f9603, f765, f760, f9575, f9571, f9567, f9503])).
fof(f760, plain, (spl58_28 <=> leq(sK45, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_28])])).
fof(f765, plain, (spl58_29 <=> leq(n0, sK45)), introduced(avatar_definition, [new_symbols(naming, [spl58_29])])).
fof(f9603, plain, ((n2 = sK45) | (n1 = sK45) | (n0 = sK45) | (n3 = sK45) | (~ spl58_28 | ~ spl58_29)), inference(subsumption_resolution, [], [f9550, f762])).
fof(f762, plain, (leq(sK45, n3) | ~ spl58_28), inference(avatar_component_clause, [], [f760])).
fof(f9550, plain, ((n2 = sK45) | (n1 = sK45) | (n0 = sK45) | ~ leq(sK45, n3) | (n3 = sK45) | ~ spl58_29), inference(resolution, [], [f767, f495])).
fof(f767, plain, (leq(n0, sK45) | ~ spl58_29), inference(avatar_component_clause, [], [f765])).
fof(f9024, plain, (spl58_23 | ~ spl58_24 | ~ spl58_25 | ~ spl58_32), inference(avatar_contradiction_clause, [], [f9023])).
fof(f9023, plain, ($false | (spl58_23 | ~ spl58_24 | ~ spl58_25 | ~ spl58_32)), inference(subsumption_resolution, [], [f9022, f8760])).
fof(f8760, plain, (~ (s_sworst7_init = a_select2(s_values7_init, sK43)) | (spl58_23 | ~ spl58_32)), inference(forward_demodulation, [], [f738, f781])).
fof(f738, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK43)) | spl58_23), inference(avatar_component_clause, [], [f736])).
fof(f736, plain, (spl58_23 <=> (s_worst7_init = a_select2(s_values7_init, sK43))), introduced(avatar_definition, [new_symbols(naming, [spl58_23])])).
fof(f9022, plain, ((s_sworst7_init = a_select2(s_values7_init, sK43)) | (~ spl58_24 | ~ spl58_25 | ~ spl58_32)), inference(subsumption_resolution, [], [f8968, f743])).
fof(f743, plain, (leq(sK43, n3) | ~ spl58_24), inference(avatar_component_clause, [], [f741])).
fof(f741, plain, (spl58_24 <=> leq(sK43, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_24])])).
fof(f8968, plain, (~ leq(sK43, n3) | (s_sworst7_init = a_select2(s_values7_init, sK43)) | (~ spl58_25 | ~ spl58_32)), inference(resolution, [], [f748, f1026])).
fof(f748, plain, (leq(n0, sK43) | ~ spl58_25), inference(avatar_component_clause, [], [f746])).
fof(f746, plain, (spl58_25 <=> leq(n0, sK43)), introduced(avatar_definition, [new_symbols(naming, [spl58_25])])).
fof(f8475, plain, (spl58_3 | ~ spl58_7 | ~ spl58_10 | ~ spl58_32), inference(avatar_contradiction_clause, [], [f8474])).
fof(f8474, plain, ($false | (spl58_3 | ~ spl58_7 | ~ spl58_10 | ~ spl58_32)), inference(subsumption_resolution, [], [f8473, f1028])).
fof(f1028, plain, (~ (s_sworst7_init = a_select2(s_values7_init, pv1388)) | (spl58_3 | ~ spl58_32)), inference(forward_demodulation, [], [f650, f781])).
fof(f650, plain, (~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | spl58_3), inference(avatar_component_clause, [], [f648])).
fof(f648, plain, (spl58_3 <=> (s_worst7_init = a_select2(s_values7_init, pv1388))), introduced(avatar_definition, [new_symbols(naming, [spl58_3])])).
fof(f8473, plain, ((s_sworst7_init = a_select2(s_values7_init, pv1388)) | (~ spl58_7 | ~ spl58_10 | ~ spl58_32)), inference(subsumption_resolution, [], [f8423, f677])).
fof(f677, plain, (leq(pv1388, n3) | ~ spl58_10), inference(avatar_component_clause, [], [f676])).
fof(f676, plain, (spl58_10 <=> leq(pv1388, n3)), introduced(avatar_definition, [new_symbols(naming, [spl58_10])])).
fof(f8423, plain, (~ leq(pv1388, n3) | (s_sworst7_init = a_select2(s_values7_init, pv1388)) | (~ spl58_7 | ~ spl58_32)), inference(resolution, [], [f665, f1026])).
fof(f665, plain, (leq(n0, pv1388) | ~ spl58_7), inference(avatar_component_clause, [], [f664])).
fof(f664, plain, (spl58_7 <=> leq(n0, pv1388)), introduced(avatar_definition, [new_symbols(naming, [spl58_7])])).
fof(f8376, plain, spl58_7, inference(avatar_split_clause, [], [f8364, f664])).
fof(f8364, plain, leq(n0, pv1388), inference(resolution, [], [f1270, f1038])).
fof(f1270, plain, ! [X62] : (~ leq(X62, n2) | leq(X62, pv1388)), inference(resolution, [], [f284, f451])).
fof(f451, plain, leq(n2, pv1388), inference(cnf_transformation, [], [f279])).
fof(f284, plain, ! [X2, X0, X1] : (~ leq(X1, X2) | leq(X0, X2) | ~ leq(X0, X1)), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ! [X0, X1, X2] : (leq(X0, X2) | ~ leq(X1, X2) | ~ leq(X0, X1)), inference(flattening, [], [f111])).
fof(f111, plain, ! [X0, X1, X2] : (leq(X0, X2) | (~ leq(X1, X2) | ~ leq(X0, X1))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0, X1, X2] : ((leq(X1, X2) & leq(X0, X1)) => leq(X0, X2)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV036+1.p', transitivity_leq)).
fof(f1011, plain, spl58_32, inference(avatar_split_clause, [], [f564, f780])).
fof(f564, plain, (s_sworst7_init = s_worst7_init), inference(definition_unfolding, [], [f446, f447])).
fof(f446, plain, (init = s_sworst7_init), inference(cnf_transformation, [], [f279])).
fof(f1010, plain, spl58_5, inference(avatar_split_clause, [], [f448, f656])).
fof(f448, plain, leq(n0, s_best7), inference(cnf_transformation, [], [f279])).
fof(f1009, plain, spl58_49, inference(avatar_split_clause, [], [f449, f864])).
fof(f449, plain, leq(n0, s_sworst7), inference(cnf_transformation, [], [f279])).
fof(f1008, plain, spl58_6, inference(avatar_split_clause, [], [f450, f660])).
fof(f450, plain, leq(n0, s_worst7), inference(cnf_transformation, [], [f279])).
fof(f1007, plain, spl58_8, inference(avatar_split_clause, [], [f452, f668])).
fof(f452, plain, leq(s_best7, n3), inference(cnf_transformation, [], [f279])).
fof(f1006, plain, spl58_50, inference(avatar_split_clause, [], [f453, f868])).
fof(f453, plain, leq(s_sworst7, n3), inference(cnf_transformation, [], [f279])).
fof(f1005, plain, spl58_9, inference(avatar_split_clause, [], [f454, f672])).
fof(f454, plain, leq(s_worst7, n3), inference(cnf_transformation, [], [f279])).
fof(f1004, plain, spl58_10, inference(avatar_split_clause, [], [f455, f676])).
fof(f455, plain, leq(pv1388, n3), inference(cnf_transformation, [], [f279])).
fof(f1003, plain, (~ spl58_19 | spl58_14), inference(avatar_split_clause, [], [f560, f692, f715])).
fof(f715, plain, (spl58_19 <=> gt(loopcounter, n1)), introduced(avatar_definition, [new_symbols(naming, [spl58_19])])).
fof(f692, plain, (spl58_14 <=> (s_worst7_init = pvar1400_init)), introduced(avatar_definition, [new_symbols(naming, [spl58_14])])).
fof(f560, plain, ((s_worst7_init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f459, f447])).
fof(f459, plain, ((init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f279])).
fof(f1002, plain, (~ spl58_19 | spl58_15), inference(avatar_split_clause, [], [f559, f696, f715])).
fof(f696, plain, (spl58_15 <=> (s_worst7_init = pvar1401_init)), introduced(avatar_definition, [new_symbols(naming, [spl58_15])])).
fof(f559, plain, ((s_worst7_init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f460, f447])).
fof(f460, plain, ((init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f279])).
fof(f1001, plain, (~ spl58_19 | spl58_16), inference(avatar_split_clause, [], [f558, f700, f715])).
fof(f700, plain, (spl58_16 <=> (s_worst7_init = pvar1402_init)), introduced(avatar_definition, [new_symbols(naming, [spl58_16])])).
fof(f558, plain, ((s_worst7_init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f461, f447])).
fof(f461, plain, ((init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f279])).
fof(f999, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | spl58_75 | spl58_19), inference(avatar_split_clause, [], [f576, f715, f993, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f640, plain, (spl58_1 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl58_1])])).
fof(f751, plain, (spl58_26 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl58_26])])).
fof(f732, plain, (spl58_22 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl58_22])])).
fof(f576, plain, (gt(loopcounter, n1) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f575])).
fof(f575, plain, (gt(loopcounter, n1) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f556])).
fof(f556, plain, (gt(loopcounter, n1) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f463, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f501, plain, (s_best7_init = s_worst7_init), inference(definition_unfolding, [], [f445, f447])).
fof(f445, plain, (s_best7_init = init), inference(cnf_transformation, [], [f279])).
fof(f463, plain, (gt(loopcounter, n1) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f998, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | spl58_74 | spl58_19), inference(avatar_split_clause, [], [f578, f715, f988, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f578, plain, (gt(loopcounter, n1) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f577])).
fof(f577, plain, (gt(loopcounter, n1) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f555])).
fof(f555, plain, (gt(loopcounter, n1) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f464, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f464, plain, (gt(loopcounter, n1) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f997, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | ~ spl58_73 | spl58_19), inference(avatar_split_clause, [], [f580, f715, f983, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f580, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f579])).
fof(f579, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f554])).
fof(f554, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f465, f447, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f465, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f996, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | spl58_75 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f582, f700, f696, f692, f993, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f582, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f581])).
fof(f581, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f553])).
fof(f553, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f466, f447, f447, f447, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f466, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK57) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f991, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | spl58_74 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f584, f700, f696, f692, f988, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f584, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f583])).
fof(f583, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f552])).
fof(f552, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f467, f447, f447, f447, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f467, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK57, n2) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f986, plain, (~ spl58_72 | ~ spl58_3 | spl58_1 | ~ spl58_32 | ~ spl58_49 | ~ spl58_6 | ~ spl58_7 | ~ spl58_50 | ~ spl58_9 | ~ spl58_10 | spl58_26 | spl58_22 | ~ spl58_73 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f586, f700, f696, f692, f983, f732, f751, f676, f672, f868, f664, f660, f864, f780, f640, f648, f979])).
fof(f586, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(trivial_inequality_removal, [], [f585])).
fof(f585, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7))), inference(duplicate_literal_removal, [], [f551])).
fof(f551, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP14 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_best7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f468, f447, f447, f447, f447, f447, f447, f447, f447, f447, f447, f501, f447, f447, f447])).
fof(f468, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_center7_init, sK57)) | sP13 | sP12 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = init) | sP14 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_best7)) | ~ (s_best7_init = init) | ~ (init = init)), inference(cnf_transformation, [], [f279])).
fof(f977, plain, (~ spl58_51 | spl58_71), inference(avatar_split_clause, [], [f440, f974, f872])).
fof(f872, plain, (spl58_51 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl58_51])])).
fof(f440, plain, (leq(n0, sK55) | ~ sP4), inference(cnf_transformation, [], [f276])).
fof(f276, plain, (((~ (init = a_select3(simplex7_init, sK56, sK55)) & leq(sK56, n3) & leq(n0, sK56)) & leq(sK55, n2) & leq(n0, sK55)) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK55, sK56])], [f273, f275, f274])).
fof(f274, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK55)) & leq(X1, n3) & leq(n0, X1)) & leq(sK55, n2) & leq(n0, sK55))), introduced(choice_axiom, [])).
fof(f275, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK55)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK56, sK55)) & leq(sK56, n3) & leq(n0, sK56))), introduced(choice_axiom, [])).
fof(f273, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP4), inference(rectify, [], [f272])).
fof(f272, plain, (? [X18] : (? [X19] : (~ (init = a_select3(simplex7_init, X19, X18)) & leq(X19, n3) & leq(n0, X19)) & leq(X18, n2) & leq(n0, X18)) | ~ sP4), inference(nnf_transformation, [], [f170])).
fof(f972, plain, (~ spl58_51 | spl58_70), inference(avatar_split_clause, [], [f441, f969, f872])).
fof(f441, plain, (leq(sK55, n2) | ~ sP4), inference(cnf_transformation, [], [f276])).
fof(f967, plain, (~ spl58_51 | spl58_69), inference(avatar_split_clause, [], [f442, f964, f872])).
fof(f442, plain, (leq(n0, sK56) | ~ sP4), inference(cnf_transformation, [], [f276])).
fof(f962, plain, (~ spl58_51 | spl58_68), inference(avatar_split_clause, [], [f443, f959, f872])).
fof(f443, plain, (leq(sK56, n3) | ~ sP4), inference(cnf_transformation, [], [f276])).
fof(f957, plain, (~ spl58_51 | ~ spl58_67), inference(avatar_split_clause, [], [f550, f954, f872])).
fof(f550, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK56, sK55)) | ~ sP4), inference(definition_unfolding, [], [f444, f447])).
fof(f444, plain, (~ (init = a_select3(simplex7_init, sK56, sK55)) | ~ sP4), inference(cnf_transformation, [], [f276])).
fof(f952, plain, (~ spl58_52 | spl58_66), inference(avatar_split_clause, [], [f437, f949, f876])).
fof(f876, plain, (spl58_52 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl58_52])])).
fof(f437, plain, (leq(n0, sK54) | ~ sP5), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ((~ (init = a_select2(s_values7_init, sK54)) & leq(sK54, n3) & leq(n0, sK54)) | ~ sP5), inference(skolemisation, [status(esa), new_symbols(skolem, [sK54])], [f269, f270])).
fof(f270, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK54)) & leq(sK54, n3) & leq(n0, sK54))), introduced(choice_axiom, [])).
fof(f269, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP5), inference(rectify, [], [f268])).
fof(f268, plain, (? [X17] : (~ (init = a_select2(s_values7_init, X17)) & leq(X17, n3) & leq(n0, X17)) | ~ sP5), inference(nnf_transformation, [], [f171])).
fof(f947, plain, (~ spl58_52 | spl58_65), inference(avatar_split_clause, [], [f438, f944, f876])).
fof(f438, plain, (leq(sK54, n3) | ~ sP5), inference(cnf_transformation, [], [f271])).
fof(f942, plain, (~ spl58_52 | ~ spl58_64), inference(avatar_split_clause, [], [f549, f939, f876])).
fof(f549, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK54)) | ~ sP5), inference(definition_unfolding, [], [f439, f447])).
fof(f439, plain, (~ (init = a_select2(s_values7_init, sK54)) | ~ sP5), inference(cnf_transformation, [], [f271])).
fof(f937, plain, (~ spl58_35 | spl58_63), inference(avatar_split_clause, [], [f432, f934, f792])).
fof(f792, plain, (spl58_35 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl58_35])])).
fof(f432, plain, (leq(n0, sK52) | ~ sP6), inference(cnf_transformation, [], [f267])).
fof(f267, plain, (((~ (init = a_select3(simplex7_init, sK53, sK52)) & leq(sK53, n3) & leq(n0, sK53)) & leq(sK52, n2) & leq(n0, sK52)) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK52, sK53])], [f264, f266, f265])).
fof(f265, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK52)) & leq(X1, n3) & leq(n0, X1)) & leq(sK52, n2) & leq(n0, sK52))), introduced(choice_axiom, [])).
fof(f266, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK52)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK53, sK52)) & leq(sK53, n3) & leq(n0, sK53))), introduced(choice_axiom, [])).
fof(f264, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP6), inference(rectify, [], [f263])).
fof(f263, plain, (? [X14] : (? [X15] : (~ (init = a_select3(simplex7_init, X15, X14)) & leq(X15, n3) & leq(n0, X15)) & leq(X14, n2) & leq(n0, X14)) | ~ sP6), inference(nnf_transformation, [], [f172])).
fof(f932, plain, (~ spl58_35 | spl58_62), inference(avatar_split_clause, [], [f433, f929, f792])).
fof(f433, plain, (leq(sK52, n2) | ~ sP6), inference(cnf_transformation, [], [f267])).
fof(f927, plain, (~ spl58_35 | spl58_61), inference(avatar_split_clause, [], [f434, f924, f792])).
fof(f434, plain, (leq(n0, sK53) | ~ sP6), inference(cnf_transformation, [], [f267])).
fof(f922, plain, (~ spl58_35 | spl58_60), inference(avatar_split_clause, [], [f435, f919, f792])).
fof(f435, plain, (leq(sK53, n3) | ~ sP6), inference(cnf_transformation, [], [f267])).
fof(f917, plain, (~ spl58_35 | ~ spl58_59), inference(avatar_split_clause, [], [f548, f914, f792])).
fof(f548, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK53, sK52)) | ~ sP6), inference(definition_unfolding, [], [f436, f447])).
fof(f436, plain, (~ (init = a_select3(simplex7_init, sK53, sK52)) | ~ sP6), inference(cnf_transformation, [], [f267])).
fof(f912, plain, (~ spl58_36 | spl58_58), inference(avatar_split_clause, [], [f429, f909, f796])).
fof(f796, plain, (spl58_36 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl58_36])])).
fof(f429, plain, (leq(n0, sK51) | ~ sP7), inference(cnf_transformation, [], [f262])).
fof(f262, plain, ((~ (init = a_select2(s_values7_init, sK51)) & leq(sK51, n3) & leq(n0, sK51)) | ~ sP7), inference(skolemisation, [status(esa), new_symbols(skolem, [sK51])], [f260, f261])).
fof(f261, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK51)) & leq(sK51, n3) & leq(n0, sK51))), introduced(choice_axiom, [])).
fof(f260, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP7), inference(rectify, [], [f259])).
fof(f259, plain, (? [X13] : (~ (init = a_select2(s_values7_init, X13)) & leq(X13, n3) & leq(n0, X13)) | ~ sP7), inference(nnf_transformation, [], [f173])).
fof(f907, plain, (~ spl58_36 | spl58_57), inference(avatar_split_clause, [], [f430, f904, f796])).
fof(f430, plain, (leq(sK51, n3) | ~ sP7), inference(cnf_transformation, [], [f262])).
fof(f902, plain, (~ spl58_36 | ~ spl58_56), inference(avatar_split_clause, [], [f547, f899, f796])).
fof(f547, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK51)) | ~ sP7), inference(definition_unfolding, [], [f431, f447])).
fof(f431, plain, (~ (init = a_select2(s_values7_init, sK51)) | ~ sP7), inference(cnf_transformation, [], [f262])).
fof(f896, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | spl58_55 | spl58_19), inference(avatar_split_clause, [], [f588, f715, f890, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f788, plain, (spl58_34 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl58_34])])).
fof(f588, plain, (gt(loopcounter, n1) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f587])).
fof(f587, plain, (gt(loopcounter, n1) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f546])).
fof(f546, plain, (gt(loopcounter, n1) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f423, f447, f447, f501, f447])).
fof(f423, plain, (gt(loopcounter, n1) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f258, plain, (((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_center7_init, sK50)) & leq(sK50, n2) & leq(n0, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ sP8), inference(skolemisation, [status(esa), new_symbols(skolem, [sK50])], [f256, f257])).
fof(f257, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK50)) & leq(sK50, n2) & leq(n0, sK50))), introduced(choice_axiom, [])).
fof(f256, plain, (((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ sP8), inference(rectify, [], [f255])).
fof(f255, plain, (((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X16] : (~ (init = a_select2(s_center7_init, X16)) & leq(X16, n2) & leq(n0, X16)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | ~ sP8), inference(nnf_transformation, [], [f174])).
fof(f895, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | spl58_54 | spl58_19), inference(avatar_split_clause, [], [f590, f715, f885, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f590, plain, (gt(loopcounter, n1) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f589])).
fof(f589, plain, (gt(loopcounter, n1) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f545])).
fof(f545, plain, (gt(loopcounter, n1) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f424, f447, f447, f501, f447])).
fof(f424, plain, (gt(loopcounter, n1) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f894, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | ~ spl58_53 | spl58_19), inference(avatar_split_clause, [], [f592, f715, f880, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f592, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f591])).
fof(f591, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f544])).
fof(f544, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f425, f447, f447, f447, f501, f447])).
fof(f425, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f893, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | spl58_55 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f594, f700, f696, f692, f890, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f594, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f593])).
fof(f593, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f543])).
fof(f543, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f426, f447, f447, f447, f447, f447, f501, f447])).
fof(f426, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK50) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f888, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | spl58_54 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f596, f700, f696, f692, f885, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f596, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f595])).
fof(f595, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f542])).
fof(f542, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f427, f447, f447, f447, f447, f447, f501, f447])).
fof(f427, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK50, n2) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f883, plain, (~ spl58_34 | ~ spl58_32 | ~ spl58_5 | ~ spl58_49 | ~ spl58_6 | ~ spl58_8 | ~ spl58_50 | ~ spl58_9 | spl58_51 | spl58_52 | ~ spl58_53 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f598, f700, f696, f692, f880, f876, f872, f672, f868, f668, f660, f864, f656, f780, f788])).
fof(f598, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(trivial_inequality_removal, [], [f597])).
fof(f597, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ sP8), inference(duplicate_literal_removal, [], [f541])).
fof(f541, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP8), inference(definition_unfolding, [], [f428, f447, f447, f447, f447, f447, f447, f501, f447])).
fof(f428, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_center7_init, sK50)) | sP5 | sP4 | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (s_best7_init = init) | ~ sP8), inference(cnf_transformation, [], [f258])).
fof(f862, plain, (~ spl58_11 | spl58_48), inference(avatar_split_clause, [], [f417, f859, f680])).
fof(f680, plain, (spl58_11 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl58_11])])).
fof(f417, plain, (leq(n0, sK48) | ~ sP9), inference(cnf_transformation, [], [f254])).
fof(f254, plain, (((~ (init = a_select3(simplex7_init, sK49, sK48)) & leq(sK49, n3) & leq(n0, sK49)) & leq(sK48, n2) & leq(n0, sK48)) | ~ sP9), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48, sK49])], [f251, f253, f252])).
fof(f252, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK48)) & leq(X1, n3) & leq(n0, X1)) & leq(sK48, n2) & leq(n0, sK48))), introduced(choice_axiom, [])).
fof(f253, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK48)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK49, sK48)) & leq(sK49, n3) & leq(n0, sK49))), introduced(choice_axiom, [])).
fof(f251, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP9), inference(rectify, [], [f250])).
fof(f250, plain, (? [X10] : (? [X11] : (~ (init = a_select3(simplex7_init, X11, X10)) & leq(X11, n3) & leq(n0, X11)) & leq(X10, n2) & leq(n0, X10)) | ~ sP9), inference(nnf_transformation, [], [f175])).
fof(f857, plain, (~ spl58_11 | spl58_47), inference(avatar_split_clause, [], [f418, f854, f680])).
fof(f418, plain, (leq(sK48, n2) | ~ sP9), inference(cnf_transformation, [], [f254])).
fof(f852, plain, (~ spl58_11 | spl58_46), inference(avatar_split_clause, [], [f419, f849, f680])).
fof(f419, plain, (leq(n0, sK49) | ~ sP9), inference(cnf_transformation, [], [f254])).
fof(f847, plain, (~ spl58_11 | spl58_45), inference(avatar_split_clause, [], [f420, f844, f680])).
fof(f420, plain, (leq(sK49, n3) | ~ sP9), inference(cnf_transformation, [], [f254])).
fof(f842, plain, (~ spl58_11 | ~ spl58_44), inference(avatar_split_clause, [], [f540, f839, f680])).
fof(f540, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK49, sK48)) | ~ sP9), inference(definition_unfolding, [], [f421, f447])).
fof(f421, plain, (~ (init = a_select3(simplex7_init, sK49, sK48)) | ~ sP9), inference(cnf_transformation, [], [f254])).
fof(f837, plain, (~ spl58_12 | spl58_43), inference(avatar_split_clause, [], [f414, f834, f684])).
fof(f684, plain, (spl58_12 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl58_12])])).
fof(f414, plain, (leq(n0, sK47) | ~ sP10), inference(cnf_transformation, [], [f249])).
fof(f249, plain, ((~ (init = a_select2(s_values7_init, sK47)) & leq(sK47, n3) & leq(n0, sK47)) | ~ sP10), inference(skolemisation, [status(esa), new_symbols(skolem, [sK47])], [f247, f248])).
fof(f248, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK47)) & leq(sK47, n3) & leq(n0, sK47))), introduced(choice_axiom, [])).
fof(f247, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP10), inference(rectify, [], [f246])).
fof(f246, plain, (? [X9] : (~ (init = a_select2(s_values7_init, X9)) & leq(X9, n3) & leq(n0, X9)) | ~ sP10), inference(nnf_transformation, [], [f176])).
fof(f832, plain, (~ spl58_12 | spl58_42), inference(avatar_split_clause, [], [f415, f829, f684])).
fof(f415, plain, (leq(sK47, n3) | ~ sP10), inference(cnf_transformation, [], [f249])).
fof(f827, plain, (~ spl58_12 | ~ spl58_41), inference(avatar_split_clause, [], [f539, f824, f684])).
fof(f539, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK47)) | ~ sP10), inference(definition_unfolding, [], [f416, f447])).
fof(f416, plain, (~ (init = a_select2(s_values7_init, sK47)) | ~ sP10), inference(cnf_transformation, [], [f249])).
fof(f816, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | spl58_39 | spl58_19), inference(avatar_split_clause, [], [f601, f715, f810, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f652, plain, (spl58_4 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl58_4])])).
fof(f601, plain, (gt(loopcounter, n1) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f600])).
fof(f600, plain, (gt(loopcounter, n1) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f537])).
fof(f537, plain, (gt(loopcounter, n1) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f408, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f408, plain, (gt(loopcounter, n1) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f245, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_center7_init, sK46)) & leq(sK46, n2) & leq(n0, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ sP11), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46])], [f243, f244])).
fof(f244, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK46)) & leq(sK46, n2) & leq(n0, sK46))), introduced(choice_axiom, [])).
fof(f243, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ sP11), inference(rectify, [], [f242])).
fof(f242, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X12] : (~ (init = a_select2(s_center7_init, X12)) & leq(X12, n2) & leq(n0, X12)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_sworst7))) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init)) & ~ leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | ~ sP11), inference(nnf_transformation, [], [f177])).
fof(f815, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | spl58_38 | spl58_19), inference(avatar_split_clause, [], [f603, f715, f805, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f603, plain, (gt(loopcounter, n1) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f602])).
fof(f602, plain, (gt(loopcounter, n1) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f536])).
fof(f536, plain, (gt(loopcounter, n1) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f409, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f409, plain, (gt(loopcounter, n1) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f814, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | ~ spl58_37 | spl58_19), inference(avatar_split_clause, [], [f605, f715, f800, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f605, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f604])).
fof(f604, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f535])).
fof(f535, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f410, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f410, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f813, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | spl58_39 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f607, f700, f696, f692, f810, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f607, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f606])).
fof(f606, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f534])).
fof(f534, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f411, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f411, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK46) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f808, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | spl58_38 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f609, f700, f696, f692, f805, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f609, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f608])).
fof(f608, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f533])).
fof(f533, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f412, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f412, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK46, n2) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f803, plain, (~ spl58_4 | ~ spl58_32 | ~ spl58_33 | ~ spl58_3 | spl58_34 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_35 | spl58_36 | ~ spl58_37 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f611, f700, f696, f692, f800, f796, f792, f676, f672, f668, f664, f660, f656, f788, f648, f784, f780, f652])).
fof(f611, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(trivial_inequality_removal, [], [f610])).
fof(f610, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ sP11), inference(duplicate_literal_removal, [], [f532])).
fof(f532, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP8 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_sworst7)) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP11), inference(definition_unfolding, [], [f413, f447, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f413, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_center7_init, sK46)) | sP7 | sP6 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP8 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_sworst7)) | ~ (init = s_sworst7_init) | ~ (init = init) | ~ sP11), inference(cnf_transformation, [], [f245])).
fof(f778, plain, (~ spl58_26 | spl58_31), inference(avatar_split_clause, [], [f401, f775, f751])).
fof(f401, plain, (leq(n0, sK44) | ~ sP12), inference(cnf_transformation, [], [f241])).
fof(f241, plain, (((~ (init = a_select3(simplex7_init, sK45, sK44)) & leq(sK45, n3) & leq(n0, sK45)) & leq(sK44, n2) & leq(n0, sK44)) | ~ sP12), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44, sK45])], [f238, f240, f239])).
fof(f239, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK44)) & leq(X1, n3) & leq(n0, X1)) & leq(sK44, n2) & leq(n0, sK44))), introduced(choice_axiom, [])).
fof(f240, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK44)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK45, sK44)) & leq(sK45, n3) & leq(n0, sK45))), introduced(choice_axiom, [])).
fof(f238, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP12), inference(rectify, [], [f237])).
fof(f237, plain, (? [X6] : (? [X7] : (~ (init = a_select3(simplex7_init, X7, X6)) & leq(X7, n3) & leq(n0, X7)) & leq(X6, n2) & leq(n0, X6)) | ~ sP12), inference(nnf_transformation, [], [f178])).
fof(f773, plain, (~ spl58_26 | spl58_30), inference(avatar_split_clause, [], [f402, f770, f751])).
fof(f402, plain, (leq(sK44, n2) | ~ sP12), inference(cnf_transformation, [], [f241])).
fof(f768, plain, (~ spl58_26 | spl58_29), inference(avatar_split_clause, [], [f403, f765, f751])).
fof(f403, plain, (leq(n0, sK45) | ~ sP12), inference(cnf_transformation, [], [f241])).
fof(f763, plain, (~ spl58_26 | spl58_28), inference(avatar_split_clause, [], [f404, f760, f751])).
fof(f404, plain, (leq(sK45, n3) | ~ sP12), inference(cnf_transformation, [], [f241])).
fof(f758, plain, (~ spl58_26 | ~ spl58_27), inference(avatar_split_clause, [], [f531, f755, f751])).
fof(f531, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK45, sK44)) | ~ sP12), inference(definition_unfolding, [], [f405, f447])).
fof(f405, plain, (~ (init = a_select3(simplex7_init, sK45, sK44)) | ~ sP12), inference(cnf_transformation, [], [f241])).
fof(f749, plain, (~ spl58_22 | spl58_25), inference(avatar_split_clause, [], [f398, f746, f732])).
fof(f398, plain, (leq(n0, sK43) | ~ sP13), inference(cnf_transformation, [], [f236])).
fof(f236, plain, ((~ (init = a_select2(s_values7_init, sK43)) & leq(sK43, n3) & leq(n0, sK43)) | ~ sP13), inference(skolemisation, [status(esa), new_symbols(skolem, [sK43])], [f234, f235])).
fof(f235, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK43)) & leq(sK43, n3) & leq(n0, sK43))), introduced(choice_axiom, [])).
fof(f234, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP13), inference(rectify, [], [f233])).
fof(f233, plain, (? [X5] : (~ (init = a_select2(s_values7_init, X5)) & leq(X5, n3) & leq(n0, X5)) | ~ sP13), inference(nnf_transformation, [], [f179])).
fof(f744, plain, (~ spl58_22 | spl58_24), inference(avatar_split_clause, [], [f399, f741, f732])).
fof(f399, plain, (leq(sK43, n3) | ~ sP13), inference(cnf_transformation, [], [f236])).
fof(f739, plain, (~ spl58_22 | ~ spl58_23), inference(avatar_split_clause, [], [f530, f736, f732])).
fof(f530, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK43)) | ~ sP13), inference(definition_unfolding, [], [f400, f447])).
fof(f400, plain, (~ (init = a_select2(s_values7_init, sK43)) | ~ sP13), inference(cnf_transformation, [], [f236])).
fof(f720, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | spl58_18 | spl58_19), inference(avatar_split_clause, [], [f615, f715, f710, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f615, plain, (gt(loopcounter, n1) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f614])).
fof(f614, plain, (gt(loopcounter, n1) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f528])).
fof(f528, plain, (gt(loopcounter, n1) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f392, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f392, plain, (gt(loopcounter, n1) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).
fof(f232, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_center7_init, sK42)) & leq(sK42, n2) & leq(n0, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ sP14), inference(skolemisation, [status(esa), new_symbols(skolem, [sK42])], [f230, f231])).
fof(f231, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK42)) & leq(sK42, n2) & leq(n0, sK42))), introduced(choice_axiom, [])).
fof(f230, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ sP14), inference(rectify, [], [f229])).
fof(f229, plain, (((((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X8] : (~ (init = a_select2(s_center7_init, X8)) & leq(X8, n2) & leq(n0, X8)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init)) & leq(a_select2(s_values7, pv1388), a_select2(s_values7, s_worst7))) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init)) & ~ gt(a_select2(s_values7, pv1388), a_select2(s_values7, s_best7))) | ~ sP14), inference(nnf_transformation, [], [f180])).
fof(f719, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | spl58_17 | spl58_19), inference(avatar_split_clause, [], [f617, f715, f705, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f617, plain, (gt(loopcounter, n1) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f616])).
fof(f616, plain, (gt(loopcounter, n1) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f527])).
fof(f527, plain, (gt(loopcounter, n1) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f393, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f393, plain, (gt(loopcounter, n1) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).
fof(f718, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | ~ spl58_13 | spl58_19), inference(avatar_split_clause, [], [f619, f715, f688, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f619, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f618])).
fof(f618, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f526])).
fof(f526, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f394, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f394, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).
fof(f713, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | spl58_18 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f621, f700, f696, f692, f710, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f621, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f620])).
fof(f620, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f525])).
fof(f525, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f395, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f395, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK42) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).
fof(f708, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | spl58_17 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f623, f700, f696, f692, f705, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f623, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f622])).
fof(f622, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f524])).
fof(f524, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f396, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f396, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK42, n2) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).
fof(f703, plain, (~ spl58_1 | ~ spl58_2 | ~ spl58_3 | spl58_4 | ~ spl58_5 | ~ spl58_6 | ~ spl58_7 | ~ spl58_8 | ~ spl58_9 | ~ spl58_10 | spl58_11 | spl58_12 | ~ spl58_13 | ~ spl58_14 | ~ spl58_15 | ~ spl58_16), inference(avatar_split_clause, [], [f625, f700, f696, f692, f688, f684, f680, f676, f672, f668, f664, f660, f656, f652, f648, f644, f640])).
fof(f625, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(trivial_inequality_removal, [], [f624])).
fof(f624, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ sP14), inference(duplicate_literal_removal, [], [f523])).
fof(f523, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | sP11 | ~ (s_worst7_init = a_select2(s_values7_init, pv1388)) | ~ (s_worst7_init = a_select2(s_values7_init, s_worst7)) | ~ (s_worst7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init) | ~ sP14), inference(definition_unfolding, [], [f397, f447, f447, f447, f447, f447, f501, f447, f447, f447, f447, f447, f447, f447, f447])).
fof(f397, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_center7_init, sK42)) | sP10 | sP9 | ~ leq(pv1388, n3) | ~ leq(s_worst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1388) | ~ leq(n0, s_worst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (s_best7_init = init) | ~ (init = init) | sP11 | ~ (init = a_select2(s_values7_init, pv1388)) | ~ (init = a_select2(s_values7_init, s_worst7)) | ~ (init = s_worst7_init) | ~ (init = init) | ~ sP14), inference(cnf_transformation, [], [f232])).