fof(f576622, plain, $false, inference(avatar_sat_refutation, [], [f512, f517, f522, f531, f536, f541, f550, f555, f560, f565, f570, f627, f632, f637, f642, f643, f644, f645, f646, f647, f648, f649, f650, f651, f652, f653, f654, f655, f656, f657, f1798, f1853, f2200, f575517, f575556, f575869, f576121, f576588])).
fof(f576588, plain, (spl39_2 | ~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16), inference(avatar_contradiction_clause, [], [f576587])).
fof(f576587, plain, ($false | (spl39_2 | ~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f576586, f576350])).
fof(f576350, plain, (~ (s_best7_init = a_select2(s_center7_init, sK34)) | (spl39_2 | ~ spl39_15)), inference(forward_demodulation, [], [f511, f573])).
fof(f573, plain, ((s_best7_init = s_worst7_init) | ~ spl39_15), inference(avatar_component_clause, [], [f572])).
fof(f572, plain, (spl39_15 <=> (s_best7_init = s_worst7_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_15])])).
fof(f511, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK34)) | spl39_2), inference(avatar_component_clause, [], [f509])).
fof(f509, plain, (spl39_2 <=> (s_worst7_init = a_select2(s_center7_init, sK34))), introduced(avatar_definition, [new_symbols(naming, [spl39_2])])).
fof(f576586, plain, ((s_best7_init = a_select2(s_center7_init, sK34)) | (~ spl39_3 | ~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f576481, f516])).
fof(f516, plain, (leq(sK34, n2) | ~ spl39_3), inference(avatar_component_clause, [], [f514])).
fof(f514, plain, (spl39_3 <=> leq(sK34, n2)), introduced(avatar_definition, [new_symbols(naming, [spl39_3])])).
fof(f576481, plain, (~ leq(sK34, n2) | (s_best7_init = a_select2(s_center7_init, sK34)) | (~ spl39_4 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f521, f677])).
fof(f677, plain, (! [X2] : (~ leq(n0, X2) | ~ leq(X2, n2) | (s_best7_init = a_select2(s_center7_init, X2))) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f673, f675])).
fof(f675, plain, ((s_best7_init = s_sworst7_init) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f577, f573])).
fof(f577, plain, ((s_sworst7_init = s_worst7_init) | ~ spl39_16), inference(avatar_component_clause, [], [f576])).
fof(f576, plain, (spl39_16 <=> (s_sworst7_init = s_worst7_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_16])])).
fof(f673, plain, (! [X2] : ((s_sworst7_init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) | ~ spl39_16), inference(backward_demodulation, [], [f465, f577])).
fof(f465, plain, ! [X2] : ((s_worst7_init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)), inference(definition_unfolding, [], [f381, f368])).
fof(f368, plain, (init = s_worst7_init), inference(cnf_transformation, [], [f243])).
fof(f243, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | (~ (init = a_select2(s_try7_init, sK38)) & leq(sK38, minus(n3, n1)) & leq(n0, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) & ! [X4] : (! [X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5)) | ~ leq(X4, n2) | ~ leq(n0, X4)) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38])], [f241, f242])).
fof(f242, plain, (? [X0] : (~ (init = a_select2(s_try7_init, X0)) & leq(X0, minus(n3, n1)) & leq(n0, X0)) => (~ (init = a_select2(s_try7_init, sK38)) & leq(sK38, minus(n3, n1)) & leq(n0, sK38))), introduced(choice_axiom, [])).
fof(f241, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X0] : (~ (init = a_select2(s_try7_init, X0)) & leq(X0, minus(n3, n1)) & leq(n0, X0)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)) & ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) & ! [X4] : (! [X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5)) | ~ leq(X4, n2) | ~ leq(n0, X4)) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(rectify, [], [f180])).
fof(f180, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & leq(X5, minus(n3, n1)) & leq(n0, X5)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | ~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(definition_folding, [], [f157, e179, e178, e177])).
fof(f177, plain, (? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ sP4), inference(usedef, [], [e177])).
fof(e177, plain, (sP4 <=> ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f178, plain, (? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ~ sP5), inference(usedef, [], [e178])).
fof(e178, plain, (sP5 <=> ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f179, plain, (? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ~ sP6), inference(usedef, [], [e179])).
fof(e179, plain, (sP6 <=> ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f157, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & leq(X5, minus(n3, n1)) & leq(n0, X5)) | ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & (((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | ~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0)) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | ~ leq(X1, n2) | ~ leq(n0, X1)) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | ~ leq(X2, n3) | ~ leq(n0, X2)) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | ~ leq(X4, n3) | ~ leq(n0, X4)) | ~ leq(X3, n2) | ~ leq(n0, X3)) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)), inference(flattening, [], [f156])).
fof(f156, plain, ((((~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init)) & gt(loopcounter, n1)) | ? [X5] : (~ (init = a_select2(s_try7_init, X5)) & (leq(X5, minus(n3, n1)) & leq(n0, X5))) | ? [X6] : (~ (init = a_select2(s_center7_init, X6)) & (leq(X6, n2) & leq(n0, X6))) | ? [X7] : (~ (init = a_select2(s_values7_init, X7)) & (leq(X7, n3) & leq(n0, X7))) | ? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & (leq(X9, n3) & leq(n0, X9))) & (leq(X8, n2) & leq(n0, X8))) | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)) & ((((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init)) | ~ gt(loopcounter, n1)) & ! [X0] : ((init = a_select2(s_try7_init, X0)) | (~ leq(X0, minus(n3, n1)) | ~ leq(n0, X0))) & ! [X1] : ((init = a_select2(s_center7_init, X1)) | (~ leq(X1, n2) | ~ leq(n0, X1))) & ! [X2] : ((init = a_select2(s_values7_init, X2)) | (~ leq(X2, n3) | ~ leq(n0, X2))) & ! [X3] : (! [X4] : ((init = a_select3(simplex7_init, X4, X3)) | (~ leq(X4, n3) | ~ leq(n0, X4))) | (~ leq(X3, n2) | ~ leq(n0, X3))) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X0] : ((leq(X0, minus(n3, n1)) & leq(n0, X0)) => (init = a_select2(s_try7_init, X0))) & ! [X1] : ((leq(X1, n2) & leq(n0, X1)) => (init = a_select2(s_center7_init, X1))) & ! [X2] : ((leq(X2, n3) & leq(n0, X2)) => (init = a_select2(s_values7_init, X2))) & ! [X3] : ((leq(X3, n2) & leq(n0, X3)) => ! [X4] : ((leq(X4, n3) & leq(n0, X4)) => (init = a_select3(simplex7_init, X4, X3)))) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X5] : ((leq(X5, minus(n3, n1)) & leq(n0, X5)) => (init = a_select2(s_try7_init, X5))) & ! [X6] : ((leq(X6, n2) & leq(n0, X6)) => (init = a_select2(s_center7_init, X6))) & ! [X7] : ((leq(X7, n3) & leq(n0, X7)) => (init = a_select2(s_values7_init, X7))) & ! [X8] : ((leq(X8, n2) & leq(n0, X8)) => ! [X9] : ((leq(X9, n3) & leq(n0, X9)) => (init = a_select3(simplex7_init, X9, X8)))) & leq(pv1376, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(rectify, [], [f54])).
fof(f54, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X20] : ((leq(X20, minus(n3, n1)) & leq(n0, X20)) => (init = a_select2(s_try7_init, X20))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X10] : ((leq(X10, minus(n3, n1)) & leq(n0, X10)) => (init = a_select2(s_try7_init, X10))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => (init = a_select2(s_center7_init, X4))) & ! [X28] : ((leq(X28, n3) & leq(n0, X28)) => (init = a_select2(s_values7_init, X28))) & ! [X21] : ((leq(X21, n2) & leq(n0, X21)) => ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select3(simplex7_init, X27, X21)))) & leq(pv1376, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), inference(negated_conjecture, [], [f53])).
fof(f53, plain, ~ (((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X20] : ((leq(X20, minus(n3, n1)) & leq(n0, X20)) => (init = a_select2(s_try7_init, X20))) & ! [X19] : ((leq(X19, n2) & leq(n0, X19)) => (init = a_select2(s_center7_init, X19))) & ! [X3] : ((leq(X3, n3) & leq(n0, X3)) => (init = a_select2(s_values7_init, X3))) & ! [X13] : ((leq(X13, n2) & leq(n0, X13)) => ! [X17] : ((leq(X17, n3) & leq(n0, X17)) => (init = a_select3(simplex7_init, X17, X13)))) & leq(pv1376, n3) & leq(pv19, minus(n410, n1)) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, pv19) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init)) => ((gt(loopcounter, n1) => ((init = pvar1402_init) & (init = pvar1401_init) & (init = pvar1400_init))) & ! [X10] : ((leq(X10, minus(n3, n1)) & leq(n0, X10)) => (init = a_select2(s_try7_init, X10))) & ! [X4] : ((leq(X4, n2) & leq(n0, X4)) => (init = a_select2(s_center7_init, X4))) & ! [X28] : ((leq(X28, n3) & leq(n0, X28)) => (init = a_select2(s_values7_init, X28))) & ! [X21] : ((leq(X21, n2) & leq(n0, X21)) => ! [X27] : ((leq(X27, n3) & leq(n0, X27)) => (init = a_select3(simplex7_init, X27, X21)))) & leq(pv1376, n3) & leq(s_worst7, n3) & leq(s_sworst7, n3) & leq(s_best7, n3) & leq(n0, pv1376) & leq(n0, s_worst7) & leq(n0, s_sworst7) & leq(n0, s_best7) & (init = s_worst7_init) & (init = s_sworst7_init) & (init = s_best7_init) & (init = init))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gauss_init_0025)).
fof(f381, plain, ! [X2] : ((init = a_select2(s_center7_init, X2)) | ~ leq(X2, n2) | ~ leq(n0, X2)), inference(cnf_transformation, [], [f243])).
fof(f521, plain, (leq(n0, sK34) | ~ spl39_4), inference(avatar_component_clause, [], [f519])).
fof(f519, plain, (spl39_4 <=> leq(n0, sK34)), introduced(avatar_definition, [new_symbols(naming, [spl39_4])])).
fof(f576121, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_65), inference(avatar_contradiction_clause, [], [f576120])).
fof(f576120, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_65)), inference(subsumption_resolution, [], [f575518, f576118])).
fof(f576118, plain, (~ (s_best7_init = a_select3(simplex7_init, n0, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_65)), inference(forward_demodulation, [], [f1104, f1153])).
fof(f1153, plain, ((n0 = sK37) | ~ spl39_65), inference(avatar_component_clause, [], [f1151])).
fof(f1151, plain, (spl39_65 <=> (n0 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_65])])).
fof(f1104, plain, (~ (s_best7_init = a_select3(simplex7_init, sK37, sK36)) | (spl39_10 | ~ spl39_15)), inference(forward_demodulation, [], [f549, f573])).
fof(f549, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK37, sK36)) | spl39_10), inference(avatar_component_clause, [], [f547])).
fof(f547, plain, (spl39_10 <=> (s_worst7_init = a_select3(simplex7_init, sK37, sK36))), introduced(avatar_definition, [new_symbols(naming, [spl39_10])])).
fof(f575518, plain, ((s_best7_init = a_select3(simplex7_init, n0, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f575420, f564])).
fof(f564, plain, (leq(sK36, n2) | ~ spl39_13), inference(avatar_component_clause, [], [f562])).
fof(f562, plain, (spl39_13 <=> leq(sK36, n2)), introduced(avatar_definition, [new_symbols(naming, [spl39_13])])).
fof(f575420, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n0, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f569, f2141])).
fof(f2141, plain, (! [X40] : (~ leq(n0, X40) | ~ leq(X40, n2) | (s_best7_init = a_select3(simplex7_init, n0, X40))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2129, f688])).
fof(f688, plain, leq(n0, n3), inference(resolution, [], [f249, f407])).
fof(f407, plain, gt(n3, n0), inference(cnf_transformation, [], [f70])).
fof(f70, plain, gt(n3, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gt_3_0)).
fof(f249, plain, ! [X0, X1] : (~ gt(X1, X0) | leq(X0, X1)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, ! [X0, X1] : (leq(X0, X1) | ~ gt(X1, X0)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (gt(X1, X0) => leq(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', leq_gt1)).
fof(f2129, plain, (! [X40] : (~ leq(n0, n3) | (s_best7_init = a_select3(simplex7_init, n0, X40)) | ~ leq(X40, n2) | ~ leq(n0, X40)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f679, f247])).
fof(f247, plain, ! [X0] : leq(X0, X0), inference(cnf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : leq(X0, X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', reflexivity_leq)).
fof(f679, plain, (! [X4, X5] : (~ leq(n0, X5) | ~ leq(X5, n3) | (s_best7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X4, n2) | ~ leq(n0, X4)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f671, f675])).
fof(f671, plain, (! [X4, X5] : ((s_sworst7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)) | ~ spl39_16), inference(backward_demodulation, [], [f467, f577])).
fof(f467, plain, ! [X4, X5] : ((s_worst7_init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)), inference(definition_unfolding, [], [f379, f368])).
fof(f379, plain, ! [X4, X5] : ((init = a_select3(simplex7_init, X5, X4)) | ~ leq(X5, n3) | ~ leq(n0, X5) | ~ leq(X4, n2) | ~ leq(n0, X4)), inference(cnf_transformation, [], [f243])).
fof(f569, plain, (leq(n0, sK36) | ~ spl39_14), inference(avatar_component_clause, [], [f567])).
fof(f567, plain, (spl39_14 <=> leq(n0, sK36)), introduced(avatar_definition, [new_symbols(naming, [spl39_14])])).
fof(f575869, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_152), inference(avatar_contradiction_clause, [], [f575868])).
fof(f575868, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_152)), inference(subsumption_resolution, [], [f575867, f575514])).
fof(f575514, plain, ((s_best7_init = a_select3(simplex7_init, n2, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f575418, f564])).
fof(f575418, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n2, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f569, f2134])).
fof(f2134, plain, (! [X4] : (~ leq(n0, X4) | ~ leq(X4, n2) | (s_best7_init = a_select3(simplex7_init, n2, X4))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2109, f704])).
fof(f704, plain, leq(n2, n3), inference(resolution, [], [f249, f416])).
fof(f416, plain, gt(n3, n2), inference(cnf_transformation, [], [f79])).
fof(f79, plain, gt(n3, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gt_3_2)).
fof(f2109, plain, (! [X4] : (~ leq(n2, n3) | (s_best7_init = a_select3(simplex7_init, n2, X4)) | ~ leq(X4, n2) | ~ leq(n0, X4)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f679, f687])).
fof(f687, plain, leq(n0, n2), inference(resolution, [], [f249, f406])).
fof(f406, plain, gt(n2, n0), inference(cnf_transformation, [], [f69])).
fof(f69, plain, gt(n2, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gt_2_0)).
fof(f575867, plain, (~ (s_best7_init = a_select3(simplex7_init, n2, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_152)), inference(forward_demodulation, [], [f1104, f2199])).
fof(f2199, plain, ((n2 = sK37) | ~ spl39_152), inference(avatar_component_clause, [], [f2197])).
fof(f2197, plain, (spl39_152 <=> (n2 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_152])])).
fof(f575556, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_143), inference(avatar_contradiction_clause, [], [f575555])).
fof(f575555, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_143)), inference(subsumption_resolution, [], [f575543, f575513])).
fof(f575513, plain, ((s_best7_init = a_select3(simplex7_init, n1, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f575417, f564])).
fof(f575417, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n1, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f569, f2133])).
fof(f2133, plain, (! [X3] : (~ leq(n0, X3) | ~ leq(X3, n2) | (s_best7_init = a_select3(simplex7_init, n1, X3))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2108, f700])).
fof(f700, plain, leq(n1, n3), inference(resolution, [], [f249, f412])).
fof(f412, plain, gt(n3, n1), inference(cnf_transformation, [], [f75])).
fof(f75, plain, gt(n3, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gt_3_1)).
fof(f2108, plain, (! [X3] : (~ leq(n1, n3) | (s_best7_init = a_select3(simplex7_init, n1, X3)) | ~ leq(X3, n2) | ~ leq(n0, X3)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f679, f686])).
fof(f686, plain, leq(n0, n1), inference(resolution, [], [f249, f405])).
fof(f405, plain, gt(n1, n0), inference(cnf_transformation, [], [f68])).
fof(f68, plain, gt(n1, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', gt_1_0)).
fof(f575543, plain, (~ (s_best7_init = a_select3(simplex7_init, n1, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_143)), inference(backward_demodulation, [], [f1104, f1987])).
fof(f1987, plain, ((n1 = sK37) | ~ spl39_143), inference(avatar_component_clause, [], [f1985])).
fof(f1985, plain, (spl39_143 <=> (n1 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_143])])).
fof(f575517, plain, (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_59), inference(avatar_contradiction_clause, [], [f575516])).
fof(f575516, plain, ($false | (spl39_10 | ~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16 | ~ spl39_59)), inference(subsumption_resolution, [], [f575515, f575341])).
fof(f575341, plain, (~ (s_best7_init = a_select3(simplex7_init, n3, sK36)) | (spl39_10 | ~ spl39_15 | ~ spl39_59)), inference(forward_demodulation, [], [f1104, f1114])).
fof(f1114, plain, ((n3 = sK37) | ~ spl39_59), inference(avatar_component_clause, [], [f1112])).
fof(f1112, plain, (spl39_59 <=> (n3 = sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_59])])).
fof(f575515, plain, ((s_best7_init = a_select3(simplex7_init, n3, sK36)) | (~ spl39_13 | ~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f575419, f564])).
fof(f575419, plain, (~ leq(sK36, n2) | (s_best7_init = a_select3(simplex7_init, n3, sK36)) | (~ spl39_14 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f569, f2135])).
fof(f2135, plain, (! [X5] : (~ leq(n0, X5) | ~ leq(X5, n2) | (s_best7_init = a_select3(simplex7_init, n3, X5))) | (~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f2110, f247])).
fof(f2110, plain, (! [X5] : (~ leq(n3, n3) | (s_best7_init = a_select3(simplex7_init, n3, X5)) | ~ leq(X5, n2) | ~ leq(n0, X5)) | (~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f679, f688])).
fof(f2200, plain, (spl39_59 | spl39_65 | spl39_143 | spl39_152 | ~ spl39_11 | ~ spl39_12), inference(avatar_split_clause, [], [f2195, f557, f552, f2197, f1985, f1151, f1112])).
fof(f552, plain, (spl39_11 <=> leq(sK37, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_11])])).
fof(f557, plain, (spl39_12 <=> leq(n0, sK37)), introduced(avatar_definition, [new_symbols(naming, [spl39_12])])).
fof(f2195, plain, ((n2 = sK37) | (n1 = sK37) | (n0 = sK37) | (n3 = sK37) | (~ spl39_11 | ~ spl39_12)), inference(subsumption_resolution, [], [f2165, f554])).
fof(f554, plain, (leq(sK37, n3) | ~ spl39_11), inference(avatar_component_clause, [], [f552])).
fof(f2165, plain, ((n2 = sK37) | (n1 = sK37) | (n0 = sK37) | ~ leq(sK37, n3) | (n3 = sK37) | ~ spl39_12), inference(resolution, [], [f425, f559])).
fof(f559, plain, (leq(n0, sK37) | ~ spl39_12), inference(avatar_component_clause, [], [f557])).
fof(f425, plain, ! [X0] : (~ leq(n0, X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | (n3 = X0)), inference(cnf_transformation, [], [f169])).
fof(f169, plain, ! [X0] : ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0) | ~ leq(X0, n3) | ~ leq(n0, X0)), inference(flattening, [], [f168])).
fof(f168, plain, ! [X0] : (((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0)) | (~ leq(X0, n3) | ~ leq(n0, X0))), inference(ennf_transformation, [], [f88])).
fof(f88, plain, ! [X0] : ((leq(X0, n3) & leq(n0, X0)) => ((n3 = X0) | (n2 = X0) | (n1 = X0) | (n0 = X0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', finite_domain_3)).
fof(f1853, plain, (~ spl39_15 | ~ spl39_16 | spl39_25 | ~ spl39_29 | ~ spl39_30), inference(avatar_contradiction_clause, [], [f1852])).
fof(f1852, plain, ($false | (~ spl39_15 | ~ spl39_16 | spl39_25 | ~ spl39_29 | ~ spl39_30)), inference(subsumption_resolution, [], [f1851, f1116])).
fof(f1116, plain, (leq(sK38, n2) | ~ spl39_29), inference(forward_demodulation, [], [f631, f730])).
fof(f730, plain, (n2 = minus(n3, n1)), inference(superposition, [], [f447, f670])).
fof(f670, plain, (n3 = plus(n2, n1)), inference(forward_demodulation, [], [f669, f665])).
fof(f665, plain, (n2 = plus(n1, n1)), inference(forward_demodulation, [], [f474, f473])).
fof(f473, plain, (n1 = plus(n0, n1)), inference(definition_unfolding, [], [f428, f324])).
fof(f324, plain, ! [X0] : (succ(X0) = plus(X0, n1)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (succ(X0) = plus(X0, n1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', succ_plus_1_r)).
fof(f428, plain, (n1 = succ(n0)), inference(cnf_transformation, [], [f91])).
fof(f91, plain, (n1 = succ(n0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', successor_1)).
fof(f474, plain, (n2 = plus(plus(n0, n1), n1)), inference(definition_unfolding, [], [f429, f324, f324])).
fof(f429, plain, (n2 = succ(succ(n0))), inference(cnf_transformation, [], [f92])).
fof(f92, plain, (n2 = succ(succ(n0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', successor_2)).
fof(f669, plain, (n3 = plus(plus(n1, n1), n1)), inference(forward_demodulation, [], [f668, f473])).
fof(f668, plain, (n3 = plus(plus(n1, plus(n0, n1)), n1)), inference(forward_demodulation, [], [f475, f438])).
fof(f438, plain, ! [X0] : (plus(X0, n1) = plus(n1, X0)), inference(definition_unfolding, [], [f325, f324])).
fof(f325, plain, ! [X0] : (succ(X0) = plus(n1, X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (succ(X0) = plus(n1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', succ_plus_1_l)).
fof(f475, plain, (n3 = plus(plus(plus(n0, n1), n1), n1)), inference(definition_unfolding, [], [f430, f324, f324, f324])).
fof(f430, plain, (n3 = succ(succ(succ(n0)))), inference(cnf_transformation, [], [f93])).
fof(f93, plain, (n3 = succ(succ(succ(n0)))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', successor_3)).
fof(f447, plain, ! [X0] : (minus(plus(X0, n1), n1) = X0), inference(definition_unfolding, [], [f335, f334, f324])).
fof(f334, plain, ! [X0] : (minus(X0, n1) = pred(X0)), inference(cnf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (minus(X0, n1) = pred(X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', pred_minus_1)).
fof(f335, plain, ! [X0] : (pred(succ(X0)) = X0), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (pred(succ(X0)) = X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV028+1.p', pred_succ)).
fof(f631, plain, (leq(sK38, minus(n3, n1)) | ~ spl39_29), inference(avatar_component_clause, [], [f629])).
fof(f629, plain, (spl39_29 <=> leq(sK38, minus(n3, n1))), introduced(avatar_definition, [new_symbols(naming, [spl39_29])])).
fof(f1851, plain, (~ leq(sK38, n2) | (~ spl39_15 | ~ spl39_16 | spl39_25 | ~ spl39_30)), inference(subsumption_resolution, [], [f1844, f1117])).
fof(f1117, plain, (~ (s_best7_init = a_select2(s_try7_init, sK38)) | (~ spl39_15 | spl39_25)), inference(forward_demodulation, [], [f614, f573])).
fof(f614, plain, (~ (s_worst7_init = a_select2(s_try7_init, sK38)) | spl39_25), inference(avatar_component_clause, [], [f612])).
fof(f612, plain, (spl39_25 <=> (s_worst7_init = a_select2(s_try7_init, sK38))), introduced(avatar_definition, [new_symbols(naming, [spl39_25])])).
fof(f1844, plain, ((s_best7_init = a_select2(s_try7_init, sK38)) | ~ leq(sK38, n2) | (~ spl39_15 | ~ spl39_16 | ~ spl39_30)), inference(resolution, [], [f733, f636])).
fof(f636, plain, (leq(n0, sK38) | ~ spl39_30), inference(avatar_component_clause, [], [f634])).
fof(f634, plain, (spl39_30 <=> leq(n0, sK38)), introduced(avatar_definition, [new_symbols(naming, [spl39_30])])).
fof(f733, plain, (! [X1] : (~ leq(n0, X1) | (s_best7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, n2)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f676, f730])).
fof(f676, plain, (! [X1] : ((s_best7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f674, f675])).
fof(f674, plain, (! [X1] : ((s_sworst7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)) | ~ spl39_16), inference(backward_demodulation, [], [f464, f577])).
fof(f464, plain, ! [X1] : ((s_worst7_init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)), inference(definition_unfolding, [], [f382, f368])).
fof(f382, plain, ! [X1] : ((init = a_select2(s_try7_init, X1)) | ~ leq(X1, minus(n3, n1)) | ~ leq(n0, X1)), inference(cnf_transformation, [], [f243])).
fof(f1798, plain, (spl39_6 | ~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16), inference(avatar_contradiction_clause, [], [f1797])).
fof(f1797, plain, ($false | (spl39_6 | ~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1796, f680])).
fof(f680, plain, (~ (s_best7_init = a_select2(s_values7_init, sK35)) | (spl39_6 | ~ spl39_15)), inference(forward_demodulation, [], [f530, f573])).
fof(f530, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK35)) | spl39_6), inference(avatar_component_clause, [], [f528])).
fof(f528, plain, (spl39_6 <=> (s_worst7_init = a_select2(s_values7_init, sK35))), introduced(avatar_definition, [new_symbols(naming, [spl39_6])])).
fof(f1796, plain, ((s_best7_init = a_select2(s_values7_init, sK35)) | (~ spl39_7 | ~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(subsumption_resolution, [], [f1747, f535])).
fof(f535, plain, (leq(sK35, n3) | ~ spl39_7), inference(avatar_component_clause, [], [f533])).
fof(f533, plain, (spl39_7 <=> leq(sK35, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_7])])).
fof(f1747, plain, (~ leq(sK35, n3) | (s_best7_init = a_select2(s_values7_init, sK35)) | (~ spl39_8 | ~ spl39_15 | ~ spl39_16)), inference(resolution, [], [f678, f540])).
fof(f540, plain, (leq(n0, sK35) | ~ spl39_8), inference(avatar_component_clause, [], [f538])).
fof(f538, plain, (spl39_8 <=> leq(n0, sK35)), introduced(avatar_definition, [new_symbols(naming, [spl39_8])])).
fof(f678, plain, (! [X3] : (~ leq(n0, X3) | ~ leq(X3, n3) | (s_best7_init = a_select2(s_values7_init, X3))) | (~ spl39_15 | ~ spl39_16)), inference(backward_demodulation, [], [f672, f675])).
fof(f672, plain, (! [X3] : ((s_sworst7_init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)) | ~ spl39_16), inference(backward_demodulation, [], [f466, f577])).
fof(f466, plain, ! [X3] : ((s_worst7_init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)), inference(definition_unfolding, [], [f380, f368])).
fof(f380, plain, ! [X3] : ((init = a_select2(s_values7_init, X3)) | ~ leq(X3, n3) | ~ leq(n0, X3)), inference(cnf_transformation, [], [f243])).
fof(f657, plain, spl39_15, inference(avatar_split_clause, [], [f469, f572])).
fof(f469, plain, (s_best7_init = s_worst7_init), inference(definition_unfolding, [], [f366, f368])).
fof(f366, plain, (init = s_best7_init), inference(cnf_transformation, [], [f243])).
fof(f656, plain, spl39_16, inference(avatar_split_clause, [], [f468, f576])).
fof(f468, plain, (s_sworst7_init = s_worst7_init), inference(definition_unfolding, [], [f367, f368])).
fof(f367, plain, (init = s_sworst7_init), inference(cnf_transformation, [], [f243])).
fof(f655, plain, spl39_17, inference(avatar_split_clause, [], [f369, f580])).
fof(f580, plain, (spl39_17 <=> leq(n0, s_best7)), introduced(avatar_definition, [new_symbols(naming, [spl39_17])])).
fof(f369, plain, leq(n0, s_best7), inference(cnf_transformation, [], [f243])).
fof(f654, plain, spl39_18, inference(avatar_split_clause, [], [f370, f584])).
fof(f584, plain, (spl39_18 <=> leq(n0, s_sworst7)), introduced(avatar_definition, [new_symbols(naming, [spl39_18])])).
fof(f370, plain, leq(n0, s_sworst7), inference(cnf_transformation, [], [f243])).
fof(f653, plain, spl39_19, inference(avatar_split_clause, [], [f371, f588])).
fof(f588, plain, (spl39_19 <=> leq(n0, s_worst7)), introduced(avatar_definition, [new_symbols(naming, [spl39_19])])).
fof(f371, plain, leq(n0, s_worst7), inference(cnf_transformation, [], [f243])).
fof(f652, plain, spl39_20, inference(avatar_split_clause, [], [f373, f592])).
fof(f592, plain, (spl39_20 <=> leq(n0, pv1376)), introduced(avatar_definition, [new_symbols(naming, [spl39_20])])).
fof(f373, plain, leq(n0, pv1376), inference(cnf_transformation, [], [f243])).
fof(f651, plain, spl39_21, inference(avatar_split_clause, [], [f374, f596])).
fof(f596, plain, (spl39_21 <=> leq(s_best7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_21])])).
fof(f374, plain, leq(s_best7, n3), inference(cnf_transformation, [], [f243])).
fof(f650, plain, spl39_22, inference(avatar_split_clause, [], [f375, f600])).
fof(f600, plain, (spl39_22 <=> leq(s_sworst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_22])])).
fof(f375, plain, leq(s_sworst7, n3), inference(cnf_transformation, [], [f243])).
fof(f649, plain, spl39_23, inference(avatar_split_clause, [], [f376, f604])).
fof(f604, plain, (spl39_23 <=> leq(s_worst7, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_23])])).
fof(f376, plain, leq(s_worst7, n3), inference(cnf_transformation, [], [f243])).
fof(f648, plain, spl39_24, inference(avatar_split_clause, [], [f378, f608])).
fof(f608, plain, (spl39_24 <=> leq(pv1376, n3)), introduced(avatar_definition, [new_symbols(naming, [spl39_24])])).
fof(f378, plain, leq(pv1376, n3), inference(cnf_transformation, [], [f243])).
fof(f647, plain, (~ spl39_31 | spl39_26), inference(avatar_split_clause, [], [f463, f616, f639])).
fof(f639, plain, (spl39_31 <=> gt(loopcounter, n1)), introduced(avatar_definition, [new_symbols(naming, [spl39_31])])).
fof(f616, plain, (spl39_26 <=> (s_worst7_init = pvar1400_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_26])])).
fof(f463, plain, ((s_worst7_init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f383, f368])).
fof(f383, plain, ((init = pvar1400_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f243])).
fof(f646, plain, (~ spl39_31 | spl39_27), inference(avatar_split_clause, [], [f462, f620, f639])).
fof(f620, plain, (spl39_27 <=> (s_worst7_init = pvar1401_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_27])])).
fof(f462, plain, ((s_worst7_init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f384, f368])).
fof(f384, plain, ((init = pvar1401_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f243])).
fof(f645, plain, (~ spl39_31 | spl39_28), inference(avatar_split_clause, [], [f461, f624, f639])).
fof(f624, plain, (spl39_28 <=> (s_worst7_init = pvar1402_init)), introduced(avatar_definition, [new_symbols(naming, [spl39_28])])).
fof(f461, plain, ((s_worst7_init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(definition_unfolding, [], [f385, f368])).
fof(f385, plain, ((init = pvar1402_init) | ~ gt(loopcounter, n1)), inference(cnf_transformation, [], [f243])).
fof(f644, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | spl39_30 | spl39_31), inference(avatar_split_clause, [], [f480, f639, f634, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f543, plain, (spl39_9 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl39_9])])).
fof(f524, plain, (spl39_5 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl39_5])])).
fof(f505, plain, (spl39_1 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl39_1])])).
fof(f480, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f479])).
fof(f479, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f460])).
fof(f460, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f386, f368, f368, f368, f368, f368])).
fof(f386, plain, (gt(loopcounter, n1) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f643, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | spl39_29 | spl39_31), inference(avatar_split_clause, [], [f482, f639, f629, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f482, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f481])).
fof(f481, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f459])).
fof(f459, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f387, f368, f368, f368, f368, f368])).
fof(f387, plain, (gt(loopcounter, n1) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f642, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | ~ spl39_25 | spl39_31), inference(avatar_split_clause, [], [f484, f639, f612, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f484, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f483])).
fof(f483, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f458])).
fof(f458, plain, (gt(loopcounter, n1) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f388, f368, f368, f368, f368, f368, f368])).
fof(f388, plain, (gt(loopcounter, n1) | ~ (init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f637, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | spl39_30 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28), inference(avatar_split_clause, [], [f486, f624, f620, f616, f634, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f486, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f485])).
fof(f485, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f457])).
fof(f457, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f389, f368, f368, f368, f368, f368, f368, f368, f368])).
fof(f389, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(n0, sK38) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f632, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | spl39_29 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28), inference(avatar_split_clause, [], [f488, f624, f620, f616, f629, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f488, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f487])).
fof(f487, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f456])).
fof(f456, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f390, f368, f368, f368, f368, f368, f368, f368, f368])).
fof(f390, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | leq(sK38, minus(n3, n1)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f627, plain, (~ spl39_15 | ~ spl39_16 | ~ spl39_17 | ~ spl39_18 | ~ spl39_19 | ~ spl39_20 | ~ spl39_21 | ~ spl39_22 | ~ spl39_23 | ~ spl39_24 | spl39_9 | spl39_5 | spl39_1 | ~ spl39_25 | ~ spl39_26 | ~ spl39_27 | ~ spl39_28), inference(avatar_split_clause, [], [f490, f624, f620, f616, f612, f505, f524, f543, f608, f604, f600, f596, f592, f588, f584, f580, f576, f572])).
fof(f490, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(trivial_inequality_removal, [], [f489])).
fof(f489, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init)), inference(duplicate_literal_removal, [], [f455])).
fof(f455, plain, (~ (s_worst7_init = pvar1402_init) | ~ (s_worst7_init = pvar1401_init) | ~ (s_worst7_init = pvar1400_init) | ~ (s_worst7_init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (s_worst7_init = s_worst7_init) | ~ (s_sworst7_init = s_worst7_init) | ~ (s_best7_init = s_worst7_init) | ~ (s_worst7_init = s_worst7_init)), inference(definition_unfolding, [], [f391, f368, f368, f368, f368, f368, f368, f368, f368, f368])).
fof(f391, plain, (~ (init = pvar1402_init) | ~ (init = pvar1401_init) | ~ (init = pvar1400_init) | ~ (init = a_select2(s_try7_init, sK38)) | sP6 | sP5 | sP4 | ~ leq(pv1376, n3) | ~ leq(s_worst7, n3) | ~ leq(s_sworst7, n3) | ~ leq(s_best7, n3) | ~ leq(n0, pv1376) | ~ leq(n0, s_worst7) | ~ leq(n0, s_sworst7) | ~ leq(n0, s_best7) | ~ (init = s_worst7_init) | ~ (init = s_sworst7_init) | ~ (init = s_best7_init) | ~ (init = init)), inference(cnf_transformation, [], [f243])).
fof(f570, plain, (~ spl39_9 | spl39_14), inference(avatar_split_clause, [], [f360, f567, f543])).
fof(f360, plain, (leq(n0, sK36) | ~ sP4), inference(cnf_transformation, [], [f240])).
fof(f240, plain, (((~ (init = a_select3(simplex7_init, sK37, sK36)) & leq(sK37, n3) & leq(n0, sK37)) & leq(sK36, n2) & leq(n0, sK36)) | ~ sP4), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f237, f239, f238])).
fof(f238, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) => (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK36)) & leq(X1, n3) & leq(n0, X1)) & leq(sK36, n2) & leq(n0, sK36))), introduced(choice_axiom, [])).
fof(f239, plain, (? [X1] : (~ (init = a_select3(simplex7_init, X1, sK36)) & leq(X1, n3) & leq(n0, X1)) => (~ (init = a_select3(simplex7_init, sK37, sK36)) & leq(sK37, n3) & leq(n0, sK37))), introduced(choice_axiom, [])).
fof(f237, plain, (? [X0] : (? [X1] : (~ (init = a_select3(simplex7_init, X1, X0)) & leq(X1, n3) & leq(n0, X1)) & leq(X0, n2) & leq(n0, X0)) | ~ sP4), inference(rectify, [], [f236])).
fof(f236, plain, (? [X8] : (? [X9] : (~ (init = a_select3(simplex7_init, X9, X8)) & leq(X9, n3) & leq(n0, X9)) & leq(X8, n2) & leq(n0, X8)) | ~ sP4), inference(nnf_transformation, [], [f177])).
fof(f565, plain, (~ spl39_9 | spl39_13), inference(avatar_split_clause, [], [f361, f562, f543])).
fof(f361, plain, (leq(sK36, n2) | ~ sP4), inference(cnf_transformation, [], [f240])).
fof(f560, plain, (~ spl39_9 | spl39_12), inference(avatar_split_clause, [], [f362, f557, f543])).
fof(f362, plain, (leq(n0, sK37) | ~ sP4), inference(cnf_transformation, [], [f240])).
fof(f555, plain, (~ spl39_9 | spl39_11), inference(avatar_split_clause, [], [f363, f552, f543])).
fof(f363, plain, (leq(sK37, n3) | ~ sP4), inference(cnf_transformation, [], [f240])).
fof(f550, plain, (~ spl39_9 | ~ spl39_10), inference(avatar_split_clause, [], [f454, f547, f543])).
fof(f454, plain, (~ (s_worst7_init = a_select3(simplex7_init, sK37, sK36)) | ~ sP4), inference(definition_unfolding, [], [f364, f368])).
fof(f364, plain, (~ (init = a_select3(simplex7_init, sK37, sK36)) | ~ sP4), inference(cnf_transformation, [], [f240])).
fof(f541, plain, (~ spl39_5 | spl39_8), inference(avatar_split_clause, [], [f357, f538, f524])).
fof(f357, plain, (leq(n0, sK35) | ~ sP5), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ((~ (init = a_select2(s_values7_init, sK35)) & leq(sK35, n3) & leq(n0, sK35)) | ~ sP5), inference(skolemisation, [status(esa), new_symbols(skolem, [sK35])], [f233, f234])).
fof(f234, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) => (~ (init = a_select2(s_values7_init, sK35)) & leq(sK35, n3) & leq(n0, sK35))), introduced(choice_axiom, [])).
fof(f233, plain, (? [X0] : (~ (init = a_select2(s_values7_init, X0)) & leq(X0, n3) & leq(n0, X0)) | ~ sP5), inference(rectify, [], [f232])).
fof(f232, plain, (? [X7] : (~ (init = a_select2(s_values7_init, X7)) & leq(X7, n3) & leq(n0, X7)) | ~ sP5), inference(nnf_transformation, [], [f178])).
fof(f536, plain, (~ spl39_5 | spl39_7), inference(avatar_split_clause, [], [f358, f533, f524])).
fof(f358, plain, (leq(sK35, n3) | ~ sP5), inference(cnf_transformation, [], [f235])).
fof(f531, plain, (~ spl39_5 | ~ spl39_6), inference(avatar_split_clause, [], [f453, f528, f524])).
fof(f453, plain, (~ (s_worst7_init = a_select2(s_values7_init, sK35)) | ~ sP5), inference(definition_unfolding, [], [f359, f368])).
fof(f359, plain, (~ (init = a_select2(s_values7_init, sK35)) | ~ sP5), inference(cnf_transformation, [], [f235])).
fof(f522, plain, (~ spl39_1 | spl39_4), inference(avatar_split_clause, [], [f354, f519, f505])).
fof(f354, plain, (leq(n0, sK34) | ~ sP6), inference(cnf_transformation, [], [f231])).
fof(f231, plain, ((~ (init = a_select2(s_center7_init, sK34)) & leq(sK34, n2) & leq(n0, sK34)) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34])], [f229, f230])).
fof(f230, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) => (~ (init = a_select2(s_center7_init, sK34)) & leq(sK34, n2) & leq(n0, sK34))), introduced(choice_axiom, [])).
fof(f229, plain, (? [X0] : (~ (init = a_select2(s_center7_init, X0)) & leq(X0, n2) & leq(n0, X0)) | ~ sP6), inference(rectify, [], [f228])).
fof(f228, plain, (? [X6] : (~ (init = a_select2(s_center7_init, X6)) & leq(X6, n2) & leq(n0, X6)) | ~ sP6), inference(nnf_transformation, [], [f179])).
fof(f517, plain, (~ spl39_1 | spl39_3), inference(avatar_split_clause, [], [f355, f514, f505])).
fof(f355, plain, (leq(sK34, n2) | ~ sP6), inference(cnf_transformation, [], [f231])).
fof(f512, plain, (~ spl39_1 | ~ spl39_2), inference(avatar_split_clause, [], [f452, f509, f505])).
fof(f452, plain, (~ (s_worst7_init = a_select2(s_center7_init, sK34)) | ~ sP6), inference(definition_unfolding, [], [f356, f368])).
fof(f356, plain, (~ (init = a_select2(s_center7_init, sK34)) | ~ sP6), inference(cnf_transformation, [], [f231])).