fof(f6395, plain, $false, inference(avatar_sat_refutation, [], [f648, f649, f650, f659, f660, f661, f669, f675, f681, f686, f687, f688, f689, f690, f708, f709, f710, f711, f712, f721, f726, f727, f728, f733, f734, f735, f837, f884, f914, f935, f954, f977, f998, f1017, f1031, f1051, f1078, f1137, f1138, f1139, f1140, f1141, f1142, f1143, f1144, f1145, f1146, f1147, f1148, f1179, f1182, f1186, f1190, f1195, f1198, f1203, f1206, f1524, f1566, f1569, f2245, f5300, f6111, f6134, f6322])).
fof(f6322, plain, ~ spl49_170, inference(avatar_contradiction_clause, [], [f6321])).
fof(f6321, plain, ($false | ~ spl49_170), inference(subsumption_resolution, [], [f6135, f295])).
fof(f295, plain, ! [X0] : ~ gt(X0, X0), inference(cnf_transformation, [], [f3])).
fof(f3, plain, ! [X0] : ~ gt(X0, X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', irreflexivity_gt)).
fof(f6135, plain, (gt(n1000, n1000) | ~ spl49_170), inference(backward_demodulation, [], [f503, f6110])).
fof(f6110, plain, ((n1000 = n588) | ~ spl49_170), inference(avatar_component_clause, [], [f6108])).
fof(f6108, plain, (spl49_170 <=> (n1000 = n588)), introduced(avatar_definition, [new_symbols(naming, [spl49_170])])).
fof(f503, plain, gt(n1000, n588), inference(cnf_transformation, [], [f55])).
fof(f55, plain, gt(n1000, n588), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_1000_588)).
fof(f6134, plain, ~ spl49_169, inference(avatar_contradiction_clause, [], [f6133])).
fof(f6133, plain, ($false | ~ spl49_169), inference(subsumption_resolution, [], [f6130, f295])).
fof(f6130, plain, (gt(n588, n588) | ~ spl49_169), inference(resolution, [], [f6106, f1328])).
fof(f1328, plain, ! [X41] : (~ gt(X41, n1000) | gt(X41, n588)), inference(resolution, [], [f294, f503])).
fof(f294, plain, ! [X2, X0, X1] : (~ gt(X1, X2) | gt(X0, X2) | ~ gt(X0, X1)), inference(cnf_transformation, [], [f146])).
fof(f146, plain, ! [X0, X1, X2] : (gt(X0, X2) | ~ gt(X1, X2) | ~ gt(X0, X1)), inference(flattening, [], [f145])).
fof(f145, plain, ! [X0, X1, X2] : (gt(X0, X2) | (~ gt(X1, X2) | ~ gt(X0, X1))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ! [X0, X1, X2] : ((gt(X1, X2) & gt(X0, X1)) => gt(X0, X2)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', transitivity_gt)).
fof(f6106, plain, (gt(n588, n1000) | ~ spl49_169), inference(avatar_component_clause, [], [f6104])).
fof(f6104, plain, (spl49_169 <=> gt(n588, n1000)), introduced(avatar_definition, [new_symbols(naming, [spl49_169])])).
fof(f6111, plain, (spl49_169 | spl49_170 | ~ spl49_3 | ~ spl49_97), inference(avatar_split_clause, [], [f6101, f2241, f645, f6108, f6104])).
fof(f645, plain, (spl49_3 <=> leq(pv5, n588)), introduced(avatar_definition, [new_symbols(naming, [spl49_3])])).
fof(f2241, plain, (spl49_97 <=> gt(pv5, n1000)), introduced(avatar_definition, [new_symbols(naming, [spl49_97])])).
fof(f6101, plain, ((n1000 = n588) | gt(n588, n1000) | (~ spl49_3 | ~ spl49_97)), inference(resolution, [], [f6086, f300])).
fof(f300, plain, ! [X0, X1] : (~ leq(X0, X1) | (X0 = X1) | gt(X1, X0)), inference(cnf_transformation, [], [f152])).
fof(f152, plain, ! [X0, X1] : (gt(X1, X0) | (X0 = X1) | ~ leq(X0, X1)), inference(flattening, [], [f151])).
fof(f151, plain, ! [X0, X1] : (gt(X1, X0) | ((X0 = X1) | ~ leq(X0, X1))), inference(ennf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1] : ((~ (X0 = X1) & leq(X0, X1)) => gt(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', leq_gt2)).
fof(f6086, plain, (leq(n1000, n588) | (~ spl49_3 | ~ spl49_97)), inference(resolution, [], [f1418, f5448])).
fof(f5448, plain, (leq(n1000, pv5) | ~ spl49_97), inference(resolution, [], [f2243, f299])).
fof(f299, plain, ! [X0, X1] : (~ gt(X1, X0) | leq(X0, X1)), inference(cnf_transformation, [], [f150])).
fof(f150, plain, ! [X0, X1] : (leq(X0, X1) | ~ gt(X1, X0)), inference(ennf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (gt(X1, X0) => leq(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', leq_gt1)).
fof(f2243, plain, (gt(pv5, n1000) | ~ spl49_97), inference(avatar_component_clause, [], [f2241])).
fof(f1418, plain, (! [X90] : (~ leq(X90, pv5) | leq(X90, n588)) | ~ spl49_3), inference(resolution, [], [f297, f646])).
fof(f646, plain, (leq(pv5, n588) | ~ spl49_3), inference(avatar_component_clause, [], [f645])).
fof(f297, plain, ! [X2, X0, X1] : (~ leq(X1, X2) | leq(X0, X2) | ~ leq(X0, X1)), inference(cnf_transformation, [], [f148])).
fof(f148, plain, ! [X0, X1, X2] : (leq(X0, X2) | ~ leq(X1, X2) | ~ leq(X0, X1)), inference(flattening, [], [f147])).
fof(f147, plain, ! [X0, X1, X2] : (leq(X0, X2) | (~ leq(X1, X2) | ~ leq(X0, X1))), inference(ennf_transformation, [], [f5])).
fof(f5, plain, ! [X0, X1, X2] : ((leq(X1, X2) & leq(X0, X1)) => leq(X0, X2)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', transitivity_leq)).
fof(f5300, plain, (~ spl49_53 | ~ spl49_96), inference(avatar_contradiction_clause, [], [f5299])).
fof(f5299, plain, ($false | (~ spl49_53 | ~ spl49_96)), inference(subsumption_resolution, [], [f5296, f295])).
fof(f5296, plain, (gt(n588, n588) | (~ spl49_53 | ~ spl49_96)), inference(resolution, [], [f1328, f2253])).
fof(f2253, plain, (gt(n588, n1000) | (~ spl49_53 | ~ spl49_96)), inference(backward_demodulation, [], [f1519, f2239])).
fof(f2239, plain, ((n1000 = pv5) | ~ spl49_96), inference(avatar_component_clause, [], [f2237])).
fof(f2237, plain, (spl49_96 <=> (n1000 = pv5)), introduced(avatar_definition, [new_symbols(naming, [spl49_96])])).
fof(f1519, plain, (gt(n588, pv5) | ~ spl49_53), inference(avatar_component_clause, [], [f1517])).
fof(f1517, plain, (spl49_53 <=> gt(n588, pv5)), introduced(avatar_definition, [new_symbols(naming, [spl49_53])])).
fof(f2245, plain, (spl49_96 | spl49_97 | spl49_41), inference(avatar_split_clause, [], [f2235, f834, f2241, f2237])).
fof(f834, plain, (spl49_41 <=> leq(pv5, minus(n1000, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_41])])).
fof(f2235, plain, (gt(pv5, n1000) | (n1000 = pv5) | spl49_41), inference(resolution, [], [f1625, f293])).
fof(f293, plain, ! [X0, X1] : (gt(X1, X0) | gt(X0, X1) | (X0 = X1)), inference(cnf_transformation, [], [f1])).
fof(f1, plain, ! [X0, X1] : ((X0 = X1) | gt(X1, X0) | gt(X0, X1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', totality)).
fof(f1625, plain, (~ gt(n1000, pv5) | spl49_41), inference(resolution, [], [f836, f571])).
fof(f571, plain, ! [X0, X1] : (leq(X0, minus(X1, n1)) | ~ gt(X1, X0)), inference(definition_unfolding, [], [f302, f384])).
fof(f384, plain, ! [X0] : (minus(X0, n1) = pred(X0)), inference(cnf_transformation, [], [f39])).
fof(f39, plain, ! [X0] : (minus(X0, n1) = pred(X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', pred_minus_1)).
fof(f302, plain, ! [X0, X1] : (leq(X0, pred(X1)) | ~ gt(X1, X0)), inference(cnf_transformation, [], [f228])).
fof(f228, plain, ! [X0, X1] : ((leq(X0, pred(X1)) | ~ gt(X1, X0)) & (gt(X1, X0) | ~ leq(X0, pred(X1)))), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ! [X0, X1] : (leq(X0, pred(X1)) <=> gt(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', leq_gt_pred)).
fof(f836, plain, (~ leq(pv5, minus(n1000, n1)) | spl49_41), inference(avatar_component_clause, [], [f834])).
fof(f1569, plain, (spl49_12 | ~ spl49_13), inference(avatar_split_clause, [], [f1568, f705, f701])).
fof(f701, plain, (spl49_12 <=> leq(pv21, n5)), introduced(avatar_definition, [new_symbols(naming, [spl49_12])])).
fof(f705, plain, (spl49_13 <=> leq(pv21, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_13])])).
fof(f1568, plain, (leq(pv21, n5) | ~ spl49_13), inference(forward_demodulation, [], [f706, f1174])).
fof(f1174, plain, (n5 = minus(n6, n1)), inference(superposition, [], [f587, f1163])).
fof(f1163, plain, (n6 = plus(n5, n1)), inference(backward_demodulation, [], [f1159, f1160])).
fof(f1160, plain, (n5 = plus(n1, n4)), inference(superposition, [], [f1158, f578])).
fof(f578, plain, ! [X0] : (plus(X0, n1) = plus(n1, X0)), inference(definition_unfolding, [], [f375, f374])).
fof(f374, plain, ! [X0] : (succ(X0) = plus(X0, n1)), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ! [X0] : (succ(X0) = plus(X0, n1)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', succ_plus_1_r)).
fof(f375, plain, ! [X0] : (succ(X0) = plus(n1, X0)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, ! [X0] : (succ(X0) = plus(n1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', succ_plus_1_l)).
fof(f1158, plain, (n5 = plus(n4, n1)), inference(backward_demodulation, [], [f1153, f1155])).
fof(f1155, plain, (n4 = plus(n1, n3)), inference(superposition, [], [f1154, f578])).
fof(f1154, plain, (n4 = plus(n3, n1)), inference(backward_demodulation, [], [f1065, f1149])).
fof(f1149, plain, (n3 = plus(n1, n2)), inference(superposition, [], [f578, f1070])).
fof(f1070, plain, (n3 = plus(n2, n1)), inference(forward_demodulation, [], [f1069, f1064])).
fof(f1064, plain, (n2 = plus(n1, n1)), inference(forward_demodulation, [], [f596, f595])).
fof(f595, plain, (n1 = plus(n0, n1)), inference(definition_unfolding, [], [f568, f374])).
fof(f568, plain, (n1 = succ(n0)), inference(cnf_transformation, [], [f120])).
fof(f120, plain, (n1 = succ(n0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_1)).
fof(f596, plain, (n2 = plus(plus(n0, n1), n1)), inference(definition_unfolding, [], [f569, f374, f374])).
fof(f569, plain, (n2 = succ(succ(n0))), inference(cnf_transformation, [], [f121])).
fof(f121, plain, (n2 = succ(succ(n0))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_2)).
fof(f1069, plain, (n3 = plus(plus(n1, n1), n1)), inference(forward_demodulation, [], [f1068, f595])).
fof(f1068, plain, (n3 = plus(plus(n1, plus(n0, n1)), n1)), inference(forward_demodulation, [], [f597, f578])).
fof(f597, plain, (n3 = plus(plus(plus(n0, n1), n1), n1)), inference(definition_unfolding, [], [f570, f374, f374, f374])).
fof(f570, plain, (n3 = succ(succ(succ(n0)))), inference(cnf_transformation, [], [f122])).
fof(f122, plain, (n3 = succ(succ(succ(n0)))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_3)).
fof(f1065, plain, (n4 = plus(plus(n1, n2), n1)), inference(backward_demodulation, [], [f1063, f1064])).
fof(f1063, plain, (n4 = plus(plus(n1, plus(n1, n1)), n1)), inference(backward_demodulation, [], [f1053, f595])).
fof(f1053, plain, (n4 = plus(plus(n1, plus(n1, plus(n0, n1))), n1)), inference(forward_demodulation, [], [f1052, f578])).
fof(f1052, plain, (n4 = plus(plus(n1, plus(plus(n0, n1), n1)), n1)), inference(forward_demodulation, [], [f592, f578])).
fof(f592, plain, (n4 = plus(plus(plus(plus(n0, n1), n1), n1), n1)), inference(definition_unfolding, [], [f565, f374, f374, f374, f374])).
fof(f565, plain, (n4 = succ(succ(succ(succ(n0))))), inference(cnf_transformation, [], [f117])).
fof(f117, plain, (n4 = succ(succ(succ(succ(n0))))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_4)).
fof(f1153, plain, (n5 = plus(plus(n1, n3), n1)), inference(backward_demodulation, [], [f1066, f1149])).
fof(f1066, plain, (n5 = plus(plus(n1, plus(n1, n2)), n1)), inference(backward_demodulation, [], [f1062, f1064])).
fof(f1062, plain, (n5 = plus(plus(n1, plus(n1, plus(n1, n1))), n1)), inference(backward_demodulation, [], [f1056, f595])).
fof(f1056, plain, (n5 = plus(plus(n1, plus(n1, plus(n1, plus(n0, n1)))), n1)), inference(forward_demodulation, [], [f1055, f578])).
fof(f1055, plain, (n5 = plus(plus(n1, plus(n1, plus(plus(n0, n1), n1))), n1)), inference(forward_demodulation, [], [f1054, f578])).
fof(f1054, plain, (n5 = plus(plus(n1, plus(plus(plus(n0, n1), n1), n1)), n1)), inference(forward_demodulation, [], [f593, f578])).
fof(f593, plain, (n5 = plus(plus(plus(plus(plus(n0, n1), n1), n1), n1), n1)), inference(definition_unfolding, [], [f566, f374, f374, f374, f374, f374])).
fof(f566, plain, (n5 = succ(succ(succ(succ(succ(n0)))))), inference(cnf_transformation, [], [f118])).
fof(f118, plain, (n5 = succ(succ(succ(succ(succ(n0)))))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_5)).
fof(f1159, plain, (n6 = plus(plus(n1, n4), n1)), inference(backward_demodulation, [], [f1152, f1155])).
fof(f1152, plain, (n6 = plus(plus(n1, plus(n1, n3)), n1)), inference(backward_demodulation, [], [f1067, f1149])).
fof(f1067, plain, (n6 = plus(plus(n1, plus(n1, plus(n1, n2))), n1)), inference(backward_demodulation, [], [f1061, f1064])).
fof(f1061, plain, (n6 = plus(plus(n1, plus(n1, plus(n1, plus(n1, n1)))), n1)), inference(backward_demodulation, [], [f1060, f595])).
fof(f1060, plain, (n6 = plus(plus(n1, plus(n1, plus(n1, plus(n1, plus(n0, n1))))), n1)), inference(forward_demodulation, [], [f1059, f578])).
fof(f1059, plain, (n6 = plus(plus(n1, plus(n1, plus(n1, plus(plus(n0, n1), n1)))), n1)), inference(forward_demodulation, [], [f1058, f578])).
fof(f1058, plain, (n6 = plus(plus(n1, plus(n1, plus(plus(plus(n0, n1), n1), n1))), n1)), inference(forward_demodulation, [], [f1057, f578])).
fof(f1057, plain, (n6 = plus(plus(n1, plus(plus(plus(plus(n0, n1), n1), n1), n1)), n1)), inference(forward_demodulation, [], [f594, f578])).
fof(f594, plain, (n6 = plus(plus(plus(plus(plus(plus(n0, n1), n1), n1), n1), n1), n1)), inference(definition_unfolding, [], [f567, f374, f374, f374, f374, f374, f374])).
fof(f567, plain, (n6 = succ(succ(succ(succ(succ(succ(n0))))))), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (n6 = succ(succ(succ(succ(succ(succ(n0))))))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', successor_6)).
fof(f587, plain, ! [X0] : (minus(plus(X0, n1), n1) = X0), inference(definition_unfolding, [], [f385, f384, f374])).
fof(f385, plain, ! [X0] : (pred(succ(X0)) = X0), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0] : (pred(succ(X0)) = X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', pred_succ)).
fof(f706, plain, (leq(pv21, minus(n6, n1)) | ~ spl49_13), inference(avatar_component_clause, [], [f705])).
fof(f1566, plain, (spl49_41 | ~ spl49_54), inference(avatar_contradiction_clause, [], [f1565])).
fof(f1565, plain, ($false | (spl49_41 | ~ spl49_54)), inference(subsumption_resolution, [], [f1526, f1215])).
fof(f1215, plain, (~ gt(n1000, pv5) | spl49_41), inference(resolution, [], [f571, f836])).
fof(f1526, plain, (gt(n1000, pv5) | ~ spl49_54), inference(backward_demodulation, [], [f503, f1523])).
fof(f1523, plain, ((pv5 = n588) | ~ spl49_54), inference(avatar_component_clause, [], [f1521])).
fof(f1521, plain, (spl49_54 <=> (pv5 = n588)), introduced(avatar_definition, [new_symbols(naming, [spl49_54])])).
fof(f1524, plain, (spl49_53 | spl49_54 | ~ spl49_3), inference(avatar_split_clause, [], [f1495, f645, f1521, f1517])).
fof(f1495, plain, ((pv5 = n588) | gt(n588, pv5) | ~ spl49_3), inference(resolution, [], [f300, f646])).
fof(f1206, plain, spl49_39, inference(avatar_contradiction_clause, [], [f1205])).
fof(f1205, plain, ($false | spl49_39), inference(subsumption_resolution, [], [f1204, f296])).
fof(f296, plain, ! [X0] : leq(X0, X0), inference(cnf_transformation, [], [f4])).
fof(f4, plain, ! [X0] : leq(X0, X0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', reflexivity_leq)).
fof(f1204, plain, (~ leq(n5, n5) | spl49_39), inference(forward_demodulation, [], [f828, f1174])).
fof(f828, plain, (~ leq(n5, minus(n6, n1)) | spl49_39), inference(avatar_component_clause, [], [f826])).
fof(f826, plain, (spl49_39 <=> leq(n5, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_39])])).
fof(f1203, plain, spl49_35, inference(avatar_contradiction_clause, [], [f1202])).
fof(f1202, plain, ($false | spl49_35), inference(subsumption_resolution, [], [f1201, f1099])).
fof(f1099, plain, leq(n3, n5), inference(resolution, [], [f299, f554])).
fof(f554, plain, gt(n5, n3), inference(cnf_transformation, [], [f106])).
fof(f106, plain, gt(n5, n3), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_5_3)).
fof(f1201, plain, (~ leq(n3, n5) | spl49_35), inference(forward_demodulation, [], [f812, f1174])).
fof(f812, plain, (~ leq(n3, minus(n6, n1)) | spl49_35), inference(avatar_component_clause, [], [f810])).
fof(f810, plain, (spl49_35 <=> leq(n3, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_35])])).
fof(f1198, plain, spl49_37, inference(avatar_contradiction_clause, [], [f1197])).
fof(f1197, plain, ($false | spl49_37), inference(subsumption_resolution, [], [f1196, f1100])).
fof(f1100, plain, leq(n4, n5), inference(resolution, [], [f299, f505])).
fof(f505, plain, gt(n5, n4), inference(cnf_transformation, [], [f57])).
fof(f57, plain, gt(n5, n4), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_5_4)).
fof(f1196, plain, (~ leq(n4, n5) | spl49_37), inference(forward_demodulation, [], [f820, f1174])).
fof(f820, plain, (~ leq(n4, minus(n6, n1)) | spl49_37), inference(avatar_component_clause, [], [f818])).
fof(f818, plain, (spl49_37 <=> leq(n4, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_37])])).
fof(f1195, plain, spl49_28, inference(avatar_contradiction_clause, [], [f1194])).
fof(f1194, plain, ($false | spl49_28), inference(subsumption_resolution, [], [f1193, f1088])).
fof(f1088, plain, leq(n1, n3), inference(resolution, [], [f299, f544])).
fof(f544, plain, gt(n3, n1), inference(cnf_transformation, [], [f96])).
fof(f96, plain, gt(n3, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_3_1)).
fof(f1193, plain, (~ leq(n1, n3) | spl49_28), inference(forward_demodulation, [], [f784, f1172])).
fof(f1172, plain, (n3 = minus(n4, n1)), inference(superposition, [], [f587, f1154])).
fof(f784, plain, (~ leq(n1, minus(n4, n1)) | spl49_28), inference(avatar_component_clause, [], [f782])).
fof(f782, plain, (spl49_28 <=> leq(n1, minus(n4, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_28])])).
fof(f1190, plain, spl49_34, inference(avatar_contradiction_clause, [], [f1189])).
fof(f1189, plain, ($false | spl49_34), inference(subsumption_resolution, [], [f1188, f296])).
fof(f1188, plain, (~ leq(n3, n3) | spl49_34), inference(forward_demodulation, [], [f808, f1172])).
fof(f808, plain, (~ leq(n3, minus(n4, n1)) | spl49_34), inference(avatar_component_clause, [], [f806])).
fof(f806, plain, (spl49_34 <=> leq(n3, minus(n4, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_34])])).
fof(f1186, plain, spl49_32, inference(avatar_contradiction_clause, [], [f1185])).
fof(f1185, plain, ($false | spl49_32), inference(subsumption_resolution, [], [f1184, f1098])).
fof(f1098, plain, leq(n2, n5), inference(resolution, [], [f299, f547])).
fof(f547, plain, gt(n5, n2), inference(cnf_transformation, [], [f99])).
fof(f99, plain, gt(n5, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_5_2)).
fof(f1184, plain, (~ leq(n2, n5) | spl49_32), inference(forward_demodulation, [], [f800, f1174])).
fof(f800, plain, (~ leq(n2, minus(n6, n1)) | spl49_32), inference(avatar_component_clause, [], [f798])).
fof(f798, plain, (spl49_32 <=> leq(n2, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_32])])).
fof(f1182, plain, spl49_31, inference(avatar_contradiction_clause, [], [f1181])).
fof(f1181, plain, ($false | spl49_31), inference(subsumption_resolution, [], [f1180, f1089])).
fof(f1089, plain, leq(n2, n3), inference(resolution, [], [f299, f551])).
fof(f551, plain, gt(n3, n2), inference(cnf_transformation, [], [f103])).
fof(f103, plain, gt(n3, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_3_2)).
fof(f1180, plain, (~ leq(n2, n3) | spl49_31), inference(forward_demodulation, [], [f796, f1172])).
fof(f796, plain, (~ leq(n2, minus(n4, n1)) | spl49_31), inference(avatar_component_clause, [], [f794])).
fof(f794, plain, (spl49_31 <=> leq(n2, minus(n4, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_31])])).
fof(f1179, plain, spl49_29, inference(avatar_contradiction_clause, [], [f1178])).
fof(f1178, plain, ($false | spl49_29), inference(subsumption_resolution, [], [f1177, f1097])).
fof(f1097, plain, leq(n1, n5), inference(resolution, [], [f299, f539])).
fof(f539, plain, gt(n5, n1), inference(cnf_transformation, [], [f91])).
fof(f91, plain, gt(n5, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_5_1)).
fof(f1177, plain, (~ leq(n1, n5) | spl49_29), inference(backward_demodulation, [], [f788, f1174])).
fof(f788, plain, (~ leq(n1, minus(n6, n1)) | spl49_29), inference(avatar_component_clause, [], [f786])).
fof(f786, plain, (spl49_29 <=> leq(n1, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_29])])).
fof(f1148, plain, spl49_21, inference(avatar_split_clause, [], [f1081, f754])).
fof(f754, plain, (spl49_21 <=> leq(n0, n1)), introduced(avatar_definition, [new_symbols(naming, [spl49_21])])).
fof(f1081, plain, leq(n0, n1), inference(resolution, [], [f299, f533])).
fof(f533, plain, gt(n1, n0), inference(cnf_transformation, [], [f85])).
fof(f85, plain, gt(n1, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_1_0)).
fof(f1147, plain, spl49_26, inference(avatar_split_clause, [], [f1119, f774])).
fof(f774, plain, (spl49_26 <=> leq(n0, n6)), introduced(avatar_definition, [new_symbols(naming, [spl49_26])])).
fof(f1119, plain, leq(n0, n6), inference(resolution, [], [f299, f531])).
fof(f531, plain, gt(n6, n0), inference(cnf_transformation, [], [f83])).
fof(f83, plain, gt(n6, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_6_0)).
fof(f1146, plain, spl49_40, inference(avatar_split_clause, [], [f1118, f830])).
fof(f830, plain, (spl49_40 <=> leq(n6, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_40])])).
fof(f1118, plain, leq(n6, n7), inference(resolution, [], [f299, f514])).
fof(f514, plain, gt(n7, n6), inference(cnf_transformation, [], [f66])).
fof(f66, plain, gt(n7, n6), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_6)).
fof(f1145, plain, spl49_38, inference(avatar_split_clause, [], [f1117, f822])).
fof(f822, plain, (spl49_38 <=> leq(n5, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_38])])).
fof(f1117, plain, leq(n5, n7), inference(resolution, [], [f299, f511])).
fof(f511, plain, gt(n7, n5), inference(cnf_transformation, [], [f63])).
fof(f63, plain, gt(n7, n5), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_5)).
fof(f1144, plain, spl49_36, inference(avatar_split_clause, [], [f1116, f814])).
fof(f814, plain, (spl49_36 <=> leq(n4, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_36])])).
fof(f1116, plain, leq(n4, n7), inference(resolution, [], [f299, f507])).
fof(f507, plain, gt(n7, n4), inference(cnf_transformation, [], [f59])).
fof(f59, plain, gt(n7, n4), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_4)).
fof(f1143, plain, spl49_33, inference(avatar_split_clause, [], [f1115, f802])).
fof(f802, plain, (spl49_33 <=> leq(n3, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_33])])).
fof(f1115, plain, leq(n3, n7), inference(resolution, [], [f299, f556])).
fof(f556, plain, gt(n7, n3), inference(cnf_transformation, [], [f108])).
fof(f108, plain, gt(n7, n3), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_3)).
fof(f1142, plain, spl49_30, inference(avatar_split_clause, [], [f1114, f790])).
fof(f790, plain, (spl49_30 <=> leq(n2, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_30])])).
fof(f1114, plain, leq(n2, n7), inference(resolution, [], [f299, f549])).
fof(f549, plain, gt(n7, n2), inference(cnf_transformation, [], [f101])).
fof(f101, plain, gt(n7, n2), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_2)).
fof(f1141, plain, spl49_27, inference(avatar_split_clause, [], [f1113, f778])).
fof(f778, plain, (spl49_27 <=> leq(n1, n7)), introduced(avatar_definition, [new_symbols(naming, [spl49_27])])).
fof(f1113, plain, leq(n1, n7), inference(resolution, [], [f299, f541])).
fof(f541, plain, gt(n7, n1), inference(cnf_transformation, [], [f93])).
fof(f93, plain, gt(n7, n1), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_7_1)).
fof(f1140, plain, spl49_25, inference(avatar_split_clause, [], [f1095, f770])).
fof(f770, plain, (spl49_25 <=> leq(n0, n5)), introduced(avatar_definition, [new_symbols(naming, [spl49_25])])).
fof(f1095, plain, leq(n0, n5), inference(resolution, [], [f299, f530])).
fof(f530, plain, gt(n5, n0), inference(cnf_transformation, [], [f82])).
fof(f82, plain, gt(n5, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_5_0)).
fof(f1139, plain, spl49_24, inference(avatar_split_clause, [], [f1090, f766])).
fof(f766, plain, (spl49_24 <=> leq(n0, n4)), introduced(avatar_definition, [new_symbols(naming, [spl49_24])])).
fof(f1090, plain, leq(n0, n4), inference(resolution, [], [f299, f529])).
fof(f529, plain, gt(n4, n0), inference(cnf_transformation, [], [f81])).
fof(f81, plain, gt(n4, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_4_0)).
fof(f1138, plain, spl49_23, inference(avatar_split_clause, [], [f1086, f762])).
fof(f762, plain, (spl49_23 <=> leq(n0, n3)), introduced(avatar_definition, [new_symbols(naming, [spl49_23])])).
fof(f1086, plain, leq(n0, n3), inference(resolution, [], [f299, f536])).
fof(f536, plain, gt(n3, n0), inference(cnf_transformation, [], [f88])).
fof(f88, plain, gt(n3, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_3_0)).
fof(f1137, plain, spl49_22, inference(avatar_split_clause, [], [f1084, f758])).
fof(f758, plain, (spl49_22 <=> leq(n0, n2)), introduced(avatar_definition, [new_symbols(naming, [spl49_22])])).
fof(f1084, plain, leq(n0, n2), inference(resolution, [], [f299, f534])).
fof(f534, plain, gt(n2, n0), inference(cnf_transformation, [], [f86])).
fof(f86, plain, gt(n2, n0), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', gt_2_0)).
fof(f1078, plain, spl49_48, inference(avatar_split_clause, [], [f1077, f1048])).
fof(f1048, plain, (spl49_48 <=> leq(n0, minus(n1000, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_48])])).
fof(f1077, plain, leq(n0, minus(n1000, n1)), inference(resolution, [], [f298, f498])).
fof(f498, plain, geq(minus(n1000, n1), n0), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ((sP17 | sP18 | sP21 | sP20 | sP19 | sP16 | sP15 | sP13 | sP14 | ((~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & true) | (~ true & true)) & geq(minus(n1000, n1), n0) & geq(n7, n0) & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)), inference(definition_folding, [], [f187, e226, e225, e224, e223, e222, e221, e220, e219, e218, e217, e216, e215, e214, e213, e212, e211, e210, e209])).
fof(f209, plain, (((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (pv31 = pv32)) | ((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & ~ (pv31 = pv32)) | ~ sP4), inference(usedef, [], [e209])).
fof(e209, plain, (sP4 <=> (((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (pv31 = pv32)) | ((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & ~ (pv31 = pv32)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f210, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP5), inference(usedef, [], [e210])).
fof(e210, plain, (sP5 <=> ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f211, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP6), inference(usedef, [], [e211])).
fof(e211, plain, (sP6 <=> ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f212, plain, (((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP5) & ~ gt(pv5, n0)) | ~ sP7), inference(usedef, [], [e212])).
fof(e212, plain, (sP7 <=> ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP5) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f213, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP8), inference(usedef, [], [e213])).
fof(e213, plain, (sP8 <=> ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f214, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP9), inference(usedef, [], [e214])).
fof(e214, plain, (sP9 <=> ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f215, plain, (((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP8) & ~ gt(pv5, n0)) | ~ sP10), inference(usedef, [], [e215])).
fof(e215, plain, (sP10 <=> ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP8) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f216, plain, (((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP6 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP7) & ~ gt(pv5, n0)) | ~ sP11), inference(usedef, [], [e216])).
fof(e216, plain, (sP11 <=> ((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP6 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP7) & ~ gt(pv5, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f217, plain, (((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP9 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP10 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP11 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP12), inference(usedef, [], [e217])).
fof(e217, plain, (sP12 <=> (((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP9 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP10 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP11 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f218, plain, ((sP4 & leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP13), inference(usedef, [], [e218])).
fof(e218, plain, (sP13 <=> (sP4 & leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f219, plain, (((~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0)) & leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) | ~ sP14), inference(usedef, [], [e219])).
fof(e219, plain, (sP14 <=> ((~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0)) & leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f220, plain, (((~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP15), inference(usedef, [], [e220])).
fof(e220, plain, (sP15 <=> ((~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f221, plain, (((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP16), inference(usedef, [], [e221])).
fof(e221, plain, (sP16 <=> ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f222, plain, ((~ true & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0) & geq(minus(n6, n1), n0) & geq(minus(n6, n1), n0)) | ~ sP17), inference(usedef, [], [e222])).
fof(e222, plain, (sP17 <=> (~ true & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0) & geq(minus(n6, n1), n0) & geq(minus(n6, n1), n0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f223, plain, ((~ true & leq(n0, a_select2(sigma, pv23)) & leq(pv23, minus(n6, n1)) & leq(n0, pv23)) | ~ sP18), inference(usedef, [], [e223])).
fof(e223, plain, (sP18 <=> (~ true & leq(n0, a_select2(sigma, pv23)) & leq(pv23, minus(n6, n1)) & leq(n0, pv23))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f224, plain, ((~ true & leq(pv5, n588) & leq(n0, pv5)) | ~ sP19), inference(usedef, [], [e224])).
fof(e224, plain, (sP19 <=> (~ true & leq(pv5, n588) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f225, plain, ((sP12 & leq(pv5, n588) & leq(n0, pv5)) | ~ sP20), inference(usedef, [], [e225])).
fof(e225, plain, (sP20 <=> (sP12 & leq(pv5, n588) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f226, plain, (((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv5, n588) & leq(n0, pv5)) | ~ sP21), inference(usedef, [], [e226])).
fof(e226, plain, (sP21 <=> ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv5, n588) & leq(n0, pv5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f187, plain, (((~ true & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0) & geq(minus(n6, n1), n0) & geq(minus(n6, n1), n0)) | (~ true & leq(n0, a_select2(sigma, pv23)) & leq(pv23, minus(n6, n1)) & leq(n0, pv23)) | ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv5, n588) & leq(n0, pv5)) | ((((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & leq(pv5, n588) & leq(n0, pv5)) | (~ true & leq(pv5, n588) & leq(n0, pv5)) | ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ((~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ((((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (pv31 = pv32)) | ((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & ~ (pv31 = pv32))) & leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) | ((~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0)) & leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) | ((~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & true) | (~ true & true)) & geq(minus(n1000, n1), n0) & geq(n7, n0) & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)), inference(flattening, [], [f186])).
fof(f186, plain, ((((((~ true & (geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0))) & geq(minus(n6, n1), n0)) & geq(minus(n6, n1), n0)) | ((~ true & leq(n0, a_select2(sigma, pv23))) & (leq(pv23, minus(n6, n1)) & leq(n0, pv23))) | ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & (leq(pv5, n588) & leq(n0, pv5))) | ((((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & (leq(pv5, n588) & leq(n0, pv5))) | (~ true & (leq(pv5, n588) & leq(n0, pv5))) | ((~ leq(pv5, n588) | ~ leq(n0, pv5)) & (leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))) | ((~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))) | ((((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (pv31 = pv32)) | ((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & ~ (pv31 = pv32))) & (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))) | ((~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0)) & (leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5))) | ((~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & true) | (~ true & true)) & (geq(minus(n1000, n1), n0) & geq(n7, n0))) & (geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0))), inference(ennf_transformation, [], [f143])).
fof(f143, plain, ~ ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => ((geq(minus(n1000, n1), n0) & geq(n7, n0)) => ((geq(minus(n6, n1), n0) => (geq(minus(n6, n1), n0) => ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => true))) & ((leq(pv23, minus(n6, n1)) & leq(n0, pv23)) => (leq(n0, a_select2(sigma, pv23)) => true)) & ((leq(pv5, n588) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv5, n588) & leq(n0, pv5)) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & ((leq(pv5, n588) & leq(n0, pv5)) => true) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))) & ((leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) => (((pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))) & (~ (pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))))) & ((leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) => (leq(pv21, minus(n6, n1)) & leq(pv21, n5) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5) & leq(n0, n0))) & (true => (leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n1000, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (true => true)))), inference(rectify, [], [f54])).
fof(f54, plain, ~ ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => ! [X13] : ((geq(minus(n1000, n1), n0) & geq(n7, n0)) => ! [X17] : ((geq(minus(n6, n1), n0) => (geq(minus(n6, n1), n0) => ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => true))) & ((leq(pv23, minus(n6, n1)) & leq(n0, pv23)) => (leq(n0, a_select2(sigma, pv23)) => true)) & ((leq(pv5, n588) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv5, n588) & leq(n0, pv5)) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & ((leq(pv5, n588) & leq(n0, pv5)) => true) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))) & ((leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) => (((pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))) & (~ (pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))))) & ((leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) => (leq(pv21, minus(n6, n1)) & leq(pv21, n5) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5) & leq(n0, n0))) & (true => (leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n1000, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (true => true)))), inference(negated_conjecture, [], [f53])).
fof(f53, plain, ~ ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => ! [X13] : ((geq(minus(n1000, n1), n0) & geq(n7, n0)) => ! [X17] : ((geq(minus(n6, n1), n0) => (geq(minus(n6, n1), n0) => ((geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0)) => true))) & ((leq(pv23, minus(n6, n1)) & leq(n0, pv23)) => (leq(n0, a_select2(sigma, pv23)) => true)) & ((leq(pv5, n588) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv5, n588) & leq(n0, pv5)) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))) & leq(pv5, minus(n1000, n1)) & leq(n7, n7) & leq(n6, n7) & leq(n5, n7) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, n7) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => ((gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (~ gt(pv5, n0) => (leq(pv5, minus(n1000, n1)) & leq(pv5, n588) & leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, n7) & leq(n0, minus(n6, n1)) & leq(n0, pv5) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))))))))) & leq(pv5, minus(n1000, n1)) & leq(n3, minus(n4, n1)) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, pv5) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & ((leq(pv5, n588) & leq(n0, pv5)) => true) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv5, n588) & leq(n0, pv5))) & ((leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) => (leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5))) & ((leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) => (((pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))) & (~ (pv31 = pv32) => (leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5))))) & ((leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) => (leq(pv21, minus(n6, n1)) & leq(pv21, n5) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5) & leq(n0, n0))) & (true => (leq(n7, n7) & leq(n6, n7) & leq(n5, minus(n6, n1)) & leq(n5, n7) & leq(n4, minus(n6, n1)) & leq(n4, n7) & leq(n3, minus(n6, n1)) & leq(n3, minus(n4, n1)) & leq(n3, n7) & leq(n2, minus(n6, n1)) & leq(n2, minus(n4, n1)) & leq(n2, n7) & leq(n1, minus(n6, n1)) & leq(n1, minus(n4, n1)) & leq(n1, n7) & leq(n0, minus(n1000, n1)) & leq(n0, minus(n6, n1)) & leq(n0, minus(n4, n1)) & leq(n0, n7) & leq(n0, n6) & leq(n0, n5) & leq(n0, n4) & leq(n0, n3) & leq(n0, n2) & leq(n0, n1) & leq(n0, n0))) & (true => true)))), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', thruster_array_0001)).
fof(f298, plain, ! [X0, X1] : (~ geq(X0, X1) | leq(X1, X0)), inference(cnf_transformation, [], [f149])).
fof(f149, plain, ! [X0, X1] : (leq(X1, X0) | ~ geq(X0, X1)), inference(ennf_transformation, [], [f144])).
fof(f144, plain, ! [X0, X1] : (geq(X0, X1) => leq(X1, X0)), inference(unused_predicate_definition_removal, [], [f7])).
fof(f7, plain, ! [X0, X1] : (geq(X0, X1) <=> leq(X1, X0)), file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', leq_geq)).
fof(f1051, plain, (~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_48 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | spl49_10 | spl49_14 | spl49_9 | spl49_6 | spl49_4 | spl49_1), inference(avatar_split_clause, [], [f1046, f637, f652, f666, f683, f714, f693, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f1048, f774, f770, f766, f762, f758, f754])).
fof(f693, plain, (spl49_10 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl49_10])])).
fof(f714, plain, (spl49_14 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl49_14])])).
fof(f683, plain, (spl49_9 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl49_9])])).
fof(f666, plain, (spl49_6 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl49_6])])).
fof(f652, plain, (spl49_4 <=> sP20), introduced(avatar_definition, [new_symbols(naming, [spl49_4])])).
fof(f637, plain, (spl49_1 <=> sP21), introduced(avatar_definition, [new_symbols(naming, [spl49_1])])).
fof(f1046, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1)), inference(subsumption_resolution, [], [f1045, f297])).
fof(f1045, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1)), inference(subsumption_resolution, [], [f1044, f297])).
fof(f1044, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1)), inference(subsumption_resolution, [], [f1043, f297])).
fof(f1043, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1)), inference(subsumption_resolution, [], [f1042, f403])).
fof(f403, plain, true, inference(cnf_transformation, [], [f51])).
fof(f51, plain, true, file('/home/ubuntu/library/tptp/Problems/SWV/SWV125+1.p', ttrue)).
fof(f1042, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ true), inference(subsumption_resolution, [], [f1041, f296])).
fof(f1041, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ true), inference(subsumption_resolution, [], [f1040, f296])).
fof(f1040, plain, (sP21 | sP20 | sP16 | sP15 | sP13 | sP14 | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ true), inference(subsumption_resolution, [], [f1039, f662])).
fof(f662, plain, ~ sP19, inference(subsumption_resolution, [], [f413, f403])).
fof(f413, plain, (~ true | ~ sP19), inference(cnf_transformation, [], [f277])).
fof(f277, plain, ((~ true & leq(pv5, n588) & leq(n0, pv5)) | ~ sP19), inference(nnf_transformation, [], [f224])).
fof(f1039, plain, (sP21 | sP20 | sP19 | sP16 | sP15 | sP13 | sP14 | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ true), inference(subsumption_resolution, [], [f1038, f663])).
fof(f663, plain, ~ sP18, inference(subsumption_resolution, [], [f417, f403])).
fof(f417, plain, (~ true | ~ sP18), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ((~ true & leq(n0, a_select2(sigma, pv23)) & leq(pv23, minus(n6, n1)) & leq(n0, pv23)) | ~ sP18), inference(nnf_transformation, [], [f223])).
fof(f1038, plain, (sP18 | sP21 | sP20 | sP19 | sP16 | sP15 | sP13 | sP14 | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ true), inference(subsumption_resolution, [], [f502, f664])).
fof(f664, plain, ~ sP17, inference(subsumption_resolution, [], [f422, f403])).
fof(f422, plain, (~ true | ~ sP17), inference(cnf_transformation, [], [f279])).
fof(f279, plain, ((~ true & geq(minus(n1000, n1), n0) & geq(minus(n4, n1), n0) & geq(minus(n6, n1), n0) & geq(minus(n6, n1), n0)) | ~ sP17), inference(nnf_transformation, [], [f222])).
fof(f502, plain, (sP17 | sP18 | sP21 | sP20 | sP19 | sP16 | sP15 | sP13 | sP14 | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n1000, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ true), inference(cnf_transformation, [], [f227])).
fof(f1031, plain, (~ spl49_15 | ~ spl49_2 | ~ spl49_8 | ~ spl49_17 | ~ spl49_3 | ~ spl49_7 | ~ spl49_16), inference(avatar_split_clause, [], [f602, f723, f671, f645, f730, f677, f641, f718])).
fof(f718, plain, (spl49_15 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl49_15])])).
fof(f641, plain, (spl49_2 <=> leq(n0, pv5)), introduced(avatar_definition, [new_symbols(naming, [spl49_2])])).
fof(f677, plain, (spl49_8 <=> leq(n0, pv31)), introduced(avatar_definition, [new_symbols(naming, [spl49_8])])).
fof(f730, plain, (spl49_17 <=> leq(n0, pv32)), introduced(avatar_definition, [new_symbols(naming, [spl49_17])])).
fof(f671, plain, (spl49_7 <=> leq(pv31, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_7])])).
fof(f723, plain, (spl49_16 <=> leq(pv32, minus(n6, n1))), introduced(avatar_definition, [new_symbols(naming, [spl49_16])])).
fof(f602, plain, (~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5) | ~ sP4), inference(duplicate_literal_removal, [], [f494])).
fof(f494, plain, (~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5) | ~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5) | ~ sP4), inference(cnf_transformation, [], [f292])).
fof(f292, plain, (((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & (pv31 = pv32)) | ((~ leq(pv32, minus(n6, n1)) | ~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv32) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & ~ (pv31 = pv32)) | ~ sP4), inference(nnf_transformation, [], [f209])).
fof(f1017, plain, (~ spl49_46 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f1016, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f974])).
fof(f974, plain, (spl49_46 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl49_46])])).
fof(f1016, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP5), inference(subsumption_resolution, [], [f1015, f297])).
fof(f1015, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP5), inference(subsumption_resolution, [], [f1014, f297])).
fof(f1014, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP5), inference(subsumption_resolution, [], [f1013, f297])).
fof(f1013, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP5), inference(subsumption_resolution, [], [f1012, f296])).
fof(f1012, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP5), inference(subsumption_resolution, [], [f603, f296])).
fof(f603, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP5), inference(duplicate_literal_removal, [], [f490])).
fof(f490, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP5), inference(cnf_transformation, [], [f291])).
fof(f291, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP5), inference(nnf_transformation, [], [f210])).
fof(f998, plain, (~ spl49_44 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f997, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f881])).
fof(f881, plain, (spl49_44 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl49_44])])).
fof(f997, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP6), inference(subsumption_resolution, [], [f996, f297])).
fof(f996, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP6), inference(subsumption_resolution, [], [f995, f297])).
fof(f995, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP6), inference(subsumption_resolution, [], [f994, f297])).
fof(f994, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP6), inference(subsumption_resolution, [], [f993, f296])).
fof(f993, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP6), inference(subsumption_resolution, [], [f604, f296])).
fof(f604, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP6), inference(duplicate_literal_removal, [], [f485])).
fof(f485, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP6), inference(cnf_transformation, [], [f290])).
fof(f290, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP6), inference(nnf_transformation, [], [f211])).
fof(f977, plain, (~ spl49_43 | spl49_46 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f972, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f974, f877])).
fof(f877, plain, (spl49_43 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl49_43])])).
fof(f972, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP5 | ~ sP7), inference(subsumption_resolution, [], [f971, f297])).
fof(f971, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP5 | ~ sP7), inference(subsumption_resolution, [], [f970, f297])).
fof(f970, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP5 | ~ sP7), inference(subsumption_resolution, [], [f969, f297])).
fof(f969, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP5 | ~ sP7), inference(subsumption_resolution, [], [f968, f296])).
fof(f968, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP5 | ~ sP7), inference(subsumption_resolution, [], [f607, f296])).
fof(f607, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP5 | ~ sP7), inference(duplicate_literal_removal, [], [f480])).
fof(f480, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP5 | ~ sP7), inference(cnf_transformation, [], [f289])).
fof(f289, plain, (((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP5) & ~ gt(pv5, n0)) | ~ sP7), inference(nnf_transformation, [], [f212])).
fof(f954, plain, (~ spl49_45 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f953, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f911])).
fof(f911, plain, (spl49_45 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl49_45])])).
fof(f953, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP8), inference(subsumption_resolution, [], [f952, f297])).
fof(f952, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP8), inference(subsumption_resolution, [], [f951, f297])).
fof(f951, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP8), inference(subsumption_resolution, [], [f950, f297])).
fof(f950, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP8), inference(subsumption_resolution, [], [f949, f296])).
fof(f949, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP8), inference(subsumption_resolution, [], [f608, f296])).
fof(f608, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP8), inference(duplicate_literal_removal, [], [f474])).
fof(f474, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP8), inference(cnf_transformation, [], [f288])).
fof(f288, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP8), inference(nnf_transformation, [], [f213])).
fof(f935, plain, (~ spl49_20 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f934, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f750])).
fof(f750, plain, (spl49_20 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl49_20])])).
fof(f934, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP9), inference(subsumption_resolution, [], [f933, f297])).
fof(f933, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP9), inference(subsumption_resolution, [], [f932, f297])).
fof(f932, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP9), inference(subsumption_resolution, [], [f931, f297])).
fof(f931, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ sP9), inference(subsumption_resolution, [], [f930, f296])).
fof(f930, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP9), inference(subsumption_resolution, [], [f609, f296])).
fof(f609, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP9), inference(duplicate_literal_removal, [], [f469])).
fof(f469, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP9), inference(cnf_transformation, [], [f287])).
fof(f287, plain, (((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0))) & ~ gt(pv5, n0)) | ~ sP9), inference(nnf_transformation, [], [f214])).
fof(f914, plain, (~ spl49_19 | spl49_45 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f909, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f911, f746])).
fof(f746, plain, (spl49_19 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl49_19])])).
fof(f909, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP8 | ~ sP10), inference(subsumption_resolution, [], [f908, f297])).
fof(f908, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP8 | ~ sP10), inference(subsumption_resolution, [], [f907, f297])).
fof(f907, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP8 | ~ sP10), inference(subsumption_resolution, [], [f906, f297])).
fof(f906, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP8 | ~ sP10), inference(subsumption_resolution, [], [f905, f296])).
fof(f905, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP8 | ~ sP10), inference(subsumption_resolution, [], [f612, f296])).
fof(f612, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP8 | ~ sP10), inference(duplicate_literal_removal, [], [f464])).
fof(f464, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP8 | ~ sP10), inference(cnf_transformation, [], [f286])).
fof(f286, plain, (((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP8) & ~ gt(pv5, n0)) | ~ sP10), inference(nnf_transformation, [], [f215])).
fof(f884, plain, (~ spl49_18 | spl49_43 | spl49_44 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f875, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f881, f877, f742])).
fof(f742, plain, (spl49_18 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl49_18])])).
fof(f875, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP6 | sP7 | ~ sP11), inference(subsumption_resolution, [], [f874, f297])).
fof(f874, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP6 | sP7 | ~ sP11), inference(subsumption_resolution, [], [f873, f297])).
fof(f873, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP6 | sP7 | ~ sP11), inference(subsumption_resolution, [], [f872, f297])).
fof(f872, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP6 | sP7 | ~ sP11), inference(subsumption_resolution, [], [f871, f296])).
fof(f871, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP6 | sP7 | ~ sP11), inference(subsumption_resolution, [], [f616, f296])).
fof(f616, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP6 | sP7 | ~ sP11), inference(duplicate_literal_removal, [], [f458])).
fof(f458, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP6 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP7 | ~ sP11), inference(cnf_transformation, [], [f285])).
fof(f285, plain, (((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP6 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP7) & ~ gt(pv5, n0)) | ~ sP11), inference(nnf_transformation, [], [f216])).
fof(f837, plain, (~ spl49_5 | spl49_18 | spl49_19 | spl49_20 | ~ spl49_21 | ~ spl49_22 | ~ spl49_23 | ~ spl49_24 | ~ spl49_25 | ~ spl49_26 | ~ spl49_2 | ~ spl49_27 | ~ spl49_28 | ~ spl49_29 | ~ spl49_30 | ~ spl49_31 | ~ spl49_32 | ~ spl49_33 | ~ spl49_34 | ~ spl49_35 | ~ spl49_36 | ~ spl49_37 | ~ spl49_38 | ~ spl49_39 | ~ spl49_40 | ~ spl49_3 | ~ spl49_41), inference(avatar_split_clause, [], [f740, f834, f645, f830, f826, f822, f818, f814, f810, f806, f802, f798, f794, f790, f786, f782, f778, f641, f774, f770, f766, f762, f758, f754, f750, f746, f742, f656])).
fof(f656, plain, (spl49_5 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl49_5])])).
fof(f740, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP9 | sP10 | sP11 | ~ sP12), inference(subsumption_resolution, [], [f739, f297])).
fof(f739, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP9 | sP10 | sP11 | ~ sP12), inference(subsumption_resolution, [], [f738, f297])).
fof(f738, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP9 | sP10 | sP11 | ~ sP12), inference(subsumption_resolution, [], [f737, f297])).
fof(f737, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | sP9 | sP10 | sP11 | ~ sP12), inference(subsumption_resolution, [], [f736, f296])).
fof(f736, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP9 | sP10 | sP11 | ~ sP12), inference(subsumption_resolution, [], [f622, f296])).
fof(f622, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP9 | sP10 | sP11 | ~ sP12), inference(duplicate_literal_removal, [], [f451])).
fof(f451, plain, (~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP9 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP10 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | sP11 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP12), inference(cnf_transformation, [], [f284])).
fof(f284, plain, (((((((((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | ((~ leq(pv5, minus(n1000, n1)) | ~ leq(pv5, n588) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & ~ gt(pv5, n0)) | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP9 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, n7) | ~ leq(n4, n7) | ~ leq(n3, minus(n6, n1)) | ~ leq(n3, n7) | ~ leq(n2, n7) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP10 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n7, n7) | ~ leq(n6, n7) | ~ leq(n5, minus(n6, n1)) | ~ leq(n5, n7) | ~ leq(n4, minus(n6, n1)) | ~ leq(n4, n7) | ~ leq(n3, n7) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, n7) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, n7) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n7) | ~ leq(n0, n6) | ~ leq(n0, n5) | ~ leq(n0, n4) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0)) & gt(pv5, n0)) | sP11 | ~ leq(pv5, minus(n1000, n1)) | ~ leq(n3, minus(n4, n1)) | ~ leq(n2, minus(n6, n1)) | ~ leq(n2, minus(n4, n1)) | ~ leq(n1, minus(n6, n1)) | ~ leq(n1, minus(n4, n1)) | ~ leq(n0, minus(n6, n1)) | ~ leq(n0, minus(n4, n1)) | ~ leq(n0, pv5) | ~ leq(n0, n3) | ~ leq(n0, n2) | ~ leq(n0, n1) | ~ leq(n0, n0) | ~ sP12), inference(nnf_transformation, [], [f217])).
fof(f735, plain, (~ spl49_14 | spl49_2), inference(avatar_split_clause, [], [f438, f641, f714])).
fof(f438, plain, (leq(n0, pv5) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f283, plain, ((sP4 & leq(pv32, minus(n6, n1)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv32) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP13), inference(nnf_transformation, [], [f218])).
fof(f734, plain, (~ spl49_14 | spl49_8), inference(avatar_split_clause, [], [f439, f677, f714])).
fof(f439, plain, (leq(n0, pv31) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f733, plain, (~ spl49_14 | spl49_17), inference(avatar_split_clause, [], [f440, f730, f714])).
fof(f440, plain, (leq(n0, pv32) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f728, plain, (~ spl49_14 | spl49_3), inference(avatar_split_clause, [], [f441, f645, f714])).
fof(f441, plain, (leq(pv5, n588) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f727, plain, (~ spl49_14 | spl49_7), inference(avatar_split_clause, [], [f442, f671, f714])).
fof(f442, plain, (leq(pv31, minus(n6, n1)) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f726, plain, (~ spl49_14 | spl49_16), inference(avatar_split_clause, [], [f443, f723, f714])).
fof(f443, plain, (leq(pv32, minus(n6, n1)) | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f721, plain, (~ spl49_14 | spl49_15), inference(avatar_split_clause, [], [f444, f718, f714])).
fof(f444, plain, (sP4 | ~ sP13), inference(cnf_transformation, [], [f283])).
fof(f712, plain, (~ spl49_10 | spl49_2), inference(avatar_split_clause, [], [f433, f641, f693])).
fof(f433, plain, (leq(n0, pv5) | ~ sP14), inference(cnf_transformation, [], [f282])).
fof(f282, plain, (((~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0)) & leq(pv21, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv21) & leq(n0, pv5)) | ~ sP14), inference(nnf_transformation, [], [f219])).
fof(f711, plain, (~ spl49_10 | spl49_11), inference(avatar_split_clause, [], [f434, f697, f693])).
fof(f697, plain, (spl49_11 <=> leq(n0, pv21)), introduced(avatar_definition, [new_symbols(naming, [spl49_11])])).
fof(f434, plain, (leq(n0, pv21) | ~ sP14), inference(cnf_transformation, [], [f282])).
fof(f710, plain, (~ spl49_10 | spl49_3), inference(avatar_split_clause, [], [f435, f645, f693])).
fof(f435, plain, (leq(pv5, n588) | ~ sP14), inference(cnf_transformation, [], [f282])).
fof(f709, plain, (~ spl49_10 | spl49_13), inference(avatar_split_clause, [], [f436, f705, f693])).
fof(f436, plain, (leq(pv21, minus(n6, n1)) | ~ sP14), inference(cnf_transformation, [], [f282])).
fof(f708, plain, (~ spl49_10 | ~ spl49_2 | ~ spl49_11 | ~ spl49_3 | ~ spl49_12 | ~ spl49_13), inference(avatar_split_clause, [], [f691, f705, f701, f645, f697, f641, f693])).
fof(f691, plain, (~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ sP14), inference(subsumption_resolution, [], [f437, f296])).
fof(f437, plain, (~ leq(pv21, minus(n6, n1)) | ~ leq(pv21, n5) | ~ leq(pv5, n588) | ~ leq(n0, pv21) | ~ leq(n0, pv5) | ~ leq(n0, n0) | ~ sP14), inference(cnf_transformation, [], [f282])).
fof(f690, plain, (~ spl49_9 | spl49_2), inference(avatar_split_clause, [], [f428, f641, f683])).
fof(f428, plain, (leq(n0, pv5) | ~ sP15), inference(cnf_transformation, [], [f281])).
fof(f281, plain, (((~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP15), inference(nnf_transformation, [], [f220])).
fof(f689, plain, (~ spl49_9 | spl49_8), inference(avatar_split_clause, [], [f429, f677, f683])).
fof(f429, plain, (leq(n0, pv31) | ~ sP15), inference(cnf_transformation, [], [f281])).
fof(f688, plain, (~ spl49_9 | spl49_3), inference(avatar_split_clause, [], [f430, f645, f683])).
fof(f430, plain, (leq(pv5, n588) | ~ sP15), inference(cnf_transformation, [], [f281])).
fof(f687, plain, (~ spl49_9 | spl49_7), inference(avatar_split_clause, [], [f431, f671, f683])).
fof(f431, plain, (leq(pv31, minus(n6, n1)) | ~ sP15), inference(cnf_transformation, [], [f281])).
fof(f686, plain, (~ spl49_9 | ~ spl49_2 | ~ spl49_8 | ~ spl49_3 | ~ spl49_7), inference(avatar_split_clause, [], [f432, f671, f645, f677, f641, f683])).
fof(f432, plain, (~ leq(pv31, minus(n6, n1)) | ~ leq(pv5, n588) | ~ leq(n0, pv31) | ~ leq(n0, pv5) | ~ sP15), inference(cnf_transformation, [], [f281])).
fof(f681, plain, (~ spl49_6 | spl49_2), inference(avatar_split_clause, [], [f423, f641, f666])).
fof(f423, plain, (leq(n0, pv5) | ~ sP16), inference(cnf_transformation, [], [f280])).
fof(f280, plain, (((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv31, minus(n6, n1)) & leq(pv5, n588) & leq(n0, pv31) & leq(n0, pv5)) | ~ sP16), inference(nnf_transformation, [], [f221])).
fof(f675, plain, (~ spl49_6 | spl49_3), inference(avatar_split_clause, [], [f425, f645, f666])).
fof(f425, plain, (leq(pv5, n588) | ~ sP16), inference(cnf_transformation, [], [f280])).
fof(f669, plain, (~ spl49_6 | ~ spl49_2 | ~ spl49_3), inference(avatar_split_clause, [], [f427, f645, f641, f666])).
fof(f427, plain, (~ leq(pv5, n588) | ~ leq(n0, pv5) | ~ sP16), inference(cnf_transformation, [], [f280])).
fof(f661, plain, (~ spl49_4 | spl49_2), inference(avatar_split_clause, [], [f408, f641, f652])).
fof(f408, plain, (leq(n0, pv5) | ~ sP20), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ((sP12 & leq(pv5, n588) & leq(n0, pv5)) | ~ sP20), inference(nnf_transformation, [], [f225])).
fof(f660, plain, (~ spl49_4 | spl49_3), inference(avatar_split_clause, [], [f409, f645, f652])).
fof(f409, plain, (leq(pv5, n588) | ~ sP20), inference(cnf_transformation, [], [f276])).
fof(f659, plain, (~ spl49_4 | spl49_5), inference(avatar_split_clause, [], [f410, f656, f652])).
fof(f410, plain, (sP12 | ~ sP20), inference(cnf_transformation, [], [f276])).
fof(f650, plain, (~ spl49_1 | spl49_2), inference(avatar_split_clause, [], [f405, f641, f637])).
fof(f405, plain, (leq(n0, pv5) | ~ sP21), inference(cnf_transformation, [], [f275])).
fof(f275, plain, (((~ leq(pv5, n588) | ~ leq(n0, pv5)) & leq(pv5, n588) & leq(n0, pv5)) | ~ sP21), inference(nnf_transformation, [], [f226])).
fof(f649, plain, (~ spl49_1 | spl49_3), inference(avatar_split_clause, [], [f406, f645, f637])).
fof(f406, plain, (leq(pv5, n588) | ~ sP21), inference(cnf_transformation, [], [f275])).
fof(f648, plain, (~ spl49_1 | ~ spl49_2 | ~ spl49_3), inference(avatar_split_clause, [], [f407, f645, f641, f637])).
fof(f407, plain, (~ leq(pv5, n588) | ~ leq(n0, pv5) | ~ sP21), inference(cnf_transformation, [], [f275])).