fof(f1137676, plain, $false, inference(avatar_sat_refutation, [], [f1056, f1059, f1114, f1119, f1168, f1178, f1203, f1358, f1539, f1544, f2127, f2136, f2219, f2228, f2247, f2300, f2313, f2318, f2347, f2495, f5395, f10035, f10739, f56665, f56723, f58879, f59001, f59021, f59030, f59595, f65761, f65878, f66059, f67344, f67421, f67757, f67758, f79053, f162693, f163116, f163124, f298237, f955782, f962830, f984228, f989490, f1012228, f1083752, f1112239, f1119178, f1122385, f1122422, f1123456, f1131017, f1133475, f1136005, f1136532, f1137167])).
fof(f1137167, plain, (spl158_5555 | ~ spl158_255 | ~ spl158_2115), inference(avatar_split_clause, [], [f1137090, f931479, f2297, f1112236])).
fof(f1112236, plain, (spl158_5555 <=> v__1(node_next(sK84), nn, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_5555])])).
fof(f2297, plain, (spl158_255 <=> (nn = sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_255])])).
fof(f931479, plain, (spl158_2115 <=> v__1(node_next(sK84), sK157, sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_2115])])).
fof(f1137090, plain, (v__1(node_next(sK84), nn, nn) | (~ spl158_255 | ~ spl158_2115)), inference(forward_demodulation, [], [f931481, f2299])).
fof(f2299, plain, ((nn = sK157) | ~ spl158_255), inference(avatar_component_clause, [], [f2297])).
fof(f931481, plain, (v__1(node_next(sK84), sK157, sK157) | ~ spl158_2115), inference(avatar_component_clause, [], [f931479])).
fof(f1136532, plain, (~ spl158_1 | ~ spl158_255 | ~ spl158_905 | spl158_2115 | ~ spl158_4475), inference(avatar_contradiction_clause, [], [f1136531])).
fof(f1136531, plain, ($false | (~ spl158_1 | ~ spl158_255 | ~ spl158_905 | spl158_2115 | ~ spl158_4475)), inference(subsumption_resolution, [], [f1136530, f1051])).
fof(f1051, plain, (v__1(nn, nn, nn) | ~ spl158_1), inference(avatar_component_clause, [], [f1049])).
fof(f1049, plain, (spl158_1 <=> v__1(nn, nn, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_1])])).
fof(f1136530, plain, (~ v__1(nn, nn, nn) | (~ spl158_255 | ~ spl158_905 | spl158_2115 | ~ spl158_4475)), inference(forward_demodulation, [], [f1136529, f994543])).
fof(f994543, plain, ((nn = node_next(nn)) | ~ spl158_4475), inference(avatar_component_clause, [], [f994541])).
fof(f994541, plain, (spl158_4475 <=> (nn = node_next(nn))), introduced(avatar_definition, [new_symbols(naming, [spl158_4475])])).
fof(f1136529, plain, (~ v__1(node_next(nn), nn, nn) | (~ spl158_255 | ~ spl158_905 | spl158_2115)), inference(forward_demodulation, [], [f1136528, f74115])).
fof(f74115, plain, ((nn = sK84) | ~ spl158_905), inference(avatar_component_clause, [], [f74113])).
fof(f74113, plain, (spl158_905 <=> (nn = sK84)), introduced(avatar_definition, [new_symbols(naming, [spl158_905])])).
fof(f1136528, plain, (~ v__1(node_next(sK84), nn, nn) | (~ spl158_255 | spl158_2115)), inference(forward_demodulation, [], [f931480, f2299])).
fof(f931480, plain, (~ v__1(node_next(sK84), sK157, sK157) | spl158_2115), inference(avatar_component_clause, [], [f931479])).
fof(f1136005, plain, (spl158_4475 | ~ spl158_24 | ~ spl158_905), inference(avatar_split_clause, [], [f1135271, f74113, f1165, f994541])).
fof(f1165, plain, (spl158_24 <=> (node_next(nn) = sK84)), introduced(avatar_definition, [new_symbols(naming, [spl158_24])])).
fof(f1135271, plain, ((nn = node_next(nn)) | (~ spl158_24 | ~ spl158_905)), inference(forward_demodulation, [], [f1167, f74115])).
fof(f1167, plain, ((node_next(nn) = sK84) | ~ spl158_24), inference(avatar_component_clause, [], [f1165])).
fof(f1133475, plain, (~ spl158_13 | ~ spl158_2 | ~ spl158_10 | ~ spl158_525), inference(avatar_split_clause, [], [f1133352, f56605, f1104, f1053, f1116])).
fof(f1116, plain, (spl158_13 <=> v__1(sK80, prev_2, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_13])])).
fof(f1053, plain, (spl158_2 <=> v__1(sortedList_first, nn, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_2])])).
fof(f1104, plain, (spl158_10 <=> (node_next(nn) = sK80)), introduced(avatar_definition, [new_symbols(naming, [spl158_10])])).
fof(f56605, plain, (spl158_525 <=> object(prev_2)), introduced(avatar_definition, [new_symbols(naming, [spl158_525])])).
fof(f1133352, plain, (~ v__1(sK80, prev_2, nn) | (~ spl158_2 | ~ spl158_10 | ~ spl158_525)), inference(forward_demodulation, [], [f545968, f1106])).
fof(f1106, plain, ((node_next(nn) = sK80) | ~ spl158_10), inference(avatar_component_clause, [], [f1104])).
fof(f545968, plain, (~ v__1(node_next(nn), prev_2, nn) | (~ spl158_2 | ~ spl158_525)), inference(subsumption_resolution, [], [f545967, f56606])).
fof(f56606, plain, (object(prev_2) | ~ spl158_525), inference(avatar_component_clause, [], [f56605])).
fof(f545967, plain, (~ v__1(node_next(nn), prev_2, nn) | ~ object(prev_2) | ~ spl158_2), inference(subsumption_resolution, [], [f545966, f496])).
fof(f496, plain, ~ (nn = null), inference(cnf_transformation, [], [f28])).
fof(f28, plain, ~ (nn = null), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_14)).
fof(f545966, plain, (~ v__1(node_next(nn), prev_2, nn) | (nn = null) | ~ object(prev_2) | ~ spl158_2), inference(subsumption_resolution, [], [f545944, f1054])).
fof(f1054, plain, (v__1(sortedList_first, nn, nn) | ~ spl158_2), inference(avatar_component_clause, [], [f1053])).
fof(f545944, plain, (~ v__1(node_next(nn), prev_2, nn) | ~ v__1(sortedList_first, nn, nn) | (nn = null) | ~ object(prev_2)), inference(superposition, [], [f72776, f1060])).
fof(f1060, plain, (nn = node_next(prev_2)), inference(subsumption_resolution, [], [f504, f490])).
fof(f490, plain, ~ (prev_2 = null), inference(cnf_transformation, [], [f26])).
fof(f26, plain, ~ (prev_2 = null), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_12)).
fof(f504, plain, ((nn = node_next(prev_2)) | (prev_2 = null)), inference(cnf_transformation, [], [f34])).
fof(f34, plain, ((nn = node_next(prev_2)) | (prev_2 = null)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_20)).
fof(f72776, plain, ! [X0] : (~ v__1(node_next(node_next(X0)), X0, node_next(X0)) | ~ v__1(sortedList_first, node_next(X0), node_next(X0)) | (null = node_next(X0)) | ~ object(X0)), inference(subsumption_resolution, [], [f72775, f475])).
fof(f475, plain, ! [X0] : object(node_next(X0)), inference(cnf_transformation, [], [f64])).
fof(f64, plain, ! [X0] : object(node_next(X0)), inference(rectify, [], [f14])).
fof(f14, plain, ! [X3] : object(node_next(X3)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_node_next)).
fof(f72775, plain, ! [X0] : ((null = node_next(X0)) | ~ v__1(sortedList_first, node_next(X0), node_next(X0)) | ~ v__1(node_next(node_next(X0)), X0, node_next(X0)) | ~ object(X0) | ~ object(node_next(node_next(X0)))), inference(subsumption_resolution, [], [f72694, f475])).
fof(f72694, plain, ! [X0] : ((null = node_next(X0)) | ~ v__1(sortedList_first, node_next(X0), node_next(X0)) | ~ object(node_next(X0)) | ~ v__1(node_next(node_next(X0)), X0, node_next(X0)) | ~ object(X0) | ~ object(node_next(node_next(X0)))), inference(resolution, [], [f61407, f7036])).
fof(f7036, plain, ! [X159, X160] : (v__1(X159, node_next(X160), node_next(X160)) | ~ v__1(X159, X160, node_next(X160)) | ~ object(X160) | ~ object(X159)), inference(subsumption_resolution, [], [f6903, f475])).
fof(f6903, plain, ! [X159, X160] : (v__1(X159, node_next(X160), node_next(X160)) | ~ v__1(X159, X160, node_next(X160)) | ~ object(node_next(X160)) | ~ object(X160) | ~ object(X159)), inference(duplicate_literal_removal, [], [f6888])).
fof(f6888, plain, ! [X159, X160] : (v__1(X159, node_next(X160), node_next(X160)) | ~ v__1(X159, X160, node_next(X160)) | ~ object(node_next(X160)) | ~ object(X160) | ~ object(X159) | ~ object(node_next(X160)) | ~ object(X160)), inference(resolution, [], [f479, f1042])).
fof(f1042, plain, ! [X0] : (v__1(X0, node_next(X0), node_next(X0)) | ~ object(X0)), inference(subsumption_resolution, [], [f1040, f475])).
fof(f1040, plain, ! [X0] : (v__1(X0, node_next(X0), node_next(X0)) | ~ object(node_next(X0)) | ~ object(X0)), inference(duplicate_literal_removal, [], [f994])).
fof(f994, plain, ! [X0] : (v__1(X0, node_next(X0), node_next(X0)) | ~ object(node_next(X0)) | ~ object(node_next(X0)) | ~ object(X0)), inference(equality_resolution, [], [f993])).
fof(f993, plain, ! [X0, X1] : (v__1(X0, X1, node_next(X0)) | ~ (node_next(X0) = X1) | ~ object(node_next(X0)) | ~ object(X1) | ~ object(X0)), inference(equality_resolution, [], [f488])).
fof(f488, plain, ! [X2, X0, X1] : (v__1(X0, X1, X2) | ~ (node_next(X0) = X2) | ~ (node_next(X0) = X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f110])).
fof(f110, plain, ! [X0] : (! [X1, X2] : (v__1(X0, X1, X2) | ~ (node_next(X0) = X2) | ~ (node_next(X0) = X1) | ~ object(X2) | ~ object(X1)) | ~ object(X0)), inference(flattening, [], [f109])).
fof(f109, plain, ! [X0] : (! [X1, X2] : (v__1(X0, X1, X2) | ~ (node_next(X0) = X2) | ~ (node_next(X0) = X1) | (~ object(X2) | ~ object(X1))) | ~ object(X0)), inference(ennf_transformation, [], [f73])).
fof(f73, plain, ! [X0] : (! [X1, X2] : (v__1(X0, X1, X2) | ~ (node_next(X0) = X2) | ~ (node_next(X0) = X1) | ~ (object(X2) & object(X1))) | ~ object(X0)), inference(rectify, [], [f24])).
fof(f24, plain, ! [X27] : (! [X28, X29] : (v__1(X27, X28, X29) | ~ (node_next(X27) = X29) | ~ (node_next(X27) = X28) | ~ (object(X29) & object(X28))) | ~ object(X27)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_10)).
fof(f479, plain, ! [X2, X0, X3, X1] : (~ v__1(X2, X0, X3) | v__1(X1, X0, X3) | ~ v__1(X1, X2, X3) | ~ object(X3) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f98])).
fof(f98, plain, ! [X0, X1, X2, X3] : ((v__1(X1, X0, X3) & v__1(X1, X2, X0)) | ~ v__1(X2, X0, X3) | ~ v__1(X1, X2, X3) | ~ object(X3) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f97])).
fof(f97, plain, ! [X0, X1, X2, X3] : ((v__1(X1, X0, X3) & v__1(X1, X2, X0)) | ~ v__1(X2, X0, X3) | ~ v__1(X1, X2, X3) | (~ object(X3) | ~ object(X2) | ~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f66])).
fof(f66, plain, ! [X0, X1, X2, X3] : ((v__1(X1, X0, X3) & v__1(X1, X2, X0)) | ~ v__1(X2, X0, X3) | ~ v__1(X1, X2, X3) | ~ (object(X3) & object(X2) & object(X1) & object(X0))), inference(rectify, [], [f17])).
fof(f17, plain, ! [X8, X9, X10, X11] : ((v__1(X9, X8, X11) & v__1(X9, X10, X8)) | ~ v__1(X10, X8, X11) | ~ v__1(X9, X10, X11) | ~ (object(X11) & object(X10) & object(X9) & object(X8))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_3)).
fof(f61407, plain, ! [X0] : (~ v__1(node_next(X0), X0, X0) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X0)), inference(duplicate_literal_removal, [], [f61406])).
fof(f61406, plain, ! [X0] : ((null = X0) | ~ v__1(node_next(X0), X0, X0) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X0) | ~ object(X0)), inference(equality_resolution, [], [f1062])).
fof(f1062, plain, ! [X0, X1] : (~ (node_key(X0) = node_key(X1)) | (null = X1) | ~ v__1(node_next(X0), X1, X1) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(subsumption_resolution, [], [f1061, f475])).
fof(f1061, plain, ! [X0, X1] : (~ (node_key(X0) = node_key(X1)) | (null = X1) | ~ v__1(node_next(X0), X1, X1) | ~ object(node_next(X0)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(subsumption_resolution, [], [f997, f471])).
fof(f471, plain, ! [X0] : integer(node_key(X0)), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ! [X0] : integer(node_key(X0)), inference(rectify, [], [f10])).
fof(f10, plain, ! [X3] : integer(node_key(X3)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_node_key)).
fof(f997, plain, ! [X0, X1] : (~ (node_key(X0) = node_key(X1)) | ~ integer(node_key(X1)) | (null = X1) | ~ v__1(node_next(X0), X1, X1) | ~ object(node_next(X0)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(equality_resolution, [], [f996])).
fof(f996, plain, ! [X0, X5, X1] : (~ (node_key(X0) = node_key(X1)) | ~ integer(node_key(X1)) | (null = X1) | ~ (node_next(X0) = X5) | ~ v__1(X5, X1, X1) | ~ object(X5) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(equality_resolution, [], [f515])).
fof(f515, plain, ! [X2, X0, X5, X1] : (~ (node_key(X1) = X2) | ~ (node_key(X0) = X2) | ~ integer(X2) | (null = X1) | ~ (node_next(X0) = X5) | ~ v__1(X5, X1, X1) | ~ object(X5) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f114])).
fof(f114, plain, ! [X0, X1] : ((! [X2] : (~ (node_key(X1) = X2) | ~ (node_key(X0) = X2) | ~ integer(X2)) & ! [X3, X4] : (lteq(X3, X4) | ~ (node_key(X1) = X4) | ~ (node_key(X0) = X3) | ~ integer(X4) | ~ integer(X3))) | (null = X1) | ! [X5] : (~ (node_next(X0) = X5) | ~ v__1(X5, X1, X1) | ~ object(X5)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f113])).
fof(f113, plain, ! [X0, X1] : ((! [X2] : (~ (node_key(X1) = X2) | ~ (node_key(X0) = X2) | ~ integer(X2)) & ! [X3, X4] : (lteq(X3, X4) | ~ (node_key(X1) = X4) | ~ (node_key(X0) = X3) | (~ integer(X4) | ~ integer(X3)))) | (null = X1) | ! [X5] : (~ (node_next(X0) = X5) | ~ v__1(X5, X1, X1) | ~ object(X5)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | (~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f77])).
fof(f77, plain, ! [X0, X1] : ((! [X2] : (~ (node_key(X1) = X2) | ~ (node_key(X0) = X2) | ~ integer(X2)) & ! [X3, X4] : (lteq(X3, X4) | ~ (node_key(X1) = X4) | ~ (node_key(X0) = X3) | ~ (integer(X4) & integer(X3)))) | (null = X1) | ! [X5] : (~ (node_next(X0) = X5) | ~ v__1(X5, X1, X1) | ~ object(X5)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ (object(X1) & object(X0))), inference(rectify, [], [f44])).
fof(f44, plain, ! [X0, X1] : ((! [X39] : (~ (node_key(X1) = X39) | ~ (node_key(X0) = X39) | ~ integer(X39)) & ! [X37, X38] : (lteq(X37, X38) | ~ (node_key(X1) = X38) | ~ (node_key(X0) = X37) | ~ (integer(X38) & integer(X37)))) | (null = X1) | ! [X36] : (~ (node_next(X0) = X36) | ~ v__1(X36, X1, X1) | ~ object(X36)) | (null = X0) | ~ v__1(sortedList_first, X0, X0) | ~ (object(X1) & object(X0))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_30)).
fof(f1131017, plain, (~ spl158_101 | ~ spl158_2 | ~ spl158_100 | ~ spl158_525), inference(avatar_split_clause, [], [f1127541, f56605, f1536, f1053, f1541])).
fof(f1541, plain, (spl158_101 <=> v__1(sK114, prev_2, prev_2)), introduced(avatar_definition, [new_symbols(naming, [spl158_101])])).
fof(f1536, plain, (spl158_100 <=> (node_next(nn) = sK114)), introduced(avatar_definition, [new_symbols(naming, [spl158_100])])).
fof(f1127541, plain, (~ v__1(sK114, prev_2, prev_2) | (~ spl158_2 | ~ spl158_100 | ~ spl158_525)), inference(backward_demodulation, [], [f487055, f1538])).
fof(f1538, plain, ((node_next(nn) = sK114) | ~ spl158_100), inference(avatar_component_clause, [], [f1536])).
fof(f487055, plain, (~ v__1(node_next(nn), prev_2, prev_2) | (~ spl158_2 | ~ spl158_525)), inference(subsumption_resolution, [], [f487054, f56606])).
fof(f487054, plain, (~ v__1(node_next(nn), prev_2, prev_2) | ~ object(prev_2) | ~ spl158_2), inference(subsumption_resolution, [], [f487053, f496])).
fof(f487053, plain, (~ v__1(node_next(nn), prev_2, prev_2) | (nn = null) | ~ object(prev_2) | ~ spl158_2), inference(subsumption_resolution, [], [f486940, f1054])).
fof(f486940, plain, (~ v__1(node_next(nn), prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn) | (nn = null) | ~ object(prev_2)), inference(superposition, [], [f72778, f1060])).
fof(f72778, plain, ! [X1] : (~ v__1(node_next(node_next(X1)), X1, X1) | ~ v__1(sortedList_first, node_next(X1), node_next(X1)) | (null = node_next(X1)) | ~ object(X1)), inference(subsumption_resolution, [], [f72777, f475])).
fof(f72777, plain, ! [X1] : ((null = node_next(X1)) | ~ v__1(sortedList_first, node_next(X1), node_next(X1)) | ~ v__1(node_next(node_next(X1)), X1, X1) | ~ object(X1) | ~ object(node_next(node_next(X1)))), inference(subsumption_resolution, [], [f72695, f475])).
fof(f72695, plain, ! [X1] : ((null = node_next(X1)) | ~ v__1(sortedList_first, node_next(X1), node_next(X1)) | ~ object(node_next(X1)) | ~ v__1(node_next(node_next(X1)), X1, X1) | ~ object(X1) | ~ object(node_next(node_next(X1)))), inference(resolution, [], [f61407, f6563])).
fof(f6563, plain, ! [X112, X113] : (v__1(X112, node_next(X113), node_next(X113)) | ~ v__1(X112, X113, X113) | ~ object(X113) | ~ object(X112)), inference(subsumption_resolution, [], [f6520, f475])).
fof(f6520, plain, ! [X112, X113] : (v__1(X112, node_next(X113), node_next(X113)) | ~ v__1(X112, X113, X113) | ~ object(node_next(X113)) | ~ object(X113) | ~ object(X112)), inference(duplicate_literal_removal, [], [f6514])).
fof(f6514, plain, ! [X112, X113] : (v__1(X112, node_next(X113), node_next(X113)) | ~ v__1(X112, X113, X113) | ~ object(node_next(X113)) | ~ object(X113) | ~ object(X112) | ~ object(X113)), inference(resolution, [], [f477, f1042])).
fof(f477, plain, ! [X2, X0, X1] : (~ v__1(X1, X2, X2) | v__1(X0, X2, X2) | ~ v__1(X0, X1, X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f96])).
fof(f96, plain, ! [X0, X1, X2] : (v__1(X0, X2, X2) | ~ v__1(X1, X2, X2) | ~ v__1(X0, X1, X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f95])).
fof(f95, plain, ! [X0, X1, X2] : (v__1(X0, X2, X2) | ~ v__1(X1, X2, X2) | ~ v__1(X0, X1, X1) | (~ object(X2) | ~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f65])).
fof(f65, plain, ! [X0, X1, X2] : (v__1(X0, X2, X2) | ~ v__1(X1, X2, X2) | ~ v__1(X0, X1, X1) | ~ (object(X2) & object(X1) & object(X0))), inference(rectify, [], [f16])).
fof(f16, plain, ! [X5, X6, X7] : (v__1(X5, X7, X7) | ~ v__1(X6, X7, X7) | ~ v__1(X5, X6, X6) | ~ (object(X7) & object(X6) & object(X5))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_2)).
fof(f1123456, plain, (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | spl158_408 | ~ spl158_4461), inference(avatar_contradiction_clause, [], [f1123455])).
fof(f1123455, plain, ($false | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | spl158_408 | ~ spl158_4461)), inference(subsumption_resolution, [], [f1123156, f8395])).
fof(f8395, plain, (~ (null = sK83(sK157)) | spl158_408), inference(avatar_component_clause, [], [f8394])).
fof(f8394, plain, (spl158_408 <=> (null = sK83(sK157))), introduced(avatar_definition, [new_symbols(naming, [spl158_408])])).
fof(f1123156, plain, ((null = sK83(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4461)), inference(backward_demodulation, [], [f1114531, f993370])).
fof(f993370, plain, ((null = sK84) | ~ spl158_4461), inference(avatar_component_clause, [], [f993368])).
fof(f993368, plain, (spl158_4461 <=> (null = sK84)), introduced(avatar_definition, [new_symbols(naming, [spl158_4461])])).
fof(f1114531, plain, ((sK84 = sK83(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376)), inference(subsumption_resolution, [], [f1114530, f5552])).
fof(f5552, plain, (~ sP52(sK157) | spl158_375), inference(avatar_component_clause, [], [f5551])).
fof(f5551, plain, (spl158_375 <=> sP52(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_375])])).
fof(f1114530, plain, ((sK84 = sK83(sK157)) | sP52(sK157) | (~ spl158_24 | ~ spl158_259 | spl158_376)), inference(subsumption_resolution, [], [f1114498, f5556])).
fof(f5556, plain, (~ sP51(sK157) | spl158_376), inference(avatar_component_clause, [], [f5555])).
fof(f5555, plain, (spl158_376 <=> sP51(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_376])])).
fof(f1114498, plain, ((sK84 = sK83(sK157)) | sP51(sK157) | sP52(sK157) | (~ spl158_24 | ~ spl158_259)), inference(resolution, [], [f1112433, f2317])).
fof(f2317, plain, (sP74(sK157) | ~ spl158_259), inference(avatar_component_clause, [], [f2315])).
fof(f2315, plain, (spl158_259 <=> sP74(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_259])])).
fof(f1112433, plain, (! [X0] : (~ sP74(X0) | (sK83(X0) = sK84) | sP51(X0) | sP52(X0)) | ~ spl158_24), inference(forward_demodulation, [], [f565, f1167])).
fof(f565, plain, ! [X0] : (~ sP74(X0) | sP51(X0) | (node_next(nn) = sK83(X0)) | sP52(X0)), inference(cnf_transformation, [], [f212])).
fof(f212, plain, ! [X0] : (sP52(X0) | sP51(X0) | ((sP47(X0) | ((node_next(nn) = sK82(X0)) & v__1(sK82(X0), X0, nn) & object(sK82(X0)))) & ((node_next(nn) = sK83(X0)) & v__1(sK83(X0), X0, X0) & object(sK83(X0)))) | ~ sP74(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK82, sK83])], [f209, f211, f210])).
fof(f210, plain, ! [X0] : (? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1)) => ((node_next(nn) = sK82(X0)) & v__1(sK82(X0), X0, nn) & object(sK82(X0)))), introduced(choice_axiom, [])).
fof(f211, plain, ! [X0] : (? [X2] : ((node_next(nn) = X2) & v__1(X2, X0, X0) & object(X2)) => ((node_next(nn) = sK83(X0)) & v__1(sK83(X0), X0, X0) & object(sK83(X0)))), introduced(choice_axiom, [])).
fof(f209, plain, ! [X0] : (sP52(X0) | sP51(X0) | ((sP47(X0) | ? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1))) & ? [X2] : ((node_next(nn) = X2) & v__1(X2, X0, X0) & object(X2))) | ~ sP74(X0)), inference(rectify, [], [f208])).
fof(f208, plain, ! [X0] : (sP52(X0) | sP51(X0) | ((sP47(X0) | ? [X43] : ((node_next(nn) = X43) & v__1(X43, X0, nn) & object(X43))) & ? [X44] : ((node_next(nn) = X44) & v__1(X44, X0, X0) & object(X44))) | ~ sP74(X0)), inference(nnf_transformation, [], [f194])).
fof(f194, plain, ! [X0] : (sP52(X0) | sP51(X0) | ((sP47(X0) | ? [X43] : ((node_next(nn) = X43) & v__1(X43, X0, nn) & object(X43))) & ? [X44] : ((node_next(nn) = X44) & v__1(X44, X0, X0) & object(X44))) | ~ sP74(X0)), inference(usedef, [], [e194])).
fof(e194, plain, ! [X0] : (sP74(X0) <=> (sP52(X0) | sP51(X0) | ((sP47(X0) | ? [X43] : ((node_next(nn) = X43) & v__1(X43, X0, nn) & object(X43))) & ? [X44] : ((node_next(nn) = X44) & v__1(X44, X0, X0) & object(X44))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP74])])).
fof(f1122422, plain, (spl158_4461 | spl158_914 | ~ spl158_24), inference(avatar_split_clause, [], [f1122421, f1165, f75186, f993368])).
fof(f75186, plain, (spl158_914 <=> v__1(sortedList_first, sK84, sK84)), introduced(avatar_definition, [new_symbols(naming, [spl158_914])])).
fof(f1122421, plain, (v__1(sortedList_first, sK84, sK84) | (null = sK84) | ~ spl158_24), inference(subsumption_resolution, [], [f1122420, f467])).
fof(f467, plain, object(nn), inference(cnf_transformation, [], [f6])).
fof(f6, plain, object(nn), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_nn)).
fof(f1122420, plain, (v__1(sortedList_first, sK84, sK84) | (null = sK84) | ~ object(nn) | ~ spl158_24), inference(subsumption_resolution, [], [f1089726, f496])).
fof(f1089726, plain, (v__1(sortedList_first, sK84, sK84) | (null = sK84) | (nn = null) | ~ object(nn) | ~ spl158_24), inference(superposition, [], [f1066, f1167])).
fof(f1066, plain, ! [X0] : (v__1(sortedList_first, node_next(X0), node_next(X0)) | (null = node_next(X0)) | (null = X0) | ~ object(X0)), inference(subsumption_resolution, [], [f1001, f475])).
fof(f1001, plain, ! [X0] : (v__1(sortedList_first, node_next(X0), node_next(X0)) | (null = node_next(X0)) | (null = X0) | ~ object(node_next(X0)) | ~ object(X0)), inference(equality_resolution, [], [f517])).
fof(f517, plain, ! [X0, X1] : (v__1(sortedList_first, X1, X1) | ~ (node_next(X0) = X1) | (null = X1) | (null = X0) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f116])).
fof(f116, plain, ! [X0, X1] : ((v__1(sortedList_first, X1, X1) & ~ (null = X1)) | ~ (node_next(X0) = X1) | (null = X1) | (null = X0) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f115])).
fof(f115, plain, ! [X0, X1] : ((v__1(sortedList_first, X1, X1) & ~ (null = X1)) | ~ (node_next(X0) = X1) | (null = X1) | (null = X0) | (~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f78])).
fof(f78, plain, ! [X0, X1] : ((v__1(sortedList_first, X1, X1) & ~ (null = X1)) | ~ (node_next(X0) = X1) | (null = X1) | (null = X0) | ~ (object(X1) & object(X0))), inference(rectify, [], [f45])).
fof(f45, plain, ! [X40, X41] : ((v__1(sortedList_first, X41, X41) & ~ (null = X41)) | ~ (node_next(X40) = X41) | (null = X41) | (null = X40) | ~ (object(X41) & object(X40))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_31)).
fof(f1122385, plain, (spl158_2115 | ~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | spl158_809 | ~ spl158_4072), inference(avatar_split_clause, [], [f1122384, f982477, f68838, f5555, f5551, f2315, f1165, f931479])).
fof(f68838, plain, (spl158_809 <=> (sK84 = sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_809])])).
fof(f982477, plain, (spl158_4072 <=> object(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_4072])])).
fof(f1122384, plain, (v__1(node_next(sK84), sK157, sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | spl158_809 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1122383, f68839])).
fof(f68839, plain, (~ (sK84 = sK157) | spl158_809), inference(avatar_component_clause, [], [f68838])).
fof(f1122383, plain, (v__1(node_next(sK84), sK157, sK157) | (sK84 = sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1122382, f2317])).
fof(f1122382, plain, (v__1(node_next(sK84), sK157, sK157) | ~ sP74(sK157) | (sK84 = sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1122381, f5552])).
fof(f1122381, plain, (v__1(node_next(sK84), sK157, sK157) | sP52(sK157) | ~ sP74(sK157) | (sK84 = sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1122380, f5556])).
fof(f1122380, plain, (v__1(node_next(sK84), sK157, sK157) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157) | (sK84 = sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1114647, f982478])).
fof(f982478, plain, (object(sK157) | ~ spl158_4072), inference(avatar_component_clause, [], [f982477])).
fof(f1114647, plain, (v__1(node_next(sK84), sK157, sK157) | ~ object(sK157) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157) | (sK84 = sK157) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376)), inference(superposition, [], [f90996, f1114531])).
fof(f90996, plain, ! [X8] : (v__1(node_next(sK83(X8)), X8, X8) | ~ object(X8) | sP51(X8) | sP52(X8) | ~ sP74(X8) | (sK83(X8) = X8)), inference(subsumption_resolution, [], [f90995, f563])).
fof(f563, plain, ! [X0] : (object(sK83(X0)) | sP51(X0) | sP52(X0) | ~ sP74(X0)), inference(cnf_transformation, [], [f212])).
fof(f90995, plain, ! [X8] : ((sK83(X8) = X8) | ~ object(X8) | sP51(X8) | sP52(X8) | ~ sP74(X8) | v__1(node_next(sK83(X8)), X8, X8) | ~ object(sK83(X8))), inference(subsumption_resolution, [], [f90982, f475])).
fof(f90982, plain, ! [X8] : ((sK83(X8) = X8) | ~ object(X8) | sP51(X8) | sP52(X8) | ~ sP74(X8) | v__1(node_next(sK83(X8)), X8, X8) | ~ object(node_next(sK83(X8))) | ~ object(sK83(X8))), inference(duplicate_literal_removal, [], [f90978])).
fof(f90978, plain, ! [X8] : ((sK83(X8) = X8) | ~ object(X8) | sP51(X8) | sP52(X8) | ~ sP74(X8) | v__1(node_next(sK83(X8)), X8, X8) | ~ object(X8) | ~ object(node_next(sK83(X8))) | ~ object(sK83(X8))), inference(resolution, [], [f4084, f484])).
fof(f484, plain, ! [X2, X0, X1] : (~ v__1(X0, X1, X2) | v__1(X1, X2, X2) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f104])).
fof(f104, plain, ! [X0, X1, X2] : ((v__1(X1, X2, X2) & v__1(X0, X1, X1)) | ~ v__1(X0, X1, X2) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f103])).
fof(f103, plain, ! [X0, X1, X2] : ((v__1(X1, X2, X2) & v__1(X0, X1, X1)) | ~ v__1(X0, X1, X2) | (~ object(X2) | ~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f69])).
fof(f69, plain, ! [X0, X1, X2] : ((v__1(X1, X2, X2) & v__1(X0, X1, X1)) | ~ v__1(X0, X1, X2) | ~ (object(X2) & object(X1) & object(X0))), inference(rectify, [], [f20])).
fof(f20, plain, ! [X19, X20, X21] : ((v__1(X20, X21, X21) & v__1(X19, X20, X20)) | ~ v__1(X19, X20, X21) | ~ (object(X21) & object(X20) & object(X19))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_6)).
fof(f4084, plain, ! [X28] : (v__1(sK83(X28), node_next(sK83(X28)), X28) | (sK83(X28) = X28) | ~ object(X28) | sP51(X28) | sP52(X28) | ~ sP74(X28)), inference(subsumption_resolution, [], [f3903, f563])).
fof(f3903, plain, ! [X28] : ((sK83(X28) = X28) | v__1(sK83(X28), node_next(sK83(X28)), X28) | ~ object(X28) | ~ object(sK83(X28)) | sP51(X28) | sP52(X28) | ~ sP74(X28)), inference(resolution, [], [f1043, f564])).
fof(f564, plain, ! [X0] : (v__1(sK83(X0), X0, X0) | sP51(X0) | sP52(X0) | ~ sP74(X0)), inference(cnf_transformation, [], [f212])).
fof(f1043, plain, ! [X0, X1] : (~ v__1(X0, X1, X1) | (X0 = X1) | v__1(X0, node_next(X0), X1) | ~ object(X1) | ~ object(X0)), inference(subsumption_resolution, [], [f995, f475])).
fof(f995, plain, ! [X0, X1] : (v__1(X0, node_next(X0), X1) | ~ object(node_next(X0)) | (X0 = X1) | ~ v__1(X0, X1, X1) | ~ object(X1) | ~ object(X0)), inference(equality_resolution, [], [f489])).
fof(f489, plain, ! [X2, X0, X1] : (v__1(X0, X2, X1) | ~ (node_next(X0) = X2) | ~ object(X2) | (X0 = X1) | ~ v__1(X0, X1, X1) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f112])).
fof(f112, plain, ! [X0, X1] : (! [X2] : (v__1(X0, X2, X1) | ~ (node_next(X0) = X2) | ~ object(X2)) | (X0 = X1) | ~ v__1(X0, X1, X1) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f111])).
fof(f111, plain, ! [X0, X1] : (! [X2] : (v__1(X0, X2, X1) | ~ (node_next(X0) = X2) | ~ object(X2)) | (X0 = X1) | ~ v__1(X0, X1, X1) | (~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f74])).
fof(f74, plain, ! [X0, X1] : (! [X2] : (v__1(X0, X2, X1) | ~ (node_next(X0) = X2) | ~ object(X2)) | (X0 = X1) | ~ v__1(X0, X1, X1) | ~ (object(X1) & object(X0))), inference(rectify, [], [f25])).
fof(f25, plain, ! [X30, X31] : (! [X32] : (v__1(X30, X32, X31) | ~ (node_next(X30) = X32) | ~ object(X32)) | (X30 = X31) | ~ v__1(X30, X31, X31) | ~ (object(X31) & object(X30))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_11)).
fof(f1119178, plain, (~ spl158_24 | spl158_254 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_536 | ~ spl158_914 | ~ spl158_4072), inference(avatar_contradiction_clause, [], [f1119177])).
fof(f1119177, plain, ($false | (~ spl158_24 | spl158_254 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_536 | ~ spl158_914 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1119176, f56679])).
fof(f56679, plain, (object(sortedList_first) | ~ spl158_536), inference(avatar_component_clause, [], [f56678])).
fof(f56678, plain, (spl158_536 <=> object(sortedList_first)), introduced(avatar_definition, [new_symbols(naming, [spl158_536])])).
fof(f1119176, plain, (~ object(sortedList_first) | (~ spl158_24 | spl158_254 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_914 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1119117, f2295])).
fof(f2295, plain, (~ v__1(sortedList_first, sK157, sK157) | spl158_254), inference(avatar_component_clause, [], [f2293])).
fof(f2293, plain, (spl158_254 <=> v__1(sortedList_first, sK157, sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_254])])).
fof(f1119117, plain, (v__1(sortedList_first, sK157, sK157) | ~ object(sortedList_first) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_914 | ~ spl158_4072)), inference(resolution, [], [f75188, f1114722])).
fof(f1114722, plain, (! [X0] : (~ v__1(X0, sK84, sK84) | v__1(X0, sK157, sK157) | ~ object(X0)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1114721, f2317])).
fof(f1114721, plain, (! [X0] : (~ v__1(X0, sK84, sK84) | v__1(X0, sK157, sK157) | ~ object(X0) | ~ sP74(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1114720, f5552])).
fof(f1114720, plain, (! [X0] : (~ v__1(X0, sK84, sK84) | v__1(X0, sK157, sK157) | ~ object(X0) | sP52(sK157) | ~ sP74(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1114719, f5556])).
fof(f1114719, plain, (! [X0] : (~ v__1(X0, sK84, sK84) | v__1(X0, sK157, sK157) | ~ object(X0) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376 | ~ spl158_4072)), inference(subsumption_resolution, [], [f1114638, f982478])).
fof(f1114638, plain, (! [X0] : (~ v__1(X0, sK84, sK84) | v__1(X0, sK157, sK157) | ~ object(sK157) | ~ object(X0) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157)) | (~ spl158_24 | ~ spl158_259 | spl158_375 | spl158_376)), inference(superposition, [], [f6548, f1114531])).
fof(f6548, plain, ! [X76, X75] : (~ v__1(X75, sK83(X76), sK83(X76)) | v__1(X75, X76, X76) | ~ object(X76) | ~ object(X75) | sP51(X76) | sP52(X76) | ~ sP74(X76)), inference(subsumption_resolution, [], [f6483, f563])).
fof(f6483, plain, ! [X76, X75] : (v__1(X75, X76, X76) | ~ v__1(X75, sK83(X76), sK83(X76)) | ~ object(X76) | ~ object(sK83(X76)) | ~ object(X75) | sP51(X76) | sP52(X76) | ~ sP74(X76)), inference(resolution, [], [f477, f564])).
fof(f75188, plain, (v__1(sortedList_first, sK84, sK84) | ~ spl158_914), inference(avatar_component_clause, [], [f75186])).
fof(f1112239, plain, (~ spl158_914 | ~ spl158_5555 | ~ spl158_24 | spl158_4461), inference(avatar_split_clause, [], [f1112234, f993368, f1165, f1112236, f75186])).
fof(f1112234, plain, (~ v__1(node_next(sK84), nn, nn) | ~ v__1(sortedList_first, sK84, sK84) | (~ spl158_24 | spl158_4461)), inference(subsumption_resolution, [], [f1112233, f467])).
fof(f1112233, plain, (~ v__1(node_next(sK84), nn, nn) | ~ v__1(sortedList_first, sK84, sK84) | ~ object(nn) | (~ spl158_24 | spl158_4461)), inference(subsumption_resolution, [], [f1089812, f993369])).
fof(f993369, plain, (~ (null = sK84) | spl158_4461), inference(avatar_component_clause, [], [f993368])).
fof(f1089812, plain, (~ v__1(node_next(sK84), nn, nn) | ~ v__1(sortedList_first, sK84, sK84) | (null = sK84) | ~ object(nn) | ~ spl158_24), inference(superposition, [], [f72778, f1167])).
fof(f1083752, plain, (spl158_254 | ~ spl158_536 | ~ spl158_3508 | ~ spl158_4364), inference(avatar_contradiction_clause, [], [f1083751])).
fof(f1083751, plain, ($false | (spl158_254 | ~ spl158_536 | ~ spl158_3508 | ~ spl158_4364)), inference(subsumption_resolution, [], [f1083750, f2295])).
fof(f1083750, plain, (v__1(sortedList_first, sK157, sK157) | (~ spl158_536 | ~ spl158_3508 | ~ spl158_4364)), inference(subsumption_resolution, [], [f1083691, f56679])).
fof(f1083691, plain, (~ object(sortedList_first) | v__1(sortedList_first, sK157, sK157) | (~ spl158_3508 | ~ spl158_4364)), inference(resolution, [], [f961761, f989489])).
fof(f989489, plain, (! [X0] : (~ v__1(X0, sK97, sK97) | ~ object(X0) | v__1(X0, sK157, sK157)) | ~ spl158_4364), inference(avatar_component_clause, [], [f989488])).
fof(f989488, plain, (spl158_4364 <=> ! [X0] : (~ v__1(X0, sK97, sK97) | ~ object(X0) | v__1(X0, sK157, sK157))), introduced(avatar_definition, [new_symbols(naming, [spl158_4364])])).
fof(f961761, plain, (v__1(sortedList_first, sK97, sK97) | ~ spl158_3508), inference(avatar_component_clause, [], [f961759])).
fof(f961759, plain, (spl158_3508 <=> v__1(sortedList_first, sK97, sK97)), introduced(avatar_definition, [new_symbols(naming, [spl158_3508])])).
fof(f1012228, plain, (~ spl158_2 | ~ spl158_62 | ~ spl158_312), inference(avatar_contradiction_clause, [], [f1012227])).
fof(f1012227, plain, ($false | (~ spl158_2 | ~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f1012226, f4069])).
fof(f4069, plain, (sP74(nn) | ~ spl158_312), inference(avatar_component_clause, [], [f4068])).
fof(f4068, plain, (spl158_312 <=> sP74(nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_312])])).
fof(f1012226, plain, (~ sP74(nn) | (~ spl158_2 | ~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f1012225, f1017])).
fof(f1017, plain, ~ sP52(nn), inference(equality_resolution, [], [f685])).
fof(f685, plain, ! [X0] : (~ (nn = X0) | ~ sP52(X0)), inference(cnf_transformation, [], [f310])).
fof(f310, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP50(X0) | ((node_next(nn) = sK115(X0)) & v__1(sK115(X0), nn, X0) & object(sK115(X0)))) & ~ (nn = X0)) | ~ sP52(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK115])], [f308, f309])).
fof(f309, plain, ! [X0] : (? [X1] : ((node_next(nn) = X1) & v__1(X1, nn, X0) & object(X1)) => ((node_next(nn) = sK115(X0)) & v__1(sK115(X0), nn, X0) & object(sK115(X0)))), introduced(choice_axiom, [])).
fof(f308, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP50(X0) | ? [X1] : ((node_next(nn) = X1) & v__1(X1, nn, X0) & object(X1))) & ~ (nn = X0)) | ~ sP52(X0)), inference(rectify, [], [f307])).
fof(f307, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP50(X0) | ? [X36] : ((node_next(nn) = X36) & v__1(X36, nn, X0) & object(X36))) & ~ (nn = X0)) | ~ sP52(X0)), inference(nnf_transformation, [], [f172])).
fof(f172, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP50(X0) | ? [X36] : ((node_next(nn) = X36) & v__1(X36, nn, X0) & object(X36))) & ~ (nn = X0)) | ~ sP52(X0)), inference(usedef, [], [e172])).
fof(e172, plain, ! [X0] : (sP52(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP50(X0) | ? [X36] : ((node_next(nn) = X36) & v__1(X36, nn, X0) & object(X36))) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP52])])).
fof(f1012225, plain, (sP52(nn) | ~ sP74(nn) | (~ spl158_2 | ~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f1012224, f1018])).
fof(f1018, plain, ~ sP51(nn), inference(equality_resolution, [], [f692])).
fof(f692, plain, ! [X0] : (~ (nn = X0) | ~ sP51(X0)), inference(cnf_transformation, [], [f314])).
fof(f314, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ((node_next(nn) = sK116(X0)) & v__1(sK116(X0), X0, nn) & object(sK116(X0))) & sP49(X0) & ~ (nn = X0)) | ~ sP51(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK116])], [f312, f313])).
fof(f313, plain, ! [X0] : (? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1)) => ((node_next(nn) = sK116(X0)) & v__1(sK116(X0), X0, nn) & object(sK116(X0)))), introduced(choice_axiom, [])).
fof(f312, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1)) & sP49(X0) & ~ (nn = X0)) | ~ sP51(X0)), inference(rectify, [], [f311])).
fof(f311, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X37] : ((node_next(nn) = X37) & v__1(X37, X0, nn) & object(X37)) & sP49(X0) & ~ (nn = X0)) | ~ sP51(X0)), inference(nnf_transformation, [], [f171])).
fof(f171, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X37] : ((node_next(nn) = X37) & v__1(X37, X0, nn) & object(X37)) & sP49(X0) & ~ (nn = X0)) | ~ sP51(X0)), inference(usedef, [], [e171])).
fof(e171, plain, ! [X0] : (sP51(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X37] : ((node_next(nn) = X37) & v__1(X37, X0, nn) & object(X37)) & sP49(X0) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP51])])).
fof(f1012224, plain, (sP51(nn) | sP52(nn) | ~ sP74(nn) | (~ spl158_2 | ~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f1012127, f935151])).
fof(f935151, plain, (~ v__1(sK97, nn, nn) | (~ spl158_2 | ~ spl158_62)), inference(subsumption_resolution, [], [f935150, f467])).
fof(f935150, plain, (~ v__1(sK97, nn, nn) | ~ object(nn) | (~ spl158_2 | ~ spl158_62)), inference(subsumption_resolution, [], [f935149, f1054])).
fof(f935149, plain, (~ v__1(sK97, nn, nn) | ~ v__1(sortedList_first, nn, nn) | ~ object(nn) | ~ spl158_62), inference(subsumption_resolution, [], [f72696, f496])).
fof(f72696, plain, (~ v__1(sK97, nn, nn) | (nn = null) | ~ v__1(sortedList_first, nn, nn) | ~ object(nn) | ~ spl158_62), inference(superposition, [], [f61407, f1357])).
fof(f1357, plain, ((node_next(nn) = sK97) | ~ spl158_62), inference(avatar_component_clause, [], [f1355])).
fof(f1355, plain, (spl158_62 <=> (node_next(nn) = sK97)), introduced(avatar_definition, [new_symbols(naming, [spl158_62])])).
fof(f1012127, plain, (v__1(sK97, nn, nn) | sP51(nn) | sP52(nn) | ~ sP74(nn) | (~ spl158_62 | ~ spl158_312)), inference(superposition, [], [f564, f992913])).
fof(f992913, plain, ((sK97 = sK83(nn)) | (~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f992912, f1017])).
fof(f992912, plain, ((sK97 = sK83(nn)) | sP52(nn) | (~ spl158_62 | ~ spl158_312)), inference(subsumption_resolution, [], [f992908, f1018])).
fof(f992908, plain, ((sK97 = sK83(nn)) | sP51(nn) | sP52(nn) | (~ spl158_62 | ~ spl158_312)), inference(resolution, [], [f72917, f4069])).
fof(f72917, plain, (! [X0] : (~ sP74(X0) | (sK83(X0) = sK97) | sP51(X0) | sP52(X0)) | ~ spl158_62), inference(forward_demodulation, [], [f565, f1357])).
fof(f989490, plain, (~ spl158_259 | spl158_375 | spl158_4364 | spl158_18 | ~ spl158_1246 | ~ spl158_4072), inference(avatar_split_clause, [], [f989486, f982477, f162713, f1137, f989488, f5551, f2315])).
fof(f1137, plain, (spl158_18 <=> v__1(null, nn, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_18])])).
fof(f162713, plain, (spl158_1246 <=> (sK97 = sK83(sK157))), introduced(avatar_definition, [new_symbols(naming, [spl158_1246])])).
fof(f989486, plain, (! [X0] : (~ v__1(X0, sK97, sK97) | v__1(X0, sK157, sK157) | ~ object(X0) | sP52(sK157) | ~ sP74(sK157)) | (spl158_18 | ~ spl158_1246 | ~ spl158_4072)), inference(subsumption_resolution, [], [f976918, f982478])).
fof(f976918, plain, (! [X0] : (~ v__1(X0, sK97, sK97) | v__1(X0, sK157, sK157) | ~ object(sK157) | ~ object(X0) | sP52(sK157) | ~ sP74(sK157)) | (spl158_18 | ~ spl158_1246)), inference(subsumption_resolution, [], [f163567, f207272])).
fof(f207272, plain, (! [X56] : (~ sP51(X56) | ~ object(X56)) | spl158_18), inference(subsumption_resolution, [], [f207271, f472])).
fof(f472, plain, object(null), inference(cnf_transformation, [], [f11])).
fof(f11, plain, object(null), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_null)).
fof(f207271, plain, (! [X56] : (~ sP51(X56) | ~ object(X56) | ~ object(null)) | spl158_18), inference(subsumption_resolution, [], [f207118, f1138])).
fof(f1138, plain, (~ v__1(null, nn, nn) | spl158_18), inference(avatar_component_clause, [], [f1137])).
fof(f207118, plain, ! [X56] : (~ sP51(X56) | v__1(null, nn, nn) | ~ object(X56) | ~ object(null)), inference(duplicate_literal_removal, [], [f206274])).
fof(f206274, plain, ! [X56] : (~ sP51(X56) | v__1(null, nn, nn) | ~ object(X56) | ~ object(null) | ~ object(X56) | ~ sP51(X56)), inference(resolution, [], [f9134, f3422])).
fof(f3422, plain, ! [X27] : (v__1(null, X27, X27) | ~ object(X27) | ~ sP51(X27)), inference(subsumption_resolution, [], [f3421, f472])).
fof(f3421, plain, ! [X27] : (v__1(null, X27, X27) | ~ object(X27) | ~ object(null) | ~ sP51(X27)), inference(subsumption_resolution, [], [f3391, f467])).
fof(f3391, plain, ! [X27] : (v__1(null, X27, X27) | ~ object(nn) | ~ object(X27) | ~ object(null) | ~ sP51(X27)), inference(duplicate_literal_removal, [], [f3312])).
fof(f3312, plain, ! [X27] : (v__1(null, X27, X27) | ~ object(nn) | ~ object(X27) | ~ object(null) | v__1(null, X27, X27) | ~ sP51(X27)), inference(resolution, [], [f483, f697])).
fof(f697, plain, ! [X0] : (v__1(null, X0, nn) | v__1(null, X0, X0) | ~ sP51(X0)), inference(cnf_transformation, [], [f314])).
fof(f483, plain, ! [X2, X0, X1] : (~ v__1(X0, X1, X2) | v__1(X0, X1, X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f104])).
fof(f9134, plain, ! [X12, X13] : (~ v__1(X13, X12, X12) | ~ sP51(X12) | v__1(X13, nn, nn) | ~ object(X12) | ~ object(X13)), inference(subsumption_resolution, [], [f9123, f467])).
fof(f9123, plain, ! [X12, X13] : (~ object(X12) | ~ sP51(X12) | v__1(X13, nn, nn) | ~ v__1(X13, X12, X12) | ~ object(nn) | ~ object(X13)), inference(duplicate_literal_removal, [], [f9086])).
fof(f9086, plain, ! [X12, X13] : (~ object(X12) | ~ sP51(X12) | v__1(X13, nn, nn) | ~ v__1(X13, X12, X12) | ~ object(nn) | ~ object(X12) | ~ object(X13)), inference(resolution, [], [f3636, f477])).
fof(f3636, plain, ! [X58] : (v__1(X58, nn, nn) | ~ object(X58) | ~ sP51(X58)), inference(subsumption_resolution, [], [f3635, f694])).
fof(f694, plain, ! [X0] : (object(sK116(X0)) | ~ sP51(X0)), inference(cnf_transformation, [], [f314])).
fof(f3635, plain, ! [X58] : (v__1(X58, nn, nn) | ~ object(X58) | ~ object(sK116(X58)) | ~ sP51(X58)), inference(subsumption_resolution, [], [f3546, f467])).
fof(f3546, plain, ! [X58] : (v__1(X58, nn, nn) | ~ object(nn) | ~ object(X58) | ~ object(sK116(X58)) | ~ sP51(X58)), inference(resolution, [], [f484, f695])).
fof(f695, plain, ! [X0] : (v__1(sK116(X0), X0, nn) | ~ sP51(X0)), inference(cnf_transformation, [], [f314])).
fof(f163567, plain, (! [X0] : (~ v__1(X0, sK97, sK97) | v__1(X0, sK157, sK157) | ~ object(sK157) | ~ object(X0) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157)) | ~ spl158_1246), inference(superposition, [], [f6548, f162715])).
fof(f162715, plain, ((sK97 = sK83(sK157)) | ~ spl158_1246), inference(avatar_component_clause, [], [f162713])).
fof(f984228, plain, spl158_4072, inference(avatar_split_clause, [], [f983, f982477])).
fof(f983, plain, object(sK157), inference(cnf_transformation, [], [f459])).
fof(f459, plain, (((nn = sK157) | ~ v__1(sortedList_first, sK157, sK157) | (null = sK157)) & ((sP75(sK157) & sP74(sK157) & (sP44(sK157) | sP46(sK157) | sP45(sK157) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, sK157, sK157)) | v__1(sortedList_first, sK157, nn)) & v__1(sortedList_first, prev_2, sK157))) & ~ (prev_2 = sK157)) | sP76(sK157) | sP77(sK157)) & ~ (null = sK157) & object(sK157)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK157])], [f198, f458])).
fof(f458, plain, (? [X0] : (((nn = X0) | ~ v__1(sortedList_first, X0, X0) | (null = X0)) & ((sP75(X0) & sP74(X0) & (sP44(X0) | sP46(X0) | sP45(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | sP76(X0) | sP77(X0)) & ~ (null = X0) & object(X0)) => (((nn = sK157) | ~ v__1(sortedList_first, sK157, sK157) | (null = sK157)) & ((sP75(sK157) & sP74(sK157) & (sP44(sK157) | sP46(sK157) | sP45(sK157) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, sK157, sK157)) | v__1(sortedList_first, sK157, nn)) & v__1(sortedList_first, prev_2, sK157))) & ~ (prev_2 = sK157)) | sP76(sK157) | sP77(sK157)) & ~ (null = sK157) & object(sK157))), introduced(choice_axiom, [])).
fof(f198, plain, ? [X0] : (((nn = X0) | ~ v__1(sortedList_first, X0, X0) | (null = X0)) & ((sP75(X0) & sP74(X0) & (sP44(X0) | sP46(X0) | sP45(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | sP76(X0) | sP77(X0)) & ~ (null = X0) & object(X0)), inference(definition_folding, [], [f119, e197, e196, e195, e194, e193, e192, e191, e190, e189, e188, e187, e186, e185, e184, e183, e182, e181, e180, e179, e178, e177, e176, e175, e174, e173, e172, e171, e170, e169, e168, e167, e166, e165, e164, e163, e162, e161, e160, e159, e158, e157, e156, e155, e154, e153, e152, e151, e150, e149, e148, e147, e146, e145, e144, e143, e142, e141, e140, e139, e138, e137, e136, e135, e134, e133, e132, e131, e130, e129, e128, e127, e126, e125, e124, e123, e122, e121, e120])).
fof(f120, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP0(X0)), inference(usedef, [], [e120])).
fof(e120, plain, ! [X0] : (sP0(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f121, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP1(X0)), inference(usedef, [], [e121])).
fof(e121, plain, ! [X0] : (sP1(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f122, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP2(X0)), inference(usedef, [], [e122])).
fof(e122, plain, ! [X0] : (sP2(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f123, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP3(X0)), inference(usedef, [], [e123])).
fof(e123, plain, ! [X0] : (sP3(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f124, plain, (((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | (nn = prev_2) | ~ sP4), inference(usedef, [], [e124])).
fof(e124, plain, (sP4 <=> (((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f125, plain, (((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2) | ~ sP5), inference(usedef, [], [e125])).
fof(e125, plain, (sP5 <=> (((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f126, plain, ! [X0] : ((sP5 & sP4 & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)) & (sP3(X0) | sP2(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)))) | ~ sP6(X0)), inference(usedef, [], [e126])).
fof(e126, plain, ! [X0] : (sP6(X0) <=> (sP5 & sP4 & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)) & (sP3(X0) | sP2(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f127, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP7(X0)), inference(usedef, [], [e127])).
fof(e127, plain, ! [X0] : (sP7(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f128, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP8(X0)), inference(usedef, [], [e128])).
fof(e128, plain, ! [X0] : (sP8(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f129, plain, ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP9), inference(usedef, [], [e129])).
fof(e129, plain, (sP9 <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f130, plain, ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP10), inference(usedef, [], [e130])).
fof(e130, plain, (sP10 <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f131, plain, ! [X0] : (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0) | ~ sP11(X0)), inference(usedef, [], [e131])).
fof(e131, plain, ! [X0] : (sP11(X0) <=> (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f132, plain, ! [X0] : (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0) | ~ sP12(X0)), inference(usedef, [], [e132])).
fof(e132, plain, ! [X0] : (sP12(X0) <=> (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f133, plain, ! [X0] : ((sP12(X0) & sP11(X0) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & (sP10 | sP9 | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)))) | ~ sP13(X0)), inference(usedef, [], [e133])).
fof(e133, plain, ! [X0] : (sP13(X0) <=> (sP12(X0) & sP11(X0) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & (sP10 | sP9 | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f134, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP14(X0)), inference(usedef, [], [e134])).
fof(e134, plain, ! [X0] : (sP14(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f135, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP15(X0)), inference(usedef, [], [e135])).
fof(e135, plain, ! [X0] : (sP15(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f136, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP16(X0)), inference(usedef, [], [e136])).
fof(e136, plain, ! [X0] : (sP16(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f137, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP17(X0)), inference(usedef, [], [e137])).
fof(e137, plain, ! [X0] : (sP17(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f138, plain, ((? [X74] : (~ v__1(X74, nn, nn) & (node_next(nn) = X74) & object(X74)) & ? [X75] : ((node_next(nn) = X75) & v__1(X75, prev_2, prev_2) & object(X75))) | ~ sP18), inference(usedef, [], [e138])).
fof(e138, plain, (sP18 <=> (? [X74] : (~ v__1(X74, nn, nn) & (node_next(nn) = X74) & object(X74)) & ? [X75] : ((node_next(nn) = X75) & v__1(X75, prev_2, prev_2) & object(X75)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f139, plain, ((? [X71] : (~ v__1(X71, prev_2, prev_2) & (node_next(nn) = X71) & object(X71)) & ? [X72] : ((node_next(nn) = X72) & v__1(X72, nn, nn) & object(X72))) | ~ sP19), inference(usedef, [], [e139])).
fof(e139, plain, (sP19 <=> (? [X71] : (~ v__1(X71, prev_2, prev_2) & (node_next(nn) = X71) & object(X71)) & ? [X72] : ((node_next(nn) = X72) & v__1(X72, nn, nn) & object(X72)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP19])])).
fof(f140, plain, (sP19 | ? [X73] : ((node_next(nn) = X73) & v__1(X73, nn, prev_2) & object(X73)) | ~ sP20), inference(usedef, [], [e140])).
fof(e140, plain, (sP20 <=> (sP19 | ? [X73] : ((node_next(nn) = X73) & v__1(X73, nn, prev_2) & object(X73)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP20])])).
fof(f141, plain, ((? [X67] : (~ v__1(X67, prev_2, prev_2) & (node_next(nn) = X67) & object(X67)) & ? [X68] : ((node_next(nn) = X68) & v__1(X68, nn, nn) & object(X68))) | ~ sP21), inference(usedef, [], [e141])).
fof(e141, plain, (sP21 <=> (? [X67] : (~ v__1(X67, prev_2, prev_2) & (node_next(nn) = X67) & object(X67)) & ? [X68] : ((node_next(nn) = X68) & v__1(X68, nn, nn) & object(X68)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP21])])).
fof(f142, plain, ! [X0] : ((? [X63] : (~ v__1(X63, nn, nn) & (node_next(nn) = X63) & object(X63)) & ? [X64] : ((node_next(nn) = X64) & v__1(X64, X0, X0) & object(X64))) | ~ sP22(X0)), inference(usedef, [], [e142])).
fof(e142, plain, ! [X0] : (sP22(X0) <=> (? [X63] : (~ v__1(X63, nn, nn) & (node_next(nn) = X63) & object(X63)) & ? [X64] : ((node_next(nn) = X64) & v__1(X64, X0, X0) & object(X64)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP22])])).
fof(f143, plain, ! [X0] : ((? [X60] : (~ v__1(X60, X0, X0) & (node_next(nn) = X60) & object(X60)) & ? [X61] : ((node_next(nn) = X61) & v__1(X61, nn, nn) & object(X61))) | ~ sP23(X0)), inference(usedef, [], [e143])).
fof(e143, plain, ! [X0] : (sP23(X0) <=> (? [X60] : (~ v__1(X60, X0, X0) & (node_next(nn) = X60) & object(X60)) & ? [X61] : ((node_next(nn) = X61) & v__1(X61, nn, nn) & object(X61)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP23])])).
fof(f144, plain, ! [X0] : (sP23(X0) | ? [X62] : ((node_next(nn) = X62) & v__1(X62, nn, X0) & object(X62)) | ~ sP24(X0)), inference(usedef, [], [e144])).
fof(e144, plain, ! [X0] : (sP24(X0) <=> (sP23(X0) | ? [X62] : ((node_next(nn) = X62) & v__1(X62, nn, X0) & object(X62)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP24])])).
fof(f145, plain, ! [X0] : ((? [X56] : (~ v__1(X56, X0, X0) & (node_next(nn) = X56) & object(X56)) & ? [X57] : ((node_next(nn) = X57) & v__1(X57, nn, nn) & object(X57))) | ~ sP25(X0)), inference(usedef, [], [e145])).
fof(e145, plain, ! [X0] : (sP25(X0) <=> (? [X56] : (~ v__1(X56, X0, X0) & (node_next(nn) = X56) & object(X56)) & ? [X57] : ((node_next(nn) = X57) & v__1(X57, nn, nn) & object(X57)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP25])])).
fof(f146, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X59] : ((node_next(nn) = X59) & v__1(X59, X0, nn) & object(X59)) & sP24(X0) & ~ (nn = X0)) | ~ sP26(X0)), inference(usedef, [], [e146])).
fof(e146, plain, ! [X0] : (sP26(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X59] : ((node_next(nn) = X59) & v__1(X59, X0, nn) & object(X59)) & sP24(X0) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP26])])).
fof(f147, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP25(X0) | ? [X58] : ((node_next(nn) = X58) & v__1(X58, nn, X0) & object(X58))) & ~ (nn = X0)) | ~ sP27(X0)), inference(usedef, [], [e147])).
fof(e147, plain, ! [X0] : (sP27(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP25(X0) | ? [X58] : ((node_next(nn) = X58) & v__1(X58, nn, X0) & object(X58))) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP27])])).
fof(f148, plain, (? [X53] : (~ v__1(X53, prev_2, prev_2) & (node_next(nn) = X53) & object(X53)) | ~ sP28), inference(usedef, [], [e148])).
fof(e148, plain, (sP28 <=> ? [X53] : (~ v__1(X53, prev_2, prev_2) & (node_next(nn) = X53) & object(X53))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP28])])).
fof(f149, plain, (((? [X52] : ((node_next(nn) = X52) & v__1(X52, nn, nn) & object(X52)) | sP28) & ? [X54] : (~ v__1(X54, prev_2, nn) & (node_next(nn) = X54) & object(X54))) | ~ sP29), inference(usedef, [], [e149])).
fof(e149, plain, (sP29 <=> ((? [X52] : ((node_next(nn) = X52) & v__1(X52, nn, nn) & object(X52)) | sP28) & ? [X54] : (~ v__1(X54, prev_2, nn) & (node_next(nn) = X54) & object(X54)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP29])])).
fof(f150, plain, (? [X49] : (~ v__1(X49, nn, nn) & (node_next(nn) = X49) & object(X49)) | ~ sP30), inference(usedef, [], [e150])).
fof(e150, plain, (sP30 <=> ? [X49] : (~ v__1(X49, nn, nn) & (node_next(nn) = X49) & object(X49))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP30])])).
fof(f151, plain, (((? [X48] : ((node_next(nn) = X48) & v__1(X48, prev_2, prev_2) & object(X48)) | sP30) & ? [X50] : (~ v__1(X50, nn, prev_2) & (node_next(nn) = X50) & object(X50))) | ~ sP31), inference(usedef, [], [e151])).
fof(e151, plain, (sP31 <=> ((? [X48] : ((node_next(nn) = X48) & v__1(X48, prev_2, prev_2) & object(X48)) | sP30) & ? [X50] : (~ v__1(X50, nn, prev_2) & (node_next(nn) = X50) & object(X50)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP31])])).
fof(f152, plain, (? [X46] : (~ v__1(X46, nn, nn) & (node_next(nn) = X46) & object(X46)) | ~ sP32), inference(usedef, [], [e152])).
fof(e152, plain, (sP32 <=> ? [X46] : (~ v__1(X46, nn, nn) & (node_next(nn) = X46) & object(X46))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP32])])).
fof(f153, plain, (((? [X45] : ((node_next(nn) = X45) & v__1(X45, prev_2, prev_2) & object(X45)) | sP32) & ? [X47] : (~ v__1(X47, nn, prev_2) & (node_next(nn) = X47) & object(X47))) | ~ sP33), inference(usedef, [], [e153])).
fof(e153, plain, (sP33 <=> ((? [X45] : ((node_next(nn) = X45) & v__1(X45, prev_2, prev_2) & object(X45)) | sP32) & ? [X47] : (~ v__1(X47, nn, prev_2) & (node_next(nn) = X47) & object(X47)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP33])])).
fof(f154, plain, ! [X0] : (sP27(X0) | sP26(X0) | ((sP22(X0) | ? [X65] : ((node_next(nn) = X65) & v__1(X65, X0, nn) & object(X65))) & ? [X66] : ((node_next(nn) = X66) & v__1(X66, X0, X0) & object(X66))) | ~ sP34(X0)), inference(usedef, [], [e154])).
fof(e154, plain, ! [X0] : (sP34(X0) <=> (sP27(X0) | sP26(X0) | ((sP22(X0) | ? [X65] : ((node_next(nn) = X65) & v__1(X65, X0, nn) & object(X65))) & ? [X66] : ((node_next(nn) = X66) & v__1(X66, X0, X0) & object(X66))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP34])])).
fof(f155, plain, (sP31 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X51] : (~ v__1(X51, prev_2, nn) & (node_next(nn) = X51) & object(X51)) | (nn = prev_2) | ~ sP35), inference(usedef, [], [e155])).
fof(e155, plain, (sP35 <=> (sP31 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X51] : (~ v__1(X51, prev_2, nn) & (node_next(nn) = X51) & object(X51)) | (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP35])])).
fof(f156, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X70] : ((node_next(nn) = X70) & v__1(X70, X0, nn) & object(X70)) & sP20 & ~ (nn = prev_2)) | ~ sP36(X0)), inference(usedef, [], [e156])).
fof(e156, plain, ! [X0] : (sP36(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X70] : ((node_next(nn) = X70) & v__1(X70, X0, nn) & object(X70)) & sP20 & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP36])])).
fof(f157, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP21 | ? [X69] : ((node_next(nn) = X69) & v__1(X69, nn, prev_2) & object(X69))) & ~ (nn = prev_2)) | ~ sP37(X0)), inference(usedef, [], [e157])).
fof(e157, plain, ! [X0] : (sP37(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP21 | ? [X69] : ((node_next(nn) = X69) & v__1(X69, nn, prev_2) & object(X69))) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP37])])).
fof(f158, plain, ! [X0] : (((sP33 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP35 & (sP29 | ? [X55] : (~ v__1(X55, prev_2, prev_2) & (node_next(nn) = X55) & object(X55))) & sP34(X0)) | ~ sP38(X0)), inference(usedef, [], [e158])).
fof(e158, plain, ! [X0] : (sP38(X0) <=> ((sP33 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP35 & (sP29 | ? [X55] : (~ v__1(X55, prev_2, prev_2) & (node_next(nn) = X55) & object(X55))) & sP34(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP38])])).
fof(f159, plain, ! [X0] : (sP38(X0) | sP37(X0) | sP36(X0) | ((sP18 | ? [X76] : ((node_next(nn) = X76) & v__1(X76, prev_2, nn) & object(X76))) & ? [X77] : ((node_next(nn) = X77) & v__1(X77, X0, prev_2) & object(X77))) | ~ sP39(X0)), inference(usedef, [], [e159])).
fof(e159, plain, ! [X0] : (sP39(X0) <=> (sP38(X0) | sP37(X0) | sP36(X0) | ((sP18 | ? [X76] : ((node_next(nn) = X76) & v__1(X76, prev_2, nn) & object(X76))) & ? [X77] : ((node_next(nn) = X77) & v__1(X77, X0, prev_2) & object(X77))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP39])])).
fof(f160, plain, ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP40), inference(usedef, [], [e160])).
fof(e160, plain, (sP40 <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP40])])).
fof(f161, plain, ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP41), inference(usedef, [], [e161])).
fof(e161, plain, (sP41 <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP41])])).
fof(f162, plain, ! [X0] : (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0) | ~ sP42(X0)), inference(usedef, [], [e162])).
fof(e162, plain, ! [X0] : (sP42(X0) <=> (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP42])])).
fof(f163, plain, ! [X0] : (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0) | ~ sP43(X0)), inference(usedef, [], [e163])).
fof(e163, plain, ! [X0] : (sP43(X0) <=> (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP43])])).
fof(f164, plain, ! [X0] : ((sP43(X0) & sP42(X0) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & (sP41 | sP40 | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)))) | ~ sP44(X0)), inference(usedef, [], [e164])).
fof(e164, plain, ! [X0] : (sP44(X0) <=> (sP43(X0) & sP42(X0) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & (sP41 | sP40 | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP44])])).
fof(f165, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP45(X0)), inference(usedef, [], [e165])).
fof(e165, plain, ! [X0] : (sP45(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP45])])).
fof(f166, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP46(X0)), inference(usedef, [], [e166])).
fof(e166, plain, ! [X0] : (sP46(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP46])])).
fof(f167, plain, ! [X0] : ((? [X41] : (~ v__1(X41, nn, nn) & (node_next(nn) = X41) & object(X41)) & ? [X42] : ((node_next(nn) = X42) & v__1(X42, X0, X0) & object(X42))) | ~ sP47(X0)), inference(usedef, [], [e167])).
fof(e167, plain, ! [X0] : (sP47(X0) <=> (? [X41] : (~ v__1(X41, nn, nn) & (node_next(nn) = X41) & object(X41)) & ? [X42] : ((node_next(nn) = X42) & v__1(X42, X0, X0) & object(X42)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP47])])).
fof(f168, plain, ! [X0] : ((? [X38] : (~ v__1(X38, X0, X0) & (node_next(nn) = X38) & object(X38)) & ? [X39] : ((node_next(nn) = X39) & v__1(X39, nn, nn) & object(X39))) | ~ sP48(X0)), inference(usedef, [], [e168])).
fof(e168, plain, ! [X0] : (sP48(X0) <=> (? [X38] : (~ v__1(X38, X0, X0) & (node_next(nn) = X38) & object(X38)) & ? [X39] : ((node_next(nn) = X39) & v__1(X39, nn, nn) & object(X39)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP48])])).
fof(f169, plain, ! [X0] : (sP48(X0) | ? [X40] : ((node_next(nn) = X40) & v__1(X40, nn, X0) & object(X40)) | ~ sP49(X0)), inference(usedef, [], [e169])).
fof(e169, plain, ! [X0] : (sP49(X0) <=> (sP48(X0) | ? [X40] : ((node_next(nn) = X40) & v__1(X40, nn, X0) & object(X40)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP49])])).
fof(f170, plain, ! [X0] : ((? [X34] : (~ v__1(X34, X0, X0) & (node_next(nn) = X34) & object(X34)) & ? [X35] : ((node_next(nn) = X35) & v__1(X35, nn, nn) & object(X35))) | ~ sP50(X0)), inference(usedef, [], [e170])).
fof(e170, plain, ! [X0] : (sP50(X0) <=> (? [X34] : (~ v__1(X34, X0, X0) & (node_next(nn) = X34) & object(X34)) & ? [X35] : ((node_next(nn) = X35) & v__1(X35, nn, nn) & object(X35)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP50])])).
fof(f173, plain, ((? [X30] : (~ v__1(X30, nn, nn) & (node_next(nn) = X30) & object(X30)) & ? [X31] : ((node_next(nn) = X31) & v__1(X31, prev_2, prev_2) & object(X31))) | ~ sP53), inference(usedef, [], [e173])).
fof(e173, plain, (sP53 <=> (? [X30] : (~ v__1(X30, nn, nn) & (node_next(nn) = X30) & object(X30)) & ? [X31] : ((node_next(nn) = X31) & v__1(X31, prev_2, prev_2) & object(X31)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP53])])).
fof(f174, plain, ((? [X27] : (~ v__1(X27, prev_2, prev_2) & (node_next(nn) = X27) & object(X27)) & ? [X28] : ((node_next(nn) = X28) & v__1(X28, nn, nn) & object(X28))) | ~ sP54), inference(usedef, [], [e174])).
fof(e174, plain, (sP54 <=> (? [X27] : (~ v__1(X27, prev_2, prev_2) & (node_next(nn) = X27) & object(X27)) & ? [X28] : ((node_next(nn) = X28) & v__1(X28, nn, nn) & object(X28)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP54])])).
fof(f175, plain, (sP54 | ? [X29] : ((node_next(nn) = X29) & v__1(X29, nn, prev_2) & object(X29)) | ~ sP55), inference(usedef, [], [e175])).
fof(e175, plain, (sP55 <=> (sP54 | ? [X29] : ((node_next(nn) = X29) & v__1(X29, nn, prev_2) & object(X29)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP55])])).
fof(f176, plain, ((? [X23] : (~ v__1(X23, prev_2, prev_2) & (node_next(nn) = X23) & object(X23)) & ? [X24] : ((node_next(nn) = X24) & v__1(X24, nn, nn) & object(X24))) | ~ sP56), inference(usedef, [], [e176])).
fof(e176, plain, (sP56 <=> (? [X23] : (~ v__1(X23, prev_2, prev_2) & (node_next(nn) = X23) & object(X23)) & ? [X24] : ((node_next(nn) = X24) & v__1(X24, nn, nn) & object(X24)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP56])])).
fof(f177, plain, ! [X0] : ((? [X19] : (~ v__1(X19, nn, nn) & (node_next(nn) = X19) & object(X19)) & ? [X20] : ((node_next(nn) = X20) & v__1(X20, X0, X0) & object(X20))) | ~ sP57(X0)), inference(usedef, [], [e177])).
fof(e177, plain, ! [X0] : (sP57(X0) <=> (? [X19] : (~ v__1(X19, nn, nn) & (node_next(nn) = X19) & object(X19)) & ? [X20] : ((node_next(nn) = X20) & v__1(X20, X0, X0) & object(X20)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP57])])).
fof(f178, plain, ! [X0] : ((? [X16] : (~ v__1(X16, X0, X0) & (node_next(nn) = X16) & object(X16)) & ? [X17] : ((node_next(nn) = X17) & v__1(X17, nn, nn) & object(X17))) | ~ sP58(X0)), inference(usedef, [], [e178])).
fof(e178, plain, ! [X0] : (sP58(X0) <=> (? [X16] : (~ v__1(X16, X0, X0) & (node_next(nn) = X16) & object(X16)) & ? [X17] : ((node_next(nn) = X17) & v__1(X17, nn, nn) & object(X17)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP58])])).
fof(f179, plain, ! [X0] : (sP58(X0) | ? [X18] : ((node_next(nn) = X18) & v__1(X18, nn, X0) & object(X18)) | ~ sP59(X0)), inference(usedef, [], [e179])).
fof(e179, plain, ! [X0] : (sP59(X0) <=> (sP58(X0) | ? [X18] : ((node_next(nn) = X18) & v__1(X18, nn, X0) & object(X18)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP59])])).
fof(f180, plain, ! [X0] : ((? [X12] : (~ v__1(X12, X0, X0) & (node_next(nn) = X12) & object(X12)) & ? [X13] : ((node_next(nn) = X13) & v__1(X13, nn, nn) & object(X13))) | ~ sP60(X0)), inference(usedef, [], [e180])).
fof(e180, plain, ! [X0] : (sP60(X0) <=> (? [X12] : (~ v__1(X12, X0, X0) & (node_next(nn) = X12) & object(X12)) & ? [X13] : ((node_next(nn) = X13) & v__1(X13, nn, nn) & object(X13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP60])])).
fof(f181, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X15] : ((node_next(nn) = X15) & v__1(X15, X0, nn) & object(X15)) & sP59(X0) & ~ (nn = X0)) | ~ sP61(X0)), inference(usedef, [], [e181])).
fof(e181, plain, ! [X0] : (sP61(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X15] : ((node_next(nn) = X15) & v__1(X15, X0, nn) & object(X15)) & sP59(X0) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP61])])).
fof(f182, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP60(X0) | ? [X14] : ((node_next(nn) = X14) & v__1(X14, nn, X0) & object(X14))) & ~ (nn = X0)) | ~ sP62(X0)), inference(usedef, [], [e182])).
fof(e182, plain, ! [X0] : (sP62(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & (sP60(X0) | ? [X14] : ((node_next(nn) = X14) & v__1(X14, nn, X0) & object(X14))) & ~ (nn = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP62])])).
fof(f183, plain, (? [X9] : (~ v__1(X9, prev_2, prev_2) & (node_next(nn) = X9) & object(X9)) | ~ sP63), inference(usedef, [], [e183])).
fof(e183, plain, (sP63 <=> ? [X9] : (~ v__1(X9, prev_2, prev_2) & (node_next(nn) = X9) & object(X9))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP63])])).
fof(f184, plain, (((? [X8] : ((node_next(nn) = X8) & v__1(X8, nn, nn) & object(X8)) | sP63) & ? [X10] : (~ v__1(X10, prev_2, nn) & (node_next(nn) = X10) & object(X10))) | ~ sP64), inference(usedef, [], [e184])).
fof(e184, plain, (sP64 <=> ((? [X8] : ((node_next(nn) = X8) & v__1(X8, nn, nn) & object(X8)) | sP63) & ? [X10] : (~ v__1(X10, prev_2, nn) & (node_next(nn) = X10) & object(X10)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP64])])).
fof(f185, plain, (? [X5] : (~ v__1(X5, nn, nn) & (node_next(nn) = X5) & object(X5)) | ~ sP65), inference(usedef, [], [e185])).
fof(e185, plain, (sP65 <=> ? [X5] : (~ v__1(X5, nn, nn) & (node_next(nn) = X5) & object(X5))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP65])])).
fof(f186, plain, (((? [X4] : ((node_next(nn) = X4) & v__1(X4, prev_2, prev_2) & object(X4)) | sP65) & ? [X6] : (~ v__1(X6, nn, prev_2) & (node_next(nn) = X6) & object(X6))) | ~ sP66), inference(usedef, [], [e186])).
fof(e186, plain, (sP66 <=> ((? [X4] : ((node_next(nn) = X4) & v__1(X4, prev_2, prev_2) & object(X4)) | sP65) & ? [X6] : (~ v__1(X6, nn, prev_2) & (node_next(nn) = X6) & object(X6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP66])])).
fof(f187, plain, (? [X2] : (~ v__1(X2, nn, nn) & (node_next(nn) = X2) & object(X2)) | ~ sP67), inference(usedef, [], [e187])).
fof(e187, plain, (sP67 <=> ? [X2] : (~ v__1(X2, nn, nn) & (node_next(nn) = X2) & object(X2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP67])])).
fof(f188, plain, (((? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, prev_2) & object(X1)) | sP67) & ? [X3] : (~ v__1(X3, nn, prev_2) & (node_next(nn) = X3) & object(X3))) | ~ sP68), inference(usedef, [], [e188])).
fof(e188, plain, (sP68 <=> ((? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, prev_2) & object(X1)) | sP67) & ? [X3] : (~ v__1(X3, nn, prev_2) & (node_next(nn) = X3) & object(X3)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP68])])).
fof(f189, plain, ! [X0] : (sP62(X0) | sP61(X0) | ((sP57(X0) | ? [X21] : ((node_next(nn) = X21) & v__1(X21, X0, nn) & object(X21))) & ? [X22] : ((node_next(nn) = X22) & v__1(X22, X0, X0) & object(X22))) | ~ sP69(X0)), inference(usedef, [], [e189])).
fof(e189, plain, ! [X0] : (sP69(X0) <=> (sP62(X0) | sP61(X0) | ((sP57(X0) | ? [X21] : ((node_next(nn) = X21) & v__1(X21, X0, nn) & object(X21))) & ? [X22] : ((node_next(nn) = X22) & v__1(X22, X0, X0) & object(X22))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP69])])).
fof(f190, plain, (sP66 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X7] : (~ v__1(X7, prev_2, nn) & (node_next(nn) = X7) & object(X7)) | (nn = prev_2) | ~ sP70), inference(usedef, [], [e190])).
fof(e190, plain, (sP70 <=> (sP66 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X7] : (~ v__1(X7, prev_2, nn) & (node_next(nn) = X7) & object(X7)) | (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP70])])).
fof(f191, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X26] : ((node_next(nn) = X26) & v__1(X26, X0, nn) & object(X26)) & sP55 & ~ (nn = prev_2)) | ~ sP71(X0)), inference(usedef, [], [e191])).
fof(e191, plain, ! [X0] : (sP71(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X26] : ((node_next(nn) = X26) & v__1(X26, X0, nn) & object(X26)) & sP55 & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP71])])).
fof(f192, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP56 | ? [X25] : ((node_next(nn) = X25) & v__1(X25, nn, prev_2) & object(X25))) & ~ (nn = prev_2)) | ~ sP72(X0)), inference(usedef, [], [e192])).
fof(e192, plain, ! [X0] : (sP72(X0) <=> (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP56 | ? [X25] : ((node_next(nn) = X25) & v__1(X25, nn, prev_2) & object(X25))) & ~ (nn = prev_2))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP72])])).
fof(f193, plain, ! [X0] : (((sP68 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP70 & (sP64 | ? [X11] : (~ v__1(X11, prev_2, prev_2) & (node_next(nn) = X11) & object(X11))) & sP69(X0)) | ~ sP73(X0)), inference(usedef, [], [e193])).
fof(e193, plain, ! [X0] : (sP73(X0) <=> ((sP68 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP70 & (sP64 | ? [X11] : (~ v__1(X11, prev_2, prev_2) & (node_next(nn) = X11) & object(X11))) & sP69(X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP73])])).
fof(f195, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | ((sP53 | ? [X32] : ((node_next(nn) = X32) & v__1(X32, prev_2, nn) & object(X32))) & ? [X33] : ((node_next(nn) = X33) & v__1(X33, X0, prev_2) & object(X33))) | ~ sP75(X0)), inference(usedef, [], [e195])).
fof(e195, plain, ! [X0] : (sP75(X0) <=> (sP73(X0) | sP72(X0) | sP71(X0) | ((sP53 | ? [X32] : ((node_next(nn) = X32) & v__1(X32, prev_2, nn) & object(X32))) & ? [X33] : ((node_next(nn) = X33) & v__1(X33, X0, prev_2) & object(X33))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP75])])).
fof(f196, plain, ! [X0] : ((sP39(X0) & (sP17(X0) | sP16(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP13(X0) | sP15(X0) | sP14(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | ~ sP76(X0)), inference(usedef, [], [e196])).
fof(e196, plain, ! [X0] : (sP76(X0) <=> (sP39(X0) & (sP17(X0) | sP16(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP13(X0) | sP15(X0) | sP14(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP76])])).
fof(f197, plain, ! [X0] : (((sP6(X0) | sP8(X0) | sP7(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP1(X0) | sP0(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)))) | ~ sP77(X0)), inference(usedef, [], [e197])).
fof(e197, plain, ! [X0] : (sP77(X0) <=> ((sP6(X0) | sP8(X0) | sP7(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP1(X0) | sP0(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0))))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP77])])).
fof(f119, plain, ? [X0] : (((nn = X0) | ~ v__1(sortedList_first, X0, X0) | (null = X0)) & (((((((? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, prev_2) & object(X1)) | ? [X2] : (~ v__1(X2, nn, nn) & (node_next(nn) = X2) & object(X2))) & ? [X3] : (~ v__1(X3, nn, prev_2) & (node_next(nn) = X3) & object(X3))) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & (((? [X4] : ((node_next(nn) = X4) & v__1(X4, prev_2, prev_2) & object(X4)) | ? [X5] : (~ v__1(X5, nn, nn) & (node_next(nn) = X5) & object(X5))) & ? [X6] : (~ v__1(X6, nn, prev_2) & (node_next(nn) = X6) & object(X6))) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X7] : (~ v__1(X7, prev_2, nn) & (node_next(nn) = X7) & object(X7)) | (nn = prev_2)) & (((? [X8] : ((node_next(nn) = X8) & v__1(X8, nn, nn) & object(X8)) | ? [X9] : (~ v__1(X9, prev_2, prev_2) & (node_next(nn) = X9) & object(X9))) & ? [X10] : (~ v__1(X10, prev_2, nn) & (node_next(nn) = X10) & object(X10))) | ? [X11] : (~ v__1(X11, prev_2, prev_2) & (node_next(nn) = X11) & object(X11))) & ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((? [X12] : (~ v__1(X12, X0, X0) & (node_next(nn) = X12) & object(X12)) & ? [X13] : ((node_next(nn) = X13) & v__1(X13, nn, nn) & object(X13))) | ? [X14] : ((node_next(nn) = X14) & v__1(X14, nn, X0) & object(X14))) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X15] : ((node_next(nn) = X15) & v__1(X15, X0, nn) & object(X15)) & ((? [X16] : (~ v__1(X16, X0, X0) & (node_next(nn) = X16) & object(X16)) & ? [X17] : ((node_next(nn) = X17) & v__1(X17, nn, nn) & object(X17))) | ? [X18] : ((node_next(nn) = X18) & v__1(X18, nn, X0) & object(X18))) & ~ (nn = X0)) | (((? [X19] : (~ v__1(X19, nn, nn) & (node_next(nn) = X19) & object(X19)) & ? [X20] : ((node_next(nn) = X20) & v__1(X20, X0, X0) & object(X20))) | ? [X21] : ((node_next(nn) = X21) & v__1(X21, X0, nn) & object(X21))) & ? [X22] : ((node_next(nn) = X22) & v__1(X22, X0, X0) & object(X22))))) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((? [X23] : (~ v__1(X23, prev_2, prev_2) & (node_next(nn) = X23) & object(X23)) & ? [X24] : ((node_next(nn) = X24) & v__1(X24, nn, nn) & object(X24))) | ? [X25] : ((node_next(nn) = X25) & v__1(X25, nn, prev_2) & object(X25))) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X26] : ((node_next(nn) = X26) & v__1(X26, X0, nn) & object(X26)) & ((? [X27] : (~ v__1(X27, prev_2, prev_2) & (node_next(nn) = X27) & object(X27)) & ? [X28] : ((node_next(nn) = X28) & v__1(X28, nn, nn) & object(X28))) | ? [X29] : ((node_next(nn) = X29) & v__1(X29, nn, prev_2) & object(X29))) & ~ (nn = prev_2)) | (((? [X30] : (~ v__1(X30, nn, nn) & (node_next(nn) = X30) & object(X30)) & ? [X31] : ((node_next(nn) = X31) & v__1(X31, prev_2, prev_2) & object(X31))) | ? [X32] : ((node_next(nn) = X32) & v__1(X32, prev_2, nn) & object(X32))) & ? [X33] : ((node_next(nn) = X33) & v__1(X33, X0, prev_2) & object(X33)))) & ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((? [X34] : (~ v__1(X34, X0, X0) & (node_next(nn) = X34) & object(X34)) & ? [X35] : ((node_next(nn) = X35) & v__1(X35, nn, nn) & object(X35))) | ? [X36] : ((node_next(nn) = X36) & v__1(X36, nn, X0) & object(X36))) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X37] : ((node_next(nn) = X37) & v__1(X37, X0, nn) & object(X37)) & ((? [X38] : (~ v__1(X38, X0, X0) & (node_next(nn) = X38) & object(X38)) & ? [X39] : ((node_next(nn) = X39) & v__1(X39, nn, nn) & object(X39))) | ? [X40] : ((node_next(nn) = X40) & v__1(X40, nn, X0) & object(X40))) & ~ (nn = X0)) | (((? [X41] : (~ v__1(X41, nn, nn) & (node_next(nn) = X41) & object(X41)) & ? [X42] : ((node_next(nn) = X42) & v__1(X42, X0, X0) & object(X42))) | ? [X43] : ((node_next(nn) = X43) & v__1(X43, X0, nn) & object(X43))) & ? [X44] : ((node_next(nn) = X44) & v__1(X44, X0, X0) & object(X44)))) & (((((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0)) & (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)))) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | ((((((? [X45] : ((node_next(nn) = X45) & v__1(X45, prev_2, prev_2) & object(X45)) | ? [X46] : (~ v__1(X46, nn, nn) & (node_next(nn) = X46) & object(X46))) & ? [X47] : (~ v__1(X47, nn, prev_2) & (node_next(nn) = X47) & object(X47))) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & (((? [X48] : ((node_next(nn) = X48) & v__1(X48, prev_2, prev_2) & object(X48)) | ? [X49] : (~ v__1(X49, nn, nn) & (node_next(nn) = X49) & object(X49))) & ? [X50] : (~ v__1(X50, nn, prev_2) & (node_next(nn) = X50) & object(X50))) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ? [X51] : (~ v__1(X51, prev_2, nn) & (node_next(nn) = X51) & object(X51)) | (nn = prev_2)) & (((? [X52] : ((node_next(nn) = X52) & v__1(X52, nn, nn) & object(X52)) | ? [X53] : (~ v__1(X53, prev_2, prev_2) & (node_next(nn) = X53) & object(X53))) & ? [X54] : (~ v__1(X54, prev_2, nn) & (node_next(nn) = X54) & object(X54))) | ? [X55] : (~ v__1(X55, prev_2, prev_2) & (node_next(nn) = X55) & object(X55))) & ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((? [X56] : (~ v__1(X56, X0, X0) & (node_next(nn) = X56) & object(X56)) & ? [X57] : ((node_next(nn) = X57) & v__1(X57, nn, nn) & object(X57))) | ? [X58] : ((node_next(nn) = X58) & v__1(X58, nn, X0) & object(X58))) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & ? [X59] : ((node_next(nn) = X59) & v__1(X59, X0, nn) & object(X59)) & ((? [X60] : (~ v__1(X60, X0, X0) & (node_next(nn) = X60) & object(X60)) & ? [X61] : ((node_next(nn) = X61) & v__1(X61, nn, nn) & object(X61))) | ? [X62] : ((node_next(nn) = X62) & v__1(X62, nn, X0) & object(X62))) & ~ (nn = X0)) | (((? [X63] : (~ v__1(X63, nn, nn) & (node_next(nn) = X63) & object(X63)) & ? [X64] : ((node_next(nn) = X64) & v__1(X64, X0, X0) & object(X64))) | ? [X65] : ((node_next(nn) = X65) & v__1(X65, X0, nn) & object(X65))) & ? [X66] : ((node_next(nn) = X66) & v__1(X66, X0, X0) & object(X66))))) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((? [X67] : (~ v__1(X67, prev_2, prev_2) & (node_next(nn) = X67) & object(X67)) & ? [X68] : ((node_next(nn) = X68) & v__1(X68, nn, nn) & object(X68))) | ? [X69] : ((node_next(nn) = X69) & v__1(X69, nn, prev_2) & object(X69))) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X70] : ((node_next(nn) = X70) & v__1(X70, X0, nn) & object(X70)) & ((? [X71] : (~ v__1(X71, prev_2, prev_2) & (node_next(nn) = X71) & object(X71)) & ? [X72] : ((node_next(nn) = X72) & v__1(X72, nn, nn) & object(X72))) | ? [X73] : ((node_next(nn) = X73) & v__1(X73, nn, prev_2) & object(X73))) & ~ (nn = prev_2)) | (((? [X74] : (~ v__1(X74, nn, nn) & (node_next(nn) = X74) & object(X74)) & ? [X75] : ((node_next(nn) = X75) & v__1(X75, prev_2, prev_2) & object(X75))) | ? [X76] : ((node_next(nn) = X76) & v__1(X76, prev_2, nn) & object(X76))) & ? [X77] : ((node_next(nn) = X77) & v__1(X77, X0, prev_2) & object(X77)))) & ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (((((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | (nn = X0)) & (((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | ((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)) & ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)))) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, prev_2, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, prev_2, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | ((((((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & (((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)) & ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)))) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0))))) & ~ (null = X0) & object(X0)), inference(ennf_transformation, [], [f86])).
fof(f86, plain, ~ ! [X0] : ((~ (nn = X0) & v__1(sortedList_first, X0, X0) & ~ (null = X0)) | (((((((! [X1] : (~ (node_next(nn) = X1) | ~ v__1(X1, prev_2, prev_2) | ~ object(X1)) & ! [X2] : (v__1(X2, nn, nn) | ~ (node_next(nn) = X2) | ~ object(X2))) | ! [X3] : (v__1(X3, nn, prev_2) | ~ (node_next(nn) = X3) | ~ object(X3))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X4] : (~ (node_next(nn) = X4) | ~ v__1(X4, prev_2, prev_2) | ~ object(X4)) & ! [X5] : (v__1(X5, nn, nn) | ~ (node_next(nn) = X5) | ~ object(X5))) | ! [X6] : (v__1(X6, nn, prev_2) | ~ (node_next(nn) = X6) | ~ object(X6))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X7] : (v__1(X7, prev_2, nn) | ~ (node_next(nn) = X7) | ~ object(X7)) & ~ (nn = prev_2)) | (((! [X8] : (~ (node_next(nn) = X8) | ~ v__1(X8, nn, nn) | ~ object(X8)) & ! [X9] : (v__1(X9, prev_2, prev_2) | ~ (node_next(nn) = X9) | ~ object(X9))) | ! [X10] : (v__1(X10, prev_2, nn) | ~ (node_next(nn) = X10) | ~ object(X10))) & ! [X11] : (v__1(X11, prev_2, prev_2) | ~ (node_next(nn) = X11) | ~ object(X11))) | ((((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | ((! [X12] : (v__1(X12, X0, X0) | ~ (node_next(nn) = X12) | ~ object(X12)) | ! [X13] : (~ (node_next(nn) = X13) | ~ v__1(X13, nn, nn) | ~ object(X13))) & ! [X14] : (~ (node_next(nn) = X14) | ~ v__1(X14, nn, X0) | ~ object(X14))) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ! [X15] : (~ (node_next(nn) = X15) | ~ v__1(X15, X0, nn) | ~ object(X15)) | ((! [X16] : (v__1(X16, X0, X0) | ~ (node_next(nn) = X16) | ~ object(X16)) | ! [X17] : (~ (node_next(nn) = X17) | ~ v__1(X17, nn, nn) | ~ object(X17))) & ! [X18] : (~ (node_next(nn) = X18) | ~ v__1(X18, nn, X0) | ~ object(X18))) | (nn = X0)) & (((! [X19] : (v__1(X19, nn, nn) | ~ (node_next(nn) = X19) | ~ object(X19)) | ! [X20] : (~ (node_next(nn) = X20) | ~ v__1(X20, X0, X0) | ~ object(X20))) & ! [X21] : (~ (node_next(nn) = X21) | ~ v__1(X21, X0, nn) | ~ object(X21))) | ! [X22] : (~ (node_next(nn) = X22) | ~ v__1(X22, X0, X0) | ~ object(X22))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X0, prev_2) | ((! [X23] : (v__1(X23, prev_2, prev_2) | ~ (node_next(nn) = X23) | ~ object(X23)) | ! [X24] : (~ (node_next(nn) = X24) | ~ v__1(X24, nn, nn) | ~ object(X24))) & ! [X25] : (~ (node_next(nn) = X25) | ~ v__1(X25, nn, prev_2) | ~ object(X25))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X26] : (~ (node_next(nn) = X26) | ~ v__1(X26, X0, nn) | ~ object(X26)) | ((! [X27] : (v__1(X27, prev_2, prev_2) | ~ (node_next(nn) = X27) | ~ object(X27)) | ! [X28] : (~ (node_next(nn) = X28) | ~ v__1(X28, nn, nn) | ~ object(X28))) & ! [X29] : (~ (node_next(nn) = X29) | ~ v__1(X29, nn, prev_2) | ~ object(X29))) | (nn = prev_2)) & (((! [X30] : (v__1(X30, nn, nn) | ~ (node_next(nn) = X30) | ~ object(X30)) | ! [X31] : (~ (node_next(nn) = X31) | ~ v__1(X31, prev_2, prev_2) | ~ object(X31))) & ! [X32] : (~ (node_next(nn) = X32) | ~ v__1(X32, prev_2, nn) | ~ object(X32))) | ! [X33] : (~ (node_next(nn) = X33) | ~ v__1(X33, X0, prev_2) | ~ object(X33)))) | ((((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | ((! [X34] : (v__1(X34, X0, X0) | ~ (node_next(nn) = X34) | ~ object(X34)) | ! [X35] : (~ (node_next(nn) = X35) | ~ v__1(X35, nn, nn) | ~ object(X35))) & ! [X36] : (~ (node_next(nn) = X36) | ~ v__1(X36, nn, X0) | ~ object(X36))) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ! [X37] : (~ (node_next(nn) = X37) | ~ v__1(X37, X0, nn) | ~ object(X37)) | ((! [X38] : (v__1(X38, X0, X0) | ~ (node_next(nn) = X38) | ~ object(X38)) | ! [X39] : (~ (node_next(nn) = X39) | ~ v__1(X39, nn, nn) | ~ object(X39))) & ! [X40] : (~ (node_next(nn) = X40) | ~ v__1(X40, nn, X0) | ~ object(X40))) | (nn = X0)) & (((! [X41] : (v__1(X41, nn, nn) | ~ (node_next(nn) = X41) | ~ object(X41)) | ! [X42] : (~ (node_next(nn) = X42) | ~ v__1(X42, X0, X0) | ~ object(X42))) & ! [X43] : (~ (node_next(nn) = X43) | ~ v__1(X43, X0, nn) | ~ object(X43))) | ! [X44] : (~ (node_next(nn) = X44) | ~ v__1(X44, X0, X0) | ~ object(X44)))) | (((((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ~ (nn = X0)) | (((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, prev_2, X0) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, prev_2, X0))) | (prev_2 = X0)) & ((((((! [X45] : (~ (node_next(nn) = X45) | ~ v__1(X45, prev_2, prev_2) | ~ object(X45)) & ! [X46] : (v__1(X46, nn, nn) | ~ (node_next(nn) = X46) | ~ object(X46))) | ! [X47] : (v__1(X47, nn, prev_2) | ~ (node_next(nn) = X47) | ~ object(X47))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X48] : (~ (node_next(nn) = X48) | ~ v__1(X48, prev_2, prev_2) | ~ object(X48)) & ! [X49] : (v__1(X49, nn, nn) | ~ (node_next(nn) = X49) | ~ object(X49))) | ! [X50] : (v__1(X50, nn, prev_2) | ~ (node_next(nn) = X50) | ~ object(X50))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X51] : (v__1(X51, prev_2, nn) | ~ (node_next(nn) = X51) | ~ object(X51)) & ~ (nn = prev_2)) | (((! [X52] : (~ (node_next(nn) = X52) | ~ v__1(X52, nn, nn) | ~ object(X52)) & ! [X53] : (v__1(X53, prev_2, prev_2) | ~ (node_next(nn) = X53) | ~ object(X53))) | ! [X54] : (v__1(X54, prev_2, nn) | ~ (node_next(nn) = X54) | ~ object(X54))) & ! [X55] : (v__1(X55, prev_2, prev_2) | ~ (node_next(nn) = X55) | ~ object(X55))) | ((((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | ((! [X56] : (v__1(X56, X0, X0) | ~ (node_next(nn) = X56) | ~ object(X56)) | ! [X57] : (~ (node_next(nn) = X57) | ~ v__1(X57, nn, nn) | ~ object(X57))) & ! [X58] : (~ (node_next(nn) = X58) | ~ v__1(X58, nn, X0) | ~ object(X58))) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ! [X59] : (~ (node_next(nn) = X59) | ~ v__1(X59, X0, nn) | ~ object(X59)) | ((! [X60] : (v__1(X60, X0, X0) | ~ (node_next(nn) = X60) | ~ object(X60)) | ! [X61] : (~ (node_next(nn) = X61) | ~ v__1(X61, nn, nn) | ~ object(X61))) & ! [X62] : (~ (node_next(nn) = X62) | ~ v__1(X62, nn, X0) | ~ object(X62))) | (nn = X0)) & (((! [X63] : (v__1(X63, nn, nn) | ~ (node_next(nn) = X63) | ~ object(X63)) | ! [X64] : (~ (node_next(nn) = X64) | ~ v__1(X64, X0, X0) | ~ object(X64))) & ! [X65] : (~ (node_next(nn) = X65) | ~ v__1(X65, X0, nn) | ~ object(X65))) | ! [X66] : (~ (node_next(nn) = X66) | ~ v__1(X66, X0, X0) | ~ object(X66))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X0, prev_2) | ((! [X67] : (v__1(X67, prev_2, prev_2) | ~ (node_next(nn) = X67) | ~ object(X67)) | ! [X68] : (~ (node_next(nn) = X68) | ~ v__1(X68, nn, nn) | ~ object(X68))) & ! [X69] : (~ (node_next(nn) = X69) | ~ v__1(X69, nn, prev_2) | ~ object(X69))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X70] : (~ (node_next(nn) = X70) | ~ v__1(X70, X0, nn) | ~ object(X70)) | ((! [X71] : (v__1(X71, prev_2, prev_2) | ~ (node_next(nn) = X71) | ~ object(X71)) | ! [X72] : (~ (node_next(nn) = X72) | ~ v__1(X72, nn, nn) | ~ object(X72))) & ! [X73] : (~ (node_next(nn) = X73) | ~ v__1(X73, nn, prev_2) | ~ object(X73))) | (nn = prev_2)) & (((! [X74] : (v__1(X74, nn, nn) | ~ (node_next(nn) = X74) | ~ object(X74)) | ! [X75] : (~ (node_next(nn) = X75) | ~ v__1(X75, prev_2, prev_2) | ~ object(X75))) & ! [X76] : (~ (node_next(nn) = X76) | ~ v__1(X76, prev_2, nn) | ~ object(X76))) | ! [X77] : (~ (node_next(nn) = X77) | ~ v__1(X77, X0, prev_2) | ~ object(X77)))) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X0, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X0, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X0, prev_2))) | (((((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ~ (nn = X0)) | (((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ~ (nn = X0)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, prev_2, X0) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, prev_2, X0))) | (prev_2 = X0)) & ((((((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)) | ((((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0)))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X0, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X0, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X0, prev_2))) | ((((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(null, X0, X0) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(null, nn, nn) | ~ v__1(null, X0, X0)) & ~ v__1(null, X0, nn)) | ~ v__1(sortedList_first, X0, nn) | ((v__1(sortedList_first, X0, X0) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X0)) | (nn = X0)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X0, X0)) & ~ v__1(sortedList_first, X0, nn)) | ~ v__1(sortedList_first, X0, X0))))) | (null = X0) | ~ object(X0)), inference(rectify, [], [f60])).
fof(f60, plain, ~ ! [X57] : ((~ (nn = X57) & v__1(sortedList_first, X57, X57) & ~ (null = X57)) | (((((((! [X134] : (~ (node_next(nn) = X134) | ~ v__1(X134, prev_2, prev_2) | ~ object(X134)) & ! [X133] : (v__1(X133, nn, nn) | ~ (node_next(nn) = X133) | ~ object(X133))) | ! [X132] : (v__1(X132, nn, prev_2) | ~ (node_next(nn) = X132) | ~ object(X132))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X131] : (~ (node_next(nn) = X131) | ~ v__1(X131, prev_2, prev_2) | ~ object(X131)) & ! [X130] : (v__1(X130, nn, nn) | ~ (node_next(nn) = X130) | ~ object(X130))) | ! [X129] : (v__1(X129, nn, prev_2) | ~ (node_next(nn) = X129) | ~ object(X129))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X128] : (v__1(X128, prev_2, nn) | ~ (node_next(nn) = X128) | ~ object(X128)) & ~ (nn = prev_2)) | (((! [X127] : (~ (node_next(nn) = X127) | ~ v__1(X127, nn, nn) | ~ object(X127)) & ! [X126] : (v__1(X126, prev_2, prev_2) | ~ (node_next(nn) = X126) | ~ object(X126))) | ! [X125] : (v__1(X125, prev_2, nn) | ~ (node_next(nn) = X125) | ~ object(X125))) & ! [X124] : (v__1(X124, prev_2, prev_2) | ~ (node_next(nn) = X124) | ~ object(X124))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X123] : (v__1(X123, X57, X57) | ~ (node_next(nn) = X123) | ~ object(X123)) | ! [X122] : (~ (node_next(nn) = X122) | ~ v__1(X122, nn, nn) | ~ object(X122))) & ! [X121] : (~ (node_next(nn) = X121) | ~ v__1(X121, nn, X57) | ~ object(X121))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X120] : (~ (node_next(nn) = X120) | ~ v__1(X120, X57, nn) | ~ object(X120)) | ((! [X119] : (v__1(X119, X57, X57) | ~ (node_next(nn) = X119) | ~ object(X119)) | ! [X118] : (~ (node_next(nn) = X118) | ~ v__1(X118, nn, nn) | ~ object(X118))) & ! [X117] : (~ (node_next(nn) = X117) | ~ v__1(X117, nn, X57) | ~ object(X117))) | (nn = X57)) & (((! [X116] : (v__1(X116, nn, nn) | ~ (node_next(nn) = X116) | ~ object(X116)) | ! [X115] : (~ (node_next(nn) = X115) | ~ v__1(X115, X57, X57) | ~ object(X115))) & ! [X114] : (~ (node_next(nn) = X114) | ~ v__1(X114, X57, nn) | ~ object(X114))) | ! [X113] : (~ (node_next(nn) = X113) | ~ v__1(X113, X57, X57) | ~ object(X113))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((! [X112] : (v__1(X112, prev_2, prev_2) | ~ (node_next(nn) = X112) | ~ object(X112)) | ! [X111] : (~ (node_next(nn) = X111) | ~ v__1(X111, nn, nn) | ~ object(X111))) & ! [X110] : (~ (node_next(nn) = X110) | ~ v__1(X110, nn, prev_2) | ~ object(X110))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X109] : (~ (node_next(nn) = X109) | ~ v__1(X109, X57, nn) | ~ object(X109)) | ((! [X108] : (v__1(X108, prev_2, prev_2) | ~ (node_next(nn) = X108) | ~ object(X108)) | ! [X107] : (~ (node_next(nn) = X107) | ~ v__1(X107, nn, nn) | ~ object(X107))) & ! [X106] : (~ (node_next(nn) = X106) | ~ v__1(X106, nn, prev_2) | ~ object(X106))) | (nn = prev_2)) & (((! [X105] : (v__1(X105, nn, nn) | ~ (node_next(nn) = X105) | ~ object(X105)) | ! [X104] : (~ (node_next(nn) = X104) | ~ v__1(X104, prev_2, prev_2) | ~ object(X104))) & ! [X103] : (~ (node_next(nn) = X103) | ~ v__1(X103, prev_2, nn) | ~ object(X103))) | ! [X102] : (~ (node_next(nn) = X102) | ~ v__1(X102, X57, prev_2) | ~ object(X102)))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X101] : (v__1(X101, X57, X57) | ~ (node_next(nn) = X101) | ~ object(X101)) | ! [X100] : (~ (node_next(nn) = X100) | ~ v__1(X100, nn, nn) | ~ object(X100))) & ! [X99] : (~ (node_next(nn) = X99) | ~ v__1(X99, nn, X57) | ~ object(X99))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X98] : (~ (node_next(nn) = X98) | ~ v__1(X98, X57, nn) | ~ object(X98)) | ((! [X97] : (v__1(X97, X57, X57) | ~ (node_next(nn) = X97) | ~ object(X97)) | ! [X96] : (~ (node_next(nn) = X96) | ~ v__1(X96, nn, nn) | ~ object(X96))) & ! [X95] : (~ (node_next(nn) = X95) | ~ v__1(X95, nn, X57) | ~ object(X95))) | (nn = X57)) & (((! [X94] : (v__1(X94, nn, nn) | ~ (node_next(nn) = X94) | ~ object(X94)) | ! [X93] : (~ (node_next(nn) = X93) | ~ v__1(X93, X57, X57) | ~ object(X93))) & ! [X92] : (~ (node_next(nn) = X92) | ~ v__1(X92, X57, nn) | ~ object(X92))) | ! [X91] : (~ (node_next(nn) = X91) | ~ v__1(X91, X57, X57) | ~ object(X91)))) | (((((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(null, X57, X57) & ~ (nn = X57)) | (((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(sortedList_first, X57, nn) & ~ (nn = X57)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X57, X57)) | v__1(sortedList_first, X57, nn)) & v__1(sortedList_first, X57, X57)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, prev_2, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, prev_2, X57))) | (prev_2 = X57)) & ((((((! [X90] : (~ (node_next(nn) = X90) | ~ v__1(X90, prev_2, prev_2) | ~ object(X90)) & ! [X89] : (v__1(X89, nn, nn) | ~ (node_next(nn) = X89) | ~ object(X89))) | ! [X88] : (v__1(X88, nn, prev_2) | ~ (node_next(nn) = X88) | ~ object(X88))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X87] : (~ (node_next(nn) = X87) | ~ v__1(X87, prev_2, prev_2) | ~ object(X87)) & ! [X86] : (v__1(X86, nn, nn) | ~ (node_next(nn) = X86) | ~ object(X86))) | ! [X85] : (v__1(X85, nn, prev_2) | ~ (node_next(nn) = X85) | ~ object(X85))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X84] : (v__1(X84, prev_2, nn) | ~ (node_next(nn) = X84) | ~ object(X84)) & ~ (nn = prev_2)) | (((! [X83] : (~ (node_next(nn) = X83) | ~ v__1(X83, nn, nn) | ~ object(X83)) & ! [X82] : (v__1(X82, prev_2, prev_2) | ~ (node_next(nn) = X82) | ~ object(X82))) | ! [X81] : (v__1(X81, prev_2, nn) | ~ (node_next(nn) = X81) | ~ object(X81))) & ! [X80] : (v__1(X80, prev_2, prev_2) | ~ (node_next(nn) = X80) | ~ object(X80))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X79] : (v__1(X79, X57, X57) | ~ (node_next(nn) = X79) | ~ object(X79)) | ! [X78] : (~ (node_next(nn) = X78) | ~ v__1(X78, nn, nn) | ~ object(X78))) & ! [X77] : (~ (node_next(nn) = X77) | ~ v__1(X77, nn, X57) | ~ object(X77))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X76] : (~ (node_next(nn) = X76) | ~ v__1(X76, X57, nn) | ~ object(X76)) | ((! [X75] : (v__1(X75, X57, X57) | ~ (node_next(nn) = X75) | ~ object(X75)) | ! [X74] : (~ (node_next(nn) = X74) | ~ v__1(X74, nn, nn) | ~ object(X74))) & ! [X73] : (~ (node_next(nn) = X73) | ~ v__1(X73, nn, X57) | ~ object(X73))) | (nn = X57)) & (((! [X72] : (v__1(X72, nn, nn) | ~ (node_next(nn) = X72) | ~ object(X72)) | ! [X71] : (~ (node_next(nn) = X71) | ~ v__1(X71, X57, X57) | ~ object(X71))) & ! [X70] : (~ (node_next(nn) = X70) | ~ v__1(X70, X57, nn) | ~ object(X70))) | ! [X69] : (~ (node_next(nn) = X69) | ~ v__1(X69, X57, X57) | ~ object(X69))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((! [X68] : (v__1(X68, prev_2, prev_2) | ~ (node_next(nn) = X68) | ~ object(X68)) | ! [X67] : (~ (node_next(nn) = X67) | ~ v__1(X67, nn, nn) | ~ object(X67))) & ! [X66] : (~ (node_next(nn) = X66) | ~ v__1(X66, nn, prev_2) | ~ object(X66))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X65] : (~ (node_next(nn) = X65) | ~ v__1(X65, X57, nn) | ~ object(X65)) | ((! [X64] : (v__1(X64, prev_2, prev_2) | ~ (node_next(nn) = X64) | ~ object(X64)) | ! [X63] : (~ (node_next(nn) = X63) | ~ v__1(X63, nn, nn) | ~ object(X63))) & ! [X62] : (~ (node_next(nn) = X62) | ~ v__1(X62, nn, prev_2) | ~ object(X62))) | (nn = prev_2)) & (((! [X61] : (v__1(X61, nn, nn) | ~ (node_next(nn) = X61) | ~ object(X61)) | ! [X60] : (~ (node_next(nn) = X60) | ~ v__1(X60, prev_2, prev_2) | ~ object(X60))) & ! [X59] : (~ (node_next(nn) = X59) | ~ v__1(X59, prev_2, nn) | ~ object(X59))) | ! [X58] : (~ (node_next(nn) = X58) | ~ v__1(X58, X57, prev_2) | ~ object(X58)))) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X57, prev_2))) | (((((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(null, X57, X57) & ~ (nn = X57)) | (((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(sortedList_first, X57, nn) & ~ (nn = X57)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X57, X57)) | v__1(sortedList_first, X57, nn)) & v__1(sortedList_first, X57, X57)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, prev_2, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, prev_2, X57))) | (prev_2 = X57)) & ((((((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, X57, X57)))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X57, prev_2))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, X57, X57))))) | (null = X57) | ~ object(X57)), inference(negated_conjecture, [], [f59])).
fof(f59, plain, ~ ! [X57] : ((~ (nn = X57) & v__1(sortedList_first, X57, X57) & ~ (null = X57)) | (((((((! [X134] : (~ (node_next(nn) = X134) | ~ v__1(X134, prev_2, prev_2) | ~ object(X134)) & ! [X133] : (v__1(X133, nn, nn) | ~ (node_next(nn) = X133) | ~ object(X133))) | ! [X132] : (v__1(X132, nn, prev_2) | ~ (node_next(nn) = X132) | ~ object(X132))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X131] : (~ (node_next(nn) = X131) | ~ v__1(X131, prev_2, prev_2) | ~ object(X131)) & ! [X130] : (v__1(X130, nn, nn) | ~ (node_next(nn) = X130) | ~ object(X130))) | ! [X129] : (v__1(X129, nn, prev_2) | ~ (node_next(nn) = X129) | ~ object(X129))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X128] : (v__1(X128, prev_2, nn) | ~ (node_next(nn) = X128) | ~ object(X128)) & ~ (nn = prev_2)) | (((! [X127] : (~ (node_next(nn) = X127) | ~ v__1(X127, nn, nn) | ~ object(X127)) & ! [X126] : (v__1(X126, prev_2, prev_2) | ~ (node_next(nn) = X126) | ~ object(X126))) | ! [X125] : (v__1(X125, prev_2, nn) | ~ (node_next(nn) = X125) | ~ object(X125))) & ! [X124] : (v__1(X124, prev_2, prev_2) | ~ (node_next(nn) = X124) | ~ object(X124))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X123] : (v__1(X123, X57, X57) | ~ (node_next(nn) = X123) | ~ object(X123)) | ! [X122] : (~ (node_next(nn) = X122) | ~ v__1(X122, nn, nn) | ~ object(X122))) & ! [X121] : (~ (node_next(nn) = X121) | ~ v__1(X121, nn, X57) | ~ object(X121))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X120] : (~ (node_next(nn) = X120) | ~ v__1(X120, X57, nn) | ~ object(X120)) | ((! [X119] : (v__1(X119, X57, X57) | ~ (node_next(nn) = X119) | ~ object(X119)) | ! [X118] : (~ (node_next(nn) = X118) | ~ v__1(X118, nn, nn) | ~ object(X118))) & ! [X117] : (~ (node_next(nn) = X117) | ~ v__1(X117, nn, X57) | ~ object(X117))) | (nn = X57)) & (((! [X116] : (v__1(X116, nn, nn) | ~ (node_next(nn) = X116) | ~ object(X116)) | ! [X115] : (~ (node_next(nn) = X115) | ~ v__1(X115, X57, X57) | ~ object(X115))) & ! [X114] : (~ (node_next(nn) = X114) | ~ v__1(X114, X57, nn) | ~ object(X114))) | ! [X113] : (~ (node_next(nn) = X113) | ~ v__1(X113, X57, X57) | ~ object(X113))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((! [X112] : (v__1(X112, prev_2, prev_2) | ~ (node_next(nn) = X112) | ~ object(X112)) | ! [X111] : (~ (node_next(nn) = X111) | ~ v__1(X111, nn, nn) | ~ object(X111))) & ! [X110] : (~ (node_next(nn) = X110) | ~ v__1(X110, nn, prev_2) | ~ object(X110))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X109] : (~ (node_next(nn) = X109) | ~ v__1(X109, X57, nn) | ~ object(X109)) | ((! [X108] : (v__1(X108, prev_2, prev_2) | ~ (node_next(nn) = X108) | ~ object(X108)) | ! [X107] : (~ (node_next(nn) = X107) | ~ v__1(X107, nn, nn) | ~ object(X107))) & ! [X106] : (~ (node_next(nn) = X106) | ~ v__1(X106, nn, prev_2) | ~ object(X106))) | (nn = prev_2)) & (((! [X105] : (v__1(X105, nn, nn) | ~ (node_next(nn) = X105) | ~ object(X105)) | ! [X104] : (~ (node_next(nn) = X104) | ~ v__1(X104, prev_2, prev_2) | ~ object(X104))) & ! [X103] : (~ (node_next(nn) = X103) | ~ v__1(X103, prev_2, nn) | ~ object(X103))) | ! [X102] : (~ (node_next(nn) = X102) | ~ v__1(X102, X57, prev_2) | ~ object(X102)))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X101] : (v__1(X101, X57, X57) | ~ (node_next(nn) = X101) | ~ object(X101)) | ! [X100] : (~ (node_next(nn) = X100) | ~ v__1(X100, nn, nn) | ~ object(X100))) & ! [X99] : (~ (node_next(nn) = X99) | ~ v__1(X99, nn, X57) | ~ object(X99))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X98] : (~ (node_next(nn) = X98) | ~ v__1(X98, X57, nn) | ~ object(X98)) | ((! [X97] : (v__1(X97, X57, X57) | ~ (node_next(nn) = X97) | ~ object(X97)) | ! [X96] : (~ (node_next(nn) = X96) | ~ v__1(X96, nn, nn) | ~ object(X96))) & ! [X95] : (~ (node_next(nn) = X95) | ~ v__1(X95, nn, X57) | ~ object(X95))) | (nn = X57)) & (((! [X94] : (v__1(X94, nn, nn) | ~ (node_next(nn) = X94) | ~ object(X94)) | ! [X93] : (~ (node_next(nn) = X93) | ~ v__1(X93, X57, X57) | ~ object(X93))) & ! [X92] : (~ (node_next(nn) = X92) | ~ v__1(X92, X57, nn) | ~ object(X92))) | ! [X91] : (~ (node_next(nn) = X91) | ~ v__1(X91, X57, X57) | ~ object(X91)))) | (((((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(null, X57, X57) & ~ (nn = X57)) | (((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(sortedList_first, X57, nn) & ~ (nn = X57)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X57, X57)) | v__1(sortedList_first, X57, nn)) & v__1(sortedList_first, X57, X57)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, prev_2, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, prev_2, X57))) | (prev_2 = X57)) & ((((((! [X90] : (~ (node_next(nn) = X90) | ~ v__1(X90, prev_2, prev_2) | ~ object(X90)) & ! [X89] : (v__1(X89, nn, nn) | ~ (node_next(nn) = X89) | ~ object(X89))) | ! [X88] : (v__1(X88, nn, prev_2) | ~ (node_next(nn) = X88) | ~ object(X88))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((! [X87] : (~ (node_next(nn) = X87) | ~ v__1(X87, prev_2, prev_2) | ~ object(X87)) & ! [X86] : (v__1(X86, nn, nn) | ~ (node_next(nn) = X86) | ~ object(X86))) | ! [X85] : (v__1(X85, nn, prev_2) | ~ (node_next(nn) = X85) | ~ object(X85))) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ! [X84] : (v__1(X84, prev_2, nn) | ~ (node_next(nn) = X84) | ~ object(X84)) & ~ (nn = prev_2)) | (((! [X83] : (~ (node_next(nn) = X83) | ~ v__1(X83, nn, nn) | ~ object(X83)) & ! [X82] : (v__1(X82, prev_2, prev_2) | ~ (node_next(nn) = X82) | ~ object(X82))) | ! [X81] : (v__1(X81, prev_2, nn) | ~ (node_next(nn) = X81) | ~ object(X81))) & ! [X80] : (v__1(X80, prev_2, prev_2) | ~ (node_next(nn) = X80) | ~ object(X80))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((! [X79] : (v__1(X79, X57, X57) | ~ (node_next(nn) = X79) | ~ object(X79)) | ! [X78] : (~ (node_next(nn) = X78) | ~ v__1(X78, nn, nn) | ~ object(X78))) & ! [X77] : (~ (node_next(nn) = X77) | ~ v__1(X77, nn, X57) | ~ object(X77))) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ! [X76] : (~ (node_next(nn) = X76) | ~ v__1(X76, X57, nn) | ~ object(X76)) | ((! [X75] : (v__1(X75, X57, X57) | ~ (node_next(nn) = X75) | ~ object(X75)) | ! [X74] : (~ (node_next(nn) = X74) | ~ v__1(X74, nn, nn) | ~ object(X74))) & ! [X73] : (~ (node_next(nn) = X73) | ~ v__1(X73, nn, X57) | ~ object(X73))) | (nn = X57)) & (((! [X72] : (v__1(X72, nn, nn) | ~ (node_next(nn) = X72) | ~ object(X72)) | ! [X71] : (~ (node_next(nn) = X71) | ~ v__1(X71, X57, X57) | ~ object(X71))) & ! [X70] : (~ (node_next(nn) = X70) | ~ v__1(X70, X57, nn) | ~ object(X70))) | ! [X69] : (~ (node_next(nn) = X69) | ~ v__1(X69, X57, X57) | ~ object(X69))))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((! [X68] : (v__1(X68, prev_2, prev_2) | ~ (node_next(nn) = X68) | ~ object(X68)) | ! [X67] : (~ (node_next(nn) = X67) | ~ v__1(X67, nn, nn) | ~ object(X67))) & ! [X66] : (~ (node_next(nn) = X66) | ~ v__1(X66, nn, prev_2) | ~ object(X66))) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ! [X65] : (~ (node_next(nn) = X65) | ~ v__1(X65, X57, nn) | ~ object(X65)) | ((! [X64] : (v__1(X64, prev_2, prev_2) | ~ (node_next(nn) = X64) | ~ object(X64)) | ! [X63] : (~ (node_next(nn) = X63) | ~ v__1(X63, nn, nn) | ~ object(X63))) & ! [X62] : (~ (node_next(nn) = X62) | ~ v__1(X62, nn, prev_2) | ~ object(X62))) | (nn = prev_2)) & (((! [X61] : (v__1(X61, nn, nn) | ~ (node_next(nn) = X61) | ~ object(X61)) | ! [X60] : (~ (node_next(nn) = X60) | ~ v__1(X60, prev_2, prev_2) | ~ object(X60))) & ! [X59] : (~ (node_next(nn) = X59) | ~ v__1(X59, prev_2, nn) | ~ object(X59))) | ! [X58] : (~ (node_next(nn) = X58) | ~ v__1(X58, X57, prev_2) | ~ object(X58)))) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X57, prev_2))) | (((((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(null, X57, X57) & ~ (nn = X57)) | (((~ v__1(sortedList_first, X57, X57) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X57)) & ((~ v__1(null, nn, nn) & v__1(null, X57, X57)) | v__1(null, X57, nn)) & v__1(sortedList_first, X57, nn) & ~ (nn = X57)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X57, X57)) | v__1(sortedList_first, X57, nn)) & v__1(sortedList_first, X57, X57)) | ((((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)))) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, prev_2, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, prev_2, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, prev_2, X57))) | (prev_2 = X57)) & ((((((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, prev_2, prev_2) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, prev_2, nn) & ~ (nn = prev_2)) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, prev_2, prev_2)) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, X57, X57)))) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, X57, prev_2) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, prev_2, prev_2) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, prev_2)) | (nn = prev_2)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, X57, prev_2))) | ((((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(null, X57, X57) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(null, nn, nn) | ~ v__1(null, X57, X57)) & ~ v__1(null, X57, nn)) | ~ v__1(sortedList_first, X57, nn) | ((v__1(sortedList_first, X57, X57) | ~ v__1(sortedList_first, nn, nn)) & ~ v__1(sortedList_first, nn, X57)) | (nn = X57)) & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, X57, X57)) & ~ v__1(sortedList_first, X57, nn)) | ~ v__1(sortedList_first, X57, X57))))) | (null = X57) | ~ object(X57)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', goal)).
fof(f962830, plain, (spl158_3508 | ~ spl158_2 | ~ spl158_62 | ~ spl158_536), inference(avatar_split_clause, [], [f962829, f56678, f1355, f1053, f961759])).
fof(f962829, plain, (v__1(sortedList_first, sK97, sK97) | (~ spl158_2 | ~ spl158_62 | ~ spl158_536)), inference(subsumption_resolution, [], [f105184, f56679])).
fof(f105184, plain, (v__1(sortedList_first, sK97, sK97) | ~ object(sortedList_first) | (~ spl158_2 | ~ spl158_62)), inference(resolution, [], [f70135, f1054])).
fof(f70135, plain, (! [X2] : (~ v__1(X2, nn, nn) | v__1(X2, sK97, sK97) | ~ object(X2)) | ~ spl158_62), inference(subsumption_resolution, [], [f70120, f467])).
fof(f70120, plain, (! [X2] : (v__1(X2, sK97, sK97) | ~ v__1(X2, nn, nn) | ~ object(nn) | ~ object(X2)) | ~ spl158_62), inference(superposition, [], [f6563, f1357])).
fof(f955782, plain, (spl158_905 | ~ spl158_255 | ~ spl158_809), inference(avatar_split_clause, [], [f955781, f68838, f2297, f74113])).
fof(f955781, plain, ((nn = sK84) | (~ spl158_255 | ~ spl158_809)), inference(forward_demodulation, [], [f68840, f2299])).
fof(f68840, plain, ((sK84 = sK157) | ~ spl158_809), inference(avatar_component_clause, [], [f68838])).
fof(f298237, plain, (~ spl158_697 | ~ spl158_525), inference(avatar_split_clause, [], [f298236, f56605, f61554])).
fof(f61554, plain, (spl158_697 <=> v__1(nn, prev_2, prev_2)), introduced(avatar_definition, [new_symbols(naming, [spl158_697])])).
fof(f298236, plain, (~ v__1(nn, prev_2, prev_2) | ~ spl158_525), inference(subsumption_resolution, [], [f298235, f56606])).
fof(f298235, plain, (~ v__1(nn, prev_2, prev_2) | ~ object(prev_2)), inference(subsumption_resolution, [], [f298234, f1057])).
fof(f1057, plain, v__1(sortedList_first, prev_2, prev_2), inference(subsumption_resolution, [], [f500, f490])).
fof(f500, plain, (v__1(sortedList_first, prev_2, prev_2) | (prev_2 = null)), inference(cnf_transformation, [], [f31])).
fof(f31, plain, ((v__1(sortedList_first, prev_2, prev_2) & ~ (prev_2 = null)) | (prev_2 = null)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_17)).
fof(f298234, plain, (~ v__1(nn, prev_2, prev_2) | ~ v__1(sortedList_first, prev_2, prev_2) | ~ object(prev_2)), inference(subsumption_resolution, [], [f72697, f490])).
fof(f72697, plain, (~ v__1(nn, prev_2, prev_2) | (prev_2 = null) | ~ v__1(sortedList_first, prev_2, prev_2) | ~ object(prev_2)), inference(superposition, [], [f61407, f1060])).
fof(f163124, plain, (~ spl158_375 | spl158_460), inference(avatar_split_clause, [], [f79079, f11740, f5551])).
fof(f11740, plain, (spl158_460 <=> v__1(null, sK157, sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_460])])).
fof(f79079, plain, (~ sP52(sK157) | spl158_460), inference(resolution, [], [f11741, f689])).
fof(f689, plain, ! [X0] : (v__1(null, X0, X0) | ~ sP52(X0)), inference(cnf_transformation, [], [f310])).
fof(f11741, plain, (~ v__1(null, sK157, sK157) | spl158_460), inference(avatar_component_clause, [], [f11740])).
fof(f163116, plain, (spl158_375 | spl158_376 | spl158_1246 | ~ spl158_62 | ~ spl158_259), inference(avatar_split_clause, [], [f77465, f2315, f1355, f162713, f5555, f5551])).
fof(f77465, plain, ((sK97 = sK83(sK157)) | sP51(sK157) | sP52(sK157) | (~ spl158_62 | ~ spl158_259)), inference(resolution, [], [f72917, f2317])).
fof(f162693, plain, (~ spl158_376 | spl158_460), inference(avatar_split_clause, [], [f162692, f11740, f5555])).
fof(f162692, plain, (~ sP51(sK157) | spl158_460), inference(subsumption_resolution, [], [f79098, f983])).
fof(f79098, plain, (~ object(sK157) | ~ sP51(sK157) | spl158_460), inference(resolution, [], [f11741, f3422])).
fof(f79053, plain, (spl158_253 | ~ spl158_460), inference(avatar_contradiction_clause, [], [f79052])).
fof(f79052, plain, ($false | (spl158_253 | ~ spl158_460)), inference(subsumption_resolution, [], [f79051, f983])).
fof(f79051, plain, (~ object(sK157) | (spl158_253 | ~ spl158_460)), inference(subsumption_resolution, [], [f78985, f2290])).
fof(f2290, plain, (~ (null = sK157) | spl158_253), inference(avatar_component_clause, [], [f2289])).
fof(f2289, plain, (spl158_253 <=> (null = sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_253])])).
fof(f78985, plain, ((null = sK157) | ~ object(sK157) | ~ spl158_460), inference(resolution, [], [f11742, f3795])).
fof(f3795, plain, ! [X0] : (~ v__1(null, X0, X0) | (null = X0) | ~ object(X0)), inference(subsumption_resolution, [], [f3794, f472])).
fof(f3794, plain, ! [X0] : (~ v__1(null, X0, X0) | (null = X0) | ~ object(X0) | ~ object(null)), inference(trivial_inequality_removal, [], [f3792])).
fof(f3792, plain, ! [X0] : (~ (null = null) | ~ v__1(null, X0, X0) | (null = X0) | ~ object(X0) | ~ object(null)), inference(superposition, [], [f487, f476])).
fof(f476, plain, (null = node_next(null)), inference(cnf_transformation, [], [f15])).
fof(f15, plain, (null = node_next(null)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_1)).
fof(f487, plain, ! [X0, X1] : (~ (node_next(X0) = X0) | ~ v__1(X0, X1, X1) | (X0 = X1) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ! [X0, X1] : ((X0 = X1) | ~ v__1(X0, X1, X1) | ~ (node_next(X0) = X0) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f107])).
fof(f107, plain, ! [X0, X1] : ((X0 = X1) | ~ v__1(X0, X1, X1) | ~ (node_next(X0) = X0) | (~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f72])).
fof(f72, plain, ! [X0, X1] : ((X0 = X1) | ~ v__1(X0, X1, X1) | ~ (node_next(X0) = X0) | ~ (object(X1) & object(X0))), inference(rectify, [], [f23])).
fof(f23, plain, ! [X25, X26] : ((X25 = X26) | ~ v__1(X25, X26, X26) | ~ (node_next(X25) = X25) | ~ (object(X26) & object(X25))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_9)).
fof(f11742, plain, (v__1(null, sK157, sK157) | ~ spl158_460), inference(avatar_component_clause, [], [f11740])).
fof(f67758, plain, (~ spl158_308 | spl158_132 | ~ spl158_222 | ~ spl158_223), inference(avatar_split_clause, [], [f67557, f2131, f2122, f1677, f4036])).
fof(f4036, plain, (spl158_308 <=> sP76(nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_308])])).
fof(f1677, plain, (spl158_132 <=> v__1(sortedList_first, nn, prev_2)), introduced(avatar_definition, [new_symbols(naming, [spl158_132])])).
fof(f2122, plain, (spl158_222 <=> ! [X0] : ~ sP17(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_222])])).
fof(f2131, plain, (spl158_223 <=> ! [X0] : ~ sP16(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_223])])).
fof(f67557, plain, (~ sP76(nn) | (spl158_132 | ~ spl158_222 | ~ spl158_223)), inference(resolution, [], [f1678, f65920])).
fof(f65920, plain, (! [X0] : (v__1(sortedList_first, X0, prev_2) | ~ sP76(X0)) | (~ spl158_222 | ~ spl158_223)), inference(subsumption_resolution, [], [f65919, f2123])).
fof(f2123, plain, (! [X0] : ~ sP17(X0) | ~ spl158_222), inference(avatar_component_clause, [], [f2122])).
fof(f65919, plain, (! [X0] : (v__1(sortedList_first, X0, prev_2) | sP17(X0) | ~ sP76(X0)) | ~ spl158_223), inference(subsumption_resolution, [], [f553, f2132])).
fof(f2132, plain, (! [X0] : ~ sP16(X0) | ~ spl158_223), inference(avatar_component_clause, [], [f2131])).
fof(f553, plain, ! [X0] : (v__1(sortedList_first, X0, prev_2) | sP16(X0) | sP17(X0) | ~ sP76(X0)), inference(cnf_transformation, [], [f202])).
fof(f202, plain, ! [X0] : ((sP39(X0) & (sP17(X0) | sP16(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP13(X0) | sP15(X0) | sP14(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, prev_2, X0))) & ~ (prev_2 = X0)) | ~ sP76(X0)), inference(nnf_transformation, [], [f196])).
fof(f1678, plain, (~ v__1(sortedList_first, nn, prev_2) | spl158_132), inference(avatar_component_clause, [], [f1677])).
fof(f67757, plain, (~ spl158_259 | spl158_460 | spl158_376 | ~ spl158_408), inference(avatar_split_clause, [], [f67756, f8394, f5555, f11740, f2315])).
fof(f67756, plain, (v__1(null, sK157, sK157) | ~ sP74(sK157) | (spl158_376 | ~ spl158_408)), inference(subsumption_resolution, [], [f56693, f5556])).
fof(f56693, plain, (v__1(null, sK157, sK157) | sP51(sK157) | ~ sP74(sK157) | ~ spl158_408), inference(subsumption_resolution, [], [f8464, f689])).
fof(f8464, plain, (v__1(null, sK157, sK157) | sP51(sK157) | sP52(sK157) | ~ sP74(sK157) | ~ spl158_408), inference(superposition, [], [f564, f8396])).
fof(f8396, plain, ((null = sK83(sK157)) | ~ spl158_408), inference(avatar_component_clause, [], [f8394])).
fof(f67421, plain, (spl158_308 | ~ spl158_255 | ~ spl158_257), inference(avatar_split_clause, [], [f67356, f2306, f2297, f4036])).
fof(f2306, plain, (spl158_257 <=> sP76(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_257])])).
fof(f67356, plain, (sP76(nn) | (~ spl158_255 | ~ spl158_257)), inference(backward_demodulation, [], [f2308, f2299])).
fof(f2308, plain, (sP76(sK157) | ~ spl158_257), inference(avatar_component_clause, [], [f2306])).
fof(f67344, plain, (~ spl158_222 | ~ spl158_223 | spl158_254 | ~ spl158_257 | ~ spl158_525 | ~ spl158_536), inference(avatar_contradiction_clause, [], [f67343])).
fof(f67343, plain, ($false | (~ spl158_222 | ~ spl158_223 | spl158_254 | ~ spl158_257 | ~ spl158_525 | ~ spl158_536)), inference(subsumption_resolution, [], [f67342, f983])).
fof(f67342, plain, (~ object(sK157) | (~ spl158_222 | ~ spl158_223 | spl158_254 | ~ spl158_257 | ~ spl158_525 | ~ spl158_536)), inference(subsumption_resolution, [], [f67297, f2308])).
fof(f67297, plain, (~ sP76(sK157) | ~ object(sK157) | (~ spl158_222 | ~ spl158_223 | spl158_254 | ~ spl158_525 | ~ spl158_536)), inference(resolution, [], [f66235, f2295])).
fof(f66235, plain, (! [X9] : (v__1(sortedList_first, X9, X9) | ~ sP76(X9) | ~ object(X9)) | (~ spl158_222 | ~ spl158_223 | ~ spl158_525 | ~ spl158_536)), inference(subsumption_resolution, [], [f66234, f56679])).
fof(f66234, plain, (! [X9] : (~ sP76(X9) | v__1(sortedList_first, X9, X9) | ~ object(X9) | ~ object(sortedList_first)) | (~ spl158_222 | ~ spl158_223 | ~ spl158_525)), inference(subsumption_resolution, [], [f66085, f56606])).
fof(f66085, plain, (! [X9] : (~ sP76(X9) | v__1(sortedList_first, X9, X9) | ~ object(prev_2) | ~ object(X9) | ~ object(sortedList_first)) | (~ spl158_222 | ~ spl158_223)), inference(resolution, [], [f65920, f483])).
fof(f66059, plain, (spl158_697 | ~ spl158_132 | ~ spl158_525 | ~ spl158_536), inference(avatar_split_clause, [], [f66058, f56678, f56605, f1677, f61554])).
fof(f66058, plain, (v__1(nn, prev_2, prev_2) | (~ spl158_132 | ~ spl158_525 | ~ spl158_536)), inference(subsumption_resolution, [], [f66057, f56679])).
fof(f66057, plain, (v__1(nn, prev_2, prev_2) | ~ object(sortedList_first) | (~ spl158_132 | ~ spl158_525)), inference(subsumption_resolution, [], [f66056, f467])).
fof(f66056, plain, (v__1(nn, prev_2, prev_2) | ~ object(nn) | ~ object(sortedList_first) | (~ spl158_132 | ~ spl158_525)), inference(subsumption_resolution, [], [f66042, f56606])).
fof(f66042, plain, (v__1(nn, prev_2, prev_2) | ~ object(prev_2) | ~ object(nn) | ~ object(sortedList_first) | ~ spl158_132), inference(resolution, [], [f1679, f484])).
fof(f1679, plain, (v__1(sortedList_first, nn, prev_2) | ~ spl158_132), inference(avatar_component_clause, [], [f1677])).
fof(f65878, plain, (spl158_132 | ~ spl158_239 | ~ spl158_240 | ~ spl158_241 | ~ spl158_534), inference(avatar_contradiction_clause, [], [f65877])).
fof(f65877, plain, ($false | (spl158_132 | ~ spl158_239 | ~ spl158_240 | ~ spl158_241 | ~ spl158_534)), inference(subsumption_resolution, [], [f65839, f56670])).
fof(f56670, plain, (sP77(nn) | ~ spl158_534), inference(avatar_component_clause, [], [f56669])).
fof(f56669, plain, (spl158_534 <=> sP77(nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_534])])).
fof(f65839, plain, (~ sP77(nn) | (spl158_132 | ~ spl158_239 | ~ spl158_240 | ~ spl158_241)), inference(resolution, [], [f65769, f1678])).
fof(f65769, plain, (! [X0] : (v__1(sortedList_first, X0, prev_2) | ~ sP77(X0)) | (~ spl158_239 | ~ spl158_240 | ~ spl158_241)), inference(subsumption_resolution, [], [f65768, f2233])).
fof(f2233, plain, (! [X0] : ~ sP6(X0) | ~ spl158_241), inference(avatar_component_clause, [], [f2232])).
fof(f2232, plain, (spl158_241 <=> ! [X0] : ~ sP6(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_241])])).
fof(f65768, plain, (! [X0] : (v__1(sortedList_first, X0, prev_2) | sP6(X0) | ~ sP77(X0)) | (~ spl158_239 | ~ spl158_240)), inference(subsumption_resolution, [], [f65767, f2224])).
fof(f2224, plain, (! [X0] : ~ sP7(X0) | ~ spl158_240), inference(avatar_component_clause, [], [f2223])).
fof(f2223, plain, (spl158_240 <=> ! [X0] : ~ sP7(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_240])])).
fof(f65767, plain, (! [X0] : (v__1(sortedList_first, X0, prev_2) | sP7(X0) | sP6(X0) | ~ sP77(X0)) | ~ spl158_239), inference(subsumption_resolution, [], [f546, f2215])).
fof(f2215, plain, (! [X0] : ~ sP8(X0) | ~ spl158_239), inference(avatar_component_clause, [], [f2214])).
fof(f2214, plain, (spl158_239 <=> ! [X0] : ~ sP8(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_239])])).
fof(f546, plain, ! [X0] : (v__1(sortedList_first, X0, prev_2) | sP8(X0) | sP7(X0) | sP6(X0) | ~ sP77(X0)), inference(cnf_transformation, [], [f201])).
fof(f201, plain, ! [X0] : (((sP6(X0) | sP8(X0) | sP7(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, prev_2, prev_2)) | v__1(sortedList_first, prev_2, nn)) & v__1(sortedList_first, X0, prev_2))) & (sP1(X0) | sP0(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)))) | ~ sP77(X0)), inference(nnf_transformation, [], [f197])).
fof(f65761, plain, (spl158_5 | spl158_132 | ~ spl158_2 | ~ spl158_525 | ~ spl158_536), inference(avatar_split_clause, [], [f65760, f56678, f56605, f1053, f1677, f1084])).
fof(f1084, plain, (spl158_5 <=> v__1(sortedList_first, prev_2, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_5])])).
fof(f65760, plain, (v__1(sortedList_first, nn, prev_2) | v__1(sortedList_first, prev_2, nn) | (~ spl158_2 | ~ spl158_525 | ~ spl158_536)), inference(subsumption_resolution, [], [f65700, f56606])).
fof(f65700, plain, (v__1(sortedList_first, nn, prev_2) | v__1(sortedList_first, prev_2, nn) | ~ object(prev_2) | (~ spl158_2 | ~ spl158_536)), inference(resolution, [], [f59640, f1057])).
fof(f59640, plain, (! [X4] : (~ v__1(sortedList_first, X4, X4) | v__1(sortedList_first, nn, X4) | v__1(sortedList_first, X4, nn) | ~ object(X4)) | (~ spl158_2 | ~ spl158_536)), inference(subsumption_resolution, [], [f59639, f56679])).
fof(f59639, plain, (! [X4] : (v__1(sortedList_first, X4, nn) | v__1(sortedList_first, nn, X4) | ~ v__1(sortedList_first, X4, X4) | ~ object(X4) | ~ object(sortedList_first)) | ~ spl158_2), inference(subsumption_resolution, [], [f59621, f467])).
fof(f59621, plain, (! [X4] : (v__1(sortedList_first, X4, nn) | v__1(sortedList_first, nn, X4) | ~ v__1(sortedList_first, X4, X4) | ~ object(nn) | ~ object(X4) | ~ object(sortedList_first)) | ~ spl158_2), inference(resolution, [], [f1054, f482])).
fof(f482, plain, ! [X2, X0, X1] : (~ v__1(X0, X2, X2) | v__1(X0, X1, X2) | v__1(X0, X2, X1) | ~ v__1(X0, X1, X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ! [X0, X1, X2] : (v__1(X0, X2, X1) | v__1(X0, X1, X2) | ~ v__1(X0, X2, X2) | ~ v__1(X0, X1, X1) | ~ object(X2) | ~ object(X1) | ~ object(X0)), inference(flattening, [], [f101])).
fof(f101, plain, ! [X0, X1, X2] : (v__1(X0, X2, X1) | v__1(X0, X1, X2) | ~ v__1(X0, X2, X2) | ~ v__1(X0, X1, X1) | (~ object(X2) | ~ object(X1) | ~ object(X0))), inference(ennf_transformation, [], [f68])).
fof(f68, plain, ! [X0, X1, X2] : (v__1(X0, X2, X1) | v__1(X0, X1, X2) | ~ v__1(X0, X2, X2) | ~ v__1(X0, X1, X1) | ~ (object(X2) & object(X1) & object(X0))), inference(rectify, [], [f19])).
fof(f19, plain, ! [X16, X17, X18] : (v__1(X16, X18, X17) | v__1(X16, X17, X18) | ~ v__1(X16, X18, X18) | ~ v__1(X16, X17, X17) | ~ (object(X18) & object(X17) & object(X16))), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_5)).
fof(f59595, plain, (spl158_534 | ~ spl158_255 | ~ spl158_256), inference(avatar_split_clause, [], [f59316, f2302, f2297, f56669])).
fof(f2302, plain, (spl158_256 <=> sP77(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_256])])).
fof(f59316, plain, (sP77(nn) | (~ spl158_255 | ~ spl158_256)), inference(backward_demodulation, [], [f2304, f2299])).
fof(f2304, plain, (sP77(sK157) | ~ spl158_256), inference(avatar_component_clause, [], [f2302])).
fof(f59030, plain, (spl158_253 | ~ spl158_269), inference(avatar_contradiction_clause, [], [f59029])).
fof(f59029, plain, ($false | (spl158_253 | ~ spl158_269)), inference(subsumption_resolution, [], [f59028, f2290])).
fof(f59028, plain, ((null = sK157) | ~ spl158_269), inference(subsumption_resolution, [], [f59023, f983])).
fof(f59023, plain, (~ object(sK157) | (null = sK157) | ~ spl158_269), inference(resolution, [], [f2490, f9992])).
fof(f9992, plain, ! [X16] : (~ sP1(X16) | ~ object(X16) | (null = X16)), inference(resolution, [], [f3795, f974])).
fof(f974, plain, ! [X0] : (v__1(null, X0, X0) | ~ sP1(X0)), inference(cnf_transformation, [], [f456])).
fof(f456, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(null, X0, X0) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP1(X0)), inference(nnf_transformation, [], [f121])).
fof(f2490, plain, (sP1(sK157) | ~ spl158_269), inference(avatar_component_clause, [], [f2488])).
fof(f2488, plain, (spl158_269 <=> sP1(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_269])])).
fof(f59021, plain, (spl158_253 | ~ spl158_270), inference(avatar_contradiction_clause, [], [f59020])).
fof(f59020, plain, ($false | (spl158_253 | ~ spl158_270)), inference(subsumption_resolution, [], [f59019, f2290])).
fof(f59019, plain, ((null = sK157) | ~ spl158_270), inference(subsumption_resolution, [], [f59014, f983])).
fof(f59014, plain, (~ object(sK157) | (null = sK157) | ~ spl158_270), inference(resolution, [], [f2494, f10022])).
fof(f10022, plain, ! [X29] : (~ sP0(X29) | ~ object(X29) | (null = X29)), inference(duplicate_literal_removal, [], [f10005])).
fof(f10005, plain, ! [X29] : ((null = X29) | ~ object(X29) | ~ object(X29) | ~ sP0(X29)), inference(resolution, [], [f3795, f3436])).
fof(f3436, plain, ! [X37] : (v__1(null, X37, X37) | ~ object(X37) | ~ sP0(X37)), inference(subsumption_resolution, [], [f3435, f472])).
fof(f3435, plain, ! [X37] : (v__1(null, X37, X37) | ~ object(X37) | ~ object(null) | ~ sP0(X37)), inference(subsumption_resolution, [], [f3381, f467])).
fof(f3381, plain, ! [X37] : (v__1(null, X37, X37) | ~ object(nn) | ~ object(X37) | ~ object(null) | ~ sP0(X37)), inference(duplicate_literal_removal, [], [f3322])).
fof(f3322, plain, ! [X37] : (v__1(null, X37, X37) | ~ object(nn) | ~ object(X37) | ~ object(null) | v__1(null, X37, X37) | ~ sP0(X37)), inference(resolution, [], [f483, f981])).
fof(f981, plain, ! [X0] : (v__1(null, X0, nn) | v__1(null, X0, X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f457])).
fof(f457, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, X0, X0)) | v__1(null, X0, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, X0, X0) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, X0)) & ~ (nn = X0)) | ~ sP0(X0)), inference(nnf_transformation, [], [f120])).
fof(f2494, plain, (sP0(sK157) | ~ spl158_270), inference(avatar_component_clause, [], [f2492])).
fof(f2492, plain, (spl158_270 <=> sP0(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_270])])).
fof(f59001, plain, spl158_536, inference(avatar_split_clause, [], [f474, f56678])).
fof(f474, plain, object(sortedList_first), inference(cnf_transformation, [], [f13])).
fof(f13, plain, object(sortedList_first), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_sortedList_first)).
fof(f58879, plain, (spl158_17 | ~ spl158_20 | ~ spl158_525), inference(avatar_contradiction_clause, [], [f58878])).
fof(f58878, plain, ($false | (spl158_17 | ~ spl158_20 | ~ spl158_525)), inference(subsumption_resolution, [], [f58877, f472])).
fof(f58877, plain, (~ object(null) | (spl158_17 | ~ spl158_20 | ~ spl158_525)), inference(subsumption_resolution, [], [f58876, f56606])).
fof(f58876, plain, (~ object(prev_2) | ~ object(null) | (spl158_17 | ~ spl158_20)), inference(subsumption_resolution, [], [f58875, f467])).
fof(f58875, plain, (~ object(nn) | ~ object(prev_2) | ~ object(null) | (spl158_17 | ~ spl158_20)), inference(subsumption_resolution, [], [f58862, f1135])).
fof(f1135, plain, (~ v__1(null, prev_2, prev_2) | spl158_17), inference(avatar_component_clause, [], [f1133])).
fof(f1133, plain, (spl158_17 <=> v__1(null, prev_2, prev_2)), introduced(avatar_definition, [new_symbols(naming, [spl158_17])])).
fof(f58862, plain, (v__1(null, prev_2, prev_2) | ~ object(nn) | ~ object(prev_2) | ~ object(null) | ~ spl158_20), inference(resolution, [], [f1147, f483])).
fof(f1147, plain, (v__1(null, prev_2, nn) | ~ spl158_20), inference(avatar_component_clause, [], [f1146])).
fof(f1146, plain, (spl158_20 <=> v__1(null, prev_2, nn)), introduced(avatar_definition, [new_symbols(naming, [spl158_20])])).
fof(f56723, plain, (~ spl158_255 | ~ spl158_259 | spl158_312), inference(avatar_contradiction_clause, [], [f56722])).
fof(f56722, plain, ($false | (~ spl158_255 | ~ spl158_259 | spl158_312)), inference(subsumption_resolution, [], [f56721, f4070])).
fof(f4070, plain, (~ sP74(nn) | spl158_312), inference(avatar_component_clause, [], [f4068])).
fof(f56721, plain, (sP74(nn) | (~ spl158_255 | ~ spl158_259)), inference(forward_demodulation, [], [f2317, f2299])).
fof(f56665, plain, spl158_525, inference(avatar_split_clause, [], [f469, f56605])).
fof(f469, plain, object(prev_2), inference(cnf_transformation, [], [f8])).
fof(f8, plain, object(prev_2), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', sort_info_prev_2)).
fof(f10739, plain, ~ spl158_18, inference(avatar_contradiction_clause, [], [f10738])).
fof(f10738, plain, ($false | ~ spl158_18), inference(subsumption_resolution, [], [f10737, f467])).
fof(f10737, plain, (~ object(nn) | ~ spl158_18), inference(subsumption_resolution, [], [f10709, f496])).
fof(f10709, plain, ((nn = null) | ~ object(nn) | ~ spl158_18), inference(resolution, [], [f1139, f3795])).
fof(f1139, plain, (v__1(null, nn, nn) | ~ spl158_18), inference(avatar_component_clause, [], [f1137])).
fof(f10035, plain, ~ spl158_17, inference(avatar_contradiction_clause, [], [f10034])).
fof(f10034, plain, ($false | ~ spl158_17), inference(subsumption_resolution, [], [f10033, f469])).
fof(f10033, plain, (~ object(prev_2) | ~ spl158_17), inference(subsumption_resolution, [], [f10014, f490])).
fof(f10014, plain, ((prev_2 = null) | ~ object(prev_2) | ~ spl158_17), inference(resolution, [], [f3795, f1134])).
fof(f1134, plain, (v__1(null, prev_2, prev_2) | ~ spl158_17), inference(avatar_component_clause, [], [f1133])).
fof(f5395, plain, (~ spl158_12 | ~ spl158_15 | ~ spl158_26 | ~ spl158_31 | ~ spl158_258), inference(avatar_contradiction_clause, [], [f5394])).
fof(f5394, plain, ($false | (~ spl158_12 | ~ spl158_15 | ~ spl158_26 | ~ spl158_31 | ~ spl158_258)), inference(subsumption_resolution, [], [f2312, f2354])).
fof(f2354, plain, (! [X0] : ~ sP75(X0) | (~ spl158_12 | ~ spl158_15 | ~ spl158_26 | ~ spl158_31)), inference(subsumption_resolution, [], [f2353, f1176])).
fof(f1176, plain, (! [X0] : ~ sP72(X0) | ~ spl158_26), inference(avatar_component_clause, [], [f1175])).
fof(f1175, plain, (spl158_26 <=> ! [X0] : ~ sP72(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_26])])).
fof(f2353, plain, (! [X0] : (~ sP75(X0) | sP72(X0)) | (~ spl158_12 | ~ spl158_15 | ~ spl158_31)), inference(subsumption_resolution, [], [f2352, f1201])).
fof(f1201, plain, (! [X0] : ~ sP71(X0) | ~ spl158_31), inference(avatar_component_clause, [], [f1200])).
fof(f1200, plain, (spl158_31 <=> ! [X0] : ~ sP71(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_31])])).
fof(f2352, plain, (! [X0] : (~ sP75(X0) | sP71(X0) | sP72(X0)) | (~ spl158_12 | ~ spl158_15)), inference(subsumption_resolution, [], [f1113, f1127])).
fof(f1127, plain, (! [X0] : ~ sP73(X0) | ~ spl158_15), inference(avatar_component_clause, [], [f1126])).
fof(f1126, plain, (spl158_15 <=> ! [X0] : ~ sP73(X0)), introduced(avatar_definition, [new_symbols(naming, [spl158_15])])).
fof(f1113, plain, (! [X0] : (sP73(X0) | ~ sP75(X0) | sP71(X0) | sP72(X0)) | ~ spl158_12), inference(avatar_component_clause, [], [f1112])).
fof(f1112, plain, (spl158_12 <=> ! [X0] : (sP73(X0) | ~ sP75(X0) | sP71(X0) | sP72(X0))), introduced(avatar_definition, [new_symbols(naming, [spl158_12])])).
fof(f2312, plain, (sP75(sK157) | ~ spl158_258), inference(avatar_component_clause, [], [f2310])).
fof(f2310, plain, (spl158_258 <=> sP75(sK157)), introduced(avatar_definition, [new_symbols(naming, [spl158_258])])).
fof(f2495, plain, (spl158_269 | spl158_270 | spl158_254 | ~ spl158_256), inference(avatar_split_clause, [], [f2486, f2302, f2293, f2492, f2488])).
fof(f2486, plain, (sP0(sK157) | sP1(sK157) | (spl158_254 | ~ spl158_256)), inference(subsumption_resolution, [], [f2485, f2304])).
fof(f2485, plain, (sP0(sK157) | sP1(sK157) | ~ sP77(sK157) | spl158_254), inference(resolution, [], [f543, f2295])).
fof(f543, plain, ! [X0] : (v__1(sortedList_first, X0, X0) | sP0(X0) | sP1(X0) | ~ sP77(X0)), inference(cnf_transformation, [], [f201])).
fof(f2347, plain, ~ spl158_253, inference(avatar_split_clause, [], [f984, f2289])).
fof(f984, plain, ~ (null = sK157), inference(cnf_transformation, [], [f459])).
fof(f2318, plain, (spl158_256 | spl158_257 | spl158_259), inference(avatar_split_clause, [], [f989, f2315, f2306, f2302])).
fof(f989, plain, (sP74(sK157) | sP76(sK157) | sP77(sK157)), inference(cnf_transformation, [], [f459])).
fof(f2313, plain, (spl158_256 | spl158_257 | spl158_258), inference(avatar_split_clause, [], [f990, f2310, f2306, f2302])).
fof(f990, plain, (sP75(sK157) | sP76(sK157) | sP77(sK157)), inference(cnf_transformation, [], [f459])).
fof(f2300, plain, (spl158_253 | ~ spl158_254 | spl158_255), inference(avatar_split_clause, [], [f991, f2297, f2293, f2289])).
fof(f991, plain, ((nn = sK157) | ~ v__1(sortedList_first, sK157, sK157) | (null = sK157)), inference(cnf_transformation, [], [f459])).
fof(f2247, plain, (spl158_241 | ~ spl158_5), inference(avatar_split_clause, [], [f2246, f1084, f2232])).
fof(f2246, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, nn) | ~ sP6(X0)), inference(subsumption_resolution, [], [f947, f1057])).
fof(f947, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, nn) | ~ v__1(sortedList_first, prev_2, prev_2) | ~ sP6(X0)), inference(cnf_transformation, [], [f451])).
fof(f451, plain, ! [X0] : ((sP5 & sP4 & (((v__1(sortedList_first, nn, nn) | ~ v__1(sortedList_first, prev_2, prev_2)) & ~ v__1(sortedList_first, prev_2, nn)) | ~ v__1(sortedList_first, prev_2, prev_2)) & (sP3(X0) | sP2(X0) | (((~ v__1(sortedList_first, nn, nn) & v__1(sortedList_first, X0, X0)) | v__1(sortedList_first, X0, nn)) & v__1(sortedList_first, X0, X0)))) | ~ sP6(X0)), inference(nnf_transformation, [], [f126])).
fof(f2228, plain, (spl158_240 | spl158_132), inference(avatar_split_clause, [], [f2227, f1677, f2223])).
fof(f2227, plain, ! [X0] : (v__1(sortedList_first, nn, prev_2) | ~ sP7(X0)), inference(subsumption_resolution, [], [f940, f1057])).
fof(f940, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, prev_2) | v__1(sortedList_first, nn, prev_2) | ~ sP7(X0)), inference(cnf_transformation, [], [f450])).
fof(f450, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP7(X0)), inference(nnf_transformation, [], [f127])).
fof(f2219, plain, (spl158_239 | spl158_132), inference(avatar_split_clause, [], [f2218, f1677, f2214])).
fof(f2218, plain, ! [X0] : (v__1(sortedList_first, nn, prev_2) | ~ sP8(X0)), inference(subsumption_resolution, [], [f934, f1057])).
fof(f934, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, prev_2) | v__1(sortedList_first, nn, prev_2) | ~ sP8(X0)), inference(cnf_transformation, [], [f449])).
fof(f449, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP8(X0)), inference(nnf_transformation, [], [f128])).
fof(f2136, plain, (spl158_223 | spl158_132), inference(avatar_split_clause, [], [f2135, f1677, f2131])).
fof(f2135, plain, ! [X0] : (v__1(sortedList_first, nn, prev_2) | ~ sP16(X0)), inference(subsumption_resolution, [], [f889, f1057])).
fof(f889, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, prev_2) | v__1(sortedList_first, nn, prev_2) | ~ sP16(X0)), inference(cnf_transformation, [], [f441])).
fof(f441, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(sortedList_first, X0, nn) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP16(X0)), inference(nnf_transformation, [], [f136])).
fof(f2127, plain, (spl158_222 | spl158_132), inference(avatar_split_clause, [], [f2126, f1677, f2122])).
fof(f2126, plain, ! [X0] : (v__1(sortedList_first, nn, prev_2) | ~ sP17(X0)), inference(subsumption_resolution, [], [f883, f1057])).
fof(f883, plain, ! [X0] : (~ v__1(sortedList_first, prev_2, prev_2) | v__1(sortedList_first, nn, prev_2) | ~ sP17(X0)), inference(cnf_transformation, [], [f440])).
fof(f440, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & ((~ v__1(sortedList_first, prev_2, prev_2) & v__1(sortedList_first, nn, nn)) | v__1(sortedList_first, nn, prev_2)) & ~ (nn = prev_2)) | ~ sP17(X0)), inference(nnf_transformation, [], [f137])).
fof(f1544, plain, (~ spl158_11 | spl158_101), inference(avatar_split_clause, [], [f680, f1541, f1108])).
fof(f1108, plain, (spl158_11 <=> sP53), introduced(avatar_definition, [new_symbols(naming, [spl158_11])])).
fof(f680, plain, (v__1(sK114, prev_2, prev_2) | ~ sP53), inference(cnf_transformation, [], [f306])).
fof(f306, plain, (((~ v__1(sK113, nn, nn) & (node_next(nn) = sK113) & object(sK113)) & ((node_next(nn) = sK114) & v__1(sK114, prev_2, prev_2) & object(sK114))) | ~ sP53), inference(skolemisation, [status(esa), new_symbols(skolem, [sK113, sK114])], [f303, f305, f304])).
fof(f304, plain, (? [X0] : (~ v__1(X0, nn, nn) & (node_next(nn) = X0) & object(X0)) => (~ v__1(sK113, nn, nn) & (node_next(nn) = sK113) & object(sK113))), introduced(choice_axiom, [])).
fof(f305, plain, (? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, prev_2) & object(X1)) => ((node_next(nn) = sK114) & v__1(sK114, prev_2, prev_2) & object(sK114))), introduced(choice_axiom, [])).
fof(f303, plain, ((? [X0] : (~ v__1(X0, nn, nn) & (node_next(nn) = X0) & object(X0)) & ? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, prev_2) & object(X1))) | ~ sP53), inference(rectify, [], [f302])).
fof(f302, plain, ((? [X30] : (~ v__1(X30, nn, nn) & (node_next(nn) = X30) & object(X30)) & ? [X31] : ((node_next(nn) = X31) & v__1(X31, prev_2, prev_2) & object(X31))) | ~ sP53), inference(nnf_transformation, [], [f173])).
fof(f1539, plain, (~ spl158_11 | spl158_100), inference(avatar_split_clause, [], [f681, f1536, f1108])).
fof(f681, plain, ((node_next(nn) = sK114) | ~ sP53), inference(cnf_transformation, [], [f306])).
fof(f1358, plain, (~ spl158_23 | spl158_62), inference(avatar_split_clause, [], [f621, f1355, f1160])).
fof(f1160, plain, (spl158_23 <=> sP64), introduced(avatar_definition, [new_symbols(naming, [spl158_23])])).
fof(f621, plain, ((node_next(nn) = sK97) | ~ sP64), inference(cnf_transformation, [], [f256])).
fof(f256, plain, (((((node_next(nn) = sK96) & v__1(sK96, nn, nn) & object(sK96)) | sP63) & (~ v__1(sK97, prev_2, nn) & (node_next(nn) = sK97) & object(sK97))) | ~ sP64), inference(skolemisation, [status(esa), new_symbols(skolem, [sK96, sK97])], [f253, f255, f254])).
fof(f254, plain, (? [X0] : ((node_next(nn) = X0) & v__1(X0, nn, nn) & object(X0)) => ((node_next(nn) = sK96) & v__1(sK96, nn, nn) & object(sK96))), introduced(choice_axiom, [])).
fof(f255, plain, (? [X1] : (~ v__1(X1, prev_2, nn) & (node_next(nn) = X1) & object(X1)) => (~ v__1(sK97, prev_2, nn) & (node_next(nn) = sK97) & object(sK97))), introduced(choice_axiom, [])).
fof(f253, plain, (((? [X0] : ((node_next(nn) = X0) & v__1(X0, nn, nn) & object(X0)) | sP63) & ? [X1] : (~ v__1(X1, prev_2, nn) & (node_next(nn) = X1) & object(X1))) | ~ sP64), inference(rectify, [], [f252])).
fof(f252, plain, (((? [X8] : ((node_next(nn) = X8) & v__1(X8, nn, nn) & object(X8)) | sP63) & ? [X10] : (~ v__1(X10, prev_2, nn) & (node_next(nn) = X10) & object(X10))) | ~ sP64), inference(nnf_transformation, [], [f184])).
fof(f1203, plain, (spl158_31 | spl158_20 | spl158_17), inference(avatar_split_clause, [], [f588, f1133, f1146, f1200])).
fof(f588, plain, ! [X0] : (v__1(null, prev_2, prev_2) | v__1(null, prev_2, nn) | ~ sP71(X0)), inference(cnf_transformation, [], [f224])).
fof(f224, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ((node_next(nn) = sK86(X0)) & v__1(sK86(X0), X0, nn) & object(sK86(X0))) & sP55 & ~ (nn = prev_2)) | ~ sP71(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK86])], [f222, f223])).
fof(f223, plain, ! [X0] : (? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1)) => ((node_next(nn) = sK86(X0)) & v__1(sK86(X0), X0, nn) & object(sK86(X0)))), introduced(choice_axiom, [])).
fof(f222, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X1] : ((node_next(nn) = X1) & v__1(X1, X0, nn) & object(X1)) & sP55 & ~ (nn = prev_2)) | ~ sP71(X0)), inference(rectify, [], [f221])).
fof(f221, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & ? [X26] : ((node_next(nn) = X26) & v__1(X26, X0, nn) & object(X26)) & sP55 & ~ (nn = prev_2)) | ~ sP71(X0)), inference(nnf_transformation, [], [f191])).
fof(f1178, plain, (spl158_26 | spl158_20 | spl158_17), inference(avatar_split_clause, [], [f581, f1133, f1146, f1175])).
fof(f581, plain, ! [X0] : (v__1(null, prev_2, prev_2) | v__1(null, prev_2, nn) | ~ sP72(X0)), inference(cnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP56 | ((node_next(nn) = sK85) & v__1(sK85, nn, prev_2) & object(sK85))) & ~ (nn = prev_2)) | ~ sP72(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK85])], [f218, f219])).
fof(f219, plain, (? [X1] : ((node_next(nn) = X1) & v__1(X1, nn, prev_2) & object(X1)) => ((node_next(nn) = sK85) & v__1(sK85, nn, prev_2) & object(sK85))), introduced(choice_axiom, [])).
fof(f218, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP56 | ? [X1] : ((node_next(nn) = X1) & v__1(X1, nn, prev_2) & object(X1))) & ~ (nn = prev_2)) | ~ sP72(X0)), inference(rectify, [], [f217])).
fof(f217, plain, ! [X0] : ((((~ v__1(null, nn, nn) & v__1(null, prev_2, prev_2)) | v__1(null, prev_2, nn)) & v__1(null, X0, prev_2) & (sP56 | ? [X25] : ((node_next(nn) = X25) & v__1(X25, nn, prev_2) & object(X25))) & ~ (nn = prev_2)) | ~ sP72(X0)), inference(nnf_transformation, [], [f192])).
fof(f1168, plain, (spl158_15 | spl158_24 | spl158_23), inference(avatar_split_clause, [], [f571, f1160, f1165, f1126])).
fof(f571, plain, ! [X0] : (sP64 | (node_next(nn) = sK84) | ~ sP73(X0)), inference(cnf_transformation, [], [f216])).
fof(f216, plain, ! [X0] : (((sP68 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP70 & (sP64 | (~ v__1(sK84, prev_2, prev_2) & (node_next(nn) = sK84) & object(sK84))) & sP69(X0)) | ~ sP73(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK84])], [f214, f215])).
fof(f215, plain, (? [X1] : (~ v__1(X1, prev_2, prev_2) & (node_next(nn) = X1) & object(X1)) => (~ v__1(sK84, prev_2, prev_2) & (node_next(nn) = sK84) & object(sK84))), introduced(choice_axiom, [])).
fof(f214, plain, ! [X0] : (((sP68 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP70 & (sP64 | ? [X1] : (~ v__1(X1, prev_2, prev_2) & (node_next(nn) = X1) & object(X1))) & sP69(X0)) | ~ sP73(X0)), inference(rectify, [], [f213])).
fof(f213, plain, ! [X0] : (((sP68 | ((v__1(null, nn, nn) | ~ v__1(null, prev_2, prev_2)) & ~ v__1(null, prev_2, nn)) | ~ v__1(null, prev_2, prev_2) | (nn = prev_2)) & sP70 & (sP64 | ? [X11] : (~ v__1(X11, prev_2, prev_2) & (node_next(nn) = X11) & object(X11))) & sP69(X0)) | ~ sP73(X0)), inference(nnf_transformation, [], [f193])).
fof(f1119, plain, (spl158_13 | spl158_11 | spl158_12), inference(avatar_split_clause, [], [f561, f1112, f1108, f1116])).
fof(f561, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | sP53 | v__1(sK80, prev_2, nn) | ~ sP75(X0)), inference(cnf_transformation, [], [f207])).
fof(f207, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | ((sP53 | ((node_next(nn) = sK80) & v__1(sK80, prev_2, nn) & object(sK80))) & ((node_next(nn) = sK81(X0)) & v__1(sK81(X0), X0, prev_2) & object(sK81(X0)))) | ~ sP75(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK80, sK81])], [f204, f206, f205])).
fof(f205, plain, (? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, nn) & object(X1)) => ((node_next(nn) = sK80) & v__1(sK80, prev_2, nn) & object(sK80))), introduced(choice_axiom, [])).
fof(f206, plain, ! [X0] : (? [X2] : ((node_next(nn) = X2) & v__1(X2, X0, prev_2) & object(X2)) => ((node_next(nn) = sK81(X0)) & v__1(sK81(X0), X0, prev_2) & object(sK81(X0)))), introduced(choice_axiom, [])).
fof(f204, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | ((sP53 | ? [X1] : ((node_next(nn) = X1) & v__1(X1, prev_2, nn) & object(X1))) & ? [X2] : ((node_next(nn) = X2) & v__1(X2, X0, prev_2) & object(X2))) | ~ sP75(X0)), inference(rectify, [], [f203])).
fof(f203, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | ((sP53 | ? [X32] : ((node_next(nn) = X32) & v__1(X32, prev_2, nn) & object(X32))) & ? [X33] : ((node_next(nn) = X33) & v__1(X33, X0, prev_2) & object(X33))) | ~ sP75(X0)), inference(nnf_transformation, [], [f195])).
fof(f1114, plain, (spl158_10 | spl158_11 | spl158_12), inference(avatar_split_clause, [], [f562, f1112, f1108, f1104])).
fof(f562, plain, ! [X0] : (sP73(X0) | sP72(X0) | sP71(X0) | sP53 | (node_next(nn) = sK80) | ~ sP75(X0)), inference(cnf_transformation, [], [f207])).
fof(f1059, plain, spl158_2, inference(avatar_split_clause, [], [f1058, f1053])).
fof(f1058, plain, v__1(sortedList_first, nn, nn), inference(subsumption_resolution, [], [f502, f496])).
fof(f502, plain, (v__1(sortedList_first, nn, nn) | (nn = null)), inference(cnf_transformation, [], [f32])).
fof(f32, plain, ((v__1(sortedList_first, nn, nn) & ~ (nn = null)) | (nn = null)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_18)).
fof(f1056, plain, (spl158_1 | ~ spl158_2), inference(avatar_split_clause, [], [f1047, f1053, f1049])).
fof(f1047, plain, (~ v__1(sortedList_first, nn, nn) | v__1(nn, nn, nn)), inference(subsumption_resolution, [], [f498, f496])).
fof(f498, plain, (~ v__1(sortedList_first, nn, nn) | (nn = null) | v__1(nn, nn, nn)), inference(cnf_transformation, [], [f30])).
fof(f30, plain, (~ v__1(sortedList_first, nn, nn) | (nn = null) | v__1(nn, nn, nn)), file('/home/ubuntu/library/tptp/Problems/SWW/SWW095+1.p', hyp_tptp_16)).