fof(f719, plain, $false, inference(avatar_sat_refutation, [], [f236, f240, f249, f250, f264, f268, f276, f281, f286, f295, f300, f304, f309, f310, f314, f323, f334, f335, f348, f351, f352, f357, f360, f369, f374, f377, f386, f391, f392, f403, f408, f419, f420, f425, f428, f445, f466, f467, f470, f478, f482, f489, f494, f503, f508, f512, f516, f521, f522, f523, f528, f537, f547, f552, f554, f562, f570, f575, f577, f578, f579, f580, f581, f583, f585, f587, f589, f591, f597, f606, f612, f614, f616, f618, f620, f622, f624, f626, f628, f632, f637, f641, f643, f645, f660, f667, f676, f680, f682, f700, f702, f704, f706, f708, f710, f712, f714, f716, f718])).
fof(f718, plain, (~ spl49_9 | ~ spl49_34), inference(avatar_contradiction_clause, [], [f717])).
fof(f717, plain, ($false | (~ spl49_9 | ~ spl49_34)), inference(subsumption_resolution, [], [f396, f267])).
fof(f267, plain, (! [X2] : ~ p(X2) | ~ spl49_9), inference(avatar_component_clause, [], [f266])).
fof(f266, plain, (spl49_9 <=> ! [X2] : ~ p(X2)), introduced(avatar_definition, [new_symbols(naming, [spl49_9])])).
fof(f396, plain, (p(sK35) | ~ spl49_34), inference(avatar_component_clause, [], [f394])).
fof(f394, plain, (spl49_34 <=> p(sK35)), introduced(avatar_definition, [new_symbols(naming, [spl49_34])])).
fof(f716, plain, (~ spl49_9 | ~ spl49_35), inference(avatar_contradiction_clause, [], [f715])).
fof(f715, plain, ($false | (~ spl49_9 | ~ spl49_35)), inference(subsumption_resolution, [], [f401, f267])).
fof(f401, plain, (p(sK34) | ~ spl49_35), inference(avatar_component_clause, [], [f399])).
fof(f399, plain, (spl49_35 <=> p(sK34)), introduced(avatar_definition, [new_symbols(naming, [spl49_35])])).
fof(f714, plain, (~ spl49_9 | ~ spl49_32), inference(avatar_contradiction_clause, [], [f713])).
fof(f713, plain, ($false | (~ spl49_9 | ~ spl49_32)), inference(subsumption_resolution, [], [f385, f267])).
fof(f385, plain, (p(sK32) | ~ spl49_32), inference(avatar_component_clause, [], [f383])).
fof(f383, plain, (spl49_32 <=> p(sK32)), introduced(avatar_definition, [new_symbols(naming, [spl49_32])])).
fof(f712, plain, (~ spl49_9 | ~ spl49_31), inference(avatar_contradiction_clause, [], [f711])).
fof(f711, plain, ($false | (~ spl49_9 | ~ spl49_31)), inference(subsumption_resolution, [], [f381, f267])).
fof(f381, plain, (p(sK33) | ~ spl49_31), inference(avatar_component_clause, [], [f379])).
fof(f379, plain, (spl49_31 <=> p(sK33)), introduced(avatar_definition, [new_symbols(naming, [spl49_31])])).
fof(f710, plain, (~ spl49_9 | ~ spl49_23), inference(avatar_contradiction_clause, [], [f709])).
fof(f709, plain, ($false | (~ spl49_9 | ~ spl49_23)), inference(subsumption_resolution, [], [f332, f267])).
fof(f332, plain, (p(sK26) | ~ spl49_23), inference(avatar_component_clause, [], [f330])).
fof(f330, plain, (spl49_23 <=> p(sK26)), introduced(avatar_definition, [new_symbols(naming, [spl49_23])])).
fof(f708, plain, (~ spl49_9 | ~ spl49_22), inference(avatar_contradiction_clause, [], [f707])).
fof(f707, plain, ($false | (~ spl49_9 | ~ spl49_22)), inference(subsumption_resolution, [], [f327, f267])).
fof(f327, plain, (p(sK27) | ~ spl49_22), inference(avatar_component_clause, [], [f325])).
fof(f325, plain, (spl49_22 <=> p(sK27)), introduced(avatar_definition, [new_symbols(naming, [spl49_22])])).
fof(f706, plain, (~ spl49_9 | ~ spl49_40), inference(avatar_contradiction_clause, [], [f705])).
fof(f705, plain, ($false | (~ spl49_9 | ~ spl49_40)), inference(subsumption_resolution, [], [f432, f267])).
fof(f432, plain, (p(sK40) | ~ spl49_40), inference(avatar_component_clause, [], [f430])).
fof(f430, plain, (spl49_40 <=> p(sK40)), introduced(avatar_definition, [new_symbols(naming, [spl49_40])])).
fof(f704, plain, (~ spl49_11 | ~ spl49_41), inference(avatar_contradiction_clause, [], [f703])).
fof(f703, plain, ($false | (~ spl49_11 | ~ spl49_41)), inference(subsumption_resolution, [], [f436, f275])).
fof(f275, plain, (! [X0] : ~ q(X0) | ~ spl49_11), inference(avatar_component_clause, [], [f274])).
fof(f274, plain, (spl49_11 <=> ! [X0] : ~ q(X0)), introduced(avatar_definition, [new_symbols(naming, [spl49_11])])).
fof(f436, plain, (q(sK40) | ~ spl49_41), inference(avatar_component_clause, [], [f434])).
fof(f434, plain, (spl49_41 <=> q(sK40)), introduced(avatar_definition, [new_symbols(naming, [spl49_41])])).
fof(f702, plain, (~ spl49_11 | ~ spl49_43), inference(avatar_contradiction_clause, [], [f701])).
fof(f701, plain, ($false | (~ spl49_11 | ~ spl49_43)), inference(subsumption_resolution, [], [f444, f275])).
fof(f444, plain, (q(sK38) | ~ spl49_43), inference(avatar_component_clause, [], [f442])).
fof(f442, plain, (spl49_43 <=> q(sK38)), introduced(avatar_definition, [new_symbols(naming, [spl49_43])])).
fof(f700, plain, (~ spl49_9 | ~ spl49_42), inference(avatar_contradiction_clause, [], [f699])).
fof(f699, plain, ($false | (~ spl49_9 | ~ spl49_42)), inference(resolution, [], [f440, f267])).
fof(f440, plain, (p(sK39) | ~ spl49_42), inference(avatar_component_clause, [], [f438])).
fof(f438, plain, (spl49_42 <=> p(sK39)), introduced(avatar_definition, [new_symbols(naming, [spl49_42])])).
fof(f682, plain, (~ spl49_3 | spl49_38), inference(avatar_contradiction_clause, [], [f681])).
fof(f681, plain, ($false | (~ spl49_3 | spl49_38)), inference(subsumption_resolution, [], [f417, f239])).
fof(f239, plain, (! [X1] : p(X1) | ~ spl49_3), inference(avatar_component_clause, [], [f238])).
fof(f238, plain, (spl49_3 <=> ! [X1] : p(X1)), introduced(avatar_definition, [new_symbols(naming, [spl49_3])])).
fof(f417, plain, (~ p(sK36) | spl49_38), inference(avatar_component_clause, [], [f415])).
fof(f415, plain, (spl49_38 <=> p(sK36)), introduced(avatar_definition, [new_symbols(naming, [spl49_38])])).
fof(f680, plain, (spl49_68 | ~ spl49_52 | ~ spl49_71), inference(avatar_split_clause, [], [f677, f572, f484, f559])).
fof(f559, plain, (spl49_68 <=> g(sK48)), introduced(avatar_definition, [new_symbols(naming, [spl49_68])])).
fof(f484, plain, (spl49_52 <=> ! [X3] : (g(X3) | ~ f(X3))), introduced(avatar_definition, [new_symbols(naming, [spl49_52])])).
fof(f572, plain, (spl49_71 <=> f(sK48)), introduced(avatar_definition, [new_symbols(naming, [spl49_71])])).
fof(f677, plain, (g(sK48) | (~ spl49_52 | ~ spl49_71)), inference(resolution, [], [f574, f485])).
fof(f485, plain, (! [X3] : (~ f(X3) | g(X3)) | ~ spl49_52), inference(avatar_component_clause, [], [f484])).
fof(f574, plain, (f(sK48) | ~ spl49_71), inference(avatar_component_clause, [], [f572])).
fof(f676, plain, (spl49_53 | ~ spl49_50 | ~ spl49_52), inference(avatar_split_clause, [], [f675, f484, f476, f487])).
fof(f487, plain, (spl49_53 <=> ! [X2] : (h(X2) | ~ f(X2))), introduced(avatar_definition, [new_symbols(naming, [spl49_53])])).
fof(f476, plain, (spl49_50 <=> ! [X0] : (h(X0) | ~ f(X0) | ~ g(X0))), introduced(avatar_definition, [new_symbols(naming, [spl49_50])])).
fof(f675, plain, (! [X0] : (h(X0) | ~ f(X0)) | (~ spl49_50 | ~ spl49_52)), inference(subsumption_resolution, [], [f477, f485])).
fof(f477, plain, (! [X0] : (h(X0) | ~ f(X0) | ~ g(X0)) | ~ spl49_50), inference(avatar_component_clause, [], [f476])).
fof(f667, plain, (~ spl49_53 | ~ spl49_67 | ~ spl49_70), inference(avatar_contradiction_clause, [], [f666])).
fof(f666, plain, ($false | (~ spl49_53 | ~ spl49_67 | ~ spl49_70)), inference(subsumption_resolution, [], [f665, f557])).
fof(f557, plain, (! [X0] : ~ h(X0) | ~ spl49_67), inference(avatar_component_clause, [], [f556])).
fof(f556, plain, (spl49_67 <=> ! [X0] : ~ h(X0)), introduced(avatar_definition, [new_symbols(naming, [spl49_67])])).
fof(f665, plain, (! [X2] : h(X2) | (~ spl49_53 | ~ spl49_70)), inference(subsumption_resolution, [], [f488, f569])).
fof(f569, plain, (! [X0] : f(X0) | ~ spl49_70), inference(avatar_component_clause, [], [f568])).
fof(f568, plain, (spl49_70 <=> ! [X0] : f(X0)), introduced(avatar_definition, [new_symbols(naming, [spl49_70])])).
fof(f488, plain, (! [X2] : (~ f(X2) | h(X2)) | ~ spl49_53), inference(avatar_component_clause, [], [f487])).
fof(f660, plain, (spl49_56 | ~ spl49_57 | ~ spl49_58 | ~ spl49_59), inference(avatar_contradiction_clause, [], [f659])).
fof(f659, plain, ($false | (spl49_56 | ~ spl49_57 | ~ spl49_58 | ~ spl49_59)), inference(subsumption_resolution, [], [f657, f502])).
fof(f502, plain, (~ r(sK44, sK44) | spl49_56), inference(avatar_component_clause, [], [f500])).
fof(f500, plain, (spl49_56 <=> r(sK44, sK44)), introduced(avatar_definition, [new_symbols(naming, [spl49_56])])).
fof(f657, plain, (r(sK44, sK44) | (~ spl49_57 | ~ spl49_58 | ~ spl49_59)), inference(resolution, [], [f653, f507])).
fof(f507, plain, (r(sK44, sK45) | ~ spl49_57), inference(avatar_component_clause, [], [f505])).
fof(f505, plain, (spl49_57 <=> r(sK44, sK45)), introduced(avatar_definition, [new_symbols(naming, [spl49_57])])).
fof(f653, plain, (! [X1] : (~ r(X1, sK45) | r(X1, sK44)) | (~ spl49_57 | ~ spl49_58 | ~ spl49_59)), inference(resolution, [], [f511, f650])).
fof(f650, plain, (r(sK45, sK44) | (~ spl49_57 | ~ spl49_59)), inference(resolution, [], [f515, f507])).
fof(f515, plain, (! [X6, X5] : (~ r(X5, X6) | r(X6, X5)) | ~ spl49_59), inference(avatar_component_clause, [], [f514])).
fof(f514, plain, (spl49_59 <=> ! [X5, X6] : (r(X6, X5) | ~ r(X5, X6))), introduced(avatar_definition, [new_symbols(naming, [spl49_59])])).
fof(f511, plain, (! [X4, X2, X3] : (~ r(X3, X4) | ~ r(X2, X3) | r(X2, X4)) | ~ spl49_58), inference(avatar_component_clause, [], [f510])).
fof(f510, plain, (spl49_58 <=> ! [X3, X2, X4] : (r(X2, X4) | ~ r(X2, X3) | ~ r(X3, X4))), introduced(avatar_definition, [new_symbols(naming, [spl49_58])])).
fof(f645, plain, (~ spl49_3 | spl49_25), inference(avatar_contradiction_clause, [], [f644])).
fof(f644, plain, ($false | (~ spl49_3 | spl49_25)), inference(subsumption_resolution, [], [f343, f239])).
fof(f343, plain, (~ p(sK29) | spl49_25), inference(avatar_component_clause, [], [f341])).
fof(f341, plain, (spl49_25 <=> p(sK29)), introduced(avatar_definition, [new_symbols(naming, [spl49_25])])).
fof(f643, plain, (~ spl49_3 | spl49_26), inference(avatar_contradiction_clause, [], [f642])).
fof(f642, plain, ($false | (~ spl49_3 | spl49_26)), inference(subsumption_resolution, [], [f347, f239])).
fof(f347, plain, (~ p(sK28) | spl49_26), inference(avatar_component_clause, [], [f345])).
fof(f345, plain, (spl49_26 <=> p(sK28)), introduced(avatar_definition, [new_symbols(naming, [spl49_26])])).
fof(f641, plain, (~ spl49_11 | ~ spl49_12), inference(avatar_contradiction_clause, [], [f640])).
fof(f640, plain, ($false | (~ spl49_11 | ~ spl49_12)), inference(resolution, [], [f275, f280])).
fof(f280, plain, (q(sK23) | ~ spl49_12), inference(avatar_component_clause, [], [f278])).
fof(f278, plain, (spl49_12 <=> q(sK23)), introduced(avatar_definition, [new_symbols(naming, [spl49_12])])).
fof(f637, plain, (~ spl49_53 | spl49_62 | ~ spl49_65), inference(avatar_contradiction_clause, [], [f636])).
fof(f636, plain, ($false | (~ spl49_53 | spl49_62 | ~ spl49_65)), inference(subsumption_resolution, [], [f635, f532])).
fof(f532, plain, (~ h(sK47) | spl49_62), inference(avatar_component_clause, [], [f530])).
fof(f530, plain, (spl49_62 <=> h(sK47)), introduced(avatar_definition, [new_symbols(naming, [spl49_62])])).
fof(f635, plain, (h(sK47) | (~ spl49_53 | ~ spl49_65)), inference(resolution, [], [f488, f546])).
fof(f546, plain, (f(sK47) | ~ spl49_65), inference(avatar_component_clause, [], [f544])).
fof(f544, plain, (spl49_65 <=> f(sK47)), introduced(avatar_definition, [new_symbols(naming, [spl49_65])])).
fof(f632, plain, (spl49_63 | ~ spl49_52 | ~ spl49_66), inference(avatar_split_clause, [], [f629, f549, f484, f534])).
fof(f534, plain, (spl49_63 <=> g(sK46)), introduced(avatar_definition, [new_symbols(naming, [spl49_63])])).
fof(f549, plain, (spl49_66 <=> f(sK46)), introduced(avatar_definition, [new_symbols(naming, [spl49_66])])).
fof(f629, plain, (g(sK46) | (~ spl49_52 | ~ spl49_66)), inference(resolution, [], [f485, f551])).
fof(f551, plain, (f(sK46) | ~ spl49_66), inference(avatar_component_clause, [], [f549])).
fof(f628, plain, (spl49_52 | ~ spl49_51 | ~ spl49_53), inference(avatar_split_clause, [], [f627, f487, f480, f484])).
fof(f480, plain, (spl49_51 <=> ! [X1] : (g(X1) | ~ f(X1) | ~ h(X1))), introduced(avatar_definition, [new_symbols(naming, [spl49_51])])).
fof(f627, plain, (! [X1] : (g(X1) | ~ f(X1)) | (~ spl49_51 | ~ spl49_53)), inference(subsumption_resolution, [], [f481, f488])).
fof(f481, plain, (! [X1] : (g(X1) | ~ f(X1) | ~ h(X1)) | ~ spl49_51), inference(avatar_component_clause, [], [f480])).
fof(f626, plain, (~ spl49_3 | spl49_29), inference(avatar_contradiction_clause, [], [f625])).
fof(f625, plain, ($false | (~ spl49_3 | spl49_29)), inference(subsumption_resolution, [], [f368, f239])).
fof(f368, plain, (~ p(sK30) | spl49_29), inference(avatar_component_clause, [], [f366])).
fof(f366, plain, (spl49_29 <=> p(sK30)), introduced(avatar_definition, [new_symbols(naming, [spl49_29])])).
fof(f624, plain, (~ spl49_17 | spl49_46), inference(avatar_contradiction_clause, [], [f623])).
fof(f623, plain, ($false | (~ spl49_17 | spl49_46)), inference(subsumption_resolution, [], [f457, f303])).
fof(f303, plain, (! [X1] : q(X1) | ~ spl49_17), inference(avatar_component_clause, [], [f302])).
fof(f302, plain, (spl49_17 <=> ! [X1] : q(X1)), introduced(avatar_definition, [new_symbols(naming, [spl49_17])])).
fof(f457, plain, (~ q(sK43) | spl49_46), inference(avatar_component_clause, [], [f455])).
fof(f455, plain, (spl49_46 <=> q(sK43)), introduced(avatar_definition, [new_symbols(naming, [spl49_46])])).
fof(f622, plain, (~ spl49_3 | spl49_47), inference(avatar_contradiction_clause, [], [f621])).
fof(f621, plain, ($false | (~ spl49_3 | spl49_47)), inference(subsumption_resolution, [], [f461, f239])).
fof(f461, plain, (~ p(sK42) | spl49_47), inference(avatar_component_clause, [], [f459])).
fof(f459, plain, (spl49_47 <=> p(sK42)), introduced(avatar_definition, [new_symbols(naming, [spl49_47])])).
fof(f620, plain, (~ spl49_3 | spl49_45), inference(avatar_contradiction_clause, [], [f619])).
fof(f619, plain, ($false | (~ spl49_3 | spl49_45)), inference(subsumption_resolution, [], [f453, f239])).
fof(f453, plain, (~ p(sK43) | spl49_45), inference(avatar_component_clause, [], [f451])).
fof(f451, plain, (spl49_45 <=> p(sK43)), introduced(avatar_definition, [new_symbols(naming, [spl49_45])])).
fof(f618, plain, (~ spl49_17 | spl49_48), inference(avatar_contradiction_clause, [], [f617])).
fof(f617, plain, ($false | (~ spl49_17 | spl49_48)), inference(resolution, [], [f465, f303])).
fof(f465, plain, (~ q(sK41) | spl49_48), inference(avatar_component_clause, [], [f463])).
fof(f463, plain, (spl49_48 <=> q(sK41)), introduced(avatar_definition, [new_symbols(naming, [spl49_48])])).
fof(f616, plain, (~ spl49_3 | spl49_37), inference(avatar_contradiction_clause, [], [f615])).
fof(f615, plain, ($false | (~ spl49_3 | spl49_37)), inference(resolution, [], [f412, f239])).
fof(f412, plain, (~ p(sK37) | spl49_37), inference(avatar_component_clause, [], [f410])).
fof(f410, plain, (spl49_37 <=> p(sK37)), introduced(avatar_definition, [new_symbols(naming, [spl49_37])])).
fof(f614, plain, (~ spl49_3 | spl49_28), inference(avatar_contradiction_clause, [], [f613])).
fof(f613, plain, ($false | (~ spl49_3 | spl49_28)), inference(resolution, [], [f364, f239])).
fof(f364, plain, (~ p(sK31) | spl49_28), inference(avatar_component_clause, [], [f362])).
fof(f362, plain, (spl49_28 <=> p(sK31)), introduced(avatar_definition, [new_symbols(naming, [spl49_28])])).
fof(f612, plain, (~ spl49_3 | spl49_5), inference(avatar_contradiction_clause, [], [f611])).
fof(f611, plain, ($false | (~ spl49_3 | spl49_5)), inference(resolution, [], [f248, f239])).
fof(f248, plain, (~ p(sK20) | spl49_5), inference(avatar_component_clause, [], [f246])).
fof(f246, plain, (spl49_5 <=> p(sK20)), introduced(avatar_definition, [new_symbols(naming, [spl49_5])])).
fof(f606, plain, (~ spl49_3 | ~ spl49_11 | ~ spl49_19), inference(avatar_contradiction_clause, [], [f605])).
fof(f605, plain, ($false | (~ spl49_3 | ~ spl49_11 | ~ spl49_19)), inference(subsumption_resolution, [], [f604, f239])).
fof(f604, plain, (! [X0] : ~ p(sK25(X0)) | (~ spl49_11 | ~ spl49_19)), inference(subsumption_resolution, [], [f313, f275])).
fof(f313, plain, (! [X0] : (q(X0) | ~ p(sK25(X0))) | ~ spl49_19), inference(avatar_component_clause, [], [f312])).
fof(f312, plain, (spl49_19 <=> ! [X0] : (q(X0) | ~ p(sK25(X0)))), introduced(avatar_definition, [new_symbols(naming, [spl49_19])])).
fof(f597, plain, (~ spl49_3 | spl49_16), inference(avatar_contradiction_clause, [], [f596])).
fof(f596, plain, ($false | (~ spl49_3 | spl49_16)), inference(subsumption_resolution, [], [f299, f239])).
fof(f299, plain, (~ p(sK24) | spl49_16), inference(avatar_component_clause, [], [f297])).
fof(f297, plain, (spl49_16 <=> p(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl49_16])])).
fof(f591, plain, (~ spl49_3 | ~ spl49_9), inference(avatar_contradiction_clause, [], [f590])).
fof(f590, plain, ($false | (~ spl49_3 | ~ spl49_9)), inference(subsumption_resolution, [], [f239, f267])).
fof(f589, plain, (spl49_15 | ~ spl49_17), inference(avatar_contradiction_clause, [], [f588])).
fof(f588, plain, ($false | (spl49_15 | ~ spl49_17)), inference(subsumption_resolution, [], [f294, f303])).
fof(f294, plain, (~ q(sK24) | spl49_15), inference(avatar_component_clause, [], [f292])).
fof(f292, plain, (spl49_15 <=> q(sK24)), introduced(avatar_definition, [new_symbols(naming, [spl49_15])])).
fof(f587, plain, (~ spl49_9 | ~ spl49_13), inference(avatar_contradiction_clause, [], [f586])).
fof(f586, plain, ($false | (~ spl49_9 | ~ spl49_13)), inference(subsumption_resolution, [], [f285, f267])).
fof(f285, plain, (p(sK23) | ~ spl49_13), inference(avatar_component_clause, [], [f283])).
fof(f283, plain, (spl49_13 <=> p(sK23)), introduced(avatar_definition, [new_symbols(naming, [spl49_13])])).
fof(f585, plain, (~ spl49_8 | ~ spl49_9), inference(avatar_contradiction_clause, [], [f584])).
fof(f584, plain, ($false | (~ spl49_8 | ~ spl49_9)), inference(subsumption_resolution, [], [f263, f267])).
fof(f263, plain, (p(sK22) | ~ spl49_8), inference(avatar_component_clause, [], [f261])).
fof(f261, plain, (spl49_8 <=> p(sK22)), introduced(avatar_definition, [new_symbols(naming, [spl49_8])])).
fof(f583, plain, (spl49_2 | ~ spl49_3), inference(avatar_contradiction_clause, [], [f582])).
fof(f582, plain, ($false | (spl49_2 | ~ spl49_3)), inference(subsumption_resolution, [], [f235, f239])).
fof(f235, plain, (~ p(sK19) | spl49_2), inference(avatar_component_clause, [], [f233])).
fof(f233, plain, (spl49_2 <=> p(sK19)), introduced(avatar_definition, [new_symbols(naming, [spl49_2])])).
fof(f581, plain, (spl49_60 | spl49_55 | spl49_49 | spl49_18 | spl49_44 | spl49_14 | spl49_39 | spl49_4 | spl49_10 | spl49_1 | spl49_3 | spl49_6 | spl49_36 | spl49_33 | spl49_21 | spl49_30 | spl49_27 | spl49_24 | spl49_20), inference(avatar_split_clause, [], [f208, f316, f337, f354, f371, f320, f388, f405, f252, f238, f229, f270, f242, f422, f288, f447, f306, f472, f496, f518])).
fof(f518, plain, (spl49_60 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl49_60])])).
fof(f496, plain, (spl49_55 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl49_55])])).
fof(f472, plain, (spl49_49 <=> sP4), introduced(avatar_definition, [new_symbols(naming, [spl49_49])])).
fof(f306, plain, (spl49_18 <=> sP13), introduced(avatar_definition, [new_symbols(naming, [spl49_18])])).
fof(f447, plain, (spl49_44 <=> sP5), introduced(avatar_definition, [new_symbols(naming, [spl49_44])])).
fof(f288, plain, (spl49_14 <=> sP14), introduced(avatar_definition, [new_symbols(naming, [spl49_14])])).
fof(f422, plain, (spl49_39 <=> sP6), introduced(avatar_definition, [new_symbols(naming, [spl49_39])])).
fof(f242, plain, (spl49_4 <=> sP17), introduced(avatar_definition, [new_symbols(naming, [spl49_4])])).
fof(f270, plain, (spl49_10 <=> sP15), introduced(avatar_definition, [new_symbols(naming, [spl49_10])])).
fof(f229, plain, (spl49_1 <=> sP18), introduced(avatar_definition, [new_symbols(naming, [spl49_1])])).
fof(f252, plain, (spl49_6 <=> sP16), introduced(avatar_definition, [new_symbols(naming, [spl49_6])])).
fof(f405, plain, (spl49_36 <=> sP7), introduced(avatar_definition, [new_symbols(naming, [spl49_36])])).
fof(f388, plain, (spl49_33 <=> sP8), introduced(avatar_definition, [new_symbols(naming, [spl49_33])])).
fof(f320, plain, (spl49_21 <=> c), introduced(avatar_definition, [new_symbols(naming, [spl49_21])])).
fof(f371, plain, (spl49_30 <=> sP9), introduced(avatar_definition, [new_symbols(naming, [spl49_30])])).
fof(f354, plain, (spl49_27 <=> sP10), introduced(avatar_definition, [new_symbols(naming, [spl49_27])])).
fof(f337, plain, (spl49_24 <=> sP11), introduced(avatar_definition, [new_symbols(naming, [spl49_24])])).
fof(f316, plain, (spl49_20 <=> sP12), introduced(avatar_definition, [new_symbols(naming, [spl49_20])])).
fof(f208, plain, ! [X1] : (sP12 | sP11 | sP10 | sP9 | c | sP8 | sP7 | sP16 | p(X1) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(duplicate_literal_removal, [], [f200])).
fof(f200, plain, ! [X1] : (sP12 | sP11 | sP10 | sP9 | c | c | c | c | sP8 | sP7 | sP16 | p(X1) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(cnf_transformation, [], [f119])).
fof(f119, plain, (sP12 | sP11 | sP10 | sP9 | ((~ c | ~ c) & (c | c)) | ((~ c | ~ c) & (c | c)) | sP8 | sP7 | sP16 | (! [X0] : ~ p(X0) & ! [X1] : p(X1)) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(rectify, [], [f118])).
fof(f118, plain, (sP12 | sP11 | sP10 | sP9 | ((~ c | ~ c) & (c | c)) | ((~ c | ~ c) & (c | c)) | sP8 | sP7 | sP16 | (! [X18] : ~ p(X18) & ! [X17] : p(X17)) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(nnf_transformation, [], [f25])).
fof(f25, plain, (sP12 | sP11 | sP10 | sP9 | ~ (c <=> c) | ~ (c <=> c) | sP8 | sP7 | sP16 | (! [X18] : ~ p(X18) & ! [X17] : p(X17)) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(definition_folding, [], [f5, e24, e23, e22, e21, e20, e19, e18, e17, e16, e15, e14, e13, e12, e11, e10, e9, e8, e7, e6])).
fof(f6, plain, (! [X52] : (? [X53] : (~ g(X53) & f(X53)) | (~ h(X52) & g(X52) & f(X52))) | ~ sP0), inference(usedef, [], [e6])).
fof(e6, plain, (sP0 <=> ! [X52] : (? [X53] : (~ g(X53) & f(X53)) | (~ h(X52) & g(X52) & f(X52)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f7, plain, (? [X40] : (~ g(X40) & f(X40)) | ? [X39] : (~ h(X39) & g(X39) & f(X39)) | ~ sP1), inference(usedef, [], [e7])).
fof(e7, plain, (sP1 <=> (? [X40] : (~ g(X40) & f(X40)) | ? [X39] : (~ h(X39) & g(X39) & f(X39)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f8, plain, ((! [X55] : (h(X55) | ~ g(X55) | ~ f(X55)) & ! [X54] : (g(X54) | ~ h(X54) | ~ f(X54)) & (! [X50] : (h(X50) | ~ f(X50)) | ! [X51] : (g(X51) | ~ f(X51))) & sP0) | ~ sP2), inference(usedef, [], [e8])).
fof(e8, plain, (sP2 <=> (! [X55] : (h(X55) | ~ g(X55) | ~ f(X55)) & ! [X54] : (g(X54) | ~ h(X54) | ~ f(X54)) & (! [X50] : (h(X50) | ~ f(X50)) | ! [X51] : (g(X51) | ~ f(X51))) & sP0)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f9, plain, ((? [X48, X49] : (~ r(X48, X48) & r(X48, X49)) & ! [X43, X44, X45] : (r(X43, X45) | ~ r(X44, X45) | ~ r(X43, X44)) & ! [X46, X47] : (r(X47, X46) | ~ r(X46, X47))) | ~ sP3), inference(usedef, [], [e9])).
fof(e9, plain, (sP3 <=> (? [X48, X49] : (~ r(X48, X48) & r(X48, X49)) & ! [X43, X44, X45] : (r(X43, X45) | ~ r(X44, X45) | ~ r(X43, X44)) & ! [X46, X47] : (r(X47, X46) | ~ r(X46, X47)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f10, plain, ((! [X42] : (h(X42) | ~ g(X42) | ~ f(X42)) & ! [X41] : (g(X41) | ~ h(X41) | ~ f(X41)) & (! [X37] : (h(X37) | ~ f(X37)) | ! [X38] : (g(X38) | ~ f(X38))) & sP1) | ~ sP4), inference(usedef, [], [e10])).
fof(e10, plain, (sP4 <=> (! [X42] : (h(X42) | ~ g(X42) | ~ f(X42)) & ! [X41] : (g(X41) | ~ h(X41) | ~ f(X41)) & (! [X37] : (h(X37) | ~ f(X37)) | ! [X38] : (g(X38) | ~ f(X38))) & sP1)), introduced(predicate_definition_introduction, [new_symbols(naming, [sP4])])).
fof(f11, plain, (~ (! [X32] : (q(X32) & p(X32)) <=> (! [X33] : q(X33) & ! [X34] : p(X34))) | ~ sP5), inference(usedef, [], [e11])).
fof(e11, plain, (sP5 <=> ~ (! [X32] : (q(X32) & p(X32)) <=> (! [X33] : q(X33) & ! [X34] : p(X34)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP5])])).
fof(f12, plain, (~ (? [X26] : (q(X26) | p(X26)) <=> (? [X27] : q(X27) | ? [X28] : p(X28))) | ~ sP6), inference(usedef, [], [e12])).
fof(e12, plain, (sP6 <=> ~ (? [X26] : (q(X26) | p(X26)) <=> (? [X27] : q(X27) | ? [X28] : p(X28)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP6])])).
fof(f13, plain, (~ (! [X12] : (c | p(X12)) <=> (c | ! [X13] : p(X13))) | ~ sP7), inference(usedef, [], [e13])).
fof(e13, plain, (sP7 <=> ~ (! [X12] : (c | p(X12)) <=> (c | ! [X13] : p(X13)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP7])])).
fof(f14, plain, (~ (? [X10] : (c & p(X10)) <=> (c & ? [X11] : p(X11))) | ~ sP8), inference(usedef, [], [e14])).
fof(e14, plain, (sP8 <=> ~ (? [X10] : (c & p(X10)) <=> (c & ? [X11] : p(X11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP8])])).
fof(f15, plain, (~ (? [X6] : (p(X6) | ~ c) <=> (? [X7] : p(X7) | ~ c)) | ~ sP9), inference(usedef, [], [e15])).
fof(e15, plain, (sP9 <=> ~ (? [X6] : (p(X6) | ~ c) <=> (? [X7] : p(X7) | ~ c))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP9])])).
fof(f16, plain, (~ (? [X4] : (c | ~ p(X4)) <=> (c | ? [X5] : ~ p(X5))) | ~ sP10), inference(usedef, [], [e16])).
fof(e16, plain, (sP10 <=> ~ (? [X4] : (c | ~ p(X4)) <=> (c | ? [X5] : ~ p(X5)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP10])])).
fof(f17, plain, (~ (! [X2] : (p(X2) | ~ c) <=> (! [X3] : p(X3) | ~ c)) | ~ sP11), inference(usedef, [], [e17])).
fof(e17, plain, (sP11 <=> ~ (! [X2] : (p(X2) | ~ c) <=> (! [X3] : p(X3) | ~ c))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP11])])).
fof(f18, plain, (~ (! [X0] : (c | ~ p(X0)) <=> (c | ! [X1] : ~ p(X1))) | ~ sP12), inference(usedef, [], [e18])).
fof(e18, plain, (sP12 <=> ~ (! [X0] : (c | ~ p(X0)) <=> (c | ! [X1] : ~ p(X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP12])])).
fof(f19, plain, (! [X35] : ? [X36] : (~ q(X35) & p(X35) & (q(X35) | ~ p(X36))) | ~ sP13), inference(usedef, [], [e19])).
fof(e19, plain, (sP13 <=> ! [X35] : ? [X36] : (~ q(X35) & p(X35) & (q(X35) | ~ p(X36)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP13])])).
fof(f20, plain, ((? [X31] : (~ q(X31) & ~ p(X31)) & (! [X29] : q(X29) | ! [X30] : p(X30))) | ~ sP14), inference(usedef, [], [e20])).
fof(e20, plain, (sP14 <=> (? [X31] : (~ q(X31) & ~ p(X31)) & (! [X29] : q(X29) | ! [X30] : p(X30)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP14])])).
fof(f21, plain, (((! [X22] : ~ q(X22) | ! [X23] : ~ p(X23)) & ? [X21] : (q(X21) & p(X21))) | ~ sP15), inference(usedef, [], [e21])).
fof(e21, plain, (sP15 <=> ((! [X22] : ~ q(X22) | ! [X23] : ~ p(X23)) & ? [X21] : (q(X21) & p(X21)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP15])])).
fof(f22, plain, ((? [X15] : (~ p(X15) & ? [X16] : p(X16)) & ! [X14] : ~ p(X14)) | ~ sP16), inference(usedef, [], [e22])).
fof(e22, plain, (sP16 <=> (? [X15] : (~ p(X15) & ? [X16] : p(X16)) & ! [X14] : ~ p(X14))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP16])])).
fof(f23, plain, (! [X24] : (? [X25] : ~ p(X25) & p(X24)) | ~ sP17), inference(usedef, [], [e23])).
fof(e23, plain, (sP17 <=> ! [X24] : (? [X25] : ~ p(X25) & p(X24))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP17])])).
fof(f24, plain, (? [X19] : (~ p(X19) & ! [X20] : p(X20)) | ~ sP18), inference(usedef, [], [e24])).
fof(e24, plain, (sP18 <=> ? [X19] : (~ p(X19) & ! [X20] : p(X20))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP18])])).
fof(f5, plain, (~ (! [X0] : (c | ~ p(X0)) <=> (c | ! [X1] : ~ p(X1))) | ~ (! [X2] : (p(X2) | ~ c) <=> (! [X3] : p(X3) | ~ c)) | ~ (? [X4] : (c | ~ p(X4)) <=> (c | ? [X5] : ~ p(X5))) | ~ (? [X6] : (p(X6) | ~ c) <=> (? [X7] : p(X7) | ~ c)) | ~ (c <=> c) | ~ (c <=> c) | ~ (? [X10] : (c & p(X10)) <=> (c & ? [X11] : p(X11))) | ~ (! [X12] : (c | p(X12)) <=> (c | ! [X13] : p(X13))) | (? [X15] : (~ p(X15) & ? [X16] : p(X16)) & ! [X14] : ~ p(X14)) | (! [X18] : ~ p(X18) & ! [X17] : p(X17)) | ? [X19] : (~ p(X19) & ! [X20] : p(X20)) | ((! [X22] : ~ q(X22) | ! [X23] : ~ p(X23)) & ? [X21] : (q(X21) & p(X21))) | ! [X24] : (? [X25] : ~ p(X25) & p(X24)) | ~ (? [X26] : (q(X26) | p(X26)) <=> (? [X27] : q(X27) | ? [X28] : p(X28))) | (? [X31] : (~ q(X31) & ~ p(X31)) & (! [X29] : q(X29) | ! [X30] : p(X30))) | ~ (! [X32] : (q(X32) & p(X32)) <=> (! [X33] : q(X33) & ! [X34] : p(X34))) | ! [X35] : ? [X36] : (~ q(X35) & p(X35) & (q(X35) | ~ p(X36))) | (! [X42] : (h(X42) | ~ g(X42) | ~ f(X42)) & ! [X41] : (g(X41) | ~ h(X41) | ~ f(X41)) & (! [X37] : (h(X37) | ~ f(X37)) | ! [X38] : (g(X38) | ~ f(X38))) & (? [X40] : (~ g(X40) & f(X40)) | ? [X39] : (~ h(X39) & g(X39) & f(X39)))) | (? [X48, X49] : (~ r(X48, X48) & r(X48, X49)) & ! [X43, X44, X45] : (r(X43, X45) | ~ r(X44, X45) | ~ r(X43, X44)) & ! [X46, X47] : (r(X47, X46) | ~ r(X46, X47))) | (! [X55] : (h(X55) | ~ g(X55) | ~ f(X55)) & ! [X54] : (g(X54) | ~ h(X54) | ~ f(X54)) & (! [X50] : (h(X50) | ~ f(X50)) | ! [X51] : (g(X51) | ~ f(X51))) & ! [X52] : (? [X53] : (~ g(X53) & f(X53)) | (~ h(X52) & g(X52) & f(X52))))), inference(flattening, [], [f4])).
fof(f4, plain, (~ (! [X0] : (c | ~ p(X0)) <=> (c | ! [X1] : ~ p(X1))) | ~ (! [X2] : (p(X2) | ~ c) <=> (! [X3] : p(X3) | ~ c)) | ~ (? [X4] : (c | ~ p(X4)) <=> (c | ? [X5] : ~ p(X5))) | ~ (? [X6] : (p(X6) | ~ c) <=> (? [X7] : p(X7) | ~ c)) | ~ (c <=> c) | ~ (c <=> c) | ~ (? [X10] : (c & p(X10)) <=> (c & ? [X11] : p(X11))) | ~ (! [X12] : (c | p(X12)) <=> (c | ! [X13] : p(X13))) | (? [X15] : (~ p(X15) & ? [X16] : p(X16)) & ! [X14] : ~ p(X14)) | (! [X18] : ~ p(X18) & ! [X17] : p(X17)) | ? [X19] : (~ p(X19) & ! [X20] : p(X20)) | ((! [X22] : ~ q(X22) | ! [X23] : ~ p(X23)) & ? [X21] : (q(X21) & p(X21))) | ! [X24] : (? [X25] : ~ p(X25) & p(X24)) | ~ (? [X26] : (q(X26) | p(X26)) <=> (? [X27] : q(X27) | ? [X28] : p(X28))) | (? [X31] : (~ q(X31) & ~ p(X31)) & (! [X29] : q(X29) | ! [X30] : p(X30))) | ~ (! [X32] : (q(X32) & p(X32)) <=> (! [X33] : q(X33) & ! [X34] : p(X34))) | ! [X35] : ? [X36] : ((~ q(X35) & p(X35)) & (q(X35) | ~ p(X36))) | ((! [X42] : (h(X42) | ~ g(X42) | ~ f(X42)) & ! [X41] : (g(X41) | (~ h(X41) | ~ f(X41)))) & ((! [X37] : (h(X37) | ~ f(X37)) | ! [X38] : (g(X38) | ~ f(X38))) & (? [X40] : (~ g(X40) & f(X40)) | ? [X39] : (~ h(X39) & (g(X39) & f(X39)))))) | (? [X48, X49] : (~ r(X48, X48) & r(X48, X49)) & (! [X43, X44, X45] : (r(X43, X45) | (~ r(X44, X45) | ~ r(X43, X44))) & ! [X46, X47] : (r(X47, X46) | ~ r(X46, X47)))) | ((! [X55] : (h(X55) | ~ g(X55) | ~ f(X55)) & ! [X54] : (g(X54) | (~ h(X54) | ~ f(X54)))) & ((! [X50] : (h(X50) | ~ f(X50)) | ! [X51] : (g(X51) | ~ f(X51))) & ! [X52] : (? [X53] : (~ g(X53) & f(X53)) | (~ h(X52) & (g(X52) & f(X52))))))), inference(ennf_transformation, [], [f3])).
fof(f3, plain, ~ ((! [X0] : (p(X0) => c) <=> (? [X1] : p(X1) => c)) & (! [X2] : (c => p(X2)) <=> (c => ! [X3] : p(X3))) & (? [X4] : (p(X4) => c) <=> (! [X5] : p(X5) => c)) & (? [X6] : (c => p(X6)) <=> (c => ? [X7] : p(X7))) & (c <=> c) & (c <=> c) & (? [X10] : (c & p(X10)) <=> (c & ? [X11] : p(X11))) & (! [X12] : (c | p(X12)) <=> (c | ! [X13] : p(X13))) & (~ ? [X14] : p(X14) => ! [X15] : (? [X16] : p(X16) => p(X15))) & (! [X17] : p(X17) => ? [X18] : p(X18)) & ! [X19] : (! [X20] : p(X20) => p(X19)) & (? [X21] : (q(X21) & p(X21)) => (? [X22] : q(X22) & ? [X23] : p(X23))) & ? [X24] : (p(X24) => ! [X25] : p(X25)) & (? [X26] : (q(X26) | p(X26)) <=> (? [X27] : q(X27) | ? [X28] : p(X28))) & ((! [X29] : q(X29) | ! [X30] : p(X30)) => ! [X31] : (q(X31) | p(X31))) & (! [X32] : (q(X32) & p(X32)) <=> (! [X33] : q(X33) & ! [X34] : p(X34))) & ? [X35] : ! [X36] : ((p(X36) => q(X35)) => (p(X35) => q(X35))) & (((! [X37] : (f(X37) => h(X37)) | ! [X38] : (f(X38) => g(X38))) & (! [X39] : ((g(X39) & f(X39)) => h(X39)) => ? [X40] : (~ g(X40) & f(X40)))) => (! [X41] : ((h(X41) & f(X41)) => g(X41)) => ? [X42] : (~ h(X42) & g(X42) & f(X42)))) & ((! [X43, X44, X45] : ((r(X44, X45) & r(X43, X44)) => r(X43, X45)) & ! [X46, X47] : (r(X46, X47) => r(X47, X46))) => ! [X48, X49] : (r(X48, X49) => r(X48, X48))) & (((! [X50] : (f(X50) => h(X50)) | ! [X51] : (f(X51) => g(X51))) & ! [X52] : (((g(X52) & f(X52)) => h(X52)) => ? [X53] : (~ g(X53) & f(X53)))) => (! [X54] : ((h(X54) & f(X54)) => g(X54)) => ? [X55] : (~ h(X55) & g(X55) & f(X55))))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ((! [X0] : (p(X0) => c) <=> (? [X0] : p(X0) => c)) & (! [X0] : (c => p(X0)) <=> (c => ! [X0] : p(X0))) & (? [X0] : (p(X0) => c) <=> (! [X0] : p(X0) => c)) & (? [X0] : (c => p(X0)) <=> (c => ? [X0] : p(X0))) & (! [X0] : c <=> c) & (? [X0] : c <=> c) & (? [X0] : (c & p(X0)) <=> (c & ? [X0] : p(X0))) & (! [X0] : (c | p(X0)) <=> (c | ! [X0] : p(X0))) & (~ ? [X1] : p(X1) => ! [X1] : (? [X0] : p(X0) => p(X1))) & (! [X0] : p(X0) => ? [X0] : p(X0)) & ! [X1] : (! [X0] : p(X0) => p(X1)) & (? [X0] : (q(X0) & p(X0)) => (? [X0] : q(X0) & ? [X0] : p(X0))) & ? [X1] : (p(X1) => ! [X0] : p(X0)) & (? [X0] : (q(X0) | p(X0)) <=> (? [X0] : q(X0) | ? [X0] : p(X0))) & ((! [X0] : q(X0) | ! [X0] : p(X0)) => ! [X0] : (q(X0) | p(X0))) & (! [X0] : (q(X0) & p(X0)) <=> (! [X0] : q(X0) & ! [X0] : p(X0))) & ? [X0] : ! [X1] : ((p(X1) => q(X0)) => (p(X0) => q(X0))) & (((! [X3] : (f(X3) => h(X3)) | ! [X2] : (f(X2) => g(X2))) & (! [X0] : ((g(X0) & f(X0)) => h(X0)) => ? [X0] : (~ g(X0) & f(X0)))) => (! [X4] : ((h(X4) & f(X4)) => g(X4)) => ? [X5] : (~ h(X5) & g(X5) & f(X5)))) & ((! [X0, X1, X3] : ((r(X1, X3) & r(X0, X1)) => r(X0, X3)) & ! [X0, X1] : (r(X0, X1) => r(X1, X0))) => ! [X0, X1] : (r(X0, X1) => r(X0, X0))) & (((! [X3] : (f(X3) => h(X3)) | ! [X2] : (f(X2) => g(X2))) & ! [X0] : (((g(X0) & f(X0)) => h(X0)) => ? [X1] : (~ g(X1) & f(X1)))) => (! [X4] : ((h(X4) & f(X4)) => g(X4)) => ? [X5] : (~ h(X5) & g(X5) & f(X5))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ((! [X0] : (p(X0) => c) <=> (? [X0] : p(X0) => c)) & (! [X0] : (c => p(X0)) <=> (c => ! [X0] : p(X0))) & (? [X0] : (p(X0) => c) <=> (! [X0] : p(X0) => c)) & (? [X0] : (c => p(X0)) <=> (c => ? [X0] : p(X0))) & (! [X0] : c <=> c) & (? [X0] : c <=> c) & (? [X0] : (c & p(X0)) <=> (c & ? [X0] : p(X0))) & (! [X0] : (c | p(X0)) <=> (c | ! [X0] : p(X0))) & (~ ? [X1] : p(X1) => ! [X1] : (? [X0] : p(X0) => p(X1))) & (! [X0] : p(X0) => ? [X0] : p(X0)) & ! [X1] : (! [X0] : p(X0) => p(X1)) & (? [X0] : (q(X0) & p(X0)) => (? [X0] : q(X0) & ? [X0] : p(X0))) & ? [X1] : (p(X1) => ! [X0] : p(X0)) & (? [X0] : (q(X0) | p(X0)) <=> (? [X0] : q(X0) | ? [X0] : p(X0))) & ((! [X0] : q(X0) | ! [X0] : p(X0)) => ! [X0] : (q(X0) | p(X0))) & (! [X0] : (q(X0) & p(X0)) <=> (! [X0] : q(X0) & ! [X0] : p(X0))) & ? [X0] : ! [X1] : ((p(X1) => q(X0)) => (p(X0) => q(X0))) & (((! [X3] : (f(X3) => h(X3)) | ! [X2] : (f(X2) => g(X2))) & (! [X0] : ((g(X0) & f(X0)) => h(X0)) => ? [X0] : (~ g(X0) & f(X0)))) => (! [X4] : ((h(X4) & f(X4)) => g(X4)) => ? [X5] : (~ h(X5) & g(X5) & f(X5)))) & ((! [X0, X1, X3] : ((r(X1, X3) & r(X0, X1)) => r(X0, X3)) & ! [X0, X1] : (r(X0, X1) => r(X1, X0))) => ! [X0, X1] : (r(X0, X1) => r(X0, X0))) & (((! [X3] : (f(X3) => h(X3)) | ! [X2] : (f(X2) => g(X2))) & ! [X0] : (((g(X0) & f(X0)) => h(X0)) => ? [X1] : (~ g(X1) & f(X1)))) => (! [X4] : ((h(X4) & f(X4)) => g(X4)) => ? [X5] : (~ h(X5) & g(X5) & f(X5))))), file('/home/ubuntu/library/tptp/Problems/SYN/SYN917+1.p', prove_this)).
fof(f580, plain, (spl49_60 | spl49_55 | spl49_49 | spl49_18 | spl49_44 | spl49_14 | spl49_39 | spl49_4 | spl49_10 | spl49_1 | spl49_9 | spl49_6 | spl49_36 | spl49_33 | spl49_21 | spl49_30 | spl49_27 | spl49_24 | spl49_20), inference(avatar_split_clause, [], [f209, f316, f337, f354, f371, f320, f388, f405, f252, f266, f229, f270, f242, f422, f288, f447, f306, f472, f496, f518])).
fof(f209, plain, ! [X0] : (sP12 | sP11 | sP10 | sP9 | c | sP8 | sP7 | sP16 | ~ p(X0) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(duplicate_literal_removal, [], [f201])).
fof(f201, plain, ! [X0] : (sP12 | sP11 | sP10 | sP9 | c | c | c | c | sP8 | sP7 | sP16 | ~ p(X0) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(cnf_transformation, [], [f119])).
fof(f579, plain, (spl49_60 | spl49_55 | spl49_49 | spl49_18 | spl49_44 | spl49_14 | spl49_39 | spl49_4 | spl49_10 | spl49_1 | spl49_3 | spl49_6 | spl49_36 | spl49_33 | ~ spl49_21 | spl49_30 | spl49_27 | spl49_24 | spl49_20), inference(avatar_split_clause, [], [f214, f316, f337, f354, f371, f320, f388, f405, f252, f238, f229, f270, f242, f422, f288, f447, f306, f472, f496, f518])).
fof(f214, plain, ! [X1] : (sP12 | sP11 | sP10 | sP9 | ~ c | sP8 | sP7 | sP16 | p(X1) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(duplicate_literal_removal, [], [f206])).
fof(f206, plain, ! [X1] : (sP12 | sP11 | sP10 | sP9 | ~ c | ~ c | ~ c | ~ c | sP8 | sP7 | sP16 | p(X1) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(cnf_transformation, [], [f119])).
fof(f578, plain, (spl49_60 | spl49_55 | spl49_49 | spl49_18 | spl49_44 | spl49_14 | spl49_39 | spl49_4 | spl49_10 | spl49_1 | spl49_9 | spl49_6 | spl49_36 | spl49_33 | ~ spl49_21 | spl49_30 | spl49_27 | spl49_24 | spl49_20), inference(avatar_split_clause, [], [f215, f316, f337, f354, f371, f320, f388, f405, f252, f266, f229, f270, f242, f422, f288, f447, f306, f472, f496, f518])).
fof(f215, plain, ! [X0] : (sP12 | sP11 | sP10 | sP9 | ~ c | sP8 | sP7 | sP16 | ~ p(X0) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(duplicate_literal_removal, [], [f207])).
fof(f207, plain, ! [X0] : (sP12 | sP11 | sP10 | sP9 | ~ c | ~ c | ~ c | ~ c | sP8 | sP7 | sP16 | ~ p(X0) | sP18 | sP15 | sP17 | sP6 | sP14 | sP5 | sP13 | sP4 | sP3 | sP2), inference(cnf_transformation, [], [f119])).
fof(f577, plain, (~ spl49_61 | spl49_70 | spl49_71), inference(avatar_split_clause, [], [f194, f572, f568, f525])).
fof(f525, plain, (spl49_61 <=> sP0), introduced(avatar_definition, [new_symbols(naming, [spl49_61])])).
fof(f194, plain, ! [X0] : (f(sK48) | f(X0) | ~ sP0), inference(cnf_transformation, [], [f117])).
fof(f117, plain, (! [X0] : ((~ g(sK48) & f(sK48)) | (~ h(X0) & g(X0) & f(X0))) | ~ sP0), inference(skolemisation, [status(esa), new_symbols(skolem, [sK48])], [f115, f116])).
fof(f116, plain, (? [X1] : (~ g(X1) & f(X1)) => (~ g(sK48) & f(sK48))), introduced(choice_axiom, [])).
fof(f115, plain, (! [X0] : (? [X1] : (~ g(X1) & f(X1)) | (~ h(X0) & g(X0) & f(X0))) | ~ sP0), inference(rectify, [], [f114])).
fof(f114, plain, (! [X52] : (? [X53] : (~ g(X53) & f(X53)) | (~ h(X52) & g(X52) & f(X52))) | ~ sP0), inference(nnf_transformation, [], [f6])).
fof(f575, plain, (~ spl49_61 | spl49_67 | spl49_71), inference(avatar_split_clause, [], [f196, f572, f556, f525])).
fof(f196, plain, ! [X0] : (f(sK48) | ~ h(X0) | ~ sP0), inference(cnf_transformation, [], [f117])).
fof(f570, plain, (~ spl49_61 | spl49_70 | ~ spl49_68), inference(avatar_split_clause, [], [f197, f559, f568, f525])).
fof(f197, plain, ! [X0] : (~ g(sK48) | f(X0) | ~ sP0), inference(cnf_transformation, [], [f117])).
fof(f562, plain, (~ spl49_61 | spl49_67 | ~ spl49_68), inference(avatar_split_clause, [], [f199, f559, f556, f525])).
fof(f199, plain, ! [X0] : (~ g(sK48) | ~ h(X0) | ~ sP0), inference(cnf_transformation, [], [f117])).
fof(f554, plain, (~ spl49_54 | spl49_65 | spl49_66), inference(avatar_split_clause, [], [f188, f549, f544, f491])).
fof(f491, plain, (spl49_54 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl49_54])])).
fof(f188, plain, (f(sK46) | f(sK47) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f113, plain, ((~ g(sK46) & f(sK46)) | (~ h(sK47) & g(sK47) & f(sK47)) | ~ sP1), inference(skolemisation, [status(esa), new_symbols(skolem, [sK46, sK47])], [f110, f112, f111])).
fof(f111, plain, (? [X0] : (~ g(X0) & f(X0)) => (~ g(sK46) & f(sK46))), introduced(choice_axiom, [])).
fof(f112, plain, (? [X1] : (~ h(X1) & g(X1) & f(X1)) => (~ h(sK47) & g(sK47) & f(sK47))), introduced(choice_axiom, [])).
fof(f110, plain, (? [X0] : (~ g(X0) & f(X0)) | ? [X1] : (~ h(X1) & g(X1) & f(X1)) | ~ sP1), inference(rectify, [], [f109])).
fof(f109, plain, (? [X40] : (~ g(X40) & f(X40)) | ? [X39] : (~ h(X39) & g(X39) & f(X39)) | ~ sP1), inference(nnf_transformation, [], [f7])).
fof(f552, plain, (~ spl49_54 | ~ spl49_62 | spl49_66), inference(avatar_split_clause, [], [f190, f549, f530, f491])).
fof(f190, plain, (f(sK46) | ~ h(sK47) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f547, plain, (~ spl49_54 | spl49_65 | ~ spl49_63), inference(avatar_split_clause, [], [f191, f534, f544, f491])).
fof(f191, plain, (~ g(sK46) | f(sK47) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f537, plain, (~ spl49_54 | ~ spl49_62 | ~ spl49_63), inference(avatar_split_clause, [], [f193, f534, f530, f491])).
fof(f193, plain, (~ g(sK46) | ~ h(sK47) | ~ sP1), inference(cnf_transformation, [], [f113])).
fof(f528, plain, (~ spl49_60 | spl49_61), inference(avatar_split_clause, [], [f184, f525, f518])).
fof(f184, plain, (sP0 | ~ sP2), inference(cnf_transformation, [], [f108])).
fof(f108, plain, ((! [X0] : (h(X0) | ~ g(X0) | ~ f(X0)) & ! [X1] : (g(X1) | ~ h(X1) | ~ f(X1)) & (! [X2] : (h(X2) | ~ f(X2)) | ! [X3] : (g(X3) | ~ f(X3))) & sP0) | ~ sP2), inference(rectify, [], [f107])).
fof(f107, plain, ((! [X55] : (h(X55) | ~ g(X55) | ~ f(X55)) & ! [X54] : (g(X54) | ~ h(X54) | ~ f(X54)) & (! [X50] : (h(X50) | ~ f(X50)) | ! [X51] : (g(X51) | ~ f(X51))) & sP0) | ~ sP2), inference(nnf_transformation, [], [f8])).
fof(f523, plain, (~ spl49_60 | spl49_52 | spl49_53), inference(avatar_split_clause, [], [f185, f487, f484, f518])).
fof(f185, plain, ! [X2, X3] : (h(X2) | ~ f(X2) | g(X3) | ~ f(X3) | ~ sP2), inference(cnf_transformation, [], [f108])).
fof(f522, plain, (~ spl49_60 | spl49_51), inference(avatar_split_clause, [], [f186, f480, f518])).
fof(f186, plain, ! [X1] : (g(X1) | ~ h(X1) | ~ f(X1) | ~ sP2), inference(cnf_transformation, [], [f108])).
fof(f521, plain, (~ spl49_60 | spl49_50), inference(avatar_split_clause, [], [f187, f476, f518])).
fof(f187, plain, ! [X0] : (h(X0) | ~ g(X0) | ~ f(X0) | ~ sP2), inference(cnf_transformation, [], [f108])).
fof(f516, plain, (~ spl49_55 | spl49_59), inference(avatar_split_clause, [], [f180, f514, f496])).
fof(f180, plain, ! [X6, X5] : (r(X6, X5) | ~ r(X5, X6) | ~ sP3), inference(cnf_transformation, [], [f106])).
fof(f106, plain, (((~ r(sK44, sK44) & r(sK44, sK45)) & ! [X2, X3, X4] : (r(X2, X4) | ~ r(X3, X4) | ~ r(X2, X3)) & ! [X5, X6] : (r(X6, X5) | ~ r(X5, X6))) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK44, sK45])], [f104, f105])).
fof(f105, plain, (? [X0, X1] : (~ r(X0, X0) & r(X0, X1)) => (~ r(sK44, sK44) & r(sK44, sK45))), introduced(choice_axiom, [])).
fof(f104, plain, ((? [X0, X1] : (~ r(X0, X0) & r(X0, X1)) & ! [X2, X3, X4] : (r(X2, X4) | ~ r(X3, X4) | ~ r(X2, X3)) & ! [X5, X6] : (r(X6, X5) | ~ r(X5, X6))) | ~ sP3), inference(rectify, [], [f103])).
fof(f103, plain, ((? [X48, X49] : (~ r(X48, X48) & r(X48, X49)) & ! [X43, X44, X45] : (r(X43, X45) | ~ r(X44, X45) | ~ r(X43, X44)) & ! [X46, X47] : (r(X47, X46) | ~ r(X46, X47))) | ~ sP3), inference(nnf_transformation, [], [f9])).
fof(f512, plain, (~ spl49_55 | spl49_58), inference(avatar_split_clause, [], [f181, f510, f496])).
fof(f181, plain, ! [X4, X2, X3] : (r(X2, X4) | ~ r(X3, X4) | ~ r(X2, X3) | ~ sP3), inference(cnf_transformation, [], [f106])).
fof(f508, plain, (~ spl49_55 | spl49_57), inference(avatar_split_clause, [], [f182, f505, f496])).
fof(f182, plain, (r(sK44, sK45) | ~ sP3), inference(cnf_transformation, [], [f106])).
fof(f503, plain, (~ spl49_55 | ~ spl49_56), inference(avatar_split_clause, [], [f183, f500, f496])).
fof(f183, plain, (~ r(sK44, sK44) | ~ sP3), inference(cnf_transformation, [], [f106])).
fof(f494, plain, (~ spl49_49 | spl49_54), inference(avatar_split_clause, [], [f176, f491, f472])).
fof(f176, plain, (sP1 | ~ sP4), inference(cnf_transformation, [], [f102])).
fof(f102, plain, ((! [X0] : (h(X0) | ~ g(X0) | ~ f(X0)) & ! [X1] : (g(X1) | ~ h(X1) | ~ f(X1)) & (! [X2] : (h(X2) | ~ f(X2)) | ! [X3] : (g(X3) | ~ f(X3))) & sP1) | ~ sP4), inference(rectify, [], [f101])).
fof(f101, plain, ((! [X42] : (h(X42) | ~ g(X42) | ~ f(X42)) & ! [X41] : (g(X41) | ~ h(X41) | ~ f(X41)) & (! [X37] : (h(X37) | ~ f(X37)) | ! [X38] : (g(X38) | ~ f(X38))) & sP1) | ~ sP4), inference(nnf_transformation, [], [f10])).
fof(f489, plain, (~ spl49_49 | spl49_52 | spl49_53), inference(avatar_split_clause, [], [f177, f487, f484, f472])).
fof(f177, plain, ! [X2, X3] : (h(X2) | ~ f(X2) | g(X3) | ~ f(X3) | ~ sP4), inference(cnf_transformation, [], [f102])).
fof(f482, plain, (~ spl49_49 | spl49_51), inference(avatar_split_clause, [], [f178, f480, f472])).
fof(f178, plain, ! [X1] : (g(X1) | ~ h(X1) | ~ f(X1) | ~ sP4), inference(cnf_transformation, [], [f102])).
fof(f478, plain, (~ spl49_49 | spl49_50), inference(avatar_split_clause, [], [f179, f476, f472])).
fof(f179, plain, ! [X0] : (h(X0) | ~ g(X0) | ~ f(X0) | ~ sP4), inference(cnf_transformation, [], [f102])).
fof(f470, plain, (~ spl49_44 | spl49_3 | spl49_3), inference(avatar_split_clause, [], [f171, f238, f238, f447])).
fof(f171, plain, ! [X4, X5] : (p(X4) | p(X5) | ~ sP5), inference(cnf_transformation, [], [f100])).
fof(f100, plain, (((~ q(sK41) | ~ p(sK42) | (~ q(sK43) | ~ p(sK43))) & ((! [X3] : q(X3) & ! [X4] : p(X4)) | ! [X5] : (q(X5) & p(X5)))) | ~ sP5), inference(skolemisation, [status(esa), new_symbols(skolem, [sK41, sK42, sK43])], [f96, f99, f98, f97])).
fof(f97, plain, (? [X0] : ~ q(X0) => ~ q(sK41)), introduced(choice_axiom, [])).
fof(f98, plain, (? [X1] : ~ p(X1) => ~ p(sK42)), introduced(choice_axiom, [])).
fof(f99, plain, (? [X2] : (~ q(X2) | ~ p(X2)) => (~ q(sK43) | ~ p(sK43))), introduced(choice_axiom, [])).
fof(f96, plain, (((? [X0] : ~ q(X0) | ? [X1] : ~ p(X1) | ? [X2] : (~ q(X2) | ~ p(X2))) & ((! [X3] : q(X3) & ! [X4] : p(X4)) | ! [X5] : (q(X5) & p(X5)))) | ~ sP5), inference(rectify, [], [f95])).
fof(f95, plain, (((? [X33] : ~ q(X33) | ? [X34] : ~ p(X34) | ? [X32] : (~ q(X32) | ~ p(X32))) & ((! [X33] : q(X33) & ! [X34] : p(X34)) | ! [X32] : (q(X32) & p(X32)))) | ~ sP5), inference(flattening, [], [f94])).
fof(f94, plain, ((((? [X33] : ~ q(X33) | ? [X34] : ~ p(X34)) | ? [X32] : (~ q(X32) | ~ p(X32))) & ((! [X33] : q(X33) & ! [X34] : p(X34)) | ! [X32] : (q(X32) & p(X32)))) | ~ sP5), inference(nnf_transformation, [], [f11])).
fof(f467, plain, (~ spl49_44 | spl49_17 | spl49_17), inference(avatar_split_clause, [], [f174, f302, f302, f447])).
fof(f174, plain, ! [X5, X3] : (q(X3) | q(X5) | ~ sP5), inference(cnf_transformation, [], [f100])).
fof(f466, plain, (~ spl49_44 | ~ spl49_45 | ~ spl49_46 | ~ spl49_47 | ~ spl49_48), inference(avatar_split_clause, [], [f175, f463, f459, f455, f451, f447])).
fof(f175, plain, (~ q(sK41) | ~ p(sK42) | ~ q(sK43) | ~ p(sK43) | ~ sP5), inference(cnf_transformation, [], [f100])).
fof(f445, plain, (~ spl49_39 | spl49_40 | spl49_41 | spl49_42 | spl49_43), inference(avatar_split_clause, [], [f166, f442, f438, f434, f430, f422])).
fof(f166, plain, (q(sK38) | p(sK39) | q(sK40) | p(sK40) | ~ sP6), inference(cnf_transformation, [], [f93])).
fof(f93, plain, ((((! [X0] : ~ q(X0) & ! [X1] : ~ p(X1)) | ! [X2] : (~ q(X2) & ~ p(X2))) & (q(sK38) | p(sK39) | (q(sK40) | p(sK40)))) | ~ sP6), inference(skolemisation, [status(esa), new_symbols(skolem, [sK38, sK39, sK40])], [f89, f92, f91, f90])).
fof(f90, plain, (? [X3] : q(X3) => q(sK38)), introduced(choice_axiom, [])).
fof(f91, plain, (? [X4] : p(X4) => p(sK39)), introduced(choice_axiom, [])).
fof(f92, plain, (? [X5] : (q(X5) | p(X5)) => (q(sK40) | p(sK40))), introduced(choice_axiom, [])).
fof(f89, plain, ((((! [X0] : ~ q(X0) & ! [X1] : ~ p(X1)) | ! [X2] : (~ q(X2) & ~ p(X2))) & (? [X3] : q(X3) | ? [X4] : p(X4) | ? [X5] : (q(X5) | p(X5)))) | ~ sP6), inference(rectify, [], [f88])).
fof(f88, plain, ((((! [X27] : ~ q(X27) & ! [X28] : ~ p(X28)) | ! [X26] : (~ q(X26) & ~ p(X26))) & (? [X27] : q(X27) | ? [X28] : p(X28) | ? [X26] : (q(X26) | p(X26)))) | ~ sP6), inference(flattening, [], [f87])).
fof(f87, plain, ((((! [X27] : ~ q(X27) & ! [X28] : ~ p(X28)) | ! [X26] : (~ q(X26) & ~ p(X26))) & ((? [X27] : q(X27) | ? [X28] : p(X28)) | ? [X26] : (q(X26) | p(X26)))) | ~ sP6), inference(nnf_transformation, [], [f12])).
fof(f428, plain, (~ spl49_39 | spl49_9 | spl49_9), inference(avatar_split_clause, [], [f167, f266, f266, f422])).
fof(f167, plain, ! [X2, X1] : (~ p(X1) | ~ p(X2) | ~ sP6), inference(cnf_transformation, [], [f93])).
fof(f425, plain, (~ spl49_39 | spl49_11 | spl49_11), inference(avatar_split_clause, [], [f170, f274, f274, f422])).
fof(f170, plain, ! [X2, X0] : (~ q(X0) | ~ q(X2) | ~ sP6), inference(cnf_transformation, [], [f93])).
fof(f420, plain, (~ spl49_36 | spl49_3 | spl49_3 | spl49_21), inference(avatar_split_clause, [], [f216, f320, f238, f238, f405])).
fof(f216, plain, ! [X2, X3] : (c | p(X2) | p(X3) | ~ sP7), inference(duplicate_literal_removal, [], [f161])).
fof(f161, plain, ! [X2, X3] : (c | p(X2) | c | p(X3) | ~ sP7), inference(cnf_transformation, [], [f86])).
fof(f86, plain, ((((~ c & ~ p(sK36)) | (~ c & ~ p(sK37))) & (c | ! [X2] : p(X2) | ! [X3] : (c | p(X3)))) | ~ sP7), inference(skolemisation, [status(esa), new_symbols(skolem, [sK36, sK37])], [f83, f85, f84])).
fof(f84, plain, (? [X0] : ~ p(X0) => ~ p(sK36)), introduced(choice_axiom, [])).
fof(f85, plain, (? [X1] : (~ c & ~ p(X1)) => (~ c & ~ p(sK37))), introduced(choice_axiom, [])).
fof(f83, plain, ((((~ c & ? [X0] : ~ p(X0)) | ? [X1] : (~ c & ~ p(X1))) & (c | ! [X2] : p(X2) | ! [X3] : (c | p(X3)))) | ~ sP7), inference(rectify, [], [f82])).
fof(f82, plain, ((((~ c & ? [X13] : ~ p(X13)) | ? [X12] : (~ c & ~ p(X12))) & (c | ! [X13] : p(X13) | ! [X12] : (c | p(X12)))) | ~ sP7), inference(flattening, [], [f81])).
fof(f81, plain, ((((~ c & ? [X13] : ~ p(X13)) | ? [X12] : (~ c & ~ p(X12))) & ((c | ! [X13] : p(X13)) | ! [X12] : (c | p(X12)))) | ~ sP7), inference(nnf_transformation, [], [f13])).
fof(f419, plain, (~ spl49_36 | ~ spl49_37 | ~ spl49_38), inference(avatar_split_clause, [], [f162, f415, f410, f405])).
fof(f162, plain, (~ p(sK36) | ~ p(sK37) | ~ sP7), inference(cnf_transformation, [], [f86])).
fof(f408, plain, (~ spl49_36 | ~ spl49_21), inference(avatar_split_clause, [], [f217, f320, f405])).
fof(f217, plain, (~ c | ~ sP7), inference(duplicate_literal_removal, [], [f165])).
fof(f165, plain, (~ c | ~ c | ~ sP7), inference(cnf_transformation, [], [f86])).
fof(f403, plain, (~ spl49_33 | spl49_34 | spl49_35), inference(avatar_split_clause, [], [f156, f399, f394, f388])).
fof(f156, plain, (p(sK34) | p(sK35) | ~ sP8), inference(cnf_transformation, [], [f80])).
fof(f80, plain, (((~ c | ! [X0] : ~ p(X0) | ! [X1] : (~ c | ~ p(X1))) & ((c & p(sK34)) | (c & p(sK35)))) | ~ sP8), inference(skolemisation, [status(esa), new_symbols(skolem, [sK34, sK35])], [f77, f79, f78])).
fof(f78, plain, (? [X2] : p(X2) => p(sK34)), introduced(choice_axiom, [])).
fof(f79, plain, (? [X3] : (c & p(X3)) => (c & p(sK35))), introduced(choice_axiom, [])).
fof(f77, plain, (((~ c | ! [X0] : ~ p(X0) | ! [X1] : (~ c | ~ p(X1))) & ((c & ? [X2] : p(X2)) | ? [X3] : (c & p(X3)))) | ~ sP8), inference(rectify, [], [f76])).
fof(f76, plain, (((~ c | ! [X11] : ~ p(X11) | ! [X10] : (~ c | ~ p(X10))) & ((c & ? [X11] : p(X11)) | ? [X10] : (c & p(X10)))) | ~ sP8), inference(flattening, [], [f75])).
fof(f75, plain, ((((~ c | ! [X11] : ~ p(X11)) | ! [X10] : (~ c | ~ p(X10))) & ((c & ? [X11] : p(X11)) | ? [X10] : (c & p(X10)))) | ~ sP8), inference(nnf_transformation, [], [f14])).
fof(f392, plain, (~ spl49_33 | spl49_21), inference(avatar_split_clause, [], [f218, f320, f388])).
fof(f218, plain, (c | ~ sP8), inference(duplicate_literal_removal, [], [f159])).
fof(f159, plain, (c | c | ~ sP8), inference(cnf_transformation, [], [f80])).
fof(f391, plain, (~ spl49_33 | spl49_9 | spl49_9 | ~ spl49_21), inference(avatar_split_clause, [], [f219, f320, f266, f266, f388])).
fof(f219, plain, ! [X0, X1] : (~ c | ~ p(X0) | ~ p(X1) | ~ sP8), inference(duplicate_literal_removal, [], [f160])).
fof(f160, plain, ! [X0, X1] : (~ c | ~ p(X0) | ~ c | ~ p(X1) | ~ sP8), inference(cnf_transformation, [], [f80])).
fof(f386, plain, (~ spl49_30 | spl49_31 | ~ spl49_21 | spl49_32), inference(avatar_split_clause, [], [f220, f383, f320, f379, f371])).
fof(f220, plain, (p(sK32) | ~ c | p(sK33) | ~ sP9), inference(duplicate_literal_removal, [], [f151])).
fof(f151, plain, (p(sK32) | ~ c | p(sK33) | ~ c | ~ sP9), inference(cnf_transformation, [], [f74])).
fof(f74, plain, ((((! [X0] : ~ p(X0) & c) | ! [X1] : (~ p(X1) & c)) & (p(sK32) | ~ c | (p(sK33) | ~ c))) | ~ sP9), inference(skolemisation, [status(esa), new_symbols(skolem, [sK32, sK33])], [f71, f73, f72])).
fof(f72, plain, (? [X2] : p(X2) => p(sK32)), introduced(choice_axiom, [])).
fof(f73, plain, (? [X3] : (p(X3) | ~ c) => (p(sK33) | ~ c)), introduced(choice_axiom, [])).
fof(f71, plain, ((((! [X0] : ~ p(X0) & c) | ! [X1] : (~ p(X1) & c)) & (? [X2] : p(X2) | ~ c | ? [X3] : (p(X3) | ~ c))) | ~ sP9), inference(rectify, [], [f70])).
fof(f70, plain, ((((! [X7] : ~ p(X7) & c) | ! [X6] : (~ p(X6) & c)) & (? [X7] : p(X7) | ~ c | ? [X6] : (p(X6) | ~ c))) | ~ sP9), inference(flattening, [], [f69])).
fof(f69, plain, ((((! [X7] : ~ p(X7) & c) | ! [X6] : (~ p(X6) & c)) & ((? [X7] : p(X7) | ~ c) | ? [X6] : (p(X6) | ~ c))) | ~ sP9), inference(nnf_transformation, [], [f15])).
fof(f377, plain, (~ spl49_30 | spl49_21), inference(avatar_split_clause, [], [f221, f320, f371])).
fof(f221, plain, (c | ~ sP9), inference(duplicate_literal_removal, [], [f152])).
fof(f152, plain, (c | c | ~ sP9), inference(cnf_transformation, [], [f74])).
fof(f374, plain, (~ spl49_30 | spl49_9 | spl49_9), inference(avatar_split_clause, [], [f155, f266, f266, f371])).
fof(f155, plain, ! [X0, X1] : (~ p(X0) | ~ p(X1) | ~ sP9), inference(cnf_transformation, [], [f74])).
fof(f369, plain, (~ spl49_27 | ~ spl49_28 | ~ spl49_29 | spl49_21), inference(avatar_split_clause, [], [f222, f320, f366, f362, f354])).
fof(f222, plain, (c | ~ p(sK30) | ~ p(sK31) | ~ sP10), inference(duplicate_literal_removal, [], [f146])).
fof(f146, plain, (c | ~ p(sK30) | c | ~ p(sK31) | ~ sP10), inference(cnf_transformation, [], [f68])).
fof(f68, plain, ((((~ c & ! [X0] : p(X0)) | ! [X1] : (~ c & p(X1))) & (c | ~ p(sK30) | (c | ~ p(sK31)))) | ~ sP10), inference(skolemisation, [status(esa), new_symbols(skolem, [sK30, sK31])], [f65, f67, f66])).
fof(f66, plain, (? [X2] : ~ p(X2) => ~ p(sK30)), introduced(choice_axiom, [])).
fof(f67, plain, (? [X3] : (c | ~ p(X3)) => (c | ~ p(sK31))), introduced(choice_axiom, [])).
fof(f65, plain, ((((~ c & ! [X0] : p(X0)) | ! [X1] : (~ c & p(X1))) & (c | ? [X2] : ~ p(X2) | ? [X3] : (c | ~ p(X3)))) | ~ sP10), inference(rectify, [], [f64])).
fof(f64, plain, ((((~ c & ! [X5] : p(X5)) | ! [X4] : (~ c & p(X4))) & (c | ? [X5] : ~ p(X5) | ? [X4] : (c | ~ p(X4)))) | ~ sP10), inference(flattening, [], [f63])).
fof(f63, plain, ((((~ c & ! [X5] : p(X5)) | ! [X4] : (~ c & p(X4))) & ((c | ? [X5] : ~ p(X5)) | ? [X4] : (c | ~ p(X4)))) | ~ sP10), inference(nnf_transformation, [], [f16])).
fof(f360, plain, (~ spl49_27 | spl49_3 | spl49_3), inference(avatar_split_clause, [], [f147, f238, f238, f354])).
fof(f147, plain, ! [X0, X1] : (p(X0) | p(X1) | ~ sP10), inference(cnf_transformation, [], [f68])).
fof(f357, plain, (~ spl49_27 | ~ spl49_21), inference(avatar_split_clause, [], [f223, f320, f354])).
fof(f223, plain, (~ c | ~ sP10), inference(duplicate_literal_removal, [], [f150])).
fof(f150, plain, (~ c | ~ c | ~ sP10), inference(cnf_transformation, [], [f68])).
fof(f352, plain, (~ spl49_24 | spl49_3 | ~ spl49_21 | spl49_3), inference(avatar_split_clause, [], [f224, f238, f320, f238, f337])).
fof(f224, plain, ! [X2, X3] : (p(X2) | ~ c | p(X3) | ~ sP11), inference(duplicate_literal_removal, [], [f141])).
fof(f141, plain, ! [X2, X3] : (p(X2) | ~ c | p(X3) | ~ c | ~ sP11), inference(cnf_transformation, [], [f62])).
fof(f62, plain, ((((~ p(sK28) & c) | (~ p(sK29) & c)) & (! [X2] : p(X2) | ~ c | ! [X3] : (p(X3) | ~ c))) | ~ sP11), inference(skolemisation, [status(esa), new_symbols(skolem, [sK28, sK29])], [f59, f61, f60])).
fof(f60, plain, (? [X0] : ~ p(X0) => ~ p(sK28)), introduced(choice_axiom, [])).
fof(f61, plain, (? [X1] : (~ p(X1) & c) => (~ p(sK29) & c)), introduced(choice_axiom, [])).
fof(f59, plain, ((((? [X0] : ~ p(X0) & c) | ? [X1] : (~ p(X1) & c)) & (! [X2] : p(X2) | ~ c | ! [X3] : (p(X3) | ~ c))) | ~ sP11), inference(rectify, [], [f58])).
fof(f58, plain, ((((? [X3] : ~ p(X3) & c) | ? [X2] : (~ p(X2) & c)) & (! [X3] : p(X3) | ~ c | ! [X2] : (p(X2) | ~ c))) | ~ sP11), inference(flattening, [], [f57])).
fof(f57, plain, ((((? [X3] : ~ p(X3) & c) | ? [X2] : (~ p(X2) & c)) & ((! [X3] : p(X3) | ~ c) | ! [X2] : (p(X2) | ~ c))) | ~ sP11), inference(nnf_transformation, [], [f17])).
fof(f351, plain, (~ spl49_24 | spl49_21), inference(avatar_split_clause, [], [f225, f320, f337])).
fof(f225, plain, (c | ~ sP11), inference(duplicate_literal_removal, [], [f142])).
fof(f142, plain, (c | c | ~ sP11), inference(cnf_transformation, [], [f62])).
fof(f348, plain, (~ spl49_24 | ~ spl49_25 | ~ spl49_26), inference(avatar_split_clause, [], [f145, f345, f341, f337])).
fof(f145, plain, (~ p(sK28) | ~ p(sK29) | ~ sP11), inference(cnf_transformation, [], [f62])).
fof(f335, plain, (~ spl49_20 | spl49_9 | spl49_9 | spl49_21), inference(avatar_split_clause, [], [f226, f320, f266, f266, f316])).
fof(f226, plain, ! [X2, X3] : (c | ~ p(X2) | ~ p(X3) | ~ sP12), inference(duplicate_literal_removal, [], [f136])).
fof(f136, plain, ! [X2, X3] : (c | ~ p(X2) | c | ~ p(X3) | ~ sP12), inference(cnf_transformation, [], [f56])).
fof(f56, plain, ((((~ c & p(sK26)) | (~ c & p(sK27))) & (c | ! [X2] : ~ p(X2) | ! [X3] : (c | ~ p(X3)))) | ~ sP12), inference(skolemisation, [status(esa), new_symbols(skolem, [sK26, sK27])], [f53, f55, f54])).
fof(f54, plain, (? [X0] : p(X0) => p(sK26)), introduced(choice_axiom, [])).
fof(f55, plain, (? [X1] : (~ c & p(X1)) => (~ c & p(sK27))), introduced(choice_axiom, [])).
fof(f53, plain, ((((~ c & ? [X0] : p(X0)) | ? [X1] : (~ c & p(X1))) & (c | ! [X2] : ~ p(X2) | ! [X3] : (c | ~ p(X3)))) | ~ sP12), inference(rectify, [], [f52])).
fof(f52, plain, ((((~ c & ? [X1] : p(X1)) | ? [X0] : (~ c & p(X0))) & (c | ! [X1] : ~ p(X1) | ! [X0] : (c | ~ p(X0)))) | ~ sP12), inference(flattening, [], [f51])).
fof(f51, plain, ((((~ c & ? [X1] : p(X1)) | ? [X0] : (~ c & p(X0))) & ((c | ! [X1] : ~ p(X1)) | ! [X0] : (c | ~ p(X0)))) | ~ sP12), inference(nnf_transformation, [], [f18])).
fof(f334, plain, (~ spl49_20 | spl49_22 | spl49_23), inference(avatar_split_clause, [], [f137, f330, f325, f316])).
fof(f137, plain, (p(sK26) | p(sK27) | ~ sP12), inference(cnf_transformation, [], [f56])).
fof(f323, plain, (~ spl49_20 | ~ spl49_21), inference(avatar_split_clause, [], [f227, f320, f316])).
fof(f227, plain, (~ c | ~ sP12), inference(duplicate_literal_removal, [], [f140])).
fof(f140, plain, (~ c | ~ c | ~ sP12), inference(cnf_transformation, [], [f56])).
fof(f314, plain, (~ spl49_18 | spl49_19), inference(avatar_split_clause, [], [f133, f312, f306])).
fof(f133, plain, ! [X0] : (q(X0) | ~ p(sK25(X0)) | ~ sP13), inference(cnf_transformation, [], [f50])).
fof(f50, plain, (! [X0] : (~ q(X0) & p(X0) & (q(X0) | ~ p(sK25(X0)))) | ~ sP13), inference(skolemisation, [status(esa), new_symbols(skolem, [sK25])], [f48, f49])).
fof(f49, plain, ! [X0] : (? [X1] : (~ q(X0) & p(X0) & (q(X0) | ~ p(X1))) => (~ q(X0) & p(X0) & (q(X0) | ~ p(sK25(X0))))), introduced(choice_axiom, [])).
fof(f48, plain, (! [X0] : ? [X1] : (~ q(X0) & p(X0) & (q(X0) | ~ p(X1))) | ~ sP13), inference(rectify, [], [f47])).
fof(f47, plain, (! [X35] : ? [X36] : (~ q(X35) & p(X35) & (q(X35) | ~ p(X36))) | ~ sP13), inference(nnf_transformation, [], [f19])).
fof(f310, plain, (~ spl49_18 | spl49_3), inference(avatar_split_clause, [], [f134, f238, f306])).
fof(f134, plain, ! [X0] : (p(X0) | ~ sP13), inference(cnf_transformation, [], [f50])).
fof(f309, plain, (~ spl49_18 | spl49_11), inference(avatar_split_clause, [], [f135, f274, f306])).
fof(f135, plain, ! [X0] : (~ q(X0) | ~ sP13), inference(cnf_transformation, [], [f50])).
fof(f304, plain, (~ spl49_14 | spl49_3 | spl49_17), inference(avatar_split_clause, [], [f130, f302, f238, f288])).
fof(f130, plain, ! [X2, X1] : (q(X1) | p(X2) | ~ sP14), inference(cnf_transformation, [], [f46])).
fof(f46, plain, (((~ q(sK24) & ~ p(sK24)) & (! [X1] : q(X1) | ! [X2] : p(X2))) | ~ sP14), inference(skolemisation, [status(esa), new_symbols(skolem, [sK24])], [f44, f45])).
fof(f45, plain, (? [X0] : (~ q(X0) & ~ p(X0)) => (~ q(sK24) & ~ p(sK24))), introduced(choice_axiom, [])).
fof(f44, plain, ((? [X0] : (~ q(X0) & ~ p(X0)) & (! [X1] : q(X1) | ! [X2] : p(X2))) | ~ sP14), inference(rectify, [], [f43])).
fof(f43, plain, ((? [X31] : (~ q(X31) & ~ p(X31)) & (! [X29] : q(X29) | ! [X30] : p(X30))) | ~ sP14), inference(nnf_transformation, [], [f20])).
fof(f300, plain, (~ spl49_14 | ~ spl49_16), inference(avatar_split_clause, [], [f131, f297, f288])).
fof(f131, plain, (~ p(sK24) | ~ sP14), inference(cnf_transformation, [], [f46])).
fof(f295, plain, (~ spl49_14 | ~ spl49_15), inference(avatar_split_clause, [], [f132, f292, f288])).
fof(f132, plain, (~ q(sK24) | ~ sP14), inference(cnf_transformation, [], [f46])).
fof(f286, plain, (~ spl49_10 | spl49_13), inference(avatar_split_clause, [], [f127, f283, f270])).
fof(f127, plain, (p(sK23) | ~ sP15), inference(cnf_transformation, [], [f42])).
fof(f42, plain, (((! [X0] : ~ q(X0) | ! [X1] : ~ p(X1)) & (q(sK23) & p(sK23))) | ~ sP15), inference(skolemisation, [status(esa), new_symbols(skolem, [sK23])], [f40, f41])).
fof(f41, plain, (? [X2] : (q(X2) & p(X2)) => (q(sK23) & p(sK23))), introduced(choice_axiom, [])).
fof(f40, plain, (((! [X0] : ~ q(X0) | ! [X1] : ~ p(X1)) & ? [X2] : (q(X2) & p(X2))) | ~ sP15), inference(rectify, [], [f39])).
fof(f39, plain, (((! [X22] : ~ q(X22) | ! [X23] : ~ p(X23)) & ? [X21] : (q(X21) & p(X21))) | ~ sP15), inference(nnf_transformation, [], [f21])).
fof(f281, plain, (~ spl49_10 | spl49_12), inference(avatar_split_clause, [], [f128, f278, f270])).
fof(f128, plain, (q(sK23) | ~ sP15), inference(cnf_transformation, [], [f42])).
fof(f276, plain, (~ spl49_10 | spl49_9 | spl49_11), inference(avatar_split_clause, [], [f129, f274, f266, f270])).
fof(f129, plain, ! [X0, X1] : (~ q(X0) | ~ p(X1) | ~ sP15), inference(cnf_transformation, [], [f42])).
fof(f268, plain, (~ spl49_6 | spl49_9), inference(avatar_split_clause, [], [f124, f266, f252])).
fof(f124, plain, ! [X2] : (~ p(X2) | ~ sP16), inference(cnf_transformation, [], [f38])).
fof(f38, plain, (((~ p(sK21) & p(sK22)) & ! [X2] : ~ p(X2)) | ~ sP16), inference(skolemisation, [status(esa), new_symbols(skolem, [sK21, sK22])], [f35, f37, f36])).
fof(f36, plain, (? [X0] : (~ p(X0) & ? [X1] : p(X1)) => (~ p(sK21) & ? [X1] : p(X1))), introduced(choice_axiom, [])).
fof(f37, plain, (? [X1] : p(X1) => p(sK22)), introduced(choice_axiom, [])).
fof(f35, plain, ((? [X0] : (~ p(X0) & ? [X1] : p(X1)) & ! [X2] : ~ p(X2)) | ~ sP16), inference(rectify, [], [f34])).
fof(f34, plain, ((? [X15] : (~ p(X15) & ? [X16] : p(X16)) & ! [X14] : ~ p(X14)) | ~ sP16), inference(nnf_transformation, [], [f22])).
fof(f264, plain, (~ spl49_6 | spl49_8), inference(avatar_split_clause, [], [f125, f261, f252])).
fof(f125, plain, (p(sK22) | ~ sP16), inference(cnf_transformation, [], [f38])).
fof(f250, plain, (~ spl49_4 | spl49_3), inference(avatar_split_clause, [], [f122, f238, f242])).
fof(f122, plain, ! [X0] : (p(X0) | ~ sP17), inference(cnf_transformation, [], [f33])).
fof(f33, plain, (! [X0] : (~ p(sK20) & p(X0)) | ~ sP17), inference(skolemisation, [status(esa), new_symbols(skolem, [sK20])], [f31, f32])).
fof(f32, plain, (? [X1] : ~ p(X1) => ~ p(sK20)), introduced(choice_axiom, [])).
fof(f31, plain, (! [X0] : (? [X1] : ~ p(X1) & p(X0)) | ~ sP17), inference(rectify, [], [f30])).
fof(f30, plain, (! [X24] : (? [X25] : ~ p(X25) & p(X24)) | ~ sP17), inference(nnf_transformation, [], [f23])).
fof(f249, plain, (~ spl49_4 | ~ spl49_5), inference(avatar_split_clause, [], [f123, f246, f242])).
fof(f123, plain, (~ p(sK20) | ~ sP17), inference(cnf_transformation, [], [f33])).
fof(f240, plain, (~ spl49_1 | spl49_3), inference(avatar_split_clause, [], [f120, f238, f229])).
fof(f120, plain, ! [X1] : (p(X1) | ~ sP18), inference(cnf_transformation, [], [f29])).
fof(f29, plain, ((~ p(sK19) & ! [X1] : p(X1)) | ~ sP18), inference(skolemisation, [status(esa), new_symbols(skolem, [sK19])], [f27, f28])).
fof(f28, plain, (? [X0] : (~ p(X0) & ! [X1] : p(X1)) => (~ p(sK19) & ! [X1] : p(X1))), introduced(choice_axiom, [])).
fof(f27, plain, (? [X0] : (~ p(X0) & ! [X1] : p(X1)) | ~ sP18), inference(rectify, [], [f26])).
fof(f26, plain, (? [X19] : (~ p(X19) & ! [X20] : p(X20)) | ~ sP18), inference(nnf_transformation, [], [f24])).
fof(f236, plain, (~ spl49_1 | ~ spl49_2), inference(avatar_split_clause, [], [f121, f233, f229])).
fof(f121, plain, (~ p(sK19) | ~ sP18), inference(cnf_transformation, [], [f29])).