fof(f24143, plain, $false, inference(avatar_sat_refutation, [], [f118, f150, f158, f162, f186, f194, f198, f206, f427, f479, f509, f685, f818, f838, f930, f977, f979, f982, f997, f2968, f3309, f4843, f5004, f5242, f5254, f5318, f5819, f6098, f10388, f15604, f23836, f23975, f23976, f23977, f24141, f24142])).
fof(f24142, plain, (~ spl12_254 | ~ spl12_239), inference(avatar_split_clause, [], [f24134, f23833, f23972])).
fof(f23972, plain, (spl12_254 <=> g_true_only(sK8(sK10(sK11), sK11), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_254])])).
fof(f23833, plain, (spl12_239 <=> g_true_only(sK8(sK10(sK11), sK11), sK10(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_239])])).
fof(f24134, plain, (~ g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ spl12_239), inference(resolution, [], [f23835, f80])).
fof(f80, plain, ! [X5, X3] : (~ g_true_only(X5, sK10(X3)) | ~ g_true_only(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f38, plain, ((sP2 | ! [X0] : (! [X2] : ((g_true_only(X2, X0) & g_false_only(X2, sK9(X0))) | (g_false_only(X2, X0) & g_true_only(X2, sK9(X0)))) | sP0(X0)) | sP3 | sP1) & ! [X3] : ! [X5] : (((~ g_both(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_false_only(X5, sK10(X3))) & ((~ g_true_only(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_both(X5, sK10(X3))) & ((~ g_true_only(X5, X3) & ~ g_both(X5, X3)) | ~ g_true_only(X5, sK10(X3)))) & ! [X7] : (((~ g_both(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_false_only(X7, sK11)) & ((~ g_false_only(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_both(X7, sK11)) & ((~ g_false_only(X7, X7) & ~ g_both(X7, X7)) | ~ g_true_only(X7, sK11)))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f34, f37, f36, f35])).
fof(f35, plain, ! [X0] : (? [X1] : ! [X2] : ((g_true_only(X2, X0) & g_false_only(X2, X1)) | (g_false_only(X2, X0) & g_true_only(X2, X1))) => ! [X2] : ((g_true_only(X2, X0) & g_false_only(X2, sK9(X0))) | (g_false_only(X2, X0) & g_true_only(X2, sK9(X0))))), introduced(choice_axiom, [])).
fof(f36, plain, ! [X3] : (? [X4] : ! [X5] : (((~ g_both(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_false_only(X5, X4)) & ((~ g_true_only(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_both(X5, X4)) & ((~ g_true_only(X5, X3) & ~ g_both(X5, X3)) | ~ g_true_only(X5, X4))) => ! [X5] : (((~ g_both(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_false_only(X5, sK10(X3))) & ((~ g_true_only(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_both(X5, sK10(X3))) & ((~ g_true_only(X5, X3) & ~ g_both(X5, X3)) | ~ g_true_only(X5, sK10(X3))))), introduced(choice_axiom, [])).
fof(f37, plain, (? [X6] : ! [X7] : (((~ g_both(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_false_only(X7, X6)) & ((~ g_false_only(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_both(X7, X6)) & ((~ g_false_only(X7, X7) & ~ g_both(X7, X7)) | ~ g_true_only(X7, X6))) => ! [X7] : (((~ g_both(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_false_only(X7, sK11)) & ((~ g_false_only(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_both(X7, sK11)) & ((~ g_false_only(X7, X7) & ~ g_both(X7, X7)) | ~ g_true_only(X7, sK11)))), introduced(choice_axiom, [])).
fof(f34, plain, ((sP2 | ! [X0] : (? [X1] : ! [X2] : ((g_true_only(X2, X0) & g_false_only(X2, X1)) | (g_false_only(X2, X0) & g_true_only(X2, X1))) | sP0(X0)) | sP3 | sP1) & ! [X3] : ? [X4] : ! [X5] : (((~ g_both(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_false_only(X5, X4)) & ((~ g_true_only(X5, X3) & ~ g_false_only(X5, X3)) | ~ g_both(X5, X4)) & ((~ g_true_only(X5, X3) & ~ g_both(X5, X3)) | ~ g_true_only(X5, X4))) & ? [X6] : ! [X7] : (((~ g_both(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_false_only(X7, X6)) & ((~ g_false_only(X7, X7) & ~ g_true_only(X7, X7)) | ~ g_both(X7, X6)) & ((~ g_false_only(X7, X7) & ~ g_both(X7, X7)) | ~ g_true_only(X7, X6)))), inference(rectify, [], [f17])).
fof(f17, plain, ((sP2 | ! [X3] : (? [X4] : ! [X5] : ((g_true_only(X5, X3) & g_false_only(X5, X4)) | (g_false_only(X5, X3) & g_true_only(X5, X4))) | sP0(X3)) | sP3 | sP1) & ! [X14] : ? [X15] : ! [X16] : (((~ g_both(X16, X14) & ~ g_false_only(X16, X14)) | ~ g_false_only(X16, X15)) & ((~ g_true_only(X16, X14) & ~ g_false_only(X16, X14)) | ~ g_both(X16, X15)) & ((~ g_true_only(X16, X14) & ~ g_both(X16, X14)) | ~ g_true_only(X16, X15))) & ? [X17] : ! [X18] : (((~ g_both(X18, X18) & ~ g_true_only(X18, X18)) | ~ g_false_only(X18, X17)) & ((~ g_false_only(X18, X18) & ~ g_true_only(X18, X18)) | ~ g_both(X18, X17)) & ((~ g_false_only(X18, X18) & ~ g_both(X18, X18)) | ~ g_true_only(X18, X17)))), inference(definition_folding, [], [f12, e16, e15, e14, e13])).
fof(f13, plain, ! [X3] : (! [X6] : (? [X7] : (((g_both(X7, X3) | g_false_only(X7, X3)) & g_false_only(X7, X6)) | ((g_true_only(X7, X3) | g_false_only(X7, X3)) & g_both(X7, X6)) | ((g_true_only(X7, X3) | g_both(X7, X3)) & g_true_only(X7, X6))) | ! [X8] : (~ g_both(X8, X3) | ~ g_both(X8, X6))) | ~ sP0(X3)), inference(usedef, [], [e13])).
fof(e13, plain, ! [X3] : (sP0(X3) <=> ! [X6] : (? [X7] : (((g_both(X7, X3) | g_false_only(X7, X3)) & g_false_only(X7, X6)) | ((g_true_only(X7, X3) | g_false_only(X7, X3)) & g_both(X7, X6)) | ((g_true_only(X7, X3) | g_both(X7, X3)) & g_true_only(X7, X6))) | ! [X8] : (~ g_both(X8, X3) | ~ g_both(X8, X6)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP0])])).
fof(f14, plain, (! [X11] : (? [X12] : (((g_both(X12, X12) | g_true_only(X12, X12)) & g_false_only(X12, X11)) | ((g_false_only(X12, X12) | g_true_only(X12, X12)) & g_both(X12, X11)) | ((g_false_only(X12, X12) | g_both(X12, X12)) & g_true_only(X12, X11))) | ! [X13] : (~ g_both(X13, X13) | ~ g_both(X13, X11))) | ~ sP1), inference(usedef, [], [e14])).
fof(e14, plain, (sP1 <=> ! [X11] : (? [X12] : (((g_both(X12, X12) | g_true_only(X12, X12)) & g_false_only(X12, X11)) | ((g_false_only(X12, X12) | g_true_only(X12, X12)) & g_both(X12, X11)) | ((g_false_only(X12, X12) | g_both(X12, X12)) & g_true_only(X12, X11))) | ! [X13] : (~ g_both(X13, X13) | ~ g_both(X13, X11)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f15, plain, (? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) | ~ sP2), inference(usedef, [], [e15])).
fof(e15, plain, (sP2 <=> ? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f16, plain, (? [X9] : ! [X10] : ((g_false_only(X10, X10) & g_false_only(X10, X9)) | (g_true_only(X10, X10) & g_true_only(X10, X9))) | ~ sP3), inference(usedef, [], [e16])).
fof(e16, plain, (sP3 <=> ? [X9] : ! [X10] : ((g_false_only(X10, X10) & g_false_only(X10, X9)) | (g_true_only(X10, X10) & g_true_only(X10, X9)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP3])])).
fof(f12, plain, ((? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) | ! [X3] : (? [X4] : ! [X5] : ((g_true_only(X5, X3) & g_false_only(X5, X4)) | (g_false_only(X5, X3) & g_true_only(X5, X4))) | ! [X6] : (? [X7] : (((g_both(X7, X3) | g_false_only(X7, X3)) & g_false_only(X7, X6)) | ((g_true_only(X7, X3) | g_false_only(X7, X3)) & g_both(X7, X6)) | ((g_true_only(X7, X3) | g_both(X7, X3)) & g_true_only(X7, X6))) | ! [X8] : (~ g_both(X8, X3) | ~ g_both(X8, X6)))) | ? [X9] : ! [X10] : ((g_false_only(X10, X10) & g_false_only(X10, X9)) | (g_true_only(X10, X10) & g_true_only(X10, X9))) | ! [X11] : (? [X12] : (((g_both(X12, X12) | g_true_only(X12, X12)) & g_false_only(X12, X11)) | ((g_false_only(X12, X12) | g_true_only(X12, X12)) & g_both(X12, X11)) | ((g_false_only(X12, X12) | g_both(X12, X12)) & g_true_only(X12, X11))) | ! [X13] : (~ g_both(X13, X13) | ~ g_both(X13, X11)))) & ! [X14] : ? [X15] : ! [X16] : (((~ g_both(X16, X14) & ~ g_false_only(X16, X14)) | ~ g_false_only(X16, X15)) & ((~ g_true_only(X16, X14) & ~ g_false_only(X16, X14)) | ~ g_both(X16, X15)) & ((~ g_true_only(X16, X14) & ~ g_both(X16, X14)) | ~ g_true_only(X16, X15))) & ? [X17] : ! [X18] : (((~ g_both(X18, X18) & ~ g_true_only(X18, X18)) | ~ g_false_only(X18, X17)) & ((~ g_false_only(X18, X18) & ~ g_true_only(X18, X18)) | ~ g_both(X18, X17)) & ((~ g_false_only(X18, X18) & ~ g_both(X18, X18)) | ~ g_true_only(X18, X17)))), inference(ennf_transformation, [], [f7])).
fof(f7, plain, ~ ((~ ? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) & ? [X3] : (~ ? [X4] : ! [X5] : ((g_true_only(X5, X3) & g_false_only(X5, X4)) | (g_false_only(X5, X3) & g_true_only(X5, X4))) & ? [X6] : (~ ? [X7] : (((g_both(X7, X3) | g_false_only(X7, X3)) & g_false_only(X7, X6)) | ((g_true_only(X7, X3) | g_false_only(X7, X3)) & g_both(X7, X6)) | ((g_true_only(X7, X3) | g_both(X7, X3)) & g_true_only(X7, X6))) & ? [X8] : (g_both(X8, X3) & g_both(X8, X6)))) & ~ ? [X9] : ! [X10] : ((g_false_only(X10, X10) & g_false_only(X10, X9)) | (g_true_only(X10, X10) & g_true_only(X10, X9))) & ? [X11] : (~ ? [X12] : (((g_both(X12, X12) | g_true_only(X12, X12)) & g_false_only(X12, X11)) | ((g_false_only(X12, X12) | g_true_only(X12, X12)) & g_both(X12, X11)) | ((g_false_only(X12, X12) | g_both(X12, X12)) & g_true_only(X12, X11))) & ? [X13] : (g_both(X13, X13) & g_both(X13, X11)))) | ? [X14] : ! [X15] : ? [X16] : (((g_both(X16, X14) | g_false_only(X16, X14)) & g_false_only(X16, X15)) | ((g_true_only(X16, X14) | g_false_only(X16, X14)) & g_both(X16, X15)) | ((g_true_only(X16, X14) | g_both(X16, X14)) & g_true_only(X16, X15))) | ! [X17] : ? [X18] : (((g_both(X18, X18) | g_true_only(X18, X18)) & g_false_only(X18, X17)) | ((g_false_only(X18, X18) | g_true_only(X18, X18)) & g_both(X18, X17)) | ((g_false_only(X18, X18) | g_both(X18, X18)) & g_true_only(X18, X17)))), inference(rectify, [], [f2])).
fof(f2, plain, ~ ((~ ? [X1] : ! [X0] : ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) & ? [X1] : (~ ? [X0] : ! [X2] : ((g_true_only(X2, X1) & g_false_only(X2, X0)) | (g_false_only(X2, X1) & g_true_only(X2, X0))) & ? [X0] : (~ ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) & ? [X2] : (g_both(X2, X1) & g_both(X2, X0)))) & ~ ? [X0] : ! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, X0)) | (g_true_only(X1, X1) & g_true_only(X1, X0))) & ? [X0] : (~ ? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0))) & ? [X1] : (g_both(X1, X1) & g_both(X1, X0)))) | ? [X1] : ! [X0] : ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) | ! [X0] : ? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0)))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ((~ ? [X1] : ! [X0] : ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) & ? [X1] : (~ ? [X0] : ! [X2] : ((g_true_only(X2, X1) & g_false_only(X2, X0)) | (g_false_only(X2, X1) & g_true_only(X2, X0))) & ? [X0] : (~ ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) & ? [X2] : (g_both(X2, X1) & g_both(X2, X0)))) & ~ ? [X0] : ! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, X0)) | (g_true_only(X1, X1) & g_true_only(X1, X0))) & ? [X0] : (~ ? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0))) & ? [X1] : (g_both(X1, X1) & g_both(X1, X0)))) | ? [X1] : ! [X0] : ? [X2] : (((g_both(X2, X1) | g_false_only(X2, X1)) & g_false_only(X2, X0)) | ((g_true_only(X2, X1) | g_false_only(X2, X1)) & g_both(X2, X0)) | ((g_true_only(X2, X1) | g_both(X2, X1)) & g_true_only(X2, X0))) | ! [X0] : ? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0)))), file('/home/ubuntu/library/tptp/Problems/SYO/SYO605+1.p', nc4)).
fof(f23835, plain, (g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | ~ spl12_239), inference(avatar_component_clause, [], [f23833])).
fof(f24141, plain, (spl12_254 | ~ spl12_24 | ~ spl12_42 | spl12_43 | spl12_238 | ~ spl12_239), inference(avatar_split_clause, [], [f24140, f23833, f23829, f424, f420, f204, f23972])).
fof(f204, plain, (spl12_24 <=> ! [X0, X2] : (g_true_only(X2, X0) | sP0(X0) | g_false_only(X2, X0))), introduced(avatar_definition, [new_symbols(naming, [spl12_24])])).
fof(f420, plain, (spl12_42 <=> g_both(sK10(sK11), sK10(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_42])])).
fof(f424, plain, (spl12_43 <=> g_true_only(sK10(sK11), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_43])])).
fof(f23829, plain, (spl12_238 <=> g_both(sK8(sK10(sK11), sK11), sK10(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_238])])).
fof(f24140, plain, (g_true_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_42 | spl12_43 | spl12_238 | ~ spl12_239)), inference(subsumption_resolution, [], [f24139, f24040])).
fof(f24040, plain, (~ g_both(sK8(sK10(sK11), sK11), sK11) | spl12_238), inference(resolution, [], [f23830, f229])).
fof(f229, plain, ! [X4, X5] : (g_both(X4, sK10(X5)) | ~ g_both(X4, X5)), inference(subsumption_resolution, [], [f222, f79])).
fof(f79, plain, ! [X5, X3] : (~ g_true_only(X5, sK10(X3)) | ~ g_both(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f222, plain, ! [X4, X5] : (g_both(X4, sK10(X5)) | g_true_only(X4, sK10(X5)) | ~ g_both(X4, X5)), inference(resolution, [], [f98, f84])).
fof(f84, plain, ! [X5, X3] : (~ g_false_only(X5, sK10(X3)) | ~ g_both(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f98, plain, ! [X0, X1] : (g_false_only(X0, X1) | g_both(X0, X1) | g_true_only(X0, X1)), inference(cnf_transformation, [], [f11])).
fof(f11, plain, ! [X0, X1] : (g_false_only(X0, X1) | g_both(X0, X1) | g_true_only(X0, X1)), inference(rectify, [], [f6])).
fof(f6, plain, ! [X3, X4] : (g_false_only(X3, X4) | g_both(X3, X4) | g_true_only(X3, X4)), file('/home/ubuntu/library/tptp/Problems/SYO/SYO605+1.p', exhaustion_g)).
fof(f23830, plain, (~ g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | spl12_238), inference(avatar_component_clause, [], [f23829])).
fof(f24139, plain, (g_true_only(sK8(sK10(sK11), sK11), sK11) | g_both(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_42 | spl12_43 | ~ spl12_239)), inference(subsumption_resolution, [], [f24138, f422])).
fof(f422, plain, (g_both(sK10(sK11), sK10(sK11)) | ~ spl12_42), inference(avatar_component_clause, [], [f420])).
fof(f24138, plain, (g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | spl12_43 | ~ spl12_239)), inference(subsumption_resolution, [], [f24137, f425])).
fof(f425, plain, (~ g_true_only(sK10(sK11), sK11) | spl12_43), inference(avatar_component_clause, [], [f424])).
fof(f24137, plain, (g_true_only(sK10(sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_239)), inference(duplicate_literal_removal, [], [f24133])).
fof(f24133, plain, (g_true_only(sK10(sK11), sK11) | g_true_only(sK10(sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_239)), inference(resolution, [], [f23835, f15952])).
fof(f15952, plain, (! [X0, X1] : (~ g_true_only(sK8(X0, X1), X0) | g_true_only(sK10(X1), X1) | g_true_only(sK10(X1), sK11) | g_true_only(sK8(X0, X1), X1) | ~ g_both(sK10(X1), X0) | g_both(sK8(X0, X1), X1)) | ~ spl12_24), inference(subsumption_resolution, [], [f15951, f6005])).
fof(f6005, plain, (! [X4, X5] : (~ g_both(X5, X4) | sP0(X4)) | ~ spl12_24), inference(subsumption_resolution, [], [f5959, f302])).
fof(f302, plain, ! [X0, X1] : (~ g_both(X0, X1) | ~ g_true_only(X0, X1)), inference(resolution, [], [f229, f82])).
fof(f82, plain, ! [X5, X3] : (~ g_both(X5, sK10(X3)) | ~ g_true_only(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f5959, plain, (! [X4, X5] : (sP0(X4) | g_true_only(X5, X4) | ~ g_both(X5, X4)) | ~ spl12_24), inference(resolution, [], [f205, f303])).
fof(f303, plain, ! [X2, X3] : (~ g_false_only(X2, X3) | ~ g_both(X2, X3)), inference(resolution, [], [f229, f81])).
fof(f81, plain, ! [X5, X3] : (~ g_both(X5, sK10(X3)) | ~ g_false_only(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f205, plain, (! [X2, X0] : (g_false_only(X2, X0) | sP0(X0) | g_true_only(X2, X0)) | ~ spl12_24), inference(avatar_component_clause, [], [f204])).
fof(f15951, plain, ! [X0, X1] : (g_both(sK8(X0, X1), X1) | g_true_only(sK10(X1), X1) | g_true_only(sK10(X1), sK11) | g_true_only(sK8(X0, X1), X1) | ~ g_both(sK10(X1), X0) | ~ sP0(X0) | ~ g_true_only(sK8(X0, X1), X0)), inference(subsumption_resolution, [], [f15901, f302])).
fof(f15901, plain, ! [X0, X1] : (g_both(sK8(X0, X1), X1) | g_true_only(sK10(X1), X1) | g_true_only(sK10(X1), sK11) | g_true_only(sK8(X0, X1), X1) | ~ g_both(sK10(X1), X0) | g_both(sK8(X0, X1), X0) | ~ sP0(X0) | ~ g_true_only(sK8(X0, X1), X0)), inference(resolution, [], [f405, f319])).
fof(f319, plain, ! [X0, X1] : (~ g_false_only(X0, X1) | ~ g_true_only(X0, X1)), inference(resolution, [], [f230, f80])).
fof(f230, plain, ! [X6, X7] : (g_true_only(X6, sK10(X7)) | ~ g_false_only(X6, X7)), inference(subsumption_resolution, [], [f223, f81])).
fof(f223, plain, ! [X6, X7] : (g_both(X6, sK10(X7)) | g_true_only(X6, sK10(X7)) | ~ g_false_only(X6, X7)), inference(resolution, [], [f98, f83])).
fof(f83, plain, ! [X5, X3] : (~ g_false_only(X5, sK10(X3)) | ~ g_false_only(X5, X3)), inference(cnf_transformation, [], [f38])).
fof(f405, plain, ! [X10, X9] : (g_false_only(sK8(X10, X9), X10) | g_both(sK8(X10, X9), X9) | g_true_only(sK10(X9), X9) | g_true_only(sK10(X9), sK11) | g_true_only(sK8(X10, X9), X9) | ~ g_both(sK10(X9), X10) | g_both(sK8(X10, X9), X10) | ~ sP0(X10)), inference(resolution, [], [f339, f69])).
fof(f69, plain, ! [X0, X3, X1] : (~ g_both(X3, X1) | g_false_only(sK8(X0, X1), X0) | g_both(sK8(X0, X1), X1) | g_true_only(sK8(X0, X1), X1) | ~ g_both(X3, X0) | g_both(sK8(X0, X1), X0) | ~ sP0(X0)), inference(cnf_transformation, [], [f33])).
fof(f33, plain, ! [X0] : (! [X1] : ((((g_both(sK8(X0, X1), X0) | g_false_only(sK8(X0, X1), X0)) & g_false_only(sK8(X0, X1), X1)) | ((g_true_only(sK8(X0, X1), X0) | g_false_only(sK8(X0, X1), X0)) & g_both(sK8(X0, X1), X1)) | ((g_true_only(sK8(X0, X1), X0) | g_both(sK8(X0, X1), X0)) & g_true_only(sK8(X0, X1), X1))) | ! [X3] : (~ g_both(X3, X0) | ~ g_both(X3, X1))) | ~ sP0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK8])], [f31, f32])).
fof(f32, plain, ! [X1, X0] : (? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) => (((g_both(sK8(X0, X1), X0) | g_false_only(sK8(X0, X1), X0)) & g_false_only(sK8(X0, X1), X1)) | ((g_true_only(sK8(X0, X1), X0) | g_false_only(sK8(X0, X1), X0)) & g_both(sK8(X0, X1), X1)) | ((g_true_only(sK8(X0, X1), X0) | g_both(sK8(X0, X1), X0)) & g_true_only(sK8(X0, X1), X1)))), introduced(choice_axiom, [])).
fof(f31, plain, ! [X0] : (! [X1] : (? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) | ! [X3] : (~ g_both(X3, X0) | ~ g_both(X3, X1))) | ~ sP0(X0)), inference(rectify, [], [f30])).
fof(f30, plain, ! [X3] : (! [X6] : (? [X7] : (((g_both(X7, X3) | g_false_only(X7, X3)) & g_false_only(X7, X6)) | ((g_true_only(X7, X3) | g_false_only(X7, X3)) & g_both(X7, X6)) | ((g_true_only(X7, X3) | g_both(X7, X3)) & g_true_only(X7, X6))) | ! [X8] : (~ g_both(X8, X3) | ~ g_both(X8, X6))) | ~ sP0(X3)), inference(nnf_transformation, [], [f13])).
fof(f339, plain, ! [X0] : (g_both(sK10(X0), X0) | g_true_only(sK10(X0), sK11) | g_true_only(sK10(X0), X0)), inference(resolution, [], [f321, f98])).
fof(f321, plain, ! [X4] : (~ g_false_only(sK10(X4), X4) | g_true_only(sK10(X4), sK11)), inference(resolution, [], [f230, f232])).
fof(f232, plain, ! [X9] : (~ g_true_only(X9, X9) | g_true_only(X9, sK11)), inference(subsumption_resolution, [], [f225, f75])).
fof(f75, plain, ! [X7] : (~ g_both(X7, sK11) | ~ g_true_only(X7, X7)), inference(cnf_transformation, [], [f38])).
fof(f225, plain, ! [X9] : (g_both(X9, sK11) | g_true_only(X9, sK11) | ~ g_true_only(X9, X9)), inference(resolution, [], [f98, f77])).
fof(f77, plain, ! [X7] : (~ g_false_only(X7, sK11) | ~ g_true_only(X7, X7)), inference(cnf_transformation, [], [f38])).
fof(f23977, plain, (~ spl12_253 | ~ spl12_238), inference(avatar_split_clause, [], [f23935, f23829, f23968])).
fof(f23968, plain, (spl12_253 <=> g_false_only(sK8(sK10(sK11), sK11), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_253])])).
fof(f23935, plain, (~ g_false_only(sK8(sK10(sK11), sK11), sK11) | ~ spl12_238), inference(resolution, [], [f23831, f81])).
fof(f23831, plain, (g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | ~ spl12_238), inference(avatar_component_clause, [], [f23829])).
fof(f23976, plain, (~ spl12_254 | ~ spl12_238), inference(avatar_split_clause, [], [f23934, f23829, f23972])).
fof(f23934, plain, (~ g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ spl12_238), inference(resolution, [], [f23831, f82])).
fof(f23975, plain, (spl12_253 | spl12_254 | ~ spl12_24 | ~ spl12_42 | spl12_43 | ~ spl12_238), inference(avatar_split_clause, [], [f23966, f23829, f424, f420, f204, f23972, f23968])).
fof(f23966, plain, (g_true_only(sK8(sK10(sK11), sK11), sK11) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_42 | spl12_43 | ~ spl12_238)), inference(subsumption_resolution, [], [f23965, f422])).
fof(f23965, plain, (g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | spl12_43 | ~ spl12_238)), inference(subsumption_resolution, [], [f23964, f425])).
fof(f23964, plain, (g_true_only(sK10(sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_238)), inference(duplicate_literal_removal, [], [f23933])).
fof(f23933, plain, (g_true_only(sK10(sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK11) | ~ g_both(sK10(sK11), sK10(sK11)) | g_true_only(sK10(sK11), sK11) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_24 | ~ spl12_238)), inference(resolution, [], [f23831, f15869])).
fof(f15869, plain, (! [X2, X3] : (~ g_both(sK8(X2, X3), X2) | g_true_only(sK10(X3), X3) | g_true_only(sK8(X2, X3), X3) | ~ g_both(sK10(X3), X2) | g_true_only(sK10(X3), sK11) | g_false_only(sK8(X2, X3), X3)) | ~ spl12_24), inference(subsumption_resolution, [], [f15868, f6005])).
fof(f15868, plain, ! [X2, X3] : (g_false_only(sK8(X2, X3), X3) | g_true_only(sK10(X3), X3) | g_true_only(sK8(X2, X3), X3) | ~ g_both(sK10(X3), X2) | g_true_only(sK10(X3), sK11) | ~ sP0(X2) | ~ g_both(sK8(X2, X3), X2)), inference(subsumption_resolution, [], [f15844, f302])).
fof(f15844, plain, ! [X2, X3] : (g_false_only(sK8(X2, X3), X3) | g_true_only(sK8(X2, X3), X2) | g_true_only(sK10(X3), X3) | g_true_only(sK8(X2, X3), X3) | ~ g_both(sK10(X3), X2) | g_true_only(sK10(X3), sK11) | ~ sP0(X2) | ~ g_both(sK8(X2, X3), X2)), inference(resolution, [], [f404, f303])).
fof(f404, plain, ! [X8, X7] : (g_false_only(sK8(X8, X7), X8) | g_false_only(sK8(X8, X7), X7) | g_true_only(sK8(X8, X7), X8) | g_true_only(sK10(X7), X7) | g_true_only(sK8(X8, X7), X7) | ~ g_both(sK10(X7), X8) | g_true_only(sK10(X7), sK11) | ~ sP0(X8)), inference(resolution, [], [f339, f67])).
fof(f67, plain, ! [X0, X3, X1] : (~ g_both(X3, X1) | g_true_only(sK8(X0, X1), X0) | g_false_only(sK8(X0, X1), X0) | g_true_only(sK8(X0, X1), X1) | ~ g_both(X3, X0) | g_false_only(sK8(X0, X1), X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f33])).
fof(f23836, plain, (spl12_238 | spl12_239 | ~ spl12_42 | spl12_43 | ~ spl12_112), inference(avatar_split_clause, [], [f23827, f3783, f424, f420, f23833, f23829])).
fof(f3783, plain, (spl12_112 <=> sP0(sK10(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_112])])).
fof(f23827, plain, (g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | (~ spl12_42 | spl12_43 | ~ spl12_112)), inference(subsumption_resolution, [], [f23826, f230])).
fof(f23826, plain, (g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_42 | spl12_43 | ~ spl12_112)), inference(subsumption_resolution, [], [f23825, f229])).
fof(f23825, plain, (g_both(sK8(sK10(sK11), sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | (~ spl12_42 | spl12_43 | ~ spl12_112)), inference(subsumption_resolution, [], [f23824, f3784])).
fof(f3784, plain, (sP0(sK10(sK11)) | ~ spl12_112), inference(avatar_component_clause, [], [f3783])).
fof(f23824, plain, (g_both(sK8(sK10(sK11), sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | ~ sP0(sK10(sK11)) | (~ spl12_42 | spl12_43)), inference(subsumption_resolution, [], [f23816, f425])).
fof(f23816, plain, (g_true_only(sK10(sK11), sK11) | g_both(sK8(sK10(sK11), sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | g_false_only(sK8(sK10(sK11), sK11), sK11) | ~ sP0(sK10(sK11)) | ~ spl12_42), inference(duplicate_literal_removal, [], [f23802])).
fof(f23802, plain, (g_true_only(sK10(sK11), sK11) | g_both(sK8(sK10(sK11), sK11), sK11) | g_true_only(sK8(sK10(sK11), sK11), sK10(sK11)) | g_both(sK8(sK10(sK11), sK11), sK10(sK11)) | g_true_only(sK10(sK11), sK11) | g_false_only(sK8(sK10(sK11), sK11), sK11) | ~ sP0(sK10(sK11)) | ~ spl12_42), inference(resolution, [], [f2934, f422])).
fof(f2934, plain, ! [X12, X11] : (~ g_both(sK10(X11), X12) | g_true_only(sK10(X11), sK11) | g_both(sK8(X12, sK11), sK11) | g_true_only(sK8(X12, sK11), X12) | g_both(sK8(X12, sK11), X12) | g_true_only(sK10(X11), X11) | g_false_only(sK8(X12, sK11), sK11) | ~ sP0(X12)), inference(resolution, [], [f400, f66])).
fof(f66, plain, ! [X0, X3, X1] : (~ g_both(X3, X1) | g_both(sK8(X0, X1), X1) | g_true_only(sK8(X0, X1), X0) | g_both(sK8(X0, X1), X0) | ~ g_both(X3, X0) | g_false_only(sK8(X0, X1), X1) | ~ sP0(X0)), inference(cnf_transformation, [], [f33])).
fof(f400, plain, ! [X1] : (g_both(sK10(X1), sK11) | g_true_only(sK10(X1), X1) | g_true_only(sK10(X1), sK11)), inference(resolution, [], [f339, f307])).
fof(f307, plain, ! [X10] : (~ g_both(sK10(X10), X10) | g_both(sK10(X10), sK11)), inference(resolution, [], [f229, f231])).
fof(f231, plain, ! [X8] : (~ g_both(X8, X8) | g_both(X8, sK11)), inference(subsumption_resolution, [], [f224, f73])).
fof(f73, plain, ! [X7] : (~ g_both(X7, X7) | ~ g_true_only(X7, sK11)), inference(cnf_transformation, [], [f38])).
fof(f224, plain, ! [X8] : (g_both(X8, sK11) | g_true_only(X8, sK11) | ~ g_both(X8, X8)), inference(resolution, [], [f98, f78])).
fof(f78, plain, ! [X7] : (~ g_false_only(X7, sK11) | ~ g_both(X7, X7)), inference(cnf_transformation, [], [f38])).
fof(f15604, plain, (spl12_40 | spl12_41 | ~ spl12_13), inference(avatar_split_clause, [], [f1232, f160, f393, f389])).
fof(f389, plain, (spl12_40 <=> g_true_only(sK6(sK10(sK5)), sK5)), introduced(avatar_definition, [new_symbols(naming, [spl12_40])])).
fof(f393, plain, (spl12_41 <=> g_both(sK6(sK10(sK5)), sK5)), introduced(avatar_definition, [new_symbols(naming, [spl12_41])])).
fof(f160, plain, (spl12_13 <=> ! [X1] : (g_false_only(sK6(X1), X1) | g_both(sK6(X1), sK5) | g_true_only(sK6(X1), sK5) | g_both(sK6(X1), X1))), introduced(avatar_definition, [new_symbols(naming, [spl12_13])])).
fof(f1232, plain, (g_both(sK6(sK10(sK5)), sK5) | g_true_only(sK6(sK10(sK5)), sK5) | ~ spl12_13), inference(duplicate_literal_removal, [], [f1231])).
fof(f1231, plain, (g_both(sK6(sK10(sK5)), sK5) | g_true_only(sK6(sK10(sK5)), sK5) | g_true_only(sK6(sK10(sK5)), sK5) | ~ spl12_13), inference(factoring, [], [f688])).
fof(f688, plain, (! [X0] : (g_both(sK6(sK10(X0)), sK5) | g_both(sK6(sK10(X0)), X0) | g_true_only(sK6(sK10(X0)), sK5) | g_true_only(sK6(sK10(X0)), X0)) | ~ spl12_13), inference(resolution, [], [f617, f98])).
fof(f617, plain, (! [X4] : (~ g_false_only(sK6(sK10(X4)), X4) | g_true_only(sK6(sK10(X4)), sK5) | g_both(sK6(sK10(X4)), sK5)) | ~ spl12_13), inference(subsumption_resolution, [], [f613, f81])).
fof(f613, plain, (! [X4] : (g_both(sK6(sK10(X4)), sK5) | g_true_only(sK6(sK10(X4)), sK5) | g_both(sK6(sK10(X4)), sK10(X4)) | ~ g_false_only(sK6(sK10(X4)), X4)) | ~ spl12_13), inference(resolution, [], [f161, f83])).
fof(f161, plain, (! [X1] : (g_false_only(sK6(X1), X1) | g_both(sK6(X1), sK5) | g_true_only(sK6(X1), sK5) | g_both(sK6(X1), X1)) | ~ spl12_13), inference(avatar_component_clause, [], [f160])).
fof(f10388, plain, (spl12_108 | ~ spl12_109), inference(avatar_contradiction_clause, [], [f10387])).
fof(f10387, plain, ($false | (spl12_108 | ~ spl12_109)), inference(subsumption_resolution, [], [f10386, f5164])).
fof(f5164, plain, (g_both(sK10(sK10(sK10(sK11))), sK11) | (spl12_108 | ~ spl12_109)), inference(subsumption_resolution, [], [f5160, f3303])).
fof(f3303, plain, (~ g_true_only(sK10(sK10(sK10(sK11))), sK11) | spl12_108), inference(avatar_component_clause, [], [f3302])).
fof(f3302, plain, (spl12_108 <=> g_true_only(sK10(sK10(sK10(sK11))), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_108])])).
fof(f5160, plain, (g_both(sK10(sK10(sK10(sK11))), sK11) | g_true_only(sK10(sK10(sK10(sK11))), sK11) | ~ spl12_109), inference(resolution, [], [f5106, f98])).
fof(f5106, plain, (~ g_false_only(sK10(sK10(sK10(sK11))), sK11) | ~ spl12_109), inference(resolution, [], [f5075, f230])).
fof(f5075, plain, (~ g_true_only(sK10(sK10(sK10(sK11))), sK10(sK11)) | ~ spl12_109), inference(resolution, [], [f3308, f80])).
fof(f3308, plain, (g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK11))) | ~ spl12_109), inference(avatar_component_clause, [], [f3306])).
fof(f3306, plain, (spl12_109 <=> g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK11)))), introduced(avatar_definition, [new_symbols(naming, [spl12_109])])).
fof(f10386, plain, (~ g_both(sK10(sK10(sK10(sK11))), sK11) | ~ spl12_109), inference(resolution, [], [f5076, f229])).
fof(f5076, plain, (~ g_both(sK10(sK10(sK10(sK11))), sK10(sK11)) | ~ spl12_109), inference(resolution, [], [f3308, f79])).
fof(f6098, plain, (spl12_112 | ~ spl12_24 | ~ spl12_42 | ~ spl12_47), inference(avatar_split_clause, [], [f6097, f461, f420, f204, f3783])).
fof(f461, plain, (spl12_47 <=> g_both(sK10(sK11), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_47])])).
fof(f6097, plain, (sP0(sK10(sK11)) | (~ spl12_24 | ~ spl12_42 | ~ spl12_47)), inference(subsumption_resolution, [], [f5996, f662])).
fof(f662, plain, (~ g_true_only(sK10(sK11), sK10(sK11)) | ~ spl12_47), inference(resolution, [], [f462, f75])).
fof(f462, plain, (g_both(sK10(sK11), sK11) | ~ spl12_47), inference(avatar_component_clause, [], [f461])).
fof(f5996, plain, (sP0(sK10(sK11)) | g_true_only(sK10(sK11), sK10(sK11)) | (~ spl12_24 | ~ spl12_42)), inference(resolution, [], [f205, f1376])).
fof(f1376, plain, (~ g_false_only(sK10(sK11), sK10(sK11)) | ~ spl12_42), inference(resolution, [], [f897, f96])).
fof(f96, plain, ! [X0, X1] : (~ g_true(X0, X1) | ~ g_false_only(X0, X1)), inference(cnf_transformation, [], [f44])).
fof(f44, plain, ! [X0, X1] : ((g_false_only(X0, X1) | g_true(X0, X1) | ~ g_false(X0, X1)) & ((~ g_true(X0, X1) & g_false(X0, X1)) | ~ g_false_only(X0, X1))), inference(flattening, [], [f43])).
fof(f43, plain, ! [X0, X1] : ((g_false_only(X0, X1) | (g_true(X0, X1) | ~ g_false(X0, X1))) & ((~ g_true(X0, X1) & g_false(X0, X1)) | ~ g_false_only(X0, X1))), inference(nnf_transformation, [], [f10])).
fof(f10, plain, ! [X0, X1] : (g_false_only(X0, X1) <=> (~ g_true(X0, X1) & g_false(X0, X1))), inference(rectify, [], [f5])).
fof(f5, plain, ! [X3, X4] : (g_false_only(X3, X4) <=> (~ g_true(X3, X4) & g_false(X3, X4))), file('/home/ubuntu/library/tptp/Problems/SYO/SYO605+1.p', false_only_g)).
fof(f897, plain, (g_true(sK10(sK11), sK10(sK11)) | ~ spl12_42), inference(resolution, [], [f422, f92])).
fof(f92, plain, ! [X0, X1] : (~ g_both(X0, X1) | g_true(X0, X1)), inference(cnf_transformation, [], [f42])).
fof(f42, plain, ! [X0, X1] : ((g_both(X0, X1) | ~ g_false(X0, X1) | ~ g_true(X0, X1)) & ((g_false(X0, X1) & g_true(X0, X1)) | ~ g_both(X0, X1))), inference(flattening, [], [f41])).
fof(f41, plain, ! [X0, X1] : ((g_both(X0, X1) | (~ g_false(X0, X1) | ~ g_true(X0, X1))) & ((g_false(X0, X1) & g_true(X0, X1)) | ~ g_both(X0, X1))), inference(nnf_transformation, [], [f9])).
fof(f9, plain, ! [X0, X1] : (g_both(X0, X1) <=> (g_false(X0, X1) & g_true(X0, X1))), inference(rectify, [], [f4])).
fof(f4, plain, ! [X3, X4] : (g_both(X3, X4) <=> (g_false(X3, X4) & g_true(X3, X4))), file('/home/ubuntu/library/tptp/Problems/SYO/SYO605+1.p', both_g)).
fof(f5819, plain, (~ spl12_55 | ~ spl12_122), inference(avatar_contradiction_clause, [], [f5818])).
fof(f5818, plain, ($false | (~ spl12_55 | ~ spl12_122)), inference(subsumption_resolution, [], [f5814, f684])).
fof(f684, plain, (g_both(sK7(sK11), sK7(sK11)) | ~ spl12_55), inference(avatar_component_clause, [], [f682])).
fof(f682, plain, (spl12_55 <=> g_both(sK7(sK11), sK7(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_55])])).
fof(f5814, plain, (~ g_both(sK7(sK11), sK7(sK11)) | ~ spl12_122), inference(resolution, [], [f5253, f303])).
fof(f5253, plain, (g_false_only(sK7(sK11), sK7(sK11)) | ~ spl12_122), inference(avatar_component_clause, [], [f5251])).
fof(f5251, plain, (spl12_122 <=> g_false_only(sK7(sK11), sK7(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_122])])).
fof(f5318, plain, (~ spl12_55 | ~ spl12_121), inference(avatar_contradiction_clause, [], [f5317])).
fof(f5317, plain, ($false | (~ spl12_55 | ~ spl12_121)), inference(subsumption_resolution, [], [f5313, f5222])).
fof(f5222, plain, (g_both(sK7(sK11), sK11) | ~ spl12_55), inference(resolution, [], [f684, f231])).
fof(f5313, plain, (~ g_both(sK7(sK11), sK11) | ~ spl12_121), inference(resolution, [], [f5249, f303])).
fof(f5249, plain, (g_false_only(sK7(sK11), sK11) | ~ spl12_121), inference(avatar_component_clause, [], [f5247])).
fof(f5247, plain, (spl12_121 <=> g_false_only(sK7(sK11), sK11)), introduced(avatar_definition, [new_symbols(naming, [spl12_121])])).
fof(f5254, plain, (spl12_92 | spl12_121 | spl12_122 | spl12_54 | ~ spl12_21 | ~ spl12_55), inference(avatar_split_clause, [], [f5245, f682, f192, f678, f5251, f5247, f2966])).
fof(f2966, plain, (spl12_92 <=> ! [X4] : ~ g_both(sK10(sK10(X4)), X4)), introduced(avatar_definition, [new_symbols(naming, [spl12_92])])).
fof(f678, plain, (spl12_54 <=> g_true_only(sK7(sK11), sK7(sK11))), introduced(avatar_definition, [new_symbols(naming, [spl12_54])])).
fof(f192, plain, (spl12_21 <=> ! [X0, X2] : (g_false_only(sK7(X0), X0) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | g_true_only(sK7(X0), X0) | g_true_only(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0)))), introduced(avatar_definition, [new_symbols(naming, [spl12_21])])).
fof(f5245, plain, (! [X5] : (g_true_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X5)), X5)) | (~ spl12_21 | ~ spl12_55)), inference(subsumption_resolution, [], [f5244, f338])).
fof(f338, plain, ! [X0] : (~ g_both(sK10(sK10(X0)), X0) | ~ g_true_only(sK10(sK10(X0)), sK11)), inference(resolution, [], [f308, f229])).
fof(f308, plain, ! [X11] : (~ g_both(sK10(X11), X11) | ~ g_true_only(sK10(X11), sK11)), inference(resolution, [], [f229, f73])).
fof(f5244, plain, (! [X5] : (g_true_only(sK10(sK10(X5)), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X5)), X5)) | (~ spl12_21 | ~ spl12_55)), inference(subsumption_resolution, [], [f5243, f79])).
fof(f5243, plain, (! [X5] : (g_true_only(sK10(sK10(X5)), sK10(X5)) | g_true_only(sK10(sK10(X5)), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X5)), X5)) | (~ spl12_21 | ~ spl12_55)), inference(subsumption_resolution, [], [f2929, f5223])).
fof(f5223, plain, (~ g_true_only(sK7(sK11), sK11) | ~ spl12_55), inference(resolution, [], [f684, f73])).
fof(f2929, plain, (! [X5] : (g_true_only(sK10(sK10(X5)), sK10(X5)) | g_true_only(sK10(sK10(X5)), sK11) | g_true_only(sK7(sK11), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X5)), X5)) | ~ spl12_21), inference(resolution, [], [f400, f1263])).
fof(f1263, plain, (! [X4, X3] : (~ g_both(sK10(sK10(X3)), X4) | g_true_only(sK7(X4), X4) | g_true_only(sK7(X4), sK7(X4)) | g_false_only(sK7(X4), sK7(X4)) | g_false_only(sK7(X4), X4) | ~ g_both(sK10(sK10(X3)), X3)) | ~ spl12_21), inference(resolution, [], [f1104, f229])).
fof(f1104, plain, (! [X0, X1] : (~ g_both(sK10(X1), X1) | ~ g_both(sK10(X1), X0) | g_true_only(sK7(X0), X0) | g_true_only(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), X0)) | ~ spl12_21), inference(resolution, [], [f193, f229])).
fof(f193, plain, (! [X2, X0] : (~ g_both(X2, X2) | g_false_only(sK7(X0), X0) | ~ g_both(X2, X0) | g_true_only(sK7(X0), X0) | g_true_only(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0))) | ~ spl12_21), inference(avatar_component_clause, [], [f192])).
fof(f5242, plain, (~ spl12_54 | ~ spl12_55), inference(avatar_split_clause, [], [f5231, f682, f678])).
fof(f5231, plain, (~ g_true_only(sK7(sK11), sK7(sK11)) | ~ spl12_55), inference(resolution, [], [f684, f302])).
fof(f5004, plain, (~ spl12_92 | ~ spl12_108 | spl12_109), inference(avatar_contradiction_clause, [], [f5003])).
fof(f5003, plain, ($false | (~ spl12_92 | ~ spl12_108 | spl12_109)), inference(subsumption_resolution, [], [f5000, f3304])).
fof(f3304, plain, (g_true_only(sK10(sK10(sK10(sK11))), sK11) | ~ spl12_108), inference(avatar_component_clause, [], [f3302])).
fof(f5000, plain, (~ g_true_only(sK10(sK10(sK10(sK11))), sK11) | (~ spl12_92 | spl12_109)), inference(resolution, [], [f4947, f80])).
fof(f4947, plain, (g_true_only(sK10(sK10(sK10(sK11))), sK10(sK11)) | (~ spl12_92 | spl12_109)), inference(subsumption_resolution, [], [f4946, f2967])).
fof(f2967, plain, (! [X4] : ~ g_both(sK10(sK10(X4)), X4) | ~ spl12_92), inference(avatar_component_clause, [], [f2966])).
fof(f4946, plain, (g_both(sK10(sK10(sK10(sK11))), sK10(sK11)) | g_true_only(sK10(sK10(sK10(sK11))), sK10(sK11)) | spl12_109), inference(resolution, [], [f4896, f98])).
fof(f4896, plain, (~ g_false_only(sK10(sK10(sK10(sK11))), sK10(sK11)) | spl12_109), inference(resolution, [], [f3307, f230])).
fof(f3307, plain, (~ g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK11))) | spl12_109), inference(avatar_component_clause, [], [f3306])).
fof(f4843, plain, (~ spl12_109 | ~ spl12_108), inference(avatar_split_clause, [], [f4838, f3302, f3306])).
fof(f4838, plain, (~ g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK11))) | ~ spl12_108), inference(resolution, [], [f3322, f80])).
fof(f3322, plain, (g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK10(sK11)))) | ~ spl12_108), inference(resolution, [], [f3304, f227])).
fof(f227, plain, ! [X1] : (~ g_true_only(X1, sK11) | g_true_only(X1, X1)), inference(subsumption_resolution, [], [f220, f73])).
fof(f220, plain, ! [X1] : (g_both(X1, X1) | g_true_only(X1, X1) | ~ g_true_only(X1, sK11)), inference(resolution, [], [f98, f74])).
fof(f74, plain, ! [X7] : (~ g_false_only(X7, X7) | ~ g_true_only(X7, sK11)), inference(cnf_transformation, [], [f38])).
fof(f3309, plain, (spl12_108 | spl12_109 | ~ spl12_92), inference(avatar_split_clause, [], [f3299, f2966, f3306, f3302])).
fof(f3299, plain, (g_true_only(sK10(sK10(sK10(sK11))), sK10(sK10(sK11))) | g_true_only(sK10(sK10(sK10(sK11))), sK11) | ~ spl12_92), inference(resolution, [], [f3008, f400])).
fof(f3008, plain, (! [X0] : ~ g_both(sK10(sK10(sK10(X0))), X0) | ~ spl12_92), inference(resolution, [], [f2967, f229])).
fof(f2968, plain, (spl12_92 | spl12_55 | ~ spl12_22 | ~ spl12_54), inference(avatar_split_clause, [], [f2964, f678, f196, f682, f2966])).
fof(f196, plain, (spl12_22 <=> ! [X0, X2] : (g_false_only(sK7(X0), X0) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | g_both(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), X0))), introduced(avatar_definition, [new_symbols(naming, [spl12_22])])).
fof(f2964, plain, (! [X4] : (g_both(sK7(sK11), sK7(sK11)) | ~ g_both(sK10(sK10(X4)), X4)) | (~ spl12_22 | ~ spl12_54)), inference(subsumption_resolution, [], [f2963, f338])).
fof(f2963, plain, (! [X4] : (g_true_only(sK10(sK10(X4)), sK11) | g_both(sK7(sK11), sK7(sK11)) | ~ g_both(sK10(sK10(X4)), X4)) | (~ spl12_22 | ~ spl12_54)), inference(subsumption_resolution, [], [f2962, f79])).
fof(f2962, plain, (! [X4] : (g_true_only(sK10(sK10(X4)), sK10(X4)) | g_true_only(sK10(sK10(X4)), sK11) | g_both(sK7(sK11), sK7(sK11)) | ~ g_both(sK10(sK10(X4)), X4)) | (~ spl12_22 | ~ spl12_54)), inference(subsumption_resolution, [], [f2961, f226])).
fof(f226, plain, ! [X0] : (~ g_both(X0, sK11) | g_both(X0, X0)), inference(subsumption_resolution, [], [f219, f75])).
fof(f219, plain, ! [X0] : (g_both(X0, X0) | g_true_only(X0, X0) | ~ g_both(X0, sK11)), inference(resolution, [], [f98, f76])).
fof(f76, plain, ! [X7] : (~ g_false_only(X7, X7) | ~ g_both(X7, sK11)), inference(cnf_transformation, [], [f38])).
fof(f2961, plain, (! [X4] : (g_true_only(sK10(sK10(X4)), sK10(X4)) | g_true_only(sK10(sK10(X4)), sK11) | g_both(sK7(sK11), sK7(sK11)) | g_both(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X4)), X4)) | (~ spl12_22 | ~ spl12_54)), inference(subsumption_resolution, [], [f2960, f1541])).
fof(f1541, plain, (~ g_false_only(sK7(sK11), sK11) | ~ spl12_54), inference(resolution, [], [f1501, f96])).
fof(f1501, plain, (g_true(sK7(sK11), sK11) | ~ spl12_54), inference(resolution, [], [f1480, f89])).
fof(f89, plain, ! [X0, X1] : (~ g_true_only(X0, X1) | g_true(X0, X1)), inference(cnf_transformation, [], [f40])).
fof(f40, plain, ! [X0, X1] : ((g_true_only(X0, X1) | g_false(X0, X1) | ~ g_true(X0, X1)) & ((~ g_false(X0, X1) & g_true(X0, X1)) | ~ g_true_only(X0, X1))), inference(flattening, [], [f39])).
fof(f39, plain, ! [X0, X1] : ((g_true_only(X0, X1) | (g_false(X0, X1) | ~ g_true(X0, X1))) & ((~ g_false(X0, X1) & g_true(X0, X1)) | ~ g_true_only(X0, X1))), inference(nnf_transformation, [], [f8])).
fof(f8, plain, ! [X0, X1] : (g_true_only(X0, X1) <=> (~ g_false(X0, X1) & g_true(X0, X1))), inference(rectify, [], [f3])).
fof(f3, plain, ! [X3, X4] : (g_true_only(X3, X4) <=> (~ g_false(X3, X4) & g_true(X3, X4))), file('/home/ubuntu/library/tptp/Problems/SYO/SYO605+1.p', true_only_g)).
fof(f1480, plain, (g_true_only(sK7(sK11), sK11) | ~ spl12_54), inference(resolution, [], [f680, f232])).
fof(f680, plain, (g_true_only(sK7(sK11), sK7(sK11)) | ~ spl12_54), inference(avatar_component_clause, [], [f678])).
fof(f2960, plain, (! [X4] : (g_true_only(sK10(sK10(X4)), sK10(X4)) | g_true_only(sK10(sK10(X4)), sK11) | g_both(sK7(sK11), sK7(sK11)) | g_both(sK7(sK11), sK11) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X4)), X4)) | (~ spl12_22 | ~ spl12_54)), inference(subsumption_resolution, [], [f2928, f1625])).
fof(f1625, plain, (~ g_false_only(sK7(sK11), sK7(sK11)) | ~ spl12_54), inference(resolution, [], [f1481, f96])).
fof(f1481, plain, (g_true(sK7(sK11), sK7(sK11)) | ~ spl12_54), inference(resolution, [], [f680, f89])).
fof(f2928, plain, (! [X4] : (g_true_only(sK10(sK10(X4)), sK10(X4)) | g_true_only(sK10(sK10(X4)), sK11) | g_both(sK7(sK11), sK7(sK11)) | g_false_only(sK7(sK11), sK7(sK11)) | g_both(sK7(sK11), sK11) | g_false_only(sK7(sK11), sK11) | ~ g_both(sK10(sK10(X4)), X4)) | ~ spl12_22), inference(resolution, [], [f400, f1278])).
fof(f1278, plain, (! [X4, X3] : (~ g_both(sK10(sK10(X3)), X4) | g_both(sK7(X4), sK7(X4)) | g_false_only(sK7(X4), sK7(X4)) | g_both(sK7(X4), X4) | g_false_only(sK7(X4), X4) | ~ g_both(sK10(sK10(X3)), X3)) | ~ spl12_22), inference(resolution, [], [f1107, f229])).
fof(f1107, plain, (! [X0, X1] : (~ g_both(sK10(X1), X1) | ~ g_both(sK10(X1), X0) | g_both(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), X0) | g_false_only(sK7(X0), X0)) | ~ spl12_22), inference(resolution, [], [f197, f229])).
fof(f197, plain, (! [X2, X0] : (~ g_both(X2, X2) | g_false_only(sK7(X0), X0) | ~ g_both(X2, X0) | g_both(sK7(X0), sK7(X0)) | g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), X0)) | ~ spl12_22), inference(avatar_component_clause, [], [f196])).
fof(f997, plain, (spl12_51 | ~ spl12_2), inference(avatar_split_clause, [], [f996, f116, f525])).
fof(f525, plain, (spl12_51 <=> ! [X12] : ~ g_both(X12, X12)), introduced(avatar_definition, [new_symbols(naming, [spl12_51])])).
fof(f116, plain, (spl12_2 <=> ! [X1] : (g_false_only(X1, X1) | g_true_only(X1, X1))), introduced(avatar_definition, [new_symbols(naming, [spl12_2])])).
fof(f996, plain, (! [X3] : ~ g_both(X3, X3) | ~ spl12_2), inference(subsumption_resolution, [], [f988, f302])).
fof(f988, plain, (! [X3] : (g_true_only(X3, X3) | ~ g_both(X3, X3)) | ~ spl12_2), inference(resolution, [], [f117, f303])).
fof(f117, plain, (! [X1] : (g_false_only(X1, X1) | g_true_only(X1, X1)) | ~ spl12_2), inference(avatar_component_clause, [], [f116])).
fof(f982, plain, (~ spl12_42 | ~ spl12_51), inference(avatar_contradiction_clause, [], [f981])).
fof(f981, plain, ($false | (~ spl12_42 | ~ spl12_51)), inference(resolution, [], [f526, f422])).
fof(f526, plain, (! [X12] : ~ g_both(X12, X12) | ~ spl12_51), inference(avatar_component_clause, [], [f525])).
fof(f979, plain, (spl12_57 | ~ spl12_12 | spl12_40 | ~ spl12_41), inference(avatar_split_clause, [], [f978, f393, f389, f156, f815])).
fof(f815, plain, (spl12_57 <=> g_false_only(sK6(sK10(sK5)), sK5)), introduced(avatar_definition, [new_symbols(naming, [spl12_57])])).
fof(f156, plain, (spl12_12 <=> ! [X1] : (g_false_only(sK6(X1), X1) | g_true_only(sK6(X1), X1) | g_false_only(sK6(X1), sK5) | g_true_only(sK6(X1), sK5))), introduced(avatar_definition, [new_symbols(naming, [spl12_12])])).
fof(f978, plain, (g_false_only(sK6(sK10(sK5)), sK5) | (~ spl12_12 | spl12_40 | ~ spl12_41)), inference(subsumption_resolution, [], [f955, f390])).
fof(f390, plain, (~ g_true_only(sK6(sK10(sK5)), sK5) | spl12_40), inference(avatar_component_clause, [], [f389])).
fof(f955, plain, (g_true_only(sK6(sK10(sK5)), sK5) | g_false_only(sK6(sK10(sK5)), sK5) | (~ spl12_12 | ~ spl12_41)), inference(resolution, [], [f395, f600])).
fof(f600, plain, (! [X3] : (~ g_both(sK6(sK10(X3)), X3) | g_true_only(sK6(sK10(X3)), sK5) | g_false_only(sK6(sK10(X3)), sK5)) | ~ spl12_12), inference(subsumption_resolution, [], [f592, f79])).
fof(f592, plain, (! [X3] : (g_false_only(sK6(sK10(X3)), sK5) | g_true_only(sK6(sK10(X3)), sK10(X3)) | g_true_only(sK6(sK10(X3)), sK5) | ~ g_both(sK6(sK10(X3)), X3)) | ~ spl12_12), inference(resolution, [], [f157, f84])).
fof(f157, plain, (! [X1] : (g_false_only(sK6(X1), sK5) | g_false_only(sK6(X1), X1) | g_true_only(sK6(X1), X1) | g_true_only(sK6(X1), sK5)) | ~ spl12_12), inference(avatar_component_clause, [], [f156])).
fof(f395, plain, (g_both(sK6(sK10(sK5)), sK5) | ~ spl12_41), inference(avatar_component_clause, [], [f393])).
fof(f977, plain, (~ spl12_41 | ~ spl12_57), inference(avatar_contradiction_clause, [], [f976])).
fof(f976, plain, ($false | (~ spl12_41 | ~ spl12_57)), inference(subsumption_resolution, [], [f974, f395])).
fof(f974, plain, (~ g_both(sK6(sK10(sK5)), sK5) | ~ spl12_57), inference(resolution, [], [f817, f303])).
fof(f817, plain, (g_false_only(sK6(sK10(sK5)), sK5) | ~ spl12_57), inference(avatar_component_clause, [], [f815])).
fof(f930, plain, (~ spl12_40 | ~ spl12_57), inference(avatar_contradiction_clause, [], [f929])).
fof(f929, plain, ($false | (~ spl12_40 | ~ spl12_57)), inference(subsumption_resolution, [], [f918, f817])).
fof(f918, plain, (~ g_false_only(sK6(sK10(sK5)), sK5) | ~ spl12_40), inference(resolution, [], [f813, f96])).
fof(f813, plain, (g_true(sK6(sK10(sK5)), sK5) | ~ spl12_40), inference(resolution, [], [f391, f89])).
fof(f391, plain, (g_true_only(sK6(sK10(sK5)), sK5) | ~ spl12_40), inference(avatar_component_clause, [], [f389])).
fof(f838, plain, (~ spl12_40 | ~ spl12_41), inference(avatar_contradiction_clause, [], [f837])).
fof(f837, plain, ($false | (~ spl12_40 | ~ spl12_41)), inference(subsumption_resolution, [], [f833, f391])).
fof(f833, plain, (~ g_true_only(sK6(sK10(sK5)), sK5) | ~ spl12_41), inference(resolution, [], [f395, f302])).
fof(f818, plain, (spl12_57 | spl12_41 | ~ spl12_10 | ~ spl12_40), inference(avatar_split_clause, [], [f812, f389, f148, f393, f815])).
fof(f148, plain, (spl12_10 <=> ! [X1] : (g_both(sK6(X1), sK5) | g_true_only(sK6(X1), X1) | g_both(sK6(X1), X1) | g_false_only(sK6(X1), sK5))), introduced(avatar_definition, [new_symbols(naming, [spl12_10])])).
fof(f812, plain, (g_both(sK6(sK10(sK5)), sK5) | g_false_only(sK6(sK10(sK5)), sK5) | (~ spl12_10 | ~ spl12_40)), inference(resolution, [], [f391, f581])).
fof(f581, plain, (! [X19] : (~ g_true_only(sK6(sK10(X19)), X19) | g_both(sK6(sK10(X19)), sK5) | g_false_only(sK6(sK10(X19)), sK5)) | ~ spl12_10), inference(subsumption_resolution, [], [f575, f80])).
fof(f575, plain, (! [X19] : (g_false_only(sK6(sK10(X19)), sK5) | g_true_only(sK6(sK10(X19)), sK10(X19)) | g_both(sK6(sK10(X19)), sK5) | ~ g_true_only(sK6(sK10(X19)), X19)) | ~ spl12_10), inference(resolution, [], [f149, f82])).
fof(f149, plain, (! [X1] : (g_false_only(sK6(X1), sK5) | g_both(sK6(X1), X1) | g_true_only(sK6(X1), X1) | g_both(sK6(X1), sK5)) | ~ spl12_10), inference(avatar_component_clause, [], [f148])).
fof(f685, plain, (spl12_54 | spl12_55 | ~ spl12_19 | ~ spl12_47), inference(avatar_split_clause, [], [f676, f461, f184, f682, f678])).
fof(f184, plain, (spl12_19 <=> ! [X0, X2] : (g_both(sK7(X0), sK7(X0)) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | g_true_only(sK7(X0), X0) | g_both(sK7(X0), X0) | g_true_only(sK7(X0), sK7(X0)))), introduced(avatar_definition, [new_symbols(naming, [spl12_19])])).
fof(f676, plain, (g_both(sK7(sK11), sK7(sK11)) | g_true_only(sK7(sK11), sK7(sK11)) | (~ spl12_19 | ~ spl12_47)), inference(subsumption_resolution, [], [f675, f227])).
fof(f675, plain, (g_both(sK7(sK11), sK7(sK11)) | g_true_only(sK7(sK11), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | (~ spl12_19 | ~ spl12_47)), inference(subsumption_resolution, [], [f674, f226])).
fof(f674, plain, (g_both(sK7(sK11), sK7(sK11)) | g_true_only(sK7(sK11), sK11) | g_both(sK7(sK11), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | (~ spl12_19 | ~ spl12_47)), inference(subsumption_resolution, [], [f660, f462])).
fof(f660, plain, (g_both(sK7(sK11), sK7(sK11)) | ~ g_both(sK10(sK11), sK11) | g_true_only(sK7(sK11), sK11) | g_both(sK7(sK11), sK11) | g_true_only(sK7(sK11), sK7(sK11)) | (~ spl12_19 | ~ spl12_47)), inference(resolution, [], [f462, f306])).
fof(f306, plain, (! [X8, X9] : (~ g_both(sK10(X8), X9) | g_both(sK7(X9), sK7(X9)) | ~ g_both(sK10(X8), X8) | g_true_only(sK7(X9), X9) | g_both(sK7(X9), X9) | g_true_only(sK7(X9), sK7(X9))) | ~ spl12_19), inference(resolution, [], [f229, f185])).
fof(f185, plain, (! [X2, X0] : (~ g_both(X2, X2) | g_both(sK7(X0), sK7(X0)) | ~ g_both(X2, X0) | g_true_only(sK7(X0), X0) | g_both(sK7(X0), X0) | g_true_only(sK7(X0), sK7(X0))) | ~ spl12_19), inference(avatar_component_clause, [], [f184])).
fof(f509, plain, ~ spl12_43, inference(avatar_contradiction_clause, [], [f508])).
fof(f508, plain, ($false | ~ spl12_43), inference(subsumption_resolution, [], [f504, f426])).
fof(f426, plain, (g_true_only(sK10(sK11), sK11) | ~ spl12_43), inference(avatar_component_clause, [], [f424])).
fof(f504, plain, (~ g_true_only(sK10(sK11), sK11) | ~ spl12_43), inference(resolution, [], [f484, f80])).
fof(f484, plain, (g_true_only(sK10(sK11), sK10(sK11)) | ~ spl12_43), inference(resolution, [], [f426, f227])).
fof(f479, plain, (spl12_47 | ~ spl12_42), inference(avatar_split_clause, [], [f438, f420, f461])).
fof(f438, plain, (g_both(sK10(sK11), sK11) | ~ spl12_42), inference(resolution, [], [f422, f231])).
fof(f427, plain, (spl12_42 | spl12_43), inference(avatar_split_clause, [], [f415, f424, f420])).
fof(f415, plain, (g_true_only(sK10(sK11), sK11) | g_both(sK10(sK11), sK10(sK11))), inference(duplicate_literal_removal, [], [f412])).
fof(f412, plain, (g_true_only(sK10(sK11), sK11) | g_true_only(sK10(sK11), sK11) | g_both(sK10(sK11), sK10(sK11))), inference(resolution, [], [f339, f226])).
fof(f206, plain, (spl12_15 | spl12_1 | spl12_24 | spl12_6), inference(avatar_split_clause, [], [f88, f132, f204, f112, f168])).
fof(f168, plain, (spl12_15 <=> sP1), introduced(avatar_definition, [new_symbols(naming, [spl12_15])])).
fof(f112, plain, (spl12_1 <=> sP3), introduced(avatar_definition, [new_symbols(naming, [spl12_1])])).
fof(f132, plain, (spl12_6 <=> sP2), introduced(avatar_definition, [new_symbols(naming, [spl12_6])])).
fof(f88, plain, ! [X2, X0] : (sP2 | g_true_only(X2, X0) | g_false_only(X2, X0) | sP0(X0) | sP3 | sP1), inference(cnf_transformation, [], [f38])).
fof(f198, plain, (~ spl12_15 | spl12_22), inference(avatar_split_clause, [], [f58, f196, f168])).
fof(f58, plain, ! [X2, X0] : (g_false_only(sK7(X0), X0) | g_both(sK7(X0), X0) | g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), sK7(X0)) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | ~ sP1), inference(cnf_transformation, [], [f29])).
fof(f29, plain, (! [X0] : ((((g_both(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0))) & g_false_only(sK7(X0), X0)) | ((g_false_only(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0))) & g_both(sK7(X0), X0)) | ((g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), sK7(X0))) & g_true_only(sK7(X0), X0))) | ! [X2] : (~ g_both(X2, X2) | ~ g_both(X2, X0))) | ~ sP1), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7])], [f27, f28])).
fof(f28, plain, ! [X0] : (? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0))) => (((g_both(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0))) & g_false_only(sK7(X0), X0)) | ((g_false_only(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0))) & g_both(sK7(X0), X0)) | ((g_false_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), sK7(X0))) & g_true_only(sK7(X0), X0)))), introduced(choice_axiom, [])).
fof(f27, plain, (! [X0] : (? [X1] : (((g_both(X1, X1) | g_true_only(X1, X1)) & g_false_only(X1, X0)) | ((g_false_only(X1, X1) | g_true_only(X1, X1)) & g_both(X1, X0)) | ((g_false_only(X1, X1) | g_both(X1, X1)) & g_true_only(X1, X0))) | ! [X2] : (~ g_both(X2, X2) | ~ g_both(X2, X0))) | ~ sP1), inference(rectify, [], [f26])).
fof(f26, plain, (! [X11] : (? [X12] : (((g_both(X12, X12) | g_true_only(X12, X12)) & g_false_only(X12, X11)) | ((g_false_only(X12, X12) | g_true_only(X12, X12)) & g_both(X12, X11)) | ((g_false_only(X12, X12) | g_both(X12, X12)) & g_true_only(X12, X11))) | ! [X13] : (~ g_both(X13, X13) | ~ g_both(X13, X11))) | ~ sP1), inference(nnf_transformation, [], [f14])).
fof(f194, plain, (~ spl12_15 | spl12_21), inference(avatar_split_clause, [], [f59, f192, f168])).
fof(f59, plain, ! [X2, X0] : (g_false_only(sK7(X0), X0) | g_false_only(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), X0) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | ~ sP1), inference(cnf_transformation, [], [f29])).
fof(f186, plain, (~ spl12_15 | spl12_19), inference(avatar_split_clause, [], [f61, f184, f168])).
fof(f61, plain, ! [X2, X0] : (g_both(sK7(X0), sK7(X0)) | g_true_only(sK7(X0), sK7(X0)) | g_both(sK7(X0), X0) | g_true_only(sK7(X0), X0) | ~ g_both(X2, X2) | ~ g_both(X2, X0) | ~ sP1), inference(cnf_transformation, [], [f29])).
fof(f162, plain, (~ spl12_6 | spl12_13), inference(avatar_split_clause, [], [f50, f160, f132])).
fof(f50, plain, ! [X1] : (g_false_only(sK6(X1), X1) | g_both(sK6(X1), X1) | g_true_only(sK6(X1), sK5) | g_both(sK6(X1), sK5) | ~ sP2), inference(cnf_transformation, [], [f25])).
fof(f25, plain, (! [X1] : (((g_both(sK6(X1), sK5) | g_false_only(sK6(X1), sK5)) & g_false_only(sK6(X1), X1)) | ((g_true_only(sK6(X1), sK5) | g_false_only(sK6(X1), sK5)) & g_both(sK6(X1), X1)) | ((g_true_only(sK6(X1), sK5) | g_both(sK6(X1), sK5)) & g_true_only(sK6(X1), X1))) | ~ sP2), inference(skolemisation, [status(esa), new_symbols(skolem, [sK5, sK6])], [f22, f24, f23])).
fof(f23, plain, (? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) => ! [X1] : ? [X2] : (((g_both(X2, sK5) | g_false_only(X2, sK5)) & g_false_only(X2, X1)) | ((g_true_only(X2, sK5) | g_false_only(X2, sK5)) & g_both(X2, X1)) | ((g_true_only(X2, sK5) | g_both(X2, sK5)) & g_true_only(X2, X1)))), introduced(choice_axiom, [])).
fof(f24, plain, ! [X1] : (? [X2] : (((g_both(X2, sK5) | g_false_only(X2, sK5)) & g_false_only(X2, X1)) | ((g_true_only(X2, sK5) | g_false_only(X2, sK5)) & g_both(X2, X1)) | ((g_true_only(X2, sK5) | g_both(X2, sK5)) & g_true_only(X2, X1))) => (((g_both(sK6(X1), sK5) | g_false_only(sK6(X1), sK5)) & g_false_only(sK6(X1), X1)) | ((g_true_only(sK6(X1), sK5) | g_false_only(sK6(X1), sK5)) & g_both(sK6(X1), X1)) | ((g_true_only(sK6(X1), sK5) | g_both(sK6(X1), sK5)) & g_true_only(sK6(X1), X1)))), introduced(choice_axiom, [])).
fof(f22, plain, (? [X0] : ! [X1] : ? [X2] : (((g_both(X2, X0) | g_false_only(X2, X0)) & g_false_only(X2, X1)) | ((g_true_only(X2, X0) | g_false_only(X2, X0)) & g_both(X2, X1)) | ((g_true_only(X2, X0) | g_both(X2, X0)) & g_true_only(X2, X1))) | ~ sP2), inference(nnf_transformation, [], [f15])).
fof(f158, plain, (~ spl12_6 | spl12_12), inference(avatar_split_clause, [], [f51, f156, f132])).
fof(f51, plain, ! [X1] : (g_false_only(sK6(X1), X1) | g_true_only(sK6(X1), sK5) | g_false_only(sK6(X1), sK5) | g_true_only(sK6(X1), X1) | ~ sP2), inference(cnf_transformation, [], [f25])).
fof(f150, plain, (~ spl12_6 | spl12_10), inference(avatar_split_clause, [], [f53, f148, f132])).
fof(f53, plain, ! [X1] : (g_both(sK6(X1), sK5) | g_false_only(sK6(X1), sK5) | g_both(sK6(X1), X1) | g_true_only(sK6(X1), X1) | ~ sP2), inference(cnf_transformation, [], [f25])).
fof(f118, plain, (~ spl12_1 | spl12_2), inference(avatar_split_clause, [], [f48, f116, f112])).
fof(f48, plain, ! [X1] : (g_false_only(X1, X1) | g_true_only(X1, X1) | ~ sP3), inference(cnf_transformation, [], [f21])).
fof(f21, plain, (! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, sK4)) | (g_true_only(X1, X1) & g_true_only(X1, sK4))) | ~ sP3), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4])], [f19, f20])).
fof(f20, plain, (? [X0] : ! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, X0)) | (g_true_only(X1, X1) & g_true_only(X1, X0))) => ! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, sK4)) | (g_true_only(X1, X1) & g_true_only(X1, sK4)))), introduced(choice_axiom, [])).
fof(f19, plain, (? [X0] : ! [X1] : ((g_false_only(X1, X1) & g_false_only(X1, X0)) | (g_true_only(X1, X1) & g_true_only(X1, X0))) | ~ sP3), inference(rectify, [], [f18])).
fof(f18, plain, (? [X9] : ! [X10] : ((g_false_only(X10, X10) & g_false_only(X10, X9)) | (g_true_only(X10, X10) & g_true_only(X10, X9))) | ~ sP3), inference(nnf_transformation, [], [f16])).