fof(f970547, plain, $false, inference(avatar_sat_refutation, [], [f12795, f12818, f12828, f12920, f13069, f132324, f132586, f970546])).
fof(f970546, plain, ~ spl31_3653, inference(avatar_contradiction_clause, [], [f970545])).
fof(f970545, plain, ($false | ~ spl31_3653), inference(subsumption_resolution, [], [f970544, f363])).
fof(f363, plain, v4_pre_topc(sK6, sK4), inference(cnf_transformation, [], [f299])).
fof(f299, plain, (((~ v4_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5) & v4_pre_topc(sK6, sK4) & m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(sK5, sK4) & v2_tsp_2(sK5, sK4) & ~ v3_struct_0(sK5)) & l1_pre_topc(sK4) & v2_pre_topc(sK4) & ~ v3_struct_0(sK4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6])], [f155, f298, f297, f296])).
fof(f296, plain, (? [X0] : (? [X1] : (? [X2] : (~ v4_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v4_pre_topc(X2, X0) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => (? [X1] : (? [X2] : (~ v4_pre_topc(k4_pre_topc(sK4, X1, k4_tsp_2(sK4, X1), X2), X1) & v4_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(X1, sK4) & v2_tsp_2(X1, sK4) & ~ v3_struct_0(X1)) & l1_pre_topc(sK4) & v2_pre_topc(sK4) & ~ v3_struct_0(sK4))), introduced(choice_axiom, [])).
fof(f297, plain, (? [X1] : (? [X2] : (~ v4_pre_topc(k4_pre_topc(sK4, X1, k4_tsp_2(sK4, X1), X2), X1) & v4_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(X1, sK4) & v2_tsp_2(X1, sK4) & ~ v3_struct_0(X1)) => (? [X2] : (~ v4_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X2), sK5) & v4_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(sK5, sK4) & v2_tsp_2(sK5, sK4) & ~ v3_struct_0(sK5))), introduced(choice_axiom, [])).
fof(f298, plain, (? [X2] : (~ v4_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X2), sK5) & v4_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) => (~ v4_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5) & v4_pre_topc(sK6, sK4) & m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4))))), introduced(choice_axiom, [])).
fof(f155, plain, ? [X0] : (? [X1] : (? [X2] : (~ v4_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v4_pre_topc(X2, X0) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f154])).
fof(f154, plain, ? [X0] : (? [X1] : (? [X2] : ((~ v4_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v4_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & (m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X2, X0) => v4_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X2, X0) => v4_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', t35_tsp_2)).
fof(f970544, plain, (~ v4_pre_topc(sK6, sK4) | ~ spl31_3653), inference(subsumption_resolution, [], [f970440, f2261])).
fof(f2261, plain, sP1(sK5, sK4), inference(subsumption_resolution, [], [f2256, f358])).
fof(f358, plain, l1_pre_topc(sK4), inference(cnf_transformation, [], [f299])).
fof(f2256, plain, (sP1(sK5, sK4) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f1162, f361])).
fof(f361, plain, m2_tsp_1(sK5, sK4), inference(cnf_transformation, [], [f299])).
fof(f1162, plain, ! [X0, X1] : (~ m2_tsp_1(X0, X1) | sP1(X0, X1) | ~ l1_pre_topc(X1)), inference(subsumption_resolution, [], [f1161, f499])).
fof(f499, plain, ! [X0, X1] : (~ m2_tsp_1(X1, X0) | l1_pre_topc(X1) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0] : (! [X1] : (l1_pre_topc(X1) | ~ m2_tsp_1(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f70])).
fof(f70, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m2_tsp_1(X1, X0) => l1_pre_topc(X1))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', dt_m2_tsp_1)).
fof(f1161, plain, ! [X0, X1] : (~ m2_tsp_1(X0, X1) | sP1(X0, X1) | ~ l1_pre_topc(X0) | ~ l1_pre_topc(X1)), inference(resolution, [], [f472, f484])).
fof(f484, plain, ! [X0, X1] : (sP2(X0, X1) | ~ l1_pre_topc(X1) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f293])).
fof(f293, plain, ! [X0] : (! [X1] : (sP2(X0, X1) | ~ l1_pre_topc(X1)) | ~ l1_pre_topc(X0)), inference(definition_folding, [], [f221, e292, e291])).
fof(f291, plain, ! [X1, X0] : (sP1(X1, X0) <=> (! [X2] : ((v4_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))), inference(usedef, [], [e291])).
fof(e291, plain, ! [X1, X0] : (sP1(X1, X0) <=> (! [X2] : ((v4_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f292, plain, ! [X0, X1] : ((m2_tsp_1(X1, X0) <=> sP1(X1, X0)) | ~ sP2(X0, X1)), inference(usedef, [], [e292])).
fof(e292, plain, ! [X0, X1] : (sP2(X0, X1) <=> (m2_tsp_1(X1, X0) <=> sP1(X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f221, plain, ! [X0] : (! [X1] : ((m2_tsp_1(X1, X0) <=> (! [X2] : ((v4_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))) | ~ l1_pre_topc(X1)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (l1_pre_topc(X1) => (m2_tsp_1(X1, X0) <=> (! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1))) => (v4_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', d2_tsp_1)).
fof(f472, plain, ! [X0, X1] : (~ sP2(X0, X1) | ~ m2_tsp_1(X1, X0) | sP1(X1, X0)), inference(cnf_transformation, [], [f306])).
fof(f306, plain, ! [X0, X1] : (((m2_tsp_1(X1, X0) | ~ sP1(X1, X0)) & (sP1(X1, X0) | ~ m2_tsp_1(X1, X0))) | ~ sP2(X0, X1)), inference(nnf_transformation, [], [f292])).
fof(f970440, plain, (~ sP1(sK5, sK4) | ~ v4_pre_topc(sK6, sK4) | ~ spl31_3653), inference(resolution, [], [f132574, f362])).
fof(f362, plain, m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4))), inference(cnf_transformation, [], [f299])).
fof(f132574, plain, (! [X0] : (~ m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(X0))) | ~ sP1(sK5, X0) | ~ v4_pre_topc(sK6, X0)) | ~ spl31_3653), inference(avatar_component_clause, [], [f132573])).
fof(f132573, plain, (spl31_3653 <=> ! [X0] : (~ v4_pre_topc(sK6, X0) | ~ sP1(sK5, X0) | ~ m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(X0))))), introduced(avatar_definition, [new_symbols(naming, [spl31_3653])])).
fof(f132586, plain, (spl31_2327 | spl31_3653 | ~ spl31_202 | ~ spl31_204 | ~ spl31_205 | ~ spl31_208), inference(avatar_split_clause, [], [f113202, f12995, f12784, f12780, f12772, f132573, f112755])).
fof(f112755, plain, (spl31_2327 <=> v4_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5)), introduced(avatar_definition, [new_symbols(naming, [spl31_2327])])).
fof(f12772, plain, (spl31_202 <=> m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl31_202])])).
fof(f12780, plain, (spl31_204 <=> v1_funct_1(k4_tsp_2(sK4, sK5))), introduced(avatar_definition, [new_symbols(naming, [spl31_204])])).
fof(f12784, plain, (spl31_205 <=> v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl31_205])])).
fof(f12995, plain, (spl31_208 <=> m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl31_208])])).
fof(f113202, plain, (! [X0] : (~ m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(X0))) | v4_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5) | ~ sP1(sK5, X0) | ~ v4_pre_topc(sK6, X0)) | (~ spl31_202 | ~ spl31_204 | ~ spl31_205 | ~ spl31_208)), inference(resolution, [], [f112763, f2000])).
fof(f2000, plain, ! [X2, X0, X1] : (~ r1_tarski(k3_xboole_0(X0, u1_struct_0(X2)), u1_struct_0(X2)) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | v4_pre_topc(k3_xboole_0(X0, u1_struct_0(X2)), X2) | ~ sP1(X2, X1) | ~ v4_pre_topc(X0, X1)), inference(resolution, [], [f618, f609])).
fof(f609, plain, ! [X0, X1] : (m1_subset_1(X0, k1_zfmisc_1(X1)) | ~ r1_tarski(X0, X1)), inference(cnf_transformation, [], [f355])).
fof(f355, plain, ! [X0, X1] : ((m1_subset_1(X0, k1_zfmisc_1(X1)) | ~ r1_tarski(X0, X1)) & (r1_tarski(X0, X1) | ~ m1_subset_1(X0, k1_zfmisc_1(X1)))), inference(nnf_transformation, [], [f117])).
fof(f117, plain, ! [X0, X1] : (m1_subset_1(X0, k1_zfmisc_1(X1)) <=> r1_tarski(X0, X1)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', t3_subset)).
fof(f618, plain, ! [X6, X0, X1] : (~ m1_subset_1(k3_xboole_0(X6, u1_struct_0(X0)), k1_zfmisc_1(u1_struct_0(X0))) | ~ v4_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))) | v4_pre_topc(k3_xboole_0(X6, u1_struct_0(X0)), X0) | ~ sP1(X0, X1)), inference(equality_resolution, [], [f478])).
fof(f478, plain, ! [X6, X0, X5, X1] : (v4_pre_topc(X5, X0) | ~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v4_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0))) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f313])).
fof(f313, plain, ! [X0, X1] : ((sP1(X0, X1) | ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = sK9(X0, X1)) | ~ v4_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(sK9(X0, X1), X0)) & (((sK9(X0, X1) = k3_xboole_0(sK10(X0, X1), u1_struct_0(X0))) & v4_pre_topc(sK10(X0, X1), X1) & m1_subset_1(sK10(X0, X1), k1_zfmisc_1(u1_struct_0(X1)))) | v4_pre_topc(sK9(X0, X1), X0)) & m1_subset_1(sK9(X0, X1), k1_zfmisc_1(u1_struct_0(X0)))) | ~ r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) & ((! [X5] : (((v4_pre_topc(X5, X0) | ! [X6] : (~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v4_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))))) & (((k3_xboole_0(sK11(X0, X1, X5), u1_struct_0(X0)) = X5) & v4_pre_topc(sK11(X0, X1, X5), X1) & m1_subset_1(sK11(X0, X1, X5), k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(X5, X0))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) & r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) | ~ sP1(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f309, f312, f311, f310])).
fof(f310, plain, ! [X1, X0] : (? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = X2) | ~ v4_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(X2, X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = X2) & v4_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v4_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) => ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = sK9(X0, X1)) | ~ v4_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(sK9(X0, X1), X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = sK9(X0, X1)) & v4_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v4_pre_topc(sK9(X0, X1), X0)) & m1_subset_1(sK9(X0, X1), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f311, plain, ! [X1, X0] : (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = sK9(X0, X1)) & v4_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) => ((sK9(X0, X1) = k3_xboole_0(sK10(X0, X1), u1_struct_0(X0))) & v4_pre_topc(sK10(X0, X1), X1) & m1_subset_1(sK10(X0, X1), k1_zfmisc_1(u1_struct_0(X1))))), introduced(choice_axiom, [])).
fof(f312, plain, ! [X5, X1, X0] : (? [X7] : ((k3_xboole_0(X7, u1_struct_0(X0)) = X5) & v4_pre_topc(X7, X1) & m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X1)))) => ((k3_xboole_0(sK11(X0, X1, X5), u1_struct_0(X0)) = X5) & v4_pre_topc(sK11(X0, X1, X5), X1) & m1_subset_1(sK11(X0, X1, X5), k1_zfmisc_1(u1_struct_0(X1))))), introduced(choice_axiom, [])).
fof(f309, plain, ! [X0, X1] : ((sP1(X0, X1) | ? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = X2) | ~ v4_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(X2, X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = X2) & v4_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v4_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) | ~ r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) & ((! [X5] : (((v4_pre_topc(X5, X0) | ! [X6] : (~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v4_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))))) & (? [X7] : ((k3_xboole_0(X7, u1_struct_0(X0)) = X5) & v4_pre_topc(X7, X1) & m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v4_pre_topc(X5, X0))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) & r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) | ~ sP1(X0, X1))), inference(rectify, [], [f308])).
fof(f308, plain, ! [X1, X0] : ((sP1(X1, X0) | ? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v4_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v4_pre_topc(X2, X1)) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | v4_pre_topc(X2, X1)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) | ~ r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) & ((! [X2] : (((v4_pre_topc(X2, X1) | ! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v4_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v4_pre_topc(X2, X1))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) | ~ sP1(X1, X0))), inference(flattening, [], [f307])).
fof(f307, plain, ! [X1, X0] : ((sP1(X1, X0) | (? [X2] : (((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v4_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v4_pre_topc(X2, X1)) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | v4_pre_topc(X2, X1))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) | ~ r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))) & ((! [X2] : (((v4_pre_topc(X2, X1) | ! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v4_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v4_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v4_pre_topc(X2, X1))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) | ~ sP1(X1, X0))), inference(nnf_transformation, [], [f291])).
fof(f112763, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | (~ spl31_202 | ~ spl31_204 | ~ spl31_205 | ~ spl31_208)), inference(subsumption_resolution, [], [f112762, f12773])).
fof(f12773, plain, (m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl31_202), inference(avatar_component_clause, [], [f12772])).
fof(f112762, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl31_204 | ~ spl31_205 | ~ spl31_208)), inference(subsumption_resolution, [], [f112761, f643])).
fof(f643, plain, l1_struct_0(sK4), inference(resolution, [], [f496, f358])).
fof(f496, plain, ! [X0] : (~ l1_pre_topc(X0) | l1_struct_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ! [X0] : (l1_struct_0(X0) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f64])).
fof(f64, plain, ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', dt_l1_pre_topc)).
fof(f112761, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl31_204 | ~ spl31_205 | ~ spl31_208)), inference(subsumption_resolution, [], [f112760, f926])).
fof(f926, plain, l1_struct_0(sK5), inference(resolution, [], [f923, f496])).
fof(f923, plain, l1_pre_topc(sK5), inference(subsumption_resolution, [], [f920, f358])).
fof(f920, plain, (l1_pre_topc(sK5) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f499, f361])).
fof(f112760, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl31_204 | ~ spl31_205 | ~ spl31_208)), inference(subsumption_resolution, [], [f112759, f12781])).
fof(f12781, plain, (v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ spl31_204), inference(avatar_component_clause, [], [f12780])).
fof(f112759, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl31_205 | ~ spl31_208)), inference(subsumption_resolution, [], [f112690, f12785])).
fof(f12785, plain, (v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl31_205), inference(avatar_component_clause, [], [f12784])).
fof(f112690, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl31_208), inference(superposition, [], [f2037, f72469])).
fof(f72469, plain, ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | ~ spl31_208), inference(forward_demodulation, [], [f72468, f68761])).
fof(f68761, plain, ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | ~ spl31_208), inference(forward_demodulation, [], [f68725, f465])).
fof(f465, plain, ! [X0, X1] : (k3_xboole_0(X0, X1) = k3_xboole_0(X1, X0)), inference(cnf_transformation, [], [f50])).
fof(f50, plain, ! [X0, X1] : (k3_xboole_0(X0, X1) = k3_xboole_0(X1, X0)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', commutativity_k3_xboole_0)).
fof(f68725, plain, ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6) = k3_xboole_0(u1_struct_0(sK5), sK6)) | ~ spl31_208), inference(resolution, [], [f1677, f12996])).
fof(f12996, plain, (m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4))) | ~ spl31_208), inference(avatar_component_clause, [], [f12995])).
fof(f1677, plain, ! [X19] : (~ m1_subset_1(X19, k1_zfmisc_1(u1_struct_0(sK4))) | (k5_subset_1(u1_struct_0(sK4), X19, sK6) = k3_xboole_0(X19, sK6))), inference(resolution, [], [f598, f362])).
fof(f598, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | (k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(cnf_transformation, [], [f275])).
fof(f275, plain, ! [X0, X1, X2] : ((k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(flattening, [], [f274])).
fof(f274, plain, ! [X0, X1, X2] : ((k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, k1_zfmisc_1(X0)) & m1_subset_1(X1, k1_zfmisc_1(X0))) => (k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', redefinition_k5_subset_1)).
fof(f72468, plain, (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6)), inference(subsumption_resolution, [], [f72467, f359])).
fof(f359, plain, ~ v3_struct_0(sK5), inference(cnf_transformation, [], [f299])).
fof(f72467, plain, (v3_struct_0(sK5) | (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6))), inference(subsumption_resolution, [], [f72464, f360])).
fof(f360, plain, v2_tsp_2(sK5, sK4), inference(cnf_transformation, [], [f299])).
fof(f72464, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6))), inference(resolution, [], [f67799, f1129])).
fof(f1129, plain, m1_pre_topc(sK5, sK4), inference(subsumption_resolution, [], [f1126, f358])).
fof(f1126, plain, (m1_pre_topc(sK5, sK4) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f601, f361])).
fof(f601, plain, ! [X0, X1] : (~ m2_tsp_1(X1, X0) | m1_pre_topc(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f354])).
fof(f354, plain, ! [X0] : (! [X1] : ((m2_tsp_1(X1, X0) | ~ m1_pre_topc(X1, X0)) & (m1_pre_topc(X1, X0) | ~ m2_tsp_1(X1, X0))) | ~ l1_pre_topc(X0)), inference(nnf_transformation, [], [f276])).
fof(f276, plain, ! [X0] : (! [X1] : (m2_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f111])).
fof(f111, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m2_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', redefinition_m2_tsp_1)).
fof(f67799, plain, ! [X28] : (~ m1_pre_topc(X28, sK4) | ~ v2_tsp_2(X28, sK4) | v3_struct_0(X28) | (k4_pre_topc(sK4, X28, k4_tsp_2(sK4, X28), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(X28), sK6))), inference(forward_demodulation, [], [f67798, f1910])).
fof(f1910, plain, (sK6 = k3_tex_4(sK4, sK6)), inference(subsumption_resolution, [], [f1909, f356])).
fof(f356, plain, ~ v3_struct_0(sK4), inference(cnf_transformation, [], [f299])).
fof(f1909, plain, ((sK6 = k3_tex_4(sK4, sK6)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1908, f357])).
fof(f357, plain, v2_pre_topc(sK4), inference(cnf_transformation, [], [f299])).
fof(f1908, plain, ((sK6 = k3_tex_4(sK4, sK6)) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1907, f358])).
fof(f1907, plain, ((sK6 = k3_tex_4(sK4, sK6)) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1884, f363])).
fof(f1884, plain, (~ v4_pre_topc(sK6, sK4) | (sK6 = k3_tex_4(sK4, sK6)) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f612, f362])).
fof(f612, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ v4_pre_topc(X1, X0) | (k3_tex_4(X0, X1) = X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f285])).
fof(f285, plain, ! [X0] : (! [X1] : ((k3_tex_4(X0, X1) = X1) | ~ v4_pre_topc(X1, X0) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f284])).
fof(f284, plain, ! [X0] : (! [X1] : (((k3_tex_4(X0, X1) = X1) | ~ v4_pre_topc(X1, X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f120])).
fof(f120, plain, ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (v4_pre_topc(X1, X0) => (k3_tex_4(X0, X1) = X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', t62_tex_4)).
fof(f67798, plain, ! [X28] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X28), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X28, k4_tsp_2(sK4, X28), sK6)) | ~ v2_tsp_2(X28, sK4) | v3_struct_0(X28) | ~ m1_pre_topc(X28, sK4)), inference(subsumption_resolution, [], [f67797, f356])).
fof(f67797, plain, ! [X28] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X28), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X28, k4_tsp_2(sK4, X28), sK6)) | ~ v2_tsp_2(X28, sK4) | v3_struct_0(X28) | v3_struct_0(sK4) | ~ m1_pre_topc(X28, sK4)), inference(subsumption_resolution, [], [f67796, f357])).
fof(f67796, plain, ! [X28] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X28), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X28, k4_tsp_2(sK4, X28), sK6)) | ~ v2_tsp_2(X28, sK4) | v3_struct_0(X28) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(X28, sK4)), inference(subsumption_resolution, [], [f67568, f358])).
fof(f67568, plain, ! [X28] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X28), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X28, k4_tsp_2(sK4, X28), sK6)) | ~ v2_tsp_2(X28, sK4) | v3_struct_0(X28) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(X28, sK4)), inference(resolution, [], [f362, f2201])).
fof(f2201, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2200, f491])).
fof(f491, plain, ! [X0, X1] : (v1_funct_1(k4_tsp_2(X0, X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f229])).
fof(f229, plain, ! [X0, X1] : ((m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1))) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f228])).
fof(f228, plain, ! [X0, X1] : ((m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1))) | (~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0, X1] : ((m1_pre_topc(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', dt_k4_tsp_2)).
fof(f2200, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2199, f492])).
fof(f492, plain, ! [X0, X1] : (v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f229])).
fof(f2199, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2198, f493])).
fof(f493, plain, ! [X0, X1] : (v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f229])).
fof(f2198, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2197, f605])).
fof(f605, plain, ! [X0, X1] : (m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_pre_topc(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f278])).
fof(f278, plain, ! [X0] : (! [X1] : (m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_pre_topc(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_pre_topc(X1, X0) => m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', t1_tsep_1)).
fof(f2197, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2196, f602])).
fof(f602, plain, ! [X0, X1] : (~ m1_pre_topc(X1, X0) | m2_tsp_1(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f354])).
fof(f2196, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ m2_tsp_1(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(duplicate_literal_removal, [], [f2195])).
fof(f2195, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ m2_tsp_1(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1)), inference(resolution, [], [f617, f494])).
fof(f494, plain, ! [X0, X1] : (m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f229])).
fof(f617, plain, ! [X6, X0, X1] : (~ m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | (k5_subset_1(u1_struct_0(X0), u1_struct_0(X1), k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X6)) | ~ v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) | ~ v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(k4_tsp_2(X0, X1)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f616])).
fof(f616, plain, ! [X6, X2, X0, X1] : ((k4_pre_topc(X0, X1, X2, X6) = k5_subset_1(u1_struct_0(X0), u1_struct_0(X1), k3_tex_4(X0, X6))) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ (k4_tsp_2(X0, X1) = X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f467])).
fof(f467, plain, ! [X6, X2, X0, X5, X1] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0))) | ~ (k4_tsp_2(X0, X1) = X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f305])).
fof(f305, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ((~ (k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, sK8(X0, X1, X2))) = k4_pre_topc(X0, X1, X2, sK8(X0, X1, X2))) & m1_subset_1(sK8(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = sK7(X0, X1, X2)) & m1_subset_1(sK7(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))) & (! [X5] : (! [X6] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f302, f304, f303])).
fof(f303, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) => (? [X4] : (~ (k4_pre_topc(X0, X1, X2, X4) = k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, X4))) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = sK7(X0, X1, X2)) & m1_subset_1(sK7(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f304, plain, ! [X2, X1, X0] : (? [X4] : (~ (k4_pre_topc(X0, X1, X2, X4) = k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, X4))) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) => (~ (k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, sK8(X0, X1, X2))) = k4_pre_topc(X0, X1, X2, sK8(X0, X1, X2))) & m1_subset_1(sK8(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f302, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (! [X5] : (! [X6] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(rectify, [], [f301])).
fof(f301, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (! [X3] : (! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f220])).
fof(f220, plain, ! [X0] : (! [X1] : (! [X2] : (((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : (! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f219])).
fof(f219, plain, ! [X0] : (! [X1] : (! [X2] : (((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : ((! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3)) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1))) | (~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => ((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : (m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))) => ((u1_struct_0(X1) = X3) => ! [X4] : (m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0))) => (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4))))))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', d12_tsp_2)).
fof(f2037, plain, ! [X111, X114, X112, X113] : (r1_tarski(k4_pre_topc(X112, X113, X111, X114), u1_struct_0(X113)) | ~ v1_funct_2(X111, u1_struct_0(X112), u1_struct_0(X113)) | ~ v1_funct_1(X111) | ~ l1_struct_0(X113) | ~ l1_struct_0(X112) | ~ m1_relset_1(X111, u1_struct_0(X112), u1_struct_0(X113))), inference(resolution, [], [f490, f608])).
fof(f608, plain, ! [X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(X1)) | r1_tarski(X0, X1)), inference(cnf_transformation, [], [f355])).
fof(f490, plain, ! [X2, X0, X3, X1] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f227])).
fof(f227, plain, ! [X0, X1, X2, X3] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(flattening, [], [f226])).
fof(f226, plain, ! [X0, X1, X2, X3] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ! [X0, X1, X2, X3] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_struct_0(X1) & l1_struct_0(X0)) => m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', dt_k4_pre_topc)).
fof(f132324, plain, (~ spl31_2327 | ~ spl31_206 | ~ spl31_208), inference(avatar_split_clause, [], [f116793, f12995, f12816, f112755])).
fof(f12816, plain, (spl31_206 <=> ! [X0] : (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0))), introduced(avatar_definition, [new_symbols(naming, [spl31_206])])).
fof(f116793, plain, (~ v4_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5) | (~ spl31_206 | ~ spl31_208)), inference(backward_demodulation, [], [f67486, f116723])).
fof(f116723, plain, ((k9_relat_1(k4_tsp_2(sK4, sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | (~ spl31_206 | ~ spl31_208)), inference(superposition, [], [f12817, f72469])).
fof(f12817, plain, (! [X0] : (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ spl31_206), inference(avatar_component_clause, [], [f12816])).
fof(f67486, plain, (~ v4_pre_topc(k9_relat_1(k4_tsp_2(sK4, sK5), sK6), sK5) | ~ spl31_206), inference(forward_demodulation, [], [f364, f12817])).
fof(f364, plain, ~ v4_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5), inference(cnf_transformation, [], [f299])).
fof(f13069, plain, spl31_208, inference(avatar_contradiction_clause, [], [f13068])).
fof(f13068, plain, ($false | spl31_208), inference(subsumption_resolution, [], [f13067, f358])).
fof(f13067, plain, (~ l1_pre_topc(sK4) | spl31_208), inference(subsumption_resolution, [], [f13065, f1129])).
fof(f13065, plain, (~ m1_pre_topc(sK5, sK4) | ~ l1_pre_topc(sK4) | spl31_208), inference(resolution, [], [f12997, f605])).
fof(f12997, plain, (~ m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4))) | spl31_208), inference(avatar_component_clause, [], [f12995])).
fof(f12920, plain, spl31_205, inference(avatar_contradiction_clause, [], [f12919])).
fof(f12919, plain, ($false | spl31_205), inference(subsumption_resolution, [], [f12918, f356])).
fof(f12918, plain, (v3_struct_0(sK4) | spl31_205), inference(subsumption_resolution, [], [f12917, f357])).
fof(f12917, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_205), inference(subsumption_resolution, [], [f12916, f358])).
fof(f12916, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_205), inference(subsumption_resolution, [], [f12915, f359])).
fof(f12915, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_205), inference(subsumption_resolution, [], [f12914, f360])).
fof(f12914, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_205), inference(subsumption_resolution, [], [f12913, f1129])).
fof(f12913, plain, (~ m1_pre_topc(sK5, sK4) | ~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_205), inference(resolution, [], [f12786, f492])).
fof(f12786, plain, (~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | spl31_205), inference(avatar_component_clause, [], [f12784])).
fof(f12828, plain, spl31_204, inference(avatar_contradiction_clause, [], [f12827])).
fof(f12827, plain, ($false | spl31_204), inference(subsumption_resolution, [], [f12826, f356])).
fof(f12826, plain, (v3_struct_0(sK4) | spl31_204), inference(subsumption_resolution, [], [f12825, f357])).
fof(f12825, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_204), inference(subsumption_resolution, [], [f12824, f358])).
fof(f12824, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_204), inference(subsumption_resolution, [], [f12823, f359])).
fof(f12823, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_204), inference(subsumption_resolution, [], [f12822, f360])).
fof(f12822, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_204), inference(subsumption_resolution, [], [f12821, f1129])).
fof(f12821, plain, (~ m1_pre_topc(sK5, sK4) | ~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl31_204), inference(resolution, [], [f12782, f491])).
fof(f12782, plain, (~ v1_funct_1(k4_tsp_2(sK4, sK5)) | spl31_204), inference(avatar_component_clause, [], [f12780])).
fof(f12818, plain, (~ spl31_204 | ~ spl31_205 | spl31_206 | ~ spl31_202), inference(avatar_split_clause, [], [f12814, f12772, f12816, f12784, f12780])).
fof(f12814, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5))) | ~ spl31_202), inference(subsumption_resolution, [], [f12813, f643])).
fof(f12813, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK4)) | ~ spl31_202), inference(subsumption_resolution, [], [f12796, f926])).
fof(f12796, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4)) | ~ spl31_202), inference(resolution, [], [f12773, f596])).
fof(f596, plain, ! [X2, X0, X3, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | (k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f271])).
fof(f271, plain, ! [X0, X1, X2, X3] : ((k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(flattening, [], [f270])).
fof(f270, plain, ! [X0, X1, X2, X3] : ((k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0))), inference(ennf_transformation, [], [f107])).
fof(f107, plain, ! [X0, X1, X2, X3] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_struct_0(X1) & l1_struct_0(X0)) => (k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', redefinition_k4_pre_topc)).
fof(f12795, plain, spl31_202, inference(avatar_contradiction_clause, [], [f12794])).
fof(f12794, plain, ($false | spl31_202), inference(subsumption_resolution, [], [f12793, f1129])).
fof(f12793, plain, (~ m1_pre_topc(sK5, sK4) | spl31_202), inference(subsumption_resolution, [], [f12792, f356])).
fof(f12792, plain, (v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl31_202), inference(subsumption_resolution, [], [f12791, f357])).
fof(f12791, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl31_202), inference(subsumption_resolution, [], [f12790, f358])).
fof(f12790, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl31_202), inference(subsumption_resolution, [], [f12789, f359])).
fof(f12789, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl31_202), inference(subsumption_resolution, [], [f12788, f360])).
fof(f12788, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl31_202), inference(resolution, [], [f12774, f1989])).
fof(f1989, plain, ! [X0, X1] : (m1_relset_1(k4_tsp_2(X1, X0), u1_struct_0(X1), u1_struct_0(X0)) | ~ v2_tsp_2(X0, X1) | v3_struct_0(X0) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X0, X1)), inference(resolution, [], [f494, f599])).
fof(f599, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f353])).
fof(f353, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f110])).
fof(f110, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP042+1.p', redefinition_m2_relset_1)).
fof(f12774, plain, (~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | spl31_202), inference(avatar_component_clause, [], [f12772])).