fof(f249940, plain, $false, inference(avatar_sat_refutation, [], [f12333, f12359, f12376, f12436, f12607, f249619, f249628])).
fof(f249628, plain, (~ spl32_1480 | ~ spl32_157 | ~ spl32_161), inference(avatar_split_clause, [], [f116403, f12526, f12357, f112216])).
fof(f112216, plain, (spl32_1480 <=> v3_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5)), introduced(avatar_definition, [new_symbols(naming, [spl32_1480])])).
fof(f12357, plain, (spl32_157 <=> ! [X0] : (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0))), introduced(avatar_definition, [new_symbols(naming, [spl32_157])])).
fof(f12526, plain, (spl32_161 <=> m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4)))), introduced(avatar_definition, [new_symbols(naming, [spl32_161])])).
fof(f116403, plain, (~ v3_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5) | (~ spl32_157 | ~ spl32_161)), inference(backward_demodulation, [], [f12447, f116330])).
fof(f116330, plain, ((k9_relat_1(k4_tsp_2(sK4, sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | (~ spl32_157 | ~ spl32_161)), inference(superposition, [], [f12358, f45632])).
fof(f45632, plain, ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | ~ spl32_161), inference(forward_demodulation, [], [f45631, f12743])).
fof(f12743, plain, ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6) = k3_xboole_0(sK6, u1_struct_0(sK5))) | ~ spl32_161), inference(forward_demodulation, [], [f12639, f474])).
fof(f474, plain, ! [X0, X1] : (k3_xboole_0(X0, X1) = k3_xboole_0(X1, X0)), inference(cnf_transformation, [], [f50])).
fof(f50, plain, ! [X0, X1] : (k3_xboole_0(X0, X1) = k3_xboole_0(X1, X0)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', commutativity_k3_xboole_0)).
fof(f12639, plain, ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6) = k3_xboole_0(u1_struct_0(sK5), sK6)) | ~ spl32_161), inference(resolution, [], [f12527, f1716])).
fof(f1716, plain, ! [X19] : (~ m1_subset_1(X19, k1_zfmisc_1(u1_struct_0(sK4))) | (k5_subset_1(u1_struct_0(sK4), X19, sK6) = k3_xboole_0(X19, sK6))), inference(resolution, [], [f609, f371])).
fof(f371, plain, m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4))), inference(cnf_transformation, [], [f305])).
fof(f305, plain, (((~ v3_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5) & v3_pre_topc(sK6, sK4) & m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(sK5, sK4) & v2_tsp_2(sK5, sK4) & ~ v3_struct_0(sK5)) & l1_pre_topc(sK4) & v2_pre_topc(sK4) & ~ v3_struct_0(sK4)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK4, sK5, sK6])], [f158, f304, f303, f302])).
fof(f302, plain, (? [X0] : (? [X1] : (? [X2] : (~ v3_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v3_pre_topc(X2, X0) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => (? [X1] : (? [X2] : (~ v3_pre_topc(k4_pre_topc(sK4, X1, k4_tsp_2(sK4, X1), X2), X1) & v3_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(X1, sK4) & v2_tsp_2(X1, sK4) & ~ v3_struct_0(X1)) & l1_pre_topc(sK4) & v2_pre_topc(sK4) & ~ v3_struct_0(sK4))), introduced(choice_axiom, [])).
fof(f303, plain, (? [X1] : (? [X2] : (~ v3_pre_topc(k4_pre_topc(sK4, X1, k4_tsp_2(sK4, X1), X2), X1) & v3_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(X1, sK4) & v2_tsp_2(X1, sK4) & ~ v3_struct_0(X1)) => (? [X2] : (~ v3_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X2), sK5) & v3_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) & m2_tsp_1(sK5, sK4) & v2_tsp_2(sK5, sK4) & ~ v3_struct_0(sK5))), introduced(choice_axiom, [])).
fof(f304, plain, (? [X2] : (~ v3_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X2), sK5) & v3_pre_topc(X2, sK4) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(sK4)))) => (~ v3_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5) & v3_pre_topc(sK6, sK4) & m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(sK4))))), introduced(choice_axiom, [])).
fof(f158, plain, ? [X0] : (? [X1] : (? [X2] : (~ v3_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v3_pre_topc(X2, X0) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)), inference(flattening, [], [f157])).
fof(f157, plain, ? [X0] : (? [X1] : (? [X2] : ((~ v3_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1) & v3_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) & (m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1))) & (l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0))), inference(ennf_transformation, [], [f2])).
fof(f2, plain, ~ ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) => (v3_pre_topc(X2, X0) => v3_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1))))), inference(negated_conjecture, [], [f1])).
fof(f1, plain, ~ ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0))) => (v3_pre_topc(X2, X0) => v3_pre_topc(k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X2), X1))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', t34_tsp_2)).
fof(f609, plain, ! [X2, X0, X1] : (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | (k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(cnf_transformation, [], [f280])).
fof(f280, plain, ! [X0, X1, X2] : ((k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | ~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0))), inference(flattening, [], [f279])).
fof(f279, plain, ! [X0, X1, X2] : ((k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2)) | (~ m1_subset_1(X2, k1_zfmisc_1(X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(X0)))), inference(ennf_transformation, [], [f111])).
fof(f111, plain, ! [X0, X1, X2] : ((m1_subset_1(X2, k1_zfmisc_1(X0)) & m1_subset_1(X1, k1_zfmisc_1(X0))) => (k5_subset_1(X0, X1, X2) = k3_xboole_0(X1, X2))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', redefinition_k5_subset_1)).
fof(f12527, plain, (m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4))) | ~ spl32_161), inference(avatar_component_clause, [], [f12526])).
fof(f45631, plain, (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6)), inference(subsumption_resolution, [], [f45630, f368])).
fof(f368, plain, ~ v3_struct_0(sK5), inference(cnf_transformation, [], [f305])).
fof(f45630, plain, (v3_struct_0(sK5) | (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6))), inference(subsumption_resolution, [], [f45626, f369])).
fof(f369, plain, v2_tsp_2(sK5, sK4), inference(cnf_transformation, [], [f305])).
fof(f45626, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6) = k5_subset_1(u1_struct_0(sK4), u1_struct_0(sK5), sK6))), inference(resolution, [], [f12254, f1151])).
fof(f1151, plain, m1_pre_topc(sK5, sK4), inference(subsumption_resolution, [], [f1148, f367])).
fof(f367, plain, l1_pre_topc(sK4), inference(cnf_transformation, [], [f305])).
fof(f1148, plain, (m1_pre_topc(sK5, sK4) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f614, f370])).
fof(f370, plain, m2_tsp_1(sK5, sK4), inference(cnf_transformation, [], [f305])).
fof(f614, plain, ! [X0, X1] : (~ m2_tsp_1(X1, X0) | m1_pre_topc(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f363])).
fof(f363, plain, ! [X0] : (! [X1] : ((m2_tsp_1(X1, X0) | ~ m1_pre_topc(X1, X0)) & (m1_pre_topc(X1, X0) | ~ m2_tsp_1(X1, X0))) | ~ l1_pre_topc(X0)), inference(nnf_transformation, [], [f282])).
fof(f282, plain, ! [X0] : (! [X1] : (m2_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f114])).
fof(f114, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m2_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', redefinition_m2_tsp_1)).
fof(f12254, plain, ! [X20] : (~ m1_pre_topc(X20, sK4) | ~ v2_tsp_2(X20, sK4) | v3_struct_0(X20) | (k5_subset_1(u1_struct_0(sK4), u1_struct_0(X20), sK6) = k4_pre_topc(sK4, X20, k4_tsp_2(sK4, X20), sK6))), inference(forward_demodulation, [], [f12253, f1950])).
fof(f1950, plain, (sK6 = k3_tex_4(sK4, sK6)), inference(subsumption_resolution, [], [f1949, f365])).
fof(f365, plain, ~ v3_struct_0(sK4), inference(cnf_transformation, [], [f305])).
fof(f1949, plain, ((sK6 = k3_tex_4(sK4, sK6)) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1948, f366])).
fof(f366, plain, v2_pre_topc(sK4), inference(cnf_transformation, [], [f305])).
fof(f1948, plain, ((sK6 = k3_tex_4(sK4, sK6)) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1947, f367])).
fof(f1947, plain, ((sK6 = k3_tex_4(sK4, sK6)) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(subsumption_resolution, [], [f1924, f372])).
fof(f372, plain, v3_pre_topc(sK6, sK4), inference(cnf_transformation, [], [f305])).
fof(f1924, plain, (~ v3_pre_topc(sK6, sK4) | (sK6 = k3_tex_4(sK4, sK6)) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4)), inference(resolution, [], [f624, f371])).
fof(f624, plain, ! [X0, X1] : (~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) | ~ v3_pre_topc(X1, X0) | (k3_tex_4(X0, X1) = X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f290])).
fof(f290, plain, ! [X0] : (! [X1] : ((k3_tex_4(X0, X1) = X1) | ~ v3_pre_topc(X1, X0) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f289])).
fof(f289, plain, ! [X0] : (! [X1] : (((k3_tex_4(X0, X1) = X1) | ~ v3_pre_topc(X1, X0)) | ~ m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0)))) | (~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f122])).
fof(f122, plain, ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1, k1_zfmisc_1(u1_struct_0(X0))) => (v3_pre_topc(X1, X0) => (k3_tex_4(X0, X1) = X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', t58_tex_4)).
fof(f12253, plain, ! [X20] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X20), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X20, k4_tsp_2(sK4, X20), sK6)) | ~ v2_tsp_2(X20, sK4) | v3_struct_0(X20) | ~ m1_pre_topc(X20, sK4)), inference(subsumption_resolution, [], [f12252, f365])).
fof(f12252, plain, ! [X20] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X20), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X20, k4_tsp_2(sK4, X20), sK6)) | ~ v2_tsp_2(X20, sK4) | v3_struct_0(X20) | v3_struct_0(sK4) | ~ m1_pre_topc(X20, sK4)), inference(subsumption_resolution, [], [f12251, f366])).
fof(f12251, plain, ! [X20] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X20), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X20, k4_tsp_2(sK4, X20), sK6)) | ~ v2_tsp_2(X20, sK4) | v3_struct_0(X20) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(X20, sK4)), inference(subsumption_resolution, [], [f12216, f367])).
fof(f12216, plain, ! [X20] : ((k5_subset_1(u1_struct_0(sK4), u1_struct_0(X20), k3_tex_4(sK4, sK6)) = k4_pre_topc(sK4, X20, k4_tsp_2(sK4, X20), sK6)) | ~ v2_tsp_2(X20, sK4) | v3_struct_0(X20) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(X20, sK4)), inference(resolution, [], [f2219, f371])).
fof(f2219, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2218, f500])).
fof(f500, plain, ! [X0, X1] : (v1_funct_1(k4_tsp_2(X0, X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f232, plain, ! [X0, X1] : ((m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1))) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f231])).
fof(f231, plain, ! [X0, X1] : ((m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1))) | (~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f61])).
fof(f61, plain, ! [X0, X1] : ((m1_pre_topc(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1) & l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => (m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) & v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(k4_tsp_2(X0, X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', dt_k4_tsp_2)).
fof(f2218, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2217, f501])).
fof(f501, plain, ! [X0, X1] : (v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f2217, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2216, f502])).
fof(f502, plain, ! [X0, X1] : (v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f2216, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2215, f618])).
fof(f618, plain, ! [X0, X1] : (m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_pre_topc(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f284])).
fof(f284, plain, ! [X0] : (! [X1] : (m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_pre_topc(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f117])).
fof(f117, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_pre_topc(X1, X0) => m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', t1_tsep_1)).
fof(f2215, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(subsumption_resolution, [], [f2214, f615])).
fof(f615, plain, ! [X0, X1] : (~ m1_pre_topc(X1, X0) | m2_tsp_1(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f363])).
fof(f2214, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ m2_tsp_1(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1)), inference(duplicate_literal_removal, [], [f2213])).
fof(f2213, plain, ! [X2, X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(u1_struct_0(X2), k1_zfmisc_1(u1_struct_0(X1))) | (k5_subset_1(u1_struct_0(X1), u1_struct_0(X2), k3_tex_4(X1, X0)) = k4_pre_topc(X1, X2, k4_tsp_2(X1, X2), X0)) | ~ v5_pre_topc(k4_tsp_2(X1, X2), X1, X2) | ~ v1_funct_2(k4_tsp_2(X1, X2), u1_struct_0(X1), u1_struct_0(X2)) | ~ v1_funct_1(k4_tsp_2(X1, X2)) | ~ m2_tsp_1(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X2, X1) | ~ v2_tsp_2(X2, X1) | v3_struct_0(X2) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1)), inference(resolution, [], [f630, f503])).
fof(f503, plain, ! [X0, X1] : (m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_pre_topc(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f232])).
fof(f630, plain, ! [X6, X0, X1] : (~ m2_relset_1(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | (k5_subset_1(u1_struct_0(X0), u1_struct_0(X1), k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, k4_tsp_2(X0, X1), X6)) | ~ v5_pre_topc(k4_tsp_2(X0, X1), X0, X1) | ~ v1_funct_2(k4_tsp_2(X0, X1), u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(k4_tsp_2(X0, X1)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f629])).
fof(f629, plain, ! [X6, X2, X0, X1] : ((k4_pre_topc(X0, X1, X2, X6) = k5_subset_1(u1_struct_0(X0), u1_struct_0(X1), k3_tex_4(X0, X6))) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ m1_subset_1(u1_struct_0(X1), k1_zfmisc_1(u1_struct_0(X0))) | ~ (k4_tsp_2(X0, X1) = X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(equality_resolution, [], [f476])).
fof(f476, plain, ! [X6, X2, X0, X5, X1] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0))) | ~ (k4_tsp_2(X0, X1) = X2) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(cnf_transformation, [], [f311])).
fof(f311, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ((~ (k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, sK8(X0, X1, X2))) = k4_pre_topc(X0, X1, X2, sK8(X0, X1, X2))) & m1_subset_1(sK8(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = sK7(X0, X1, X2)) & m1_subset_1(sK7(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))) & (! [X5] : (! [X6] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(skolemisation, [status(esa), new_symbols(skolem, [sK7, sK8])], [f308, f310, f309])).
fof(f309, plain, ! [X2, X1, X0] : (? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) => (? [X4] : (~ (k4_pre_topc(X0, X1, X2, X4) = k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, X4))) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = sK7(X0, X1, X2)) & m1_subset_1(sK7(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f310, plain, ! [X2, X1, X0] : (? [X4] : (~ (k4_pre_topc(X0, X1, X2, X4) = k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, X4))) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) => (~ (k5_subset_1(u1_struct_0(X0), sK7(X0, X1, X2), k3_tex_4(X0, sK8(X0, X1, X2))) = k4_pre_topc(X0, X1, X2, sK8(X0, X1, X2))) & m1_subset_1(sK8(X0, X1, X2), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f308, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (! [X5] : (! [X6] : ((k5_subset_1(u1_struct_0(X0), X5, k3_tex_4(X0, X6)) = k4_pre_topc(X0, X1, X2, X6)) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X5) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(rectify, [], [f307])).
fof(f307, plain, ! [X0] : (! [X1] : (! [X2] : ((((k4_tsp_2(X0, X1) = X2) | ? [X3] : (? [X4] : (~ (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) & (u1_struct_0(X1) = X3) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (! [X3] : (! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (k4_tsp_2(X0, X1) = X2))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(nnf_transformation, [], [f223])).
fof(f223, plain, ! [X0] : (! [X1] : (! [X2] : (((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : (! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2)) | ~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1)) | ~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0)), inference(flattening, [], [f222])).
fof(f222, plain, ! [X0] : (! [X1] : (! [X2] : (((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : ((! [X4] : ((k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4)) | ~ m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0)))) | ~ (u1_struct_0(X1) = X3)) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | (~ m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v5_pre_topc(X2, X0, X1) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2))) | (~ m2_tsp_1(X1, X0) | ~ v2_tsp_2(X1, X0) | v3_struct_0(X1))) | (~ l1_pre_topc(X0) | ~ v2_pre_topc(X0) | v3_struct_0(X0))), inference(ennf_transformation, [], [f52])).
fof(f52, plain, ! [X0] : ((l1_pre_topc(X0) & v2_pre_topc(X0) & ~ v3_struct_0(X0)) => ! [X1] : ((m2_tsp_1(X1, X0) & v2_tsp_2(X1, X0) & ~ v3_struct_0(X1)) => ! [X2] : ((m2_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v5_pre_topc(X2, X0, X1) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2)) => ((k4_tsp_2(X0, X1) = X2) <=> ! [X3] : (m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))) => ((u1_struct_0(X1) = X3) => ! [X4] : (m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X0))) => (k5_subset_1(u1_struct_0(X0), X3, k3_tex_4(X0, X4)) = k4_pre_topc(X0, X1, X2, X4))))))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', d12_tsp_2)).
fof(f12358, plain, (! [X0] : (k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ spl32_157), inference(avatar_component_clause, [], [f12357])).
fof(f12447, plain, (~ v3_pre_topc(k9_relat_1(k4_tsp_2(sK4, sK5), sK6), sK5) | ~ spl32_157), inference(backward_demodulation, [], [f373, f12358])).
fof(f373, plain, ~ v3_pre_topc(k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), sK6), sK5), inference(cnf_transformation, [], [f305])).
fof(f249619, plain, (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161 | spl32_1480), inference(avatar_contradiction_clause, [], [f249618])).
fof(f249618, plain, ($false | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161 | spl32_1480)), inference(subsumption_resolution, [], [f249617, f372])).
fof(f249617, plain, (~ v3_pre_topc(sK6, sK4) | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161 | spl32_1480)), inference(subsumption_resolution, [], [f249615, f2328])).
fof(f2328, plain, sP1(sK5, sK4), inference(subsumption_resolution, [], [f2321, f367])).
fof(f2321, plain, (sP1(sK5, sK4) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f1187, f1154])).
fof(f1154, plain, m1_tsp_1(sK5, sK4), inference(subsumption_resolution, [], [f1152, f367])).
fof(f1152, plain, (m1_tsp_1(sK5, sK4) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f1151, f611])).
fof(f611, plain, ! [X0, X1] : (~ m1_pre_topc(X1, X0) | m1_tsp_1(X1, X0) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f361])).
fof(f361, plain, ! [X0] : (! [X1] : ((m1_tsp_1(X1, X0) | ~ m1_pre_topc(X1, X0)) & (m1_pre_topc(X1, X0) | ~ m1_tsp_1(X1, X0))) | ~ l1_pre_topc(X0)), inference(nnf_transformation, [], [f281])).
fof(f281, plain, ! [X0] : (! [X1] : (m1_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f112])).
fof(f112, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_tsp_1(X1, X0) <=> m1_pre_topc(X1, X0))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', redefinition_m1_tsp_1)).
fof(f1187, plain, ! [X0, X1] : (~ m1_tsp_1(X0, X1) | sP1(X0, X1) | ~ l1_pre_topc(X1)), inference(subsumption_resolution, [], [f1186, f507])).
fof(f507, plain, ! [X0, X1] : (~ m1_tsp_1(X1, X0) | l1_pre_topc(X1) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f237])).
fof(f237, plain, ! [X0] : (! [X1] : (l1_pre_topc(X1) | ~ m1_tsp_1(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f69])).
fof(f69, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m1_tsp_1(X1, X0) => l1_pre_topc(X1))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', dt_m1_tsp_1)).
fof(f1186, plain, ! [X0, X1] : (~ m1_tsp_1(X0, X1) | sP1(X0, X1) | ~ l1_pre_topc(X0) | ~ l1_pre_topc(X1)), inference(resolution, [], [f481, f493])).
fof(f493, plain, ! [X0, X1] : (sP2(X0, X1) | ~ l1_pre_topc(X1) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f299])).
fof(f299, plain, ! [X0] : (! [X1] : (sP2(X0, X1) | ~ l1_pre_topc(X1)) | ~ l1_pre_topc(X0)), inference(definition_folding, [], [f224, e298, e297])).
fof(f297, plain, ! [X1, X0] : (sP1(X1, X0) <=> (! [X2] : ((v3_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))), inference(usedef, [], [e297])).
fof(e297, plain, ! [X1, X0] : (sP1(X1, X0) <=> (! [X2] : ((v3_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP1])])).
fof(f298, plain, ! [X0, X1] : ((m1_tsp_1(X1, X0) <=> sP1(X1, X0)) | ~ sP2(X0, X1)), inference(usedef, [], [e298])).
fof(e298, plain, ! [X0, X1] : (sP2(X0, X1) <=> (m1_tsp_1(X1, X0) <=> sP1(X1, X0))), introduced(predicate_definition_introduction, [new_symbols(naming, [sP2])])).
fof(f224, plain, ! [X0] : (! [X1] : ((m1_tsp_1(X1, X0) <=> (! [X2] : ((v3_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))) | ~ l1_pre_topc(X1)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f53])).
fof(f53, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (l1_pre_topc(X1) => (m1_tsp_1(X1, X0) <=> (! [X2] : (m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1))) => (v3_pre_topc(X2, X1) <=> ? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', d1_tsp_1)).
fof(f481, plain, ! [X0, X1] : (~ sP2(X0, X1) | ~ m1_tsp_1(X1, X0) | sP1(X1, X0)), inference(cnf_transformation, [], [f312])).
fof(f312, plain, ! [X0, X1] : (((m1_tsp_1(X1, X0) | ~ sP1(X1, X0)) & (sP1(X1, X0) | ~ m1_tsp_1(X1, X0))) | ~ sP2(X0, X1)), inference(nnf_transformation, [], [f298])).
fof(f249615, plain, (~ sP1(sK5, sK4) | ~ v3_pre_topc(sK6, sK4) | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161 | spl32_1480)), inference(resolution, [], [f112685, f371])).
fof(f112685, plain, (! [X0] : (~ m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(X0))) | ~ sP1(sK5, X0) | ~ v3_pre_topc(sK6, X0)) | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161 | spl32_1480)), inference(subsumption_resolution, [], [f112597, f112218])).
fof(f112218, plain, (~ v3_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5) | spl32_1480), inference(avatar_component_clause, [], [f112216])).
fof(f112597, plain, (! [X0] : (~ m1_subset_1(sK6, k1_zfmisc_1(u1_struct_0(X0))) | v3_pre_topc(k3_xboole_0(sK6, u1_struct_0(sK5)), sK5) | ~ sP1(sK5, X0) | ~ v3_pre_topc(sK6, X0)) | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161)), inference(resolution, [], [f112224, f2038])).
fof(f2038, plain, ! [X2, X0, X1] : (~ r1_tarski(k3_xboole_0(X0, u1_struct_0(X2)), u1_struct_0(X2)) | ~ m1_subset_1(X0, k1_zfmisc_1(u1_struct_0(X1))) | v3_pre_topc(k3_xboole_0(X0, u1_struct_0(X2)), X2) | ~ sP1(X2, X1) | ~ v3_pre_topc(X0, X1)), inference(resolution, [], [f631, f622])).
fof(f622, plain, ! [X0, X1] : (m1_subset_1(X0, k1_zfmisc_1(X1)) | ~ r1_tarski(X0, X1)), inference(cnf_transformation, [], [f364])).
fof(f364, plain, ! [X0, X1] : ((m1_subset_1(X0, k1_zfmisc_1(X1)) | ~ r1_tarski(X0, X1)) & (r1_tarski(X0, X1) | ~ m1_subset_1(X0, k1_zfmisc_1(X1)))), inference(nnf_transformation, [], [f120])).
fof(f120, plain, ! [X0, X1] : (m1_subset_1(X0, k1_zfmisc_1(X1)) <=> r1_tarski(X0, X1)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', t3_subset)).
fof(f631, plain, ! [X6, X0, X1] : (~ m1_subset_1(k3_xboole_0(X6, u1_struct_0(X0)), k1_zfmisc_1(u1_struct_0(X0))) | ~ v3_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))) | v3_pre_topc(k3_xboole_0(X6, u1_struct_0(X0)), X0) | ~ sP1(X0, X1)), inference(equality_resolution, [], [f487])).
fof(f487, plain, ! [X6, X0, X5, X1] : (v3_pre_topc(X5, X0) | ~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v3_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0))) | ~ sP1(X0, X1)), inference(cnf_transformation, [], [f319])).
fof(f319, plain, ! [X0, X1] : ((sP1(X0, X1) | ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = sK9(X0, X1)) | ~ v3_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(sK9(X0, X1), X0)) & (((sK9(X0, X1) = k3_xboole_0(sK10(X0, X1), u1_struct_0(X0))) & v3_pre_topc(sK10(X0, X1), X1) & m1_subset_1(sK10(X0, X1), k1_zfmisc_1(u1_struct_0(X1)))) | v3_pre_topc(sK9(X0, X1), X0)) & m1_subset_1(sK9(X0, X1), k1_zfmisc_1(u1_struct_0(X0)))) | ~ r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) & ((! [X5] : (((v3_pre_topc(X5, X0) | ! [X6] : (~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v3_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))))) & (((k3_xboole_0(sK11(X0, X1, X5), u1_struct_0(X0)) = X5) & v3_pre_topc(sK11(X0, X1, X5), X1) & m1_subset_1(sK11(X0, X1, X5), k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(X5, X0))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) & r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) | ~ sP1(X0, X1))), inference(skolemisation, [status(esa), new_symbols(skolem, [sK9, sK10, sK11])], [f315, f318, f317, f316])).
fof(f316, plain, ! [X1, X0] : (? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = X2) | ~ v3_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(X2, X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = X2) & v3_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v3_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) => ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = sK9(X0, X1)) | ~ v3_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(sK9(X0, X1), X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = sK9(X0, X1)) & v3_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v3_pre_topc(sK9(X0, X1), X0)) & m1_subset_1(sK9(X0, X1), k1_zfmisc_1(u1_struct_0(X0))))), introduced(choice_axiom, [])).
fof(f317, plain, ! [X1, X0] : (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = sK9(X0, X1)) & v3_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) => ((sK9(X0, X1) = k3_xboole_0(sK10(X0, X1), u1_struct_0(X0))) & v3_pre_topc(sK10(X0, X1), X1) & m1_subset_1(sK10(X0, X1), k1_zfmisc_1(u1_struct_0(X1))))), introduced(choice_axiom, [])).
fof(f318, plain, ! [X5, X1, X0] : (? [X7] : ((k3_xboole_0(X7, u1_struct_0(X0)) = X5) & v3_pre_topc(X7, X1) & m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X1)))) => ((k3_xboole_0(sK11(X0, X1, X5), u1_struct_0(X0)) = X5) & v3_pre_topc(sK11(X0, X1, X5), X1) & m1_subset_1(sK11(X0, X1, X5), k1_zfmisc_1(u1_struct_0(X1))))), introduced(choice_axiom, [])).
fof(f315, plain, ! [X0, X1] : ((sP1(X0, X1) | ? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X0)) = X2) | ~ v3_pre_topc(X3, X1) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(X2, X0)) & (? [X4] : ((k3_xboole_0(X4, u1_struct_0(X0)) = X2) & v3_pre_topc(X4, X1) & m1_subset_1(X4, k1_zfmisc_1(u1_struct_0(X1)))) | v3_pre_topc(X2, X0)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X0)))) | ~ r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) & ((! [X5] : (((v3_pre_topc(X5, X0) | ! [X6] : (~ (k3_xboole_0(X6, u1_struct_0(X0)) = X5) | ~ v3_pre_topc(X6, X1) | ~ m1_subset_1(X6, k1_zfmisc_1(u1_struct_0(X1))))) & (? [X7] : ((k3_xboole_0(X7, u1_struct_0(X0)) = X5) & v3_pre_topc(X7, X1) & m1_subset_1(X7, k1_zfmisc_1(u1_struct_0(X1)))) | ~ v3_pre_topc(X5, X0))) | ~ m1_subset_1(X5, k1_zfmisc_1(u1_struct_0(X0)))) & r1_tarski(u1_struct_0(X0), u1_struct_0(X1))) | ~ sP1(X0, X1))), inference(rectify, [], [f314])).
fof(f314, plain, ! [X1, X0] : ((sP1(X1, X0) | ? [X2] : ((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v3_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v3_pre_topc(X2, X1)) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | v3_pre_topc(X2, X1)) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) | ~ r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) & ((! [X2] : (((v3_pre_topc(X2, X1) | ! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v3_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v3_pre_topc(X2, X1))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) | ~ sP1(X1, X0))), inference(flattening, [], [f313])).
fof(f313, plain, ! [X1, X0] : ((sP1(X1, X0) | (? [X2] : (((! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v3_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v3_pre_topc(X2, X1)) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | v3_pre_topc(X2, X1))) & m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) | ~ r1_tarski(u1_struct_0(X1), u1_struct_0(X0)))) & ((! [X2] : (((v3_pre_topc(X2, X1) | ! [X3] : (~ (k3_xboole_0(X3, u1_struct_0(X1)) = X2) | ~ v3_pre_topc(X3, X0) | ~ m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0))))) & (? [X3] : ((k3_xboole_0(X3, u1_struct_0(X1)) = X2) & v3_pre_topc(X3, X0) & m1_subset_1(X3, k1_zfmisc_1(u1_struct_0(X0)))) | ~ v3_pre_topc(X2, X1))) | ~ m1_subset_1(X2, k1_zfmisc_1(u1_struct_0(X1)))) & r1_tarski(u1_struct_0(X1), u1_struct_0(X0))) | ~ sP1(X1, X0))), inference(nnf_transformation, [], [f297])).
fof(f112224, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | (~ spl32_148 | ~ spl32_150 | ~ spl32_151 | ~ spl32_161)), inference(subsumption_resolution, [], [f112223, f12074])).
fof(f12074, plain, (m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl32_148), inference(avatar_component_clause, [], [f12073])).
fof(f12073, plain, (spl32_148 <=> m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl32_148])])).
fof(f112223, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl32_150 | ~ spl32_151 | ~ spl32_161)), inference(subsumption_resolution, [], [f112222, f656])).
fof(f656, plain, l1_struct_0(sK4), inference(resolution, [], [f505, f367])).
fof(f505, plain, ! [X0] : (~ l1_pre_topc(X0) | l1_struct_0(X0)), inference(cnf_transformation, [], [f235])).
fof(f235, plain, ! [X0] : (l1_struct_0(X0) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f64])).
fof(f64, plain, ! [X0] : (l1_pre_topc(X0) => l1_struct_0(X0)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', dt_l1_pre_topc)).
fof(f112222, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl32_150 | ~ spl32_151 | ~ spl32_161)), inference(subsumption_resolution, [], [f112221, f934])).
fof(f934, plain, l1_struct_0(sK5), inference(resolution, [], [f932, f505])).
fof(f932, plain, l1_pre_topc(sK5), inference(subsumption_resolution, [], [f929, f367])).
fof(f929, plain, (l1_pre_topc(sK5) | ~ l1_pre_topc(sK4)), inference(resolution, [], [f509, f370])).
fof(f509, plain, ! [X0, X1] : (~ m2_tsp_1(X1, X0) | l1_pre_topc(X1) | ~ l1_pre_topc(X0)), inference(cnf_transformation, [], [f239])).
fof(f239, plain, ! [X0] : (! [X1] : (l1_pre_topc(X1) | ~ m2_tsp_1(X1, X0)) | ~ l1_pre_topc(X0)), inference(ennf_transformation, [], [f71])).
fof(f71, plain, ! [X0] : (l1_pre_topc(X0) => ! [X1] : (m2_tsp_1(X1, X0) => l1_pre_topc(X1))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', dt_m2_tsp_1)).
fof(f112221, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl32_150 | ~ spl32_151 | ~ spl32_161)), inference(subsumption_resolution, [], [f112220, f12082])).
fof(f12082, plain, (v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ spl32_150), inference(avatar_component_clause, [], [f12081])).
fof(f12081, plain, (spl32_150 <=> v1_funct_1(k4_tsp_2(sK4, sK5))), introduced(avatar_definition, [new_symbols(naming, [spl32_150])])).
fof(f112220, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | (~ spl32_151 | ~ spl32_161)), inference(subsumption_resolution, [], [f112151, f12086])).
fof(f12086, plain, (v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl32_151), inference(avatar_component_clause, [], [f12085])).
fof(f12085, plain, (spl32_151 <=> v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5))), introduced(avatar_definition, [new_symbols(naming, [spl32_151])])).
fof(f112151, plain, (r1_tarski(k3_xboole_0(sK6, u1_struct_0(sK5)), u1_struct_0(sK5)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4) | ~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ spl32_161), inference(superposition, [], [f2078, f45632])).
fof(f2078, plain, ! [X111, X114, X112, X113] : (r1_tarski(k4_pre_topc(X112, X113, X111, X114), u1_struct_0(X113)) | ~ v1_funct_2(X111, u1_struct_0(X112), u1_struct_0(X113)) | ~ v1_funct_1(X111) | ~ l1_struct_0(X113) | ~ l1_struct_0(X112) | ~ m1_relset_1(X111, u1_struct_0(X112), u1_struct_0(X113))), inference(resolution, [], [f499, f621])).
fof(f621, plain, ! [X0, X1] : (~ m1_subset_1(X0, k1_zfmisc_1(X1)) | r1_tarski(X0, X1)), inference(cnf_transformation, [], [f364])).
fof(f499, plain, ! [X2, X0, X3, X1] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f230])).
fof(f230, plain, ! [X0, X1, X2, X3] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(flattening, [], [f229])).
fof(f229, plain, ! [X0, X1, X2, X3] : (m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1))) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0))), inference(ennf_transformation, [], [f60])).
fof(f60, plain, ! [X0, X1, X2, X3] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_struct_0(X1) & l1_struct_0(X0)) => m1_subset_1(k4_pre_topc(X0, X1, X2, X3), k1_zfmisc_1(u1_struct_0(X1)))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', dt_k4_pre_topc)).
fof(f12607, plain, spl32_161, inference(avatar_contradiction_clause, [], [f12606])).
fof(f12606, plain, ($false | spl32_161), inference(subsumption_resolution, [], [f12605, f367])).
fof(f12605, plain, (~ l1_pre_topc(sK4) | spl32_161), inference(subsumption_resolution, [], [f12603, f1151])).
fof(f12603, plain, (~ m1_pre_topc(sK5, sK4) | ~ l1_pre_topc(sK4) | spl32_161), inference(resolution, [], [f12528, f618])).
fof(f12528, plain, (~ m1_subset_1(u1_struct_0(sK5), k1_zfmisc_1(u1_struct_0(sK4))) | spl32_161), inference(avatar_component_clause, [], [f12526])).
fof(f12436, plain, spl32_151, inference(avatar_contradiction_clause, [], [f12435])).
fof(f12435, plain, ($false | spl32_151), inference(subsumption_resolution, [], [f12434, f365])).
fof(f12434, plain, (v3_struct_0(sK4) | spl32_151), inference(subsumption_resolution, [], [f12433, f366])).
fof(f12433, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_151), inference(subsumption_resolution, [], [f12432, f367])).
fof(f12432, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_151), inference(subsumption_resolution, [], [f12431, f368])).
fof(f12431, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_151), inference(subsumption_resolution, [], [f12430, f369])).
fof(f12430, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_151), inference(subsumption_resolution, [], [f12429, f1151])).
fof(f12429, plain, (~ m1_pre_topc(sK5, sK4) | ~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_151), inference(resolution, [], [f12087, f501])).
fof(f12087, plain, (~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | spl32_151), inference(avatar_component_clause, [], [f12085])).
fof(f12376, plain, spl32_150, inference(avatar_contradiction_clause, [], [f12375])).
fof(f12375, plain, ($false | spl32_150), inference(subsumption_resolution, [], [f12374, f365])).
fof(f12374, plain, (v3_struct_0(sK4) | spl32_150), inference(subsumption_resolution, [], [f12373, f366])).
fof(f12373, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_150), inference(subsumption_resolution, [], [f12372, f367])).
fof(f12372, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_150), inference(subsumption_resolution, [], [f12371, f368])).
fof(f12371, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_150), inference(subsumption_resolution, [], [f12370, f369])).
fof(f12370, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_150), inference(subsumption_resolution, [], [f12369, f1151])).
fof(f12369, plain, (~ m1_pre_topc(sK5, sK4) | ~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | spl32_150), inference(resolution, [], [f12083, f500])).
fof(f12083, plain, (~ v1_funct_1(k4_tsp_2(sK4, sK5)) | spl32_150), inference(avatar_component_clause, [], [f12081])).
fof(f12359, plain, (~ spl32_150 | ~ spl32_151 | spl32_157 | ~ spl32_148), inference(avatar_split_clause, [], [f12355, f12073, f12357, f12085, f12081])).
fof(f12355, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5))) | ~ spl32_148), inference(subsumption_resolution, [], [f12354, f656])).
fof(f12354, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK4)) | ~ spl32_148), inference(subsumption_resolution, [], [f12336, f934])).
fof(f12336, plain, (! [X0] : ((k4_pre_topc(sK4, sK5, k4_tsp_2(sK4, sK5), X0) = k9_relat_1(k4_tsp_2(sK4, sK5), X0)) | ~ v1_funct_2(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | ~ v1_funct_1(k4_tsp_2(sK4, sK5)) | ~ l1_struct_0(sK5) | ~ l1_struct_0(sK4)) | ~ spl32_148), inference(resolution, [], [f12074, f607])).
fof(f607, plain, ! [X2, X0, X3, X1] : (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | (k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(cnf_transformation, [], [f276])).
fof(f276, plain, ! [X0, X1, X2, X3] : ((k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | ~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0)), inference(flattening, [], [f275])).
fof(f275, plain, ! [X0, X1, X2, X3] : ((k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3)) | (~ m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) | ~ v1_funct_1(X2) | ~ l1_struct_0(X1) | ~ l1_struct_0(X0))), inference(ennf_transformation, [], [f109])).
fof(f109, plain, ! [X0, X1, X2, X3] : ((m1_relset_1(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_2(X2, u1_struct_0(X0), u1_struct_0(X1)) & v1_funct_1(X2) & l1_struct_0(X1) & l1_struct_0(X0)) => (k4_pre_topc(X0, X1, X2, X3) = k9_relat_1(X2, X3))), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', redefinition_k4_pre_topc)).
fof(f12333, plain, spl32_148, inference(avatar_contradiction_clause, [], [f12332])).
fof(f12332, plain, ($false | spl32_148), inference(subsumption_resolution, [], [f12331, f1151])).
fof(f12331, plain, (~ m1_pre_topc(sK5, sK4) | spl32_148), inference(subsumption_resolution, [], [f12330, f365])).
fof(f12330, plain, (v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl32_148), inference(subsumption_resolution, [], [f12329, f366])).
fof(f12329, plain, (~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl32_148), inference(subsumption_resolution, [], [f12328, f367])).
fof(f12328, plain, (~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl32_148), inference(subsumption_resolution, [], [f12327, f368])).
fof(f12327, plain, (v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl32_148), inference(subsumption_resolution, [], [f12326, f369])).
fof(f12326, plain, (~ v2_tsp_2(sK5, sK4) | v3_struct_0(sK5) | ~ l1_pre_topc(sK4) | ~ v2_pre_topc(sK4) | v3_struct_0(sK4) | ~ m1_pre_topc(sK5, sK4) | spl32_148), inference(resolution, [], [f12075, f2029])).
fof(f2029, plain, ! [X0, X1] : (m1_relset_1(k4_tsp_2(X1, X0), u1_struct_0(X1), u1_struct_0(X0)) | ~ v2_tsp_2(X0, X1) | v3_struct_0(X0) | ~ l1_pre_topc(X1) | ~ v2_pre_topc(X1) | v3_struct_0(X1) | ~ m1_pre_topc(X0, X1)), inference(resolution, [], [f503, f612])).
fof(f612, plain, ! [X2, X0, X1] : (~ m2_relset_1(X2, X0, X1) | m1_relset_1(X2, X0, X1)), inference(cnf_transformation, [], [f362])).
fof(f362, plain, ! [X0, X1, X2] : ((m2_relset_1(X2, X0, X1) | ~ m1_relset_1(X2, X0, X1)) & (m1_relset_1(X2, X0, X1) | ~ m2_relset_1(X2, X0, X1))), inference(nnf_transformation, [], [f113])).
fof(f113, plain, ! [X0, X1, X2] : (m2_relset_1(X2, X0, X1) <=> m1_relset_1(X2, X0, X1)), file('/home/ubuntu/library/tptp/Problems/TOP/TOP041+1.p', redefinition_m2_relset_1)).
fof(f12075, plain, (~ m1_relset_1(k4_tsp_2(sK4, sK5), u1_struct_0(sK4), u1_struct_0(sK5)) | spl32_148), inference(avatar_component_clause, [], [f12073])).